# Preprocessing time       : 0.027 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t40_waybel30, conjecture, ![X1]:((((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v2_waybel19(X1))&v2_waybel_2(X1))&l1_waybel_9(X1))=>(v2_waybel_8(X1)<=>(v3_waybel30(X1)&v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t40_waybel30)).
fof(dt_o_1_14_waybel30, axiom, ![X1]:((((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v2_waybel19(X1))&v2_waybel_2(X1))&l1_waybel_9(X1))=>(v4_waybel11(o_1_14_waybel30(X1))&m1_yellow_9(o_1_14_waybel30(X1),X1))), file('/scratch/mptp/minsubpreds3/t40_waybel30', dt_o_1_14_waybel30)).
fof(dt_l1_waybel_9, axiom, ![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1))), file('/scratch/mptp/minsubpreds3/t40_waybel30', dt_l1_waybel_9)).
fof(dt_m1_yellow_9, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_9(X2,X1)=>l1_waybel_9(X2))), file('/scratch/mptp/minsubpreds3/t40_waybel30', dt_m1_yellow_9)).
fof(d4_yellow_9, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_waybel_9(X2)=>(m1_yellow_9(X2,X1)<=>g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', d4_yellow_9)).
fof(t52_yellow_9, axiom, ![X1]:(((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&l1_orders_2(X1))=>![X2]:(((((((v3_orders_2(X2)&v4_orders_2(X2))&v5_orders_2(X2))&v1_lattice3(X2))&v2_lattice3(X2))&v3_lattice3(X2))&l1_orders_2(X2))=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>k5_waybel11(X1)=k5_waybel11(X2)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t52_yellow_9)).
fof(cc4_yellow_9, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_9(X2,X1)=>v5_orders_2(X2))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc4_yellow_9)).
fof(cc3_yellow_9, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_9(X2,X1)=>v4_orders_2(X2))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc3_yellow_9)).
fof(cc2_yellow_9, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_9(X2,X1)=>v3_orders_2(X2))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc2_yellow_9)).
fof(cc1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>((~(v2_struct_0(X1))&v3_lattice3(X1))=>((~(v2_struct_0(X1))&v1_lattice3(X1))&v2_lattice3(X1)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc1_yellow_0)).
fof(cc5_yellow_9, axiom, ![X1]:(((~(v2_struct_0(X1))&v3_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_9(X2,X1)=>v3_lattice3(X2))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc5_yellow_9)).
fof(t39_waybel14, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v4_waybel11(X1))&l1_waybel_9(X1))=>(![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>?[X3]:(((v1_tops_2(X3,X1)&v1_yellow_8(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)=>(v3_pre_topc(X4,X1)&v2_waybel_0(X4,X1))))))<=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>?[X3]:((m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))&X2=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3))&![X4]:(m1_subset_1(X4,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>(r2_hidden(X4,X3)=>v6_waybel_6(X4,k2_yellow_1(k5_waybel11(X1))))))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t39_waybel14)).
fof(cc1_yellow_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_9(X2,X1)=>~(v2_struct_0(X2)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc1_yellow_9)).
fof(cc6_yellow_9, axiom, ![X1]:(((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_9(X2,X1)=>(v4_waybel11(X2)=>v2_pre_topc(X2)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc6_yellow_9)).
fof(t44_waybel14, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v4_waybel11(X1))&l1_waybel_9(X1))=>~(((v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))&![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>?[X3]:((m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))&X2=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3))&![X4]:(m1_subset_1(X4,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>(r2_hidden(X4,X3)=>v6_waybel_6(X4,k2_yellow_1(k5_waybel11(X1))))))))&![X2]:(((v1_tops_2(X2,X1)&v1_cantor_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>X2!=a_1_0_waybel14(X1))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t44_waybel14)).
fof(t45_waybel14, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v4_waybel11(X1))&l1_waybel_9(X1))=>(?[X2]:(((v1_tops_2(X2,X1)&v1_cantor_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))&X2=a_1_0_waybel14(X1))=>v2_waybel_8(X1))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t45_waybel14)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('/scratch/mptp/minsubpreds3/t40_waybel30', dt_k2_yellow_1)).
fof(fc5_waybel30, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v2_waybel19(X1))&l1_waybel_9(X1))=>((~(v7_struct_0(k2_yellow_1(k5_waybel11(X1))))&v1_orders_2(k2_yellow_1(k5_waybel11(X1))))&v3_lattice3(k2_yellow_1(k5_waybel11(X1))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', fc5_waybel30)).
fof(cc1_waybel_8, axiom, ![X1]:(l1_orders_2(X1)=>((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v2_waybel_8(X1))=>(((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_waybel_3(X1)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc1_waybel_8)).
fof(fc5_yellow_1, axiom, ![X1]:(((v1_orders_2(k2_yellow_1(X1))&v3_orders_2(k2_yellow_1(X1)))&v4_orders_2(k2_yellow_1(X1)))&v5_orders_2(k2_yellow_1(X1))), file('/scratch/mptp/minsubpreds3/t40_waybel30', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', fc6_yellow_1)).
fof(fc3_waybel14, axiom, ![X1]:(((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&l1_orders_2(X1))=>~(v1_xboole_0(k5_waybel11(X1)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', fc3_waybel14)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc1_lattice3)).
fof(t35_waybel30, axiom, ![X1]:((((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v2_waybel19(X1))&v2_waybel_2(X1))&l1_waybel_9(X1))=>(v3_waybel_3(X1)<=>(v3_waybel30(X1)&v3_waybel_3(k2_yellow_1(k5_waybel11(X1)))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t35_waybel30)).
fof(cc4_waybel19, axiom, ![X1]:((((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v3_waybel_3(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_9(X2,X1)=>v3_waybel_3(X2))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc4_waybel19)).
fof(t35_waybel14, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v4_waybel11(X1))&l1_waybel_9(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~((v3_waybel_3(X1)&![X3]:(((v1_tops_2(X3,X1)&v1_yellow_8(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>?[X4]:((m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&r2_hidden(X4,X3))&~((v3_pre_topc(X4,X1)&v2_waybel_0(X4,X1))))))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t35_waybel14)).
fof(t17_waybel_8, axiom, ![X1]:((((~(v2_struct_0(X1))&v3_orders_2(X1))&v5_orders_2(X1))&l1_orders_2(X1))=>![X2]:((((~(v2_struct_0(X2))&v3_orders_2(X2))&v5_orders_2(X2))&l1_orders_2(X2))=>((g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))&v2_waybel_8(X1))=>v2_waybel_8(X2)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t17_waybel_8)).
fof(t42_waybel14, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v4_waybel11(X1))&l1_waybel_9(X1))=>~((v2_waybel_8(X1)&![X2]:(((v1_tops_2(X2,X1)&v1_cantor_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>X2!=a_1_0_waybel14(X1))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t42_waybel14)).
fof(cc10_waybel30, axiom, ![X1]:(l1_waybel_9(X1)=>(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v2_waybel19(X1))&v3_waybel_3(X1))=>((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v2_waybel19(X1))&v3_waybel30(X1)))), file('/scratch/mptp/minsubpreds3/t40_waybel30', cc10_waybel30)).
fof(t43_waybel14, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v4_waybel11(X1))&l1_waybel_9(X1))=>(?[X2]:(((v1_tops_2(X2,X1)&v1_cantor_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))&X2=a_1_0_waybel14(X1))=>(v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))&![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>?[X3]:((m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))&X2=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3))&![X4]:(m1_subset_1(X4,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>(r2_hidden(X4,X3)=>v6_waybel_6(X4,k2_yellow_1(k5_waybel11(X1)))))))))), file('/scratch/mptp/minsubpreds3/t40_waybel30', t43_waybel14)).
fof(c_0_30, negated_conjecture, ~(![X1]:((((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v2_waybel19(X1))&v2_waybel_2(X1))&l1_waybel_9(X1))=>(v2_waybel_8(X1)<=>(v3_waybel30(X1)&v2_waybel_8(k2_yellow_1(k5_waybel11(X1))))))), inference(assume_negation,[status(cth)],[t40_waybel30])).
fof(c_0_31, plain, ![X31]:((v4_waybel11(o_1_14_waybel30(X31))|(~v2_pre_topc(X31)|~v3_orders_2(X31)|~v4_orders_2(X31)|~v5_orders_2(X31)|~v1_lattice3(X31)|~v2_lattice3(X31)|~v3_lattice3(X31)|~v2_waybel19(X31)|~v2_waybel_2(X31)|~l1_waybel_9(X31)))&(m1_yellow_9(o_1_14_waybel30(X31),X31)|(~v2_pre_topc(X31)|~v3_orders_2(X31)|~v4_orders_2(X31)|~v5_orders_2(X31)|~v1_lattice3(X31)|~v2_lattice3(X31)|~v3_lattice3(X31)|~v2_waybel19(X31)|~v2_waybel_2(X31)|~l1_waybel_9(X31)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_o_1_14_waybel30])])])).
fof(c_0_32, negated_conjecture, ((((((((((v2_pre_topc(esk1_0)&v3_orders_2(esk1_0))&v4_orders_2(esk1_0))&v5_orders_2(esk1_0))&v1_lattice3(esk1_0))&v2_lattice3(esk1_0))&v3_lattice3(esk1_0))&v2_waybel19(esk1_0))&v2_waybel_2(esk1_0))&l1_waybel_9(esk1_0))&((~v2_waybel_8(esk1_0)|(~v3_waybel30(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))))&((v3_waybel30(esk1_0)|v2_waybel_8(esk1_0))&(v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])])).
fof(c_0_33, plain, ![X28]:((l1_pre_topc(X28)|~l1_waybel_9(X28))&(l1_orders_2(X28)|~l1_waybel_9(X28))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
fof(c_0_34, plain, ![X29, X30]:(~l1_orders_2(X29)|(~m1_yellow_9(X30,X29)|l1_waybel_9(X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_9])])])).
cnf(c_0_35, plain, (m1_yellow_9(o_1_14_waybel30(X1),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, negated_conjecture, (l1_waybel_9(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (v2_waybel_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (v2_waybel19(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46, plain, (l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_47, plain, ![X25, X26]:((~m1_yellow_9(X26,X25)|g1_orders_2(u1_struct_0(X26),u1_orders_2(X26))=g1_orders_2(u1_struct_0(X25),u1_orders_2(X25))|~l1_waybel_9(X26)|~l1_orders_2(X25))&(g1_orders_2(u1_struct_0(X26),u1_orders_2(X26))!=g1_orders_2(u1_struct_0(X25),u1_orders_2(X25))|m1_yellow_9(X26,X25)|~l1_waybel_9(X26)|~l1_orders_2(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_9])])])])).
cnf(c_0_48, plain, (l1_waybel_9(X2)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_49, plain, (m1_yellow_9(o_1_14_waybel30(X1),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_50, negated_conjecture, (l1_waybel_9(esk1_0)), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_51, negated_conjecture, (v2_waybel_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_37])).
cnf(c_0_52, negated_conjecture, (v2_waybel19(esk1_0)), inference(evalgc,[status(thm)],[c_0_38])).
cnf(c_0_53, negated_conjecture, (v3_lattice3(esk1_0)), inference(evalgc,[status(thm)],[c_0_39])).
cnf(c_0_54, negated_conjecture, (v2_lattice3(esk1_0)), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_55, negated_conjecture, (v1_lattice3(esk1_0)), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_56, negated_conjecture, (v5_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_57, negated_conjecture, (v4_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_58, negated_conjecture, (v3_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_59, negated_conjecture, (v2_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_60, plain, (l1_orders_2(X1)|~l1_waybel_9(X1)), inference(evalgc,[status(thm)],[c_0_46])).
fof(c_0_61, plain, ![X1]:(epred1_1(X1)<=>(![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>?[X3]:(((v1_tops_2(X3,X1)&v1_yellow_8(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)=>(v3_pre_topc(X4,X1)&v2_waybel_0(X4,X1))))))<=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>?[X3]:((m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))&X2=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3))&![X4]:(m1_subset_1(X4,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>(r2_hidden(X4,X3)=>v6_waybel_6(X4,k2_yellow_1(k5_waybel11(X1))))))))), introduced(definition)).
fof(c_0_62, plain, ![X62, X63]:(~v3_orders_2(X62)|~v4_orders_2(X62)|~v5_orders_2(X62)|~v1_lattice3(X62)|~v2_lattice3(X62)|~v3_lattice3(X62)|~l1_orders_2(X62)|(~v3_orders_2(X63)|~v4_orders_2(X63)|~v5_orders_2(X63)|~v1_lattice3(X63)|~v2_lattice3(X63)|~v3_lattice3(X63)|~l1_orders_2(X63)|(g1_orders_2(u1_struct_0(X62),u1_orders_2(X62))!=g1_orders_2(u1_struct_0(X63),u1_orders_2(X63))|k5_waybel11(X62)=k5_waybel11(X63)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_yellow_9])])])).
cnf(c_0_63, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~m1_yellow_9(X1,X2)|~l1_waybel_9(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_64, plain, ![X19, X20]:(~v5_orders_2(X19)|~l1_orders_2(X19)|(~m1_yellow_9(X20,X19)|v5_orders_2(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_yellow_9])])])).
fof(c_0_65, plain, ![X15, X16]:(~v4_orders_2(X15)|~l1_orders_2(X15)|(~m1_yellow_9(X16,X15)|v4_orders_2(X16))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_yellow_9])])])).
fof(c_0_66, plain, ![X13, X14]:(~v3_orders_2(X13)|~l1_orders_2(X13)|(~m1_yellow_9(X14,X13)|v3_orders_2(X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_yellow_9])])])).
fof(c_0_67, plain, ![X9]:(((~v2_struct_0(X9)|(v2_struct_0(X9)|~v3_lattice3(X9))|~l1_orders_2(X9))&(v1_lattice3(X9)|(v2_struct_0(X9)|~v3_lattice3(X9))|~l1_orders_2(X9)))&(v2_lattice3(X9)|(v2_struct_0(X9)|~v3_lattice3(X9))|~l1_orders_2(X9))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_yellow_0])])])])).
cnf(c_0_68, plain, (l1_waybel_9(X2)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_69, negated_conjecture, (m1_yellow_9(o_1_14_waybel30(esk1_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59])])])).
cnf(c_0_70, negated_conjecture, (l1_orders_2(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_50])])).
fof(c_0_71, plain, ![X21, X22]:(v2_struct_0(X21)|~v3_lattice3(X21)|~l1_orders_2(X21)|(~m1_yellow_9(X22,X21)|v3_lattice3(X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_yellow_9])])])])).
fof(c_0_72, axiom, ![X1]:(((((((((v2_pre_topc(X1)&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v2_lattice3(X1))&v3_lattice3(X1))&v4_waybel11(X1))&l1_waybel_9(X1))=>epred1_1(X1)), inference(apply_def,[status(thm)],[t39_waybel14, c_0_61])).
cnf(c_0_73, plain, (k5_waybel11(X1)=k5_waybel11(X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~l1_orders_2(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_74, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~l1_waybel_9(X1)|~l1_orders_2(X2)|~m1_yellow_9(X1,X2)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_75, plain, (v5_orders_2(X2)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_76, plain, (v4_orders_2(X2)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_77, plain, (v3_orders_2(X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_78, plain, (v1_lattice3(X1)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_79, negated_conjecture, (l1_waybel_9(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])])])).
cnf(c_0_80, plain, (v2_struct_0(X1)|v3_lattice3(X2)|~v3_lattice3(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_81, plain, ![X10, X11]:(v2_struct_0(X10)|~l1_orders_2(X10)|(~m1_yellow_9(X11,X10)|~v2_struct_0(X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_yellow_9])])])])).
cnf(c_0_82, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_83, plain, ![X47]:(~v2_pre_topc(X47)|~v3_orders_2(X47)|~v4_orders_2(X47)|~v5_orders_2(X47)|~v1_lattice3(X47)|~v2_lattice3(X47)|~v3_lattice3(X47)|~v4_waybel11(X47)|~l1_waybel_9(X47)|epred1_1(X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_72])])).
cnf(c_0_84, plain, (v4_waybel11(o_1_14_waybel30(X1))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_85, plain, (k5_waybel11(X1)=k5_waybel11(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~v3_orders_2(X2)|~v3_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v5_orders_2(X2)|~v5_orders_2(X1)|~v1_lattice3(X2)|~v1_lattice3(X1)|~v2_lattice3(X2)|~v2_lattice3(X1)|~v3_lattice3(X2)|~v3_lattice3(X1)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_86, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~m1_yellow_9(X1,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[c_0_74, c_0_68])).
cnf(c_0_87, plain, (v5_orders_2(X2)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_88, plain, (v4_orders_2(X2)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_89, plain, (v3_orders_2(X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_90, plain, (v1_lattice3(X1)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_91, negated_conjecture, (l1_orders_2(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_79])])).
cnf(c_0_92, plain, (v3_lattice3(X2)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_93, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)|~v2_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_94, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_95, plain, (epred1_1(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_96, plain, (v4_waybel11(o_1_14_waybel30(X1))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)), inference(evalgc,[status(thm)],[c_0_84])).
fof(c_0_97, plain, ![X23, X24]:(~v3_orders_2(X23)|~v4_orders_2(X23)|~v5_orders_2(X23)|~v1_lattice3(X23)|~v2_lattice3(X23)|~v3_lattice3(X23)|~l1_orders_2(X23)|(~m1_yellow_9(X24,X23)|(~v4_waybel11(X24)|v2_pre_topc(X24)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_yellow_9])])])).
cnf(c_0_98, negated_conjecture, (k5_waybel11(X1)=k5_waybel11(esk1_0)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_orders_2(X1)|~v3_lattice3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_70]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58])])])).
cnf(c_0_99, negated_conjecture, (g1_orders_2(u1_struct_0(o_1_14_waybel30(esk1_0)),u1_orders_2(o_1_14_waybel30(esk1_0)))=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_69]), c_0_70])])])).
cnf(c_0_100, negated_conjecture, (v5_orders_2(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_69]), c_0_70]), c_0_56])])])).
cnf(c_0_101, negated_conjecture, (v4_orders_2(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_69]), c_0_70]), c_0_57])])])).
cnf(c_0_102, negated_conjecture, (v3_orders_2(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_69]), c_0_70]), c_0_58])])])).
cnf(c_0_103, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v1_lattice3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
cnf(c_0_104, negated_conjecture, (v2_struct_0(esk1_0)|v3_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_69]), c_0_70]), c_0_53])])])).
cnf(c_0_105, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~v2_struct_0(X2)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_106, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_lattice3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_91])])).
cnf(c_0_107, plain, (epred1_1(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_108, negated_conjecture, (v4_waybel11(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59])])])).
cnf(c_0_109, plain, (v2_pre_topc(X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)|~v4_waybel11(X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
fof(c_0_110, plain, ![X55, X57]:(((((v1_tops_2(esk8_1(X55),X55)|(m1_subset_1(esk6_1(X55),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55)))&(v1_cantor_1(esk8_1(X55),X55)|(m1_subset_1(esk6_1(X55),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(m1_subset_1(esk8_1(X55),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X55))))|(m1_subset_1(esk6_1(X55),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(esk8_1(X55)=a_1_0_waybel14(X55)|(m1_subset_1(esk6_1(X55),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(((((v1_tops_2(esk8_1(X55),X55)|(m1_subset_1(esk7_2(X55,X57),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55)))&(v1_cantor_1(esk8_1(X55),X55)|(m1_subset_1(esk7_2(X55,X57),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(m1_subset_1(esk8_1(X55),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X55))))|(m1_subset_1(esk7_2(X55,X57),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(esk8_1(X55)=a_1_0_waybel14(X55)|(m1_subset_1(esk7_2(X55,X57),u1_struct_0(k2_yellow_1(k5_waybel11(X55))))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(((((v1_tops_2(esk8_1(X55),X55)|(r2_hidden(esk7_2(X55,X57),X57)|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55)))&(v1_cantor_1(esk8_1(X55),X55)|(r2_hidden(esk7_2(X55,X57),X57)|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(m1_subset_1(esk8_1(X55),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X55))))|(r2_hidden(esk7_2(X55,X57),X57)|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(esk8_1(X55)=a_1_0_waybel14(X55)|(r2_hidden(esk7_2(X55,X57),X57)|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&((((v1_tops_2(esk8_1(X55),X55)|(~v6_waybel_6(esk7_2(X55,X57),k2_yellow_1(k5_waybel11(X55)))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55)))&(v1_cantor_1(esk8_1(X55),X55)|(~v6_waybel_6(esk7_2(X55,X57),k2_yellow_1(k5_waybel11(X55)))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(m1_subset_1(esk8_1(X55),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X55))))|(~v6_waybel_6(esk7_2(X55,X57),k2_yellow_1(k5_waybel11(X55)))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))&(esk8_1(X55)=a_1_0_waybel14(X55)|(~v6_waybel_6(esk7_2(X55,X57),k2_yellow_1(k5_waybel11(X55)))|(~m1_subset_1(X57,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X55)))))|esk6_1(X55)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X55)),X57))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X55))))|(~v2_pre_topc(X55)|~v3_orders_2(X55)|~v4_orders_2(X55)|~v5_orders_2(X55)|~v1_lattice3(X55)|~v2_lattice3(X55)|~v3_lattice3(X55)|~v4_waybel11(X55)|~l1_waybel_9(X55))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_waybel14])])])])])).
cnf(c_0_111, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_91]), c_0_100]), c_0_101]), c_0_102])])])).
cnf(c_0_112, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v1_lattice3(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_113, negated_conjecture, (v2_struct_0(esk1_0)|~v2_struct_0(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_69]), c_0_70])])])).
cnf(c_0_114, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_lattice3(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_104])])).
cnf(c_0_115, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108])])).
cnf(c_0_116, plain, (v2_pre_topc(X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~v4_waybel11(X2)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_117, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_118, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(esk1_0)|v2_struct_0(esk1_0)|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_104])])).
cnf(c_0_119, negated_conjecture, (v2_struct_0(esk1_0)|v1_lattice3(o_1_14_waybel30(esk1_0))), inference(csr,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_120, negated_conjecture, (v2_struct_0(esk1_0)|v2_lattice3(o_1_14_waybel30(esk1_0))), inference(csr,[status(thm)],[c_0_114, c_0_113])).
fof(c_0_121, plain, ![X1]:(epred1_1(X1)=>(![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>?[X3]:(((v1_tops_2(X3,X1)&v1_yellow_8(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X4,X3)=>(v3_pre_topc(X4,X1)&v2_waybel_0(X4,X1))))))<=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>?[X3]:((m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))&X2=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3))&![X4]:(m1_subset_1(X4,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))=>(r2_hidden(X4,X3)=>v6_waybel_6(X4,k2_yellow_1(k5_waybel11(X1))))))))), inference(split_equiv,[status(thm)],[c_0_61])).
cnf(c_0_122, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_79]), c_0_100]), c_0_101]), c_0_102])])).
cnf(c_0_123, negated_conjecture, (v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_69]), c_0_108]), c_0_70]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58])])])).
cnf(c_0_124, plain, (m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_125, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(esk1_0)|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_118, c_0_119]), c_0_120])).
cnf(c_0_126, plain, (v1_tops_2(esk8_1(X1),X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_127, plain, (v1_cantor_1(esk8_1(X1),X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_128, plain, (esk8_1(X1)=a_1_0_waybel14(X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
fof(c_0_129, plain, ![X64, X66, X68, X70, X72, X74, X76]:(((((m1_subset_1(esk11_2(X64,X68),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|m1_subset_1(esk9_1(X64),u1_struct_0(X64))|~epred1_1(X64))&(X68=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),esk11_2(X64,X68))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|m1_subset_1(esk9_1(X64),u1_struct_0(X64))|~epred1_1(X64)))&(~m1_subset_1(X70,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~r2_hidden(X70,esk11_2(X64,X68))|v6_waybel_6(X70,k2_yellow_1(k5_waybel11(X64))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|m1_subset_1(esk9_1(X64),u1_struct_0(X64))|~epred1_1(X64)))&((((m1_subset_1(esk11_2(X64,X68),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(m1_subset_1(esk10_2(X64,X66),k1_zfmisc_1(u1_struct_0(X64)))|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64))&(X68=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),esk11_2(X64,X68))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(m1_subset_1(esk10_2(X64,X66),k1_zfmisc_1(u1_struct_0(X64)))|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64)))&(~m1_subset_1(X70,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~r2_hidden(X70,esk11_2(X64,X68))|v6_waybel_6(X70,k2_yellow_1(k5_waybel11(X64))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(m1_subset_1(esk10_2(X64,X66),k1_zfmisc_1(u1_struct_0(X64)))|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64)))&((((m1_subset_1(esk11_2(X64,X68),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(r2_hidden(esk10_2(X64,X66),X66)|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64))&(X68=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),esk11_2(X64,X68))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(r2_hidden(esk10_2(X64,X66),X66)|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64)))&(~m1_subset_1(X70,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~r2_hidden(X70,esk11_2(X64,X68))|v6_waybel_6(X70,k2_yellow_1(k5_waybel11(X64))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(r2_hidden(esk10_2(X64,X66),X66)|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64)))&(((m1_subset_1(esk11_2(X64,X68),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~v3_pre_topc(esk10_2(X64,X66),X64)|~v2_waybel_0(esk10_2(X64,X66),X64)|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64))&(X68=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),esk11_2(X64,X68))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~v3_pre_topc(esk10_2(X64,X66),X64)|~v2_waybel_0(esk10_2(X64,X66),X64)|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64)))&(~m1_subset_1(X70,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~r2_hidden(X70,esk11_2(X64,X68))|v6_waybel_6(X70,k2_yellow_1(k5_waybel11(X64))))|~m1_subset_1(X68,u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~v3_pre_topc(esk10_2(X64,X66),X64)|~v2_waybel_0(esk10_2(X64,X66),X64)|(~v1_tops_2(X66,X64)|~v1_yellow_8(X66,X64,esk9_1(X64))|~m1_subset_1(X66,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))))|~epred1_1(X64))))))&(((((v1_tops_2(esk14_2(X64,X74),X64)|~m1_subset_1(X74,u1_struct_0(X64))|m1_subset_1(esk12_1(X64),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|~epred1_1(X64))&(v1_yellow_8(esk14_2(X64,X74),X64,X74)|~m1_subset_1(X74,u1_struct_0(X64))|m1_subset_1(esk12_1(X64),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|~epred1_1(X64)))&(m1_subset_1(esk14_2(X64,X74),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))|~m1_subset_1(X74,u1_struct_0(X64))|m1_subset_1(esk12_1(X64),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|~epred1_1(X64)))&((v3_pre_topc(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|m1_subset_1(esk12_1(X64),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|~epred1_1(X64))&(v2_waybel_0(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|m1_subset_1(esk12_1(X64),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|~epred1_1(X64))))&(((((v1_tops_2(esk14_2(X64,X74),X64)|~m1_subset_1(X74,u1_struct_0(X64))|(m1_subset_1(esk13_2(X64,X72),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))&(v1_yellow_8(esk14_2(X64,X74),X64,X74)|~m1_subset_1(X74,u1_struct_0(X64))|(m1_subset_1(esk13_2(X64,X72),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64)))&(m1_subset_1(esk14_2(X64,X74),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))|~m1_subset_1(X74,u1_struct_0(X64))|(m1_subset_1(esk13_2(X64,X72),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64)))&((v3_pre_topc(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|(m1_subset_1(esk13_2(X64,X72),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))&(v2_waybel_0(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|(m1_subset_1(esk13_2(X64,X72),u1_struct_0(k2_yellow_1(k5_waybel11(X64))))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))))&(((((v1_tops_2(esk14_2(X64,X74),X64)|~m1_subset_1(X74,u1_struct_0(X64))|(r2_hidden(esk13_2(X64,X72),X72)|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))&(v1_yellow_8(esk14_2(X64,X74),X64,X74)|~m1_subset_1(X74,u1_struct_0(X64))|(r2_hidden(esk13_2(X64,X72),X72)|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64)))&(m1_subset_1(esk14_2(X64,X74),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))|~m1_subset_1(X74,u1_struct_0(X64))|(r2_hidden(esk13_2(X64,X72),X72)|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64)))&((v3_pre_topc(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|(r2_hidden(esk13_2(X64,X72),X72)|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))&(v2_waybel_0(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|(r2_hidden(esk13_2(X64,X72),X72)|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))))&((((v1_tops_2(esk14_2(X64,X74),X64)|~m1_subset_1(X74,u1_struct_0(X64))|(~v6_waybel_6(esk13_2(X64,X72),k2_yellow_1(k5_waybel11(X64)))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))&(v1_yellow_8(esk14_2(X64,X74),X64,X74)|~m1_subset_1(X74,u1_struct_0(X64))|(~v6_waybel_6(esk13_2(X64,X72),k2_yellow_1(k5_waybel11(X64)))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64)))&(m1_subset_1(esk14_2(X64,X74),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X64))))|~m1_subset_1(X74,u1_struct_0(X64))|(~v6_waybel_6(esk13_2(X64,X72),k2_yellow_1(k5_waybel11(X64)))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64)))&((v3_pre_topc(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|(~v6_waybel_6(esk13_2(X64,X72),k2_yellow_1(k5_waybel11(X64)))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64))&(v2_waybel_0(X76,X64)|~r2_hidden(X76,esk14_2(X64,X74))|~m1_subset_1(X76,k1_zfmisc_1(u1_struct_0(X64)))|~m1_subset_1(X74,u1_struct_0(X64))|(~v6_waybel_6(esk13_2(X64,X72),k2_yellow_1(k5_waybel11(X64)))|(~m1_subset_1(X72,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X64)))))|esk12_1(X64)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X64)),X72)))|~epred1_1(X64)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_121])])])])])).
cnf(c_0_130, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123])])])).
fof(c_0_131, plain, ![X60, X61]:(~v2_pre_topc(X60)|~v3_orders_2(X60)|~v4_orders_2(X60)|~v5_orders_2(X60)|~v1_lattice3(X60)|~v2_lattice3(X60)|~v3_lattice3(X60)|~v4_waybel11(X60)|~l1_waybel_9(X60)|(~v1_tops_2(X61,X60)|~v1_cantor_1(X61,X60)|~m1_subset_1(X61,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X60))))|X61!=a_1_0_waybel14(X60)|v2_waybel_8(X60))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_waybel14])])])).
cnf(c_0_132, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_133, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_134, plain, (v1_tops_2(esk8_1(X1),X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_135, plain, (v1_cantor_1(esk8_1(X1),X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_136, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_137, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|r2_hidden(esk7_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_138, plain, (m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|m1_subset_1(esk9_1(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_139, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_104])])).
cnf(c_0_140, plain, (v2_waybel_8(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|X2!=a_1_0_waybel14(X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_141, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_132, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_142, negated_conjecture, (v2_waybel_8(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_143, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_144, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_145, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_146, plain, (r2_hidden(esk7_2(X1,X2),X2)|m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_147, plain, (m1_subset_1(esk9_1(X1),u1_struct_0(X1))|m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_148, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_120])])).
cnf(c_0_149, plain, (v2_waybel_8(X1)|X2!=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_150, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_151, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_143, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_152, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_144, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_153, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_145, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_154, plain, (X1=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))|m1_subset_1(esk9_1(X2),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_155, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_156, plain, (v1_tops_2(esk8_1(X1),X1)|r2_hidden(esk7_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_157, plain, (v1_cantor_1(esk8_1(X1),X1)|r2_hidden(esk7_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_158, plain, (esk8_1(X1)=a_1_0_waybel14(X1)|r2_hidden(esk7_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_159, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|m1_subset_1(esk9_1(X2),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~r2_hidden(X1,esk11_2(X2,X3))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_160, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_161, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_125])])).
cnf(c_0_162, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[c_0_148, c_0_119])).
cnf(c_0_163, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_164, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_142])])).
cnf(c_0_165, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_142])])).
cnf(c_0_166, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_142])])).
cnf(c_0_167, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|m1_subset_1(esk9_1(X2),u1_struct_0(X2))|~epred1_1(X2)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_168, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_169, plain, (v1_tops_2(esk8_1(X1),X1)|r2_hidden(esk7_2(X1,X2),X2)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_170, plain, (v1_tops_2(esk8_1(X1),X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_171, plain, (v1_cantor_1(esk8_1(X1),X1)|r2_hidden(esk7_2(X1,X2),X2)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_172, plain, (v1_cantor_1(esk8_1(X1),X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_173, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|r2_hidden(esk7_2(X1,X2),X2)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_174, plain, (esk8_1(X1)=a_1_0_waybel14(X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
fof(c_0_175, plain, ![X27]:(v1_orders_2(k2_yellow_1(X27))&l1_orders_2(k2_yellow_1(X27))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_176, plain, ![X33]:(((~v7_struct_0(k2_yellow_1(k5_waybel11(X33)))|(~v2_pre_topc(X33)|~v3_orders_2(X33)|~v4_orders_2(X33)|~v5_orders_2(X33)|~v1_lattice3(X33)|~v2_lattice3(X33)|~v3_lattice3(X33)|~v2_waybel19(X33)|~l1_waybel_9(X33)))&(v1_orders_2(k2_yellow_1(k5_waybel11(X33)))|(~v2_pre_topc(X33)|~v3_orders_2(X33)|~v4_orders_2(X33)|~v5_orders_2(X33)|~v1_lattice3(X33)|~v2_lattice3(X33)|~v3_lattice3(X33)|~v2_waybel19(X33)|~l1_waybel_9(X33))))&(v3_lattice3(k2_yellow_1(k5_waybel11(X33)))|(~v2_pre_topc(X33)|~v3_orders_2(X33)|~v4_orders_2(X33)|~v5_orders_2(X33)|~v1_lattice3(X33)|~v2_lattice3(X33)|~v3_lattice3(X33)|~v2_waybel19(X33)|~l1_waybel_9(X33)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_waybel30])])])])).
cnf(c_0_177, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|m1_subset_1(esk9_1(X2),u1_struct_0(X2))|~epred1_1(X2)|~r2_hidden(X1,esk11_2(X2,X3))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_178, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_160, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_179, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(csr,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_180, negated_conjecture, (m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_163, c_0_119]), c_0_120]), c_0_104]), c_0_164]), c_0_165]), c_0_166])).
cnf(c_0_181, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_125])])).
cnf(c_0_182, negated_conjecture, (m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_183, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_184, plain, (v1_tops_2(esk8_1(X1),X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_185, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_186, plain, (v1_cantor_1(esk8_1(X1),X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_187, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_188, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_189, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_190, plain, (v3_lattice3(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_191, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_192, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_125])])).
cnf(c_0_193, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_142])])).
cnf(c_0_194, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180])])).
cnf(c_0_195, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(csr,[status(thm)],[c_0_181, c_0_162])).
cnf(c_0_196, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_182, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_197, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_183, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_198, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_199, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_185, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_200, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_201, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_187, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_202, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
fof(c_0_203, plain, ![X8]:((((((v3_orders_2(X8)|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v2_waybel_8(X8))|~l1_orders_2(X8))&(v4_orders_2(X8)|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v2_waybel_8(X8))|~l1_orders_2(X8)))&(v5_orders_2(X8)|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v2_waybel_8(X8))|~l1_orders_2(X8)))&(v1_lattice3(X8)|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v2_waybel_8(X8))|~l1_orders_2(X8)))&(v2_lattice3(X8)|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v2_waybel_8(X8))|~l1_orders_2(X8)))&(v3_waybel_3(X8)|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v2_waybel_8(X8))|~l1_orders_2(X8))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_waybel_8])])])).
fof(c_0_204, plain, ![X34]:(((v1_orders_2(k2_yellow_1(X34))&v3_orders_2(k2_yellow_1(X34)))&v4_orders_2(k2_yellow_1(X34)))&v5_orders_2(k2_yellow_1(X34))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_205, plain, ![X35]:((~v2_struct_0(k2_yellow_1(X35))|v1_xboole_0(X35))&(v1_orders_2(k2_yellow_1(X35))|v1_xboole_0(X35))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
cnf(c_0_206, plain, (l1_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_207, plain, (v3_lattice3(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~l1_waybel_9(X1)), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_208, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_209, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[c_0_192, c_0_162])).
cnf(c_0_210, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_211, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_180])])).
cnf(c_0_212, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_142])])).
cnf(c_0_213, plain, (v1_tops_2(esk8_1(X1),X1)|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_214, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_142])])).
cnf(c_0_215, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_198, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_216, plain, (v1_cantor_1(esk8_1(X1),X1)|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_217, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_142])])).
cnf(c_0_218, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_200, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_219, plain, (esk8_1(X1)=a_1_0_waybel14(X1)|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_220, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_142])])).
cnf(c_0_221, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_166])])).
cnf(c_0_222, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_202, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_223, plain, (v3_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v2_waybel_8(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_203])).
cnf(c_0_224, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_225, plain, (v4_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_226, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_204])).
fof(c_0_227, plain, ![X32]:(~v3_orders_2(X32)|~v4_orders_2(X32)|~v5_orders_2(X32)|~v1_lattice3(X32)|~v2_lattice3(X32)|~v3_lattice3(X32)|~l1_orders_2(X32)|~v1_xboole_0(k5_waybel11(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_waybel14])])])).
cnf(c_0_228, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_229, plain, (v2_struct_0(k2_yellow_1(X1))|v1_lattice3(k2_yellow_1(X1))|~v3_lattice3(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_206])])).
cnf(c_0_230, negated_conjecture, (v3_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_50]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59])])])).
cnf(c_0_231, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_232, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_180])])).
cnf(c_0_233, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_210, c_0_211])).
cnf(c_0_234, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_194])])).
cnf(c_0_235, plain, (v1_tops_2(esk8_1(X1),X1)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_236, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_194])])).
cnf(c_0_237, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_142])])).
cnf(c_0_238, plain, (v1_cantor_1(esk8_1(X1),X1)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_216])).
cnf(c_0_239, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_194])])).
cnf(c_0_240, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_142])])).
cnf(c_0_241, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_242, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221])])).
cnf(c_0_243, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_166])])).
cnf(c_0_244, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_142])])).
cnf(c_0_245, plain, (v3_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v2_waybel_8(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_246, plain, (v5_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_247, plain, (v4_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_248, plain, (v3_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_249, plain, (~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~v1_xboole_0(k5_waybel11(X1))), inference(split_conjunct,[status(thm)],[c_0_227])).
cnf(c_0_250, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_251, negated_conjecture, (v2_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))|v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230])])).
fof(c_0_252, plain, ![X7]:(~l1_orders_2(X7)|(~v1_lattice3(X7)|~v2_struct_0(X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_253, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_231, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_254, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233])])).
cnf(c_0_255, negated_conjecture, (m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_234, c_0_211])).
cnf(c_0_256, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_257, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_236, c_0_211])).
cnf(c_0_258, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_194])])).
cnf(c_0_259, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_260, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_239, c_0_211])).
cnf(c_0_261, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_194])])).
cnf(c_0_262, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_125]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_263, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_166])])).
cnf(c_0_264, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_242, c_0_243])).
cnf(c_0_265, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_221])])).
cnf(c_0_266, plain, (v3_waybel_3(k2_yellow_1(X1))|~v2_waybel_8(k2_yellow_1(X1))|~v2_lattice3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_206]), c_0_246]), c_0_247]), c_0_248])])])).
cnf(c_0_267, plain, (~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~v1_xboole_0(k5_waybel11(X1))), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_268, negated_conjecture, (v1_xboole_0(k5_waybel11(esk1_0))|v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251])])).
cnf(c_0_269, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_252])).
cnf(c_0_270, plain, (v2_struct_0(k2_yellow_1(X1))|v2_lattice3(k2_yellow_1(X1))|~v3_lattice3(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_206])])).
cnf(c_0_271, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_142])])).
cnf(c_0_272, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_254, c_0_255])).
cnf(c_0_273, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_256, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_274, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_257])])).
cnf(c_0_275, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_258, c_0_211])).
cnf(c_0_276, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_259, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_277, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_260])])).
cnf(c_0_278, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_261, c_0_211])).
cnf(c_0_279, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_262, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_280, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264])])).
cnf(c_0_281, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_265, c_0_243])).
fof(c_0_282, plain, ![X46]:(((v3_waybel30(X46)|~v3_waybel_3(X46)|(~v2_pre_topc(X46)|~v3_orders_2(X46)|~v4_orders_2(X46)|~v5_orders_2(X46)|~v1_lattice3(X46)|~v2_lattice3(X46)|~v3_lattice3(X46)|~v2_waybel19(X46)|~v2_waybel_2(X46)|~l1_waybel_9(X46)))&(v3_waybel_3(k2_yellow_1(k5_waybel11(X46)))|~v3_waybel_3(X46)|(~v2_pre_topc(X46)|~v3_orders_2(X46)|~v4_orders_2(X46)|~v5_orders_2(X46)|~v1_lattice3(X46)|~v2_lattice3(X46)|~v3_lattice3(X46)|~v2_waybel19(X46)|~v2_waybel_2(X46)|~l1_waybel_9(X46))))&(~v3_waybel30(X46)|~v3_waybel_3(k2_yellow_1(k5_waybel11(X46)))|v3_waybel_3(X46)|(~v2_pre_topc(X46)|~v3_orders_2(X46)|~v4_orders_2(X46)|~v5_orders_2(X46)|~v1_lattice3(X46)|~v2_lattice3(X46)|~v3_lattice3(X46)|~v2_waybel19(X46)|~v2_waybel_2(X46)|~l1_waybel_9(X46)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_waybel30])])])).
cnf(c_0_283, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0)|~v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))|~v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_142])])).
cnf(c_0_284, negated_conjecture, (v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_268]), c_0_70]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58])])])).
cnf(c_0_285, plain, (~v1_lattice3(X1)|~l1_orders_2(X1)|~v2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_286, negated_conjecture, (v2_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))|v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_230])])).
cnf(c_0_287, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272])])).
cnf(c_0_288, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_142])])).
cnf(c_0_289, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_274, c_0_275])).
cnf(c_0_290, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_142])])).
cnf(c_0_291, plain, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_277, c_0_278])).
cnf(c_0_292, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_142])])).
cnf(c_0_293, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_280, c_0_281])).
fof(c_0_294, plain, ![X17, X18]:(v2_struct_0(X17)|~v3_orders_2(X17)|~v4_orders_2(X17)|~v5_orders_2(X17)|~v3_waybel_3(X17)|~l1_orders_2(X17)|(~m1_yellow_9(X18,X17)|v3_waybel_3(X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc4_waybel19])])])])).
cnf(c_0_295, plain, (v3_waybel_3(X1)|~v3_waybel30(X1)|~v3_waybel_3(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_282])).
cnf(c_0_296, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0)|~v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_283, c_0_284])])])).
cnf(c_0_297, negated_conjecture, (v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_206]), c_0_284])])])).
fof(c_0_298, plain, ![X42, X43, X45]:((((v1_tops_2(esk3_2(X42,X43),X42)|~v3_waybel_3(X42)|~m1_subset_1(X43,u1_struct_0(X42))|(~v2_pre_topc(X42)|~v3_orders_2(X42)|~v4_orders_2(X42)|~v5_orders_2(X42)|~v1_lattice3(X42)|~v2_lattice3(X42)|~v3_lattice3(X42)|~v4_waybel11(X42)|~l1_waybel_9(X42)))&(v1_yellow_8(esk3_2(X42,X43),X42,X43)|~v3_waybel_3(X42)|~m1_subset_1(X43,u1_struct_0(X42))|(~v2_pre_topc(X42)|~v3_orders_2(X42)|~v4_orders_2(X42)|~v5_orders_2(X42)|~v1_lattice3(X42)|~v2_lattice3(X42)|~v3_lattice3(X42)|~v4_waybel11(X42)|~l1_waybel_9(X42))))&(m1_subset_1(esk3_2(X42,X43),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X42))))|~v3_waybel_3(X42)|~m1_subset_1(X43,u1_struct_0(X42))|(~v2_pre_topc(X42)|~v3_orders_2(X42)|~v4_orders_2(X42)|~v5_orders_2(X42)|~v1_lattice3(X42)|~v2_lattice3(X42)|~v3_lattice3(X42)|~v4_waybel11(X42)|~l1_waybel_9(X42))))&((v3_pre_topc(X45,X42)|(~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X42)))|~r2_hidden(X45,esk3_2(X42,X43)))|~v3_waybel_3(X42)|~m1_subset_1(X43,u1_struct_0(X42))|(~v2_pre_topc(X42)|~v3_orders_2(X42)|~v4_orders_2(X42)|~v5_orders_2(X42)|~v1_lattice3(X42)|~v2_lattice3(X42)|~v3_lattice3(X42)|~v4_waybel11(X42)|~l1_waybel_9(X42)))&(v2_waybel_0(X45,X42)|(~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X42)))|~r2_hidden(X45,esk3_2(X42,X43)))|~v3_waybel_3(X42)|~m1_subset_1(X43,u1_struct_0(X42))|(~v2_pre_topc(X42)|~v3_orders_2(X42)|~v4_orders_2(X42)|~v5_orders_2(X42)|~v1_lattice3(X42)|~v2_lattice3(X42)|~v3_lattice3(X42)|~v4_waybel11(X42)|~l1_waybel_9(X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_waybel14])])])])])).
cnf(c_0_299, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_287, c_0_194]), c_0_211])).
cnf(c_0_300, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289])])).
cnf(c_0_301, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_291])])).
cnf(c_0_302, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293])])).
cnf(c_0_303, plain, (v2_struct_0(X1)|v3_waybel_3(X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_294])).
cnf(c_0_304, plain, (v3_waybel_3(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)|~v3_waybel30(X1)|~v3_waybel_3(k2_yellow_1(k5_waybel11(X1)))), inference(evalgc,[status(thm)],[c_0_295])).
cnf(c_0_305, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_296, c_0_297])])])).
cnf(c_0_306, negated_conjecture, (v3_waybel30(esk1_0)|v2_waybel_8(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_307, plain, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v3_waybel_3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_308, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_299]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_309, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_300, c_0_194]), c_0_211])).
cnf(c_0_310, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_301, c_0_194]), c_0_211])).
cnf(c_0_311, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_302, c_0_221]), c_0_243])).
cnf(c_0_312, plain, (v3_waybel_3(X2)|v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1)), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_313, negated_conjecture, (v3_waybel_3(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel30(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_305]), c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59])])])).
cnf(c_0_314, negated_conjecture, (v2_waybel_8(esk1_0)|v3_waybel30(esk1_0)), inference(evalgc,[status(thm)],[c_0_306])).
cnf(c_0_315, negated_conjecture, (v3_waybel_3(esk1_0)|~v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_70]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58])])])).
cnf(c_0_316, plain, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1)|~v4_waybel11(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_317, negated_conjecture, (m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_308, c_0_119]), c_0_120]), c_0_104]), c_0_309]), c_0_310]), c_0_311])).
cnf(c_0_318, negated_conjecture, (v2_struct_0(esk1_0)|v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_waybel_3(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_69]), c_0_70]), c_0_56]), c_0_57]), c_0_58])])])).
cnf(c_0_319, negated_conjecture, (v3_waybel_3(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_313, c_0_314]), c_0_315])).
cnf(c_0_320, plain, (v1_tops_2(esk3_2(X1,X2),X1)|~v3_waybel_3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_321, plain, (v1_yellow_8(esk3_2(X1,X2),X1,X2)|~v3_waybel_3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_322, plain, (m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v1_tops_2(X3,X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_323, negated_conjecture, (m1_subset_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_317]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_324, negated_conjecture, (v2_struct_0(esk1_0)|v3_waybel_3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_318, c_0_319])])])).
cnf(c_0_325, plain, (v1_tops_2(esk3_2(X1,X2),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1)|~v4_waybel11(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_320])).
cnf(c_0_326, plain, (v1_yellow_8(esk3_2(X1,X2),X1,X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1)|~v4_waybel11(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_321])).
cnf(c_0_327, plain, (m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~v1_tops_2(X3,X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))), inference(evalgc,[status(thm)],[c_0_322])).
cnf(c_0_328, negated_conjecture, (m1_subset_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_323, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_329, negated_conjecture, (v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_317]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_330, negated_conjecture, (v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_317]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_331, plain, (X1=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))|m1_subset_1(esk10_2(X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~v1_tops_2(X3,X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_332, negated_conjecture, (m1_subset_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_311]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_333, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_328])])).
cnf(c_0_334, negated_conjecture, (v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_329, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_335, negated_conjecture, (v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_330, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_336, plain, (m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|r2_hidden(esk10_2(X1,X3),X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v1_tops_2(X3,X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_337, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|m1_subset_1(esk10_2(X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~epred1_1(X2)|~v1_tops_2(X3,X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_331])).
cnf(c_0_338, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_332, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_339, negated_conjecture, (v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_311]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_340, negated_conjecture, (v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_311]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_341, plain, (m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_333, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_342, plain, (r2_hidden(esk10_2(X1,X3),X3)|m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~v1_tops_2(X3,X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))), inference(evalgc,[status(thm)],[c_0_336])).
cnf(c_0_343, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_328])])).
cnf(c_0_344, plain, (X1=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))|r2_hidden(esk10_2(X2,X3),X3)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~v1_tops_2(X3,X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_345, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_338])])).
cnf(c_0_346, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_339, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_347, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_340, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_348, plain, (v2_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,esk3_2(X2,X3))|~v3_waybel_3(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~v4_waybel11(X2)|~l1_waybel_9(X2)), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_349, negated_conjecture, (m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_125])])).
cnf(c_0_350, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_328])])).
cnf(c_0_351, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_343, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_352, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|r2_hidden(esk10_2(X2,X3),X3)|~epred1_1(X2)|~v1_tops_2(X3,X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_344])).
cnf(c_0_353, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_345, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_354, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_338])])).
cnf(c_0_355, plain, (v2_waybel_0(X1,X2)|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~l1_waybel_9(X2)|~v3_waybel_3(X2)|~v4_waybel11(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk3_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_348])).
cnf(c_0_356, negated_conjecture, (m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_180])])).
cnf(c_0_357, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_350, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_358, plain, (v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,esk3_2(X2,X3))|~v3_waybel_3(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~v4_waybel11(X2)|~l1_waybel_9(X2)), inference(split_conjunct,[status(thm)],[c_0_298])).
cnf(c_0_359, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_125])])).
cnf(c_0_360, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_328])])).
cnf(c_0_361, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_125])])).
cnf(c_0_362, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_338])])).
cnf(c_0_363, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_354, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_364, plain, (m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v3_pre_topc(esk10_2(X1,X3),X1)|~v2_waybel_0(esk10_2(X1,X3),X1)|~v1_tops_2(X3,X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~epred1_1(X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_365, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_356]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_366, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_125])])).
cnf(c_0_367, plain, (v3_pre_topc(X1,X2)|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~l1_waybel_9(X2)|~v3_waybel_3(X2)|~v4_waybel11(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk3_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_358])).
cnf(c_0_368, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_180])])).
cnf(c_0_369, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_360, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_370, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_166])])).
cnf(c_0_371, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_362, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_372, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_125])])).
cnf(c_0_373, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_338])])).
cnf(c_0_374, plain, (m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~v1_tops_2(X3,X1)|~v3_pre_topc(esk10_2(X1,X3),X1)|~v2_waybel_0(esk10_2(X1,X3),X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1))))), inference(evalgc,[status(thm)],[c_0_364])).
cnf(c_0_375, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_365, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_376, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_180])])).
cnf(c_0_377, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_356]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_378, plain, (X1=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~v3_pre_topc(esk10_2(X2,X3),X2)|~v2_waybel_0(esk10_2(X2,X3),X2)|~v1_tops_2(X3,X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_379, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_368]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_380, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_125])])).
cnf(c_0_381, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_370]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_382, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_125])])).
cnf(c_0_383, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_166])])).
cnf(c_0_384, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_373, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_385, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_328])])).
cnf(c_0_386, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_376])])).
cnf(c_0_387, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_377, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_388, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|~epred1_1(X2)|~v1_tops_2(X3,X2)|~v3_pre_topc(esk10_2(X2,X3),X2)|~v2_waybel_0(esk10_2(X2,X3),X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_378])).
cnf(c_0_389, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_379, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_390, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_180])])).
cnf(c_0_391, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_368]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_392, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_381, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_393, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_166])])).
cnf(c_0_394, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_370]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_395, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_383]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_396, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_125])])).
cnf(c_0_397, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_385, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_398, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_386, c_0_317])).
cnf(c_0_399, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_376])])).
cnf(c_0_400, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_328])])).
cnf(c_0_401, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_390])])).
cnf(c_0_402, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_391, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_403, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_338])])).
cnf(c_0_404, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_393])])).
cnf(c_0_405, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_394, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_406, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_395, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_407, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_166])])).
cnf(c_0_408, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_383]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_409, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|m1_subset_1(esk10_2(X2,X4),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~r2_hidden(X1,esk11_2(X2,X3))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~v1_tops_2(X4,X2)|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_410, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_398])])).
cnf(c_0_411, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_399, c_0_317])).
cnf(c_0_412, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_400, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_413, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_401, c_0_317])).
cnf(c_0_414, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_390])])).
cnf(c_0_415, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_403, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_416, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_404, c_0_311])).
cnf(c_0_417, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_405, c_0_393])])).
cnf(c_0_418, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_338])])).
cnf(c_0_419, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407])])).
cnf(c_0_420, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_408, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_421, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|m1_subset_1(esk10_2(X2,X4),k1_zfmisc_1(u1_struct_0(X2)))|~epred1_1(X2)|~v1_tops_2(X4,X2)|~r2_hidden(X1,esk11_2(X2,X3))|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_409])).
cnf(c_0_422, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[c_0_410, c_0_411])).
cnf(c_0_423, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_413])])).
cnf(c_0_424, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_414, c_0_317])).
cnf(c_0_425, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_416])])).
cnf(c_0_426, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_417, c_0_311])).
cnf(c_0_427, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_418, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_428, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_419, c_0_311])).
cnf(c_0_429, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_407])])).
cnf(c_0_430, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_328])])).
cnf(c_0_431, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_125])])).
cnf(c_0_432, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[c_0_423, c_0_424])).
cnf(c_0_433, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|r2_hidden(esk10_2(X2,X4),X4)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~r2_hidden(X1,esk11_2(X2,X3))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~v1_tops_2(X4,X2)|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_434, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[c_0_425, c_0_426])).
cnf(c_0_435, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_428])])).
cnf(c_0_436, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_429, c_0_311])).
cnf(c_0_437, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_430, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_438, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_180])])).
cnf(c_0_439, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_125])])).
cnf(c_0_440, plain, (r2_hidden(esk10_2(X2,X4),X4)|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|~epred1_1(X2)|~v1_tops_2(X4,X2)|~r2_hidden(X1,esk11_2(X2,X3))|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_433])).
cnf(c_0_441, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_338])])).
cnf(c_0_442, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_125])])).
cnf(c_0_443, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))), inference(csr,[status(thm)],[c_0_435, c_0_436])).
cnf(c_0_444, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_125])])).
cnf(c_0_445, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_438])])).
cnf(c_0_446, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_439, c_0_180])])).
cnf(c_0_447, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_328])])).
cnf(c_0_448, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_441, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_449, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_166])])).
cnf(c_0_450, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_125])])).
cnf(c_0_451, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_180])])).
cnf(c_0_452, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_445, c_0_446])).
cnf(c_0_453, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_438])])).
cnf(c_0_454, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_447, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_455, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_438])])).
cnf(c_0_456, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_438])])).
cnf(c_0_457, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_448, c_0_125])])).
cnf(c_0_458, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_449])])).
cnf(c_0_459, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_166])])).
cnf(c_0_460, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_338])])).
cnf(c_0_461, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_452])])).
cnf(c_0_462, negated_conjecture, (m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_453, c_0_446])).
cnf(c_0_463, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_125])])).
cnf(c_0_464, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_455, c_0_446])).
cnf(c_0_465, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_438])])).
cnf(c_0_466, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_456, c_0_446])).
cnf(c_0_467, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_438])])).
cnf(c_0_468, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_166])])).
cnf(c_0_469, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_458, c_0_459])).
cnf(c_0_470, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_449])])).
cnf(c_0_471, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_460, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_472, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_461, c_0_462])).
cnf(c_0_473, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_463, c_0_180])])).
cnf(c_0_474, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_464])])).
cnf(c_0_475, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_465, c_0_446])).
cnf(c_0_476, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_466])])).
cnf(c_0_477, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_467, c_0_446])).
cnf(c_0_478, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_469])])).
cnf(c_0_479, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_470, c_0_459])).
cnf(c_0_480, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_125])])).
cnf(c_0_481, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~r2_hidden(X1,esk11_2(X2,X3))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~v3_pre_topc(esk10_2(X2,X4),X2)|~v2_waybel_0(esk10_2(X2,X4),X2)|~v1_tops_2(X4,X2)|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~epred1_1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_482, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_472]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_483, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_452])])).
cnf(c_0_484, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_474, c_0_475])).
cnf(c_0_485, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_476, c_0_477])).
cnf(c_0_486, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_478, c_0_479])).
cnf(c_0_487, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_166])])).
cnf(c_0_488, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|~epred1_1(X2)|~v1_tops_2(X4,X2)|~r2_hidden(X1,esk11_2(X2,X3))|~v3_pre_topc(esk10_2(X2,X4),X2)|~v2_waybel_0(esk10_2(X2,X4),X2)|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))), inference(evalgc,[status(thm)],[c_0_481])).
cnf(c_0_489, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_482, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_490, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_483, c_0_462])).
cnf(c_0_491, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_472]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_492, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_484]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_493, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_464])])).
cnf(c_0_494, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_485]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_495, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_466])])).
cnf(c_0_496, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_486]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_497, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_469])])).
cnf(c_0_498, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_488, c_0_328])])).
cnf(c_0_499, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_490])])).
cnf(c_0_500, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_491, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_501, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_492, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_502, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_493, c_0_475])).
cnf(c_0_503, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_484]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_504, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_494, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_505, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_495, c_0_477])).
cnf(c_0_506, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_485]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_507, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_496, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_508, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_497, c_0_479])).
cnf(c_0_509, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_486]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_510, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_498, c_0_334]), c_0_335]), c_0_162])).
cnf(c_0_511, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_499, c_0_317])).
cnf(c_0_512, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_490])])).
cnf(c_0_513, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_501, c_0_502])])).
cnf(c_0_514, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_503, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_515, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_504, c_0_505])])).
cnf(c_0_516, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_506, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_517, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_488, c_0_338])])).
cnf(c_0_518, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_508])])).
cnf(c_0_519, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_509, c_0_119]), c_0_120]), c_0_104]), c_0_324])).
cnf(c_0_520, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_511])])).
cnf(c_0_521, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_512, c_0_317])).
cnf(c_0_522, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_513, c_0_317])).
cnf(c_0_523, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_514, c_0_502])])).
cnf(c_0_524, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_515, c_0_317])).
cnf(c_0_525, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_516, c_0_505])])).
cnf(c_0_526, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_517, c_0_346]), c_0_347]), c_0_162])).
cnf(c_0_527, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_518, c_0_311])).
cnf(c_0_528, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_519, c_0_508])])).
cnf(c_0_529, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[c_0_520, c_0_521])).
cnf(c_0_530, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_522])])).
cnf(c_0_531, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_523, c_0_317])).
cnf(c_0_532, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_524])])).
cnf(c_0_533, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_525, c_0_317])).
cnf(c_0_534, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_527])])).
cnf(c_0_535, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_528, c_0_311])).
cnf(c_0_536, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_529, c_0_125])])).
cnf(c_0_537, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[c_0_530, c_0_531])).
cnf(c_0_538, plain, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[c_0_532, c_0_533])).
cnf(c_0_539, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(csr,[status(thm)],[c_0_534, c_0_535])).
cnf(c_0_540, plain, (m1_yellow_9(X1,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~l1_waybel_9(X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_541, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_536, c_0_180])])).
cnf(c_0_542, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_125])])).
cnf(c_0_543, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_538, c_0_125])])).
cnf(c_0_544, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_539, c_0_125])])).
fof(c_0_545, plain, ![X36, X37]:(v2_struct_0(X36)|~v3_orders_2(X36)|~v5_orders_2(X36)|~l1_orders_2(X36)|(v2_struct_0(X37)|~v3_orders_2(X37)|~v5_orders_2(X37)|~l1_orders_2(X37)|(g1_orders_2(u1_struct_0(X36),u1_orders_2(X36))!=g1_orders_2(u1_struct_0(X37),u1_orders_2(X37))|~v2_waybel_8(X36)|v2_waybel_8(X37)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_8])])])])).
cnf(c_0_546, plain, (m1_yellow_9(X1,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~l1_waybel_9(X1)|~l1_orders_2(X2)), inference(evalgc,[status(thm)],[c_0_540])).
cnf(c_0_547, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_541, c_0_452])])).
cnf(c_0_548, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_180])])).
cnf(c_0_549, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_180])])).
cnf(c_0_550, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_544, c_0_166])])).
cnf(c_0_551, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_waybel_8(X2)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v3_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~v2_waybel_8(X1)), inference(split_conjunct,[status(thm)],[c_0_545])).
cnf(c_0_552, negated_conjecture, (m1_yellow_9(X1,o_1_14_waybel30(esk1_0))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(o_1_14_waybel30(esk1_0)),u1_orders_2(o_1_14_waybel30(esk1_0)))|~l1_waybel_9(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_546, c_0_91])])).
cnf(c_0_553, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_547, c_0_462])).
cnf(c_0_554, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_548, c_0_464])])).
cnf(c_0_555, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_549, c_0_466])])).
cnf(c_0_556, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_469])])).
cnf(c_0_557, plain, (v2_waybel_8(X2)|v2_struct_0(X2)|v2_struct_0(X1)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~v3_orders_2(X2)|~v3_orders_2(X1)|~v5_orders_2(X2)|~v5_orders_2(X1)|~v2_waybel_8(X1)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_551])).
cnf(c_0_558, negated_conjecture, (m1_yellow_9(X1,o_1_14_waybel30(esk1_0))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_waybel_9(X1)), inference(rw,[status(thm)],[c_0_552, c_0_99])).
cnf(c_0_559, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_553])])).
cnf(c_0_560, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_554, c_0_475])).
cnf(c_0_561, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_555, c_0_477])).
cnf(c_0_562, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_556, c_0_479])).
cnf(c_0_563, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_orders_2(X1)|~v2_waybel_8(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_557, c_0_70]), c_0_56]), c_0_58])])])).
cnf(c_0_564, negated_conjecture, (m1_yellow_9(esk1_0,o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_558]), c_0_50])])])).
cnf(c_0_565, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_559, c_0_438]), c_0_446])).
cnf(c_0_566, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_560])])).
cnf(c_0_567, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_561])])).
cnf(c_0_568, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_562])])).
cnf(c_0_569, negated_conjecture, (v2_struct_0(esk1_0)|v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v2_waybel_8(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_563, c_0_99]), c_0_91]), c_0_100]), c_0_102])])])).
cnf(c_0_570, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_564]), c_0_91])])])).
cnf(c_0_571, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_565]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_572, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_566, c_0_438]), c_0_446])).
cnf(c_0_573, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_567, c_0_438]), c_0_446])).
cnf(c_0_574, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_568, c_0_449]), c_0_459])).
cnf(c_0_575, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v2_waybel_8(o_1_14_waybel30(esk1_0))), inference(csr,[status(thm)],[c_0_569, c_0_570])).
cnf(c_0_576, negated_conjecture, (v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_571, c_0_119]), c_0_120]), c_0_104]), c_0_572]), c_0_573]), c_0_574])).
cnf(c_0_577, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(o_1_14_waybel30(esk1_0))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(o_1_14_waybel30(esk1_0)),u1_orders_2(o_1_14_waybel30(esk1_0)))|~l1_orders_2(X1)|~v2_waybel_8(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_557, c_0_91]), c_0_100]), c_0_102])])])).
cnf(c_0_578, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_575, c_0_576])])).
fof(c_0_579, plain, ![X48]:((((v1_tops_2(esk4_1(X48),X48)|~v2_waybel_8(X48)|(~v2_pre_topc(X48)|~v3_orders_2(X48)|~v4_orders_2(X48)|~v5_orders_2(X48)|~v1_lattice3(X48)|~v2_lattice3(X48)|~v3_lattice3(X48)|~v4_waybel11(X48)|~l1_waybel_9(X48)))&(v1_cantor_1(esk4_1(X48),X48)|~v2_waybel_8(X48)|(~v2_pre_topc(X48)|~v3_orders_2(X48)|~v4_orders_2(X48)|~v5_orders_2(X48)|~v1_lattice3(X48)|~v2_lattice3(X48)|~v3_lattice3(X48)|~v4_waybel11(X48)|~l1_waybel_9(X48))))&(m1_subset_1(esk4_1(X48),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X48))))|~v2_waybel_8(X48)|(~v2_pre_topc(X48)|~v3_orders_2(X48)|~v4_orders_2(X48)|~v5_orders_2(X48)|~v1_lattice3(X48)|~v2_lattice3(X48)|~v3_lattice3(X48)|~v4_waybel11(X48)|~l1_waybel_9(X48))))&(esk4_1(X48)=a_1_0_waybel14(X48)|~v2_waybel_8(X48)|(~v2_pre_topc(X48)|~v3_orders_2(X48)|~v4_orders_2(X48)|~v5_orders_2(X48)|~v1_lattice3(X48)|~v2_lattice3(X48)|~v3_lattice3(X48)|~v4_waybel11(X48)|~l1_waybel_9(X48)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t42_waybel14])])])])).
cnf(c_0_580, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_struct_0(X1)|v2_waybel_8(o_1_14_waybel30(esk1_0))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_orders_2(X1)|~v2_waybel_8(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)), inference(rw,[status(thm)],[c_0_577, c_0_99])).
cnf(c_0_581, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_578, c_0_570])).
cnf(c_0_582, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_waybel_8(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_579])).
cnf(c_0_583, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|~v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_580]), c_0_70]), c_0_56]), c_0_58])])])).
cnf(c_0_584, negated_conjecture, (v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_581])])).
cnf(c_0_585, plain, (esk4_1(X1)=a_1_0_waybel14(X1)|~v2_waybel_8(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_579])).
cnf(c_0_586, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1)), inference(evalgc,[status(thm)],[c_0_582])).
cnf(c_0_587, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(o_1_14_waybel30(esk1_0))|~v2_waybel_8(esk1_0)), inference(csr,[status(thm)],[c_0_583, c_0_570])).
cnf(c_0_588, negated_conjecture, (v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_584]), c_0_70]), c_0_55])])])).
cnf(c_0_589, plain, (v1_tops_2(esk4_1(X1),X1)|~v2_waybel_8(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_579])).
cnf(c_0_590, plain, (v1_cantor_1(esk4_1(X1),X1)|~v2_waybel_8(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_579])).
cnf(c_0_591, plain, (esk4_1(X1)=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1)), inference(evalgc,[status(thm)],[c_0_585])).
cnf(c_0_592, negated_conjecture, (m1_subset_1(esk4_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_108])])).
cnf(c_0_593, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_587, c_0_588])])])).
cnf(c_0_594, plain, (v1_tops_2(esk4_1(X1),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1)), inference(evalgc,[status(thm)],[c_0_589])).
cnf(c_0_595, plain, (v1_cantor_1(esk4_1(X1),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1)), inference(evalgc,[status(thm)],[c_0_590])).
cnf(c_0_596, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_591, c_0_108])])).
fof(c_0_597, plain, ![X6]:(((((((((v2_pre_topc(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6))&(v3_orders_2(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6)))&(v4_orders_2(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6)))&(v5_orders_2(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6)))&(v1_lattice3(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6)))&(v2_lattice3(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6)))&(v3_lattice3(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6)))&(v2_waybel19(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6)))&(v3_waybel30(X6)|(~v2_pre_topc(X6)|~v3_orders_2(X6)|~v4_orders_2(X6)|~v5_orders_2(X6)|~v1_lattice3(X6)|~v2_lattice3(X6)|~v3_lattice3(X6)|~v2_waybel19(X6)|~v3_waybel_3(X6))|~l1_waybel_9(X6))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc10_waybel30])])])).
fof(c_0_598, plain, ![X50, X51, X52, X54]:((v2_waybel_8(k2_yellow_1(k5_waybel11(X50)))|(~v1_tops_2(X51,X50)|~v1_cantor_1(X51,X50)|~m1_subset_1(X51,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X50))))|X51!=a_1_0_waybel14(X50))|(~v2_pre_topc(X50)|~v3_orders_2(X50)|~v4_orders_2(X50)|~v5_orders_2(X50)|~v1_lattice3(X50)|~v2_lattice3(X50)|~v3_lattice3(X50)|~v4_waybel11(X50)|~l1_waybel_9(X50)))&(((m1_subset_1(esk5_2(X50,X52),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X50)))))|~m1_subset_1(X52,u1_struct_0(k2_yellow_1(k5_waybel11(X50))))|(~v1_tops_2(X51,X50)|~v1_cantor_1(X51,X50)|~m1_subset_1(X51,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X50))))|X51!=a_1_0_waybel14(X50))|(~v2_pre_topc(X50)|~v3_orders_2(X50)|~v4_orders_2(X50)|~v5_orders_2(X50)|~v1_lattice3(X50)|~v2_lattice3(X50)|~v3_lattice3(X50)|~v4_waybel11(X50)|~l1_waybel_9(X50)))&(X52=k1_yellow_0(k2_yellow_1(k5_waybel11(X50)),esk5_2(X50,X52))|~m1_subset_1(X52,u1_struct_0(k2_yellow_1(k5_waybel11(X50))))|(~v1_tops_2(X51,X50)|~v1_cantor_1(X51,X50)|~m1_subset_1(X51,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X50))))|X51!=a_1_0_waybel14(X50))|(~v2_pre_topc(X50)|~v3_orders_2(X50)|~v4_orders_2(X50)|~v5_orders_2(X50)|~v1_lattice3(X50)|~v2_lattice3(X50)|~v3_lattice3(X50)|~v4_waybel11(X50)|~l1_waybel_9(X50))))&(~m1_subset_1(X54,u1_struct_0(k2_yellow_1(k5_waybel11(X50))))|(~r2_hidden(X54,esk5_2(X50,X52))|v6_waybel_6(X54,k2_yellow_1(k5_waybel11(X50))))|~m1_subset_1(X52,u1_struct_0(k2_yellow_1(k5_waybel11(X50))))|(~v1_tops_2(X51,X50)|~v1_cantor_1(X51,X50)|~m1_subset_1(X51,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X50))))|X51!=a_1_0_waybel14(X50))|(~v2_pre_topc(X50)|~v3_orders_2(X50)|~v4_orders_2(X50)|~v5_orders_2(X50)|~v1_lattice3(X50)|~v2_lattice3(X50)|~v3_lattice3(X50)|~v4_waybel11(X50)|~l1_waybel_9(X50))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_waybel14])])])])])).
cnf(c_0_599, negated_conjecture, (m1_subset_1(esk4_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_592, c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])).
cnf(c_0_600, negated_conjecture, (v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_593])])).
cnf(c_0_601, negated_conjecture, (v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_594, c_0_108])])).
cnf(c_0_602, negated_conjecture, (v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_595, c_0_108])])).
cnf(c_0_603, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_596, c_0_79]), c_0_100]), c_0_101]), c_0_102])])).
cnf(c_0_604, plain, (v3_waybel30(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v3_waybel_3(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_597])).
cnf(c_0_605, plain, (v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|X2!=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v4_waybel11(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_598])).
cnf(c_0_606, negated_conjecture, (m1_subset_1(esk4_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_599, c_0_600])])).
cnf(c_0_607, negated_conjecture, (v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_601, c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])).
cnf(c_0_608, negated_conjecture, (v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_602, c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])).
cnf(c_0_609, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_603, c_0_123])])])).
cnf(c_0_610, negated_conjecture, (~v2_waybel_8(esk1_0)|~v3_waybel30(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_611, plain, (v3_waybel30(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1)), inference(evalgc,[status(thm)],[c_0_604])).
cnf(c_0_612, plain, (v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|X2!=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_605])).
cnf(c_0_613, negated_conjecture, (m1_subset_1(esk4_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_606, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_614, negated_conjecture, (v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_607, c_0_600])])).
cnf(c_0_615, negated_conjecture, (v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_608, c_0_600])])).
cnf(c_0_616, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_609, c_0_600])])).
cnf(c_0_617, negated_conjecture, (~v2_waybel_8(esk1_0)|~v3_waybel30(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[c_0_610])).
cnf(c_0_618, negated_conjecture, (v3_waybel30(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_611, c_0_319]), c_0_50]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58]), c_0_59])])])).
cnf(c_0_619, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|esk4_1(o_1_14_waybel30(esk1_0))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_612, c_0_613]), c_0_108]), c_0_79]), c_0_100]), c_0_101]), c_0_102]), c_0_123])])])).
cnf(c_0_620, negated_conjecture, (v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_614, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_621, negated_conjecture, (v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_615, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_622, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_616, c_0_119]), c_0_120]), c_0_104])).
cnf(c_0_623, negated_conjecture, (~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_waybel_8(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_617, c_0_618])])])).
cnf(c_0_624, negated_conjecture, (v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_619, c_0_119]), c_0_120]), c_0_104]), c_0_620]), c_0_621]), c_0_622])).
cnf(c_0_625, negated_conjecture, (~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_623, c_0_588])])])).
cnf(c_0_626, negated_conjecture, (v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_624, c_0_125]), c_0_625])])).
cnf(c_0_627, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_626]), c_0_70]), c_0_55])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 628
# Proof object clause steps            : 563
# Proof object formula steps           : 65
# Proof object conjectures             : 354
# Proof object clause conjectures      : 351
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 81
# Proof object initial formulas used   : 30
# Proof object generating inferences   : 246
# Proof object simplifying inferences  : 792
# Training examples: 334 positive, 567 negative
# Training: Positive examples begin
cnf(c_0_617, negated_conjecture, (~v2_waybel_8(esk1_0)|~v3_waybel30(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_304, plain, (v3_waybel_3(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)|~v3_waybel30(X1)|~v3_waybel_3(k2_yellow_1(k5_waybel11(X1))))).# trainpos
cnf(c_0_314, negated_conjecture, (v2_waybel_8(esk1_0)|v3_waybel30(esk1_0))).# trainpos
cnf(c_0_59, negated_conjecture, (v2_pre_topc(esk1_0))).# trainpos
cnf(c_0_58, negated_conjecture, (v3_orders_2(esk1_0))).# trainpos
cnf(c_0_57, negated_conjecture, (v4_orders_2(esk1_0))).# trainpos
cnf(c_0_56, negated_conjecture, (v5_orders_2(esk1_0))).# trainpos
cnf(c_0_55, negated_conjecture, (v1_lattice3(esk1_0))).# trainpos
cnf(c_0_54, negated_conjecture, (v2_lattice3(esk1_0))).# trainpos
cnf(c_0_53, negated_conjecture, (v3_lattice3(esk1_0))).# trainpos
cnf(c_0_52, negated_conjecture, (v2_waybel19(esk1_0))).# trainpos
cnf(c_0_51, negated_conjecture, (v2_waybel_2(esk1_0))).# trainpos
cnf(c_0_50, negated_conjecture, (l1_waybel_9(esk1_0))).# trainpos
cnf(c_0_285, plain, (~v1_lattice3(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).# trainpos
cnf(c_0_207, plain, (v3_lattice3(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~l1_waybel_9(X1))).# trainpos
cnf(c_0_142, negated_conjecture, (v2_waybel_8(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_68, plain, (l1_waybel_9(X2)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_89, plain, (v3_orders_2(X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_88, plain, (v4_orders_2(X2)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_87, plain, (v5_orders_2(X2)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_105, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~v2_struct_0(X2)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_60, plain, (l1_orders_2(X1)|~l1_waybel_9(X1))).# trainpos
cnf(c_0_250, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1)))).# trainpos
cnf(c_0_267, plain, (~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~v1_xboole_0(k5_waybel11(X1)))).# trainpos
cnf(c_0_248, plain, (v3_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_247, plain, (v4_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_246, plain, (v5_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_206, plain, (l1_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_92, plain, (v3_lattice3(X2)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_90, plain, (v1_lattice3(X1)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_94, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_107, plain, (epred1_1(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1))).# trainpos
cnf(c_0_116, plain, (v2_pre_topc(X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)|~v4_waybel11(X2)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_149, plain, (v2_waybel_8(X1)|X2!=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_312, plain, (v3_waybel_3(X2)|v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~m1_yellow_9(X2,X1))).# trainpos
cnf(c_0_245, plain, (v3_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v2_waybel_8(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_591, plain, (esk4_1(X1)=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1))).# trainpos
cnf(c_0_611, plain, (v3_waybel30(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1))).# trainpos
cnf(c_0_96, plain, (v4_waybel11(o_1_14_waybel30(X1))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1))).# trainpos
cnf(c_0_612, plain, (v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|X2!=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_367, plain, (v3_pre_topc(X1,X2)|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~l1_waybel_9(X2)|~v3_waybel_3(X2)|~v4_waybel11(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk3_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_241, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_355, plain, (v2_waybel_0(X1,X2)|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~l1_waybel_9(X2)|~v3_waybel_3(X2)|~v4_waybel11(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk3_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_235, plain, (v1_tops_2(esk8_1(X1),X1)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_49, plain, (m1_yellow_9(o_1_14_waybel30(X1),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1))).# trainpos
cnf(c_0_594, plain, (v1_tops_2(esk4_1(X1),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1))).# trainpos
cnf(c_0_595, plain, (v1_cantor_1(esk4_1(X1),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1))).# trainpos
cnf(c_0_557, plain, (v2_waybel_8(X2)|v2_struct_0(X2)|v2_struct_0(X1)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~v3_orders_2(X2)|~v3_orders_2(X1)|~v5_orders_2(X2)|~v5_orders_2(X1)|~v2_waybel_8(X1)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_238, plain, (v1_cantor_1(esk8_1(X1),X1)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_546, plain, (m1_yellow_9(X1,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~l1_waybel_9(X1)|~l1_orders_2(X2))).# trainpos
cnf(c_0_85, plain, (k5_waybel11(X1)=k5_waybel11(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~v3_orders_2(X2)|~v3_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v5_orders_2(X2)|~v5_orders_2(X1)|~v1_lattice3(X2)|~v1_lattice3(X1)|~v2_lattice3(X2)|~v2_lattice3(X1)|~v3_lattice3(X2)|~v3_lattice3(X1)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_167, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|m1_subset_1(esk9_1(X2),u1_struct_0(X2))|~epred1_1(X2)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_208, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~v6_waybel_6(esk7_2(X1,X2),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_586, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v2_waybel_8(X1)|~v4_waybel11(X1))).# trainpos
cnf(c_0_134, plain, (v1_tops_2(esk8_1(X1),X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1))))).# trainpos
cnf(c_0_135, plain, (v1_cantor_1(esk8_1(X1),X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1))))).# trainpos
cnf(c_0_325, plain, (v1_tops_2(esk3_2(X1,X2),X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1)|~v4_waybel11(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_136, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1))))).# trainpos
cnf(c_0_488, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|~epred1_1(X2)|~v1_tops_2(X4,X2)|~r2_hidden(X1,esk11_2(X2,X3))|~v3_pre_topc(esk10_2(X2,X4),X2)|~v2_waybel_0(esk10_2(X2,X4),X2)|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_421, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|m1_subset_1(esk10_2(X2,X4),k1_zfmisc_1(u1_struct_0(X2)))|~epred1_1(X2)|~v1_tops_2(X4,X2)|~r2_hidden(X1,esk11_2(X2,X3))|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_177, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|m1_subset_1(esk9_1(X2),u1_struct_0(X2))|~epred1_1(X2)|~r2_hidden(X1,esk11_2(X2,X3))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_173, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|r2_hidden(esk7_2(X1,X2),X2)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_147, plain, (m1_subset_1(esk9_1(X1),u1_struct_0(X1))|m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))).# trainpos
cnf(c_0_326, plain, (v1_yellow_8(esk3_2(X1,X2),X1,X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1)|~v4_waybel11(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_124, plain, (m1_subset_1(esk6_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1))))).# trainpos
cnf(c_0_316, plain, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1)|~v4_waybel11(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_74, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))|~l1_waybel_9(X1)|~l1_orders_2(X2)|~m1_yellow_9(X1,X2))).# trainpos
cnf(c_0_388, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|~epred1_1(X2)|~v1_tops_2(X3,X2)|~v3_pre_topc(esk10_2(X2,X3),X2)|~v2_waybel_0(esk10_2(X2,X3),X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_440, plain, (r2_hidden(esk10_2(X2,X4),X4)|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|~epred1_1(X2)|~v1_tops_2(X4,X2)|~r2_hidden(X1,esk11_2(X2,X3))|~v1_yellow_8(X4,X2,esk9_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_188, plain, (a_1_0_waybel14(X1)=esk8_1(X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_337, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|m1_subset_1(esk10_2(X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~epred1_1(X2)|~v1_tops_2(X3,X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_374, plain, (m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~v1_tops_2(X3,X1)|~v3_pre_topc(esk10_2(X1,X3),X1)|~v2_waybel_0(esk10_2(X1,X3),X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))).# trainpos
cnf(c_0_327, plain, (m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~v1_tops_2(X3,X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))).# trainpos
cnf(c_0_169, plain, (v1_tops_2(esk8_1(X1),X1)|r2_hidden(esk7_2(X1,X2),X2)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_171, plain, (v1_cantor_1(esk8_1(X1),X1)|r2_hidden(esk7_2(X1,X2),X2)|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_146, plain, (r2_hidden(esk7_2(X1,X2),X2)|m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_184, plain, (v1_tops_2(esk8_1(X1),X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_186, plain, (v1_cantor_1(esk8_1(X1),X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_352, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk11_2(X2,X1))=X1|r2_hidden(esk10_2(X2,X3),X3)|~epred1_1(X2)|~v1_tops_2(X3,X2)|~v1_yellow_8(X3,X2,esk9_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).# trainpos
cnf(c_0_168, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|m1_subset_1(esk7_2(X1,X2),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk6_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X2)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v2_waybel_8(k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).# trainpos
cnf(c_0_342, plain, (r2_hidden(esk10_2(X1,X3),X3)|m1_subset_1(esk11_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|~epred1_1(X1)|~v1_tops_2(X3,X1)|~v1_yellow_8(X3,X1,esk9_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))).# trainpos
cnf(c_0_70, negated_conjecture, (l1_orders_2(esk1_0))).# trainpos
cnf(c_0_315, negated_conjecture, (v3_waybel_3(esk1_0)|~v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_230, negated_conjecture, (v3_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_108, negated_conjecture, (v4_waybel11(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_69, negated_conjecture, (m1_yellow_9(o_1_14_waybel30(esk1_0),esk1_0))).# trainpos
cnf(c_0_113, negated_conjecture, (v2_struct_0(esk1_0)|~v2_struct_0(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_318, negated_conjecture, (v2_struct_0(esk1_0)|v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_waybel_3(esk1_0))).# trainpos
cnf(c_0_79, negated_conjecture, (l1_waybel_9(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_102, negated_conjecture, (v3_orders_2(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_101, negated_conjecture, (v4_orders_2(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_100, negated_conjecture, (v5_orders_2(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_596, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_115, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_123, negated_conjecture, (v2_pre_topc(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_130, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_609, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_91, negated_conjecture, (l1_orders_2(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_103, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v1_lattice3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_106, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_lattice3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_104, negated_conjecture, (v2_struct_0(esk1_0)|v3_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_139, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_266, plain, (v3_waybel_3(k2_yellow_1(X1))|~v2_waybel_8(k2_yellow_1(X1))|~v2_lattice3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1)))).# trainpos
cnf(c_0_283, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0)|~v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))|~v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_112, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v1_lattice3(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_114, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_lattice3(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_148, negated_conjecture, (epred1_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_229, plain, (v2_struct_0(k2_yellow_1(X1))|v1_lattice3(k2_yellow_1(X1))|~v3_lattice3(k2_yellow_1(X1)))).# trainpos
cnf(c_0_251, negated_conjecture, (v2_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))|v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_268, negated_conjecture, (v1_xboole_0(k5_waybel11(esk1_0))|v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_284, negated_conjecture, (v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_296, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0)|~v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_270, plain, (v2_struct_0(k2_yellow_1(X1))|v2_lattice3(k2_yellow_1(X1))|~v3_lattice3(k2_yellow_1(X1)))).# trainpos
cnf(c_0_601, negated_conjecture, (v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_602, negated_conjecture, (v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_286, negated_conjecture, (v2_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))|v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_297, negated_conjecture, (v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_305, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_313, negated_conjecture, (v3_waybel_3(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel30(esk1_0))).# trainpos
cnf(c_0_618, negated_conjecture, (v3_waybel30(esk1_0))).# trainpos
cnf(c_0_623, negated_conjecture, (~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_324, negated_conjecture, (v2_struct_0(esk1_0)|v3_waybel_3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_99, negated_conjecture, (g1_orders_2(u1_struct_0(o_1_14_waybel30(esk1_0)),u1_orders_2(o_1_14_waybel30(esk1_0)))=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0)))).# trainpos
cnf(c_0_592, negated_conjecture, (m1_subset_1(esk4_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~l1_waybel_9(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_552, negated_conjecture, (m1_yellow_9(X1,o_1_14_waybel30(esk1_0))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(o_1_14_waybel30(esk1_0)),u1_orders_2(o_1_14_waybel30(esk1_0)))|~l1_waybel_9(X1))).# trainpos
cnf(c_0_564, negated_conjecture, (m1_yellow_9(esk1_0,o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_570, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_98, negated_conjecture, (k5_waybel11(X1)=k5_waybel11(esk1_0)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_orders_2(X1)|~v3_lattice3(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_111, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_118, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(esk1_0)|v2_struct_0(esk1_0)|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_145, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_262, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_292, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_143, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_144, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_181, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_256, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_259, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_231, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_288, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_290, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_132, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_161, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_202, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_244, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_187, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_220, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_271, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),X1),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_183, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_192, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_166, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_198, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_200, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_237, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_240, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_185, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_214, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_217, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_164, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_165, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_563, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_orders_2(X1)|~v2_waybel_8(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_569, negated_conjecture, (v2_struct_0(esk1_0)|v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v2_waybel_8(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_182, negated_conjecture, (m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_212, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),X1),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_263, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_160, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0)))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_243, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_193, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),X1),X1)|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_577, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(o_1_14_waybel30(esk1_0))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(o_1_14_waybel30(esk1_0)),u1_orders_2(o_1_14_waybel30(esk1_0)))|~l1_orders_2(X1)|~v2_waybel_8(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_583, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|~v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_150, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_163, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_232, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_211, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_221, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_265, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_242, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_280, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_302, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_339, negated_conjecture, (v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_340, negated_conjecture, (v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_332, negated_conjecture, (m1_subset_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_517, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_418, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_403, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_354, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_372, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_373, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_396, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_345, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_361, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_383, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_408, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_395, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_362, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_382, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_441, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_460, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_457, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_468, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_480, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_487, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_194, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_370, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_258, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_261, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_394, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_381, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_407, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_429, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_419, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_435, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_450, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_459, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_234, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_236, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_239, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_210, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_274, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_300, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_277, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_301, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_393, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_417, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_404, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_425, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_442, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_449, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_470, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_458, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_478, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_509, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_496, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_497, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_528, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_518, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_254, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_287, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_308, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_329, negated_conjecture, (v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_330, negated_conjecture, (v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_323, negated_conjecture, (m1_subset_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_498, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_400, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_385, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_343, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_359, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_360, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_380, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_333, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_349, negated_conjecture, (m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_350, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_366, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_534, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_544, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_550, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_556, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_568, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_430, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_444, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_451, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_447, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_368, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_391, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_379, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_463, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_473, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_390, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_414, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_401, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_423, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_439, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_446, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_356, negated_conjecture, (m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_377, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_365, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_376, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_386, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_399, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_410, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).# trainpos
cnf(c_0_431, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_438, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_465, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_467, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_455, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_453, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_456, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_445, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_474, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_503, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_492, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_476, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_506, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_494, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_493, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_523, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_513, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_495, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_525, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_515, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_461, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_491, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_482, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_483, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_512, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_499, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).# trainpos
cnf(c_0_530, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_542, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_548, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_554, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_566, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_532, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_543, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_549, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_555, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_567, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_520, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_536, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).# trainpos
cnf(c_0_541, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).# trainpos
cnf(c_0_547, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).# trainpos
cnf(c_0_559, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).# trainpos
cnf(c_0_571, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_578, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_584, negated_conjecture, (v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_588, negated_conjecture, (v2_waybel_8(esk1_0))).# trainpos
cnf(c_0_625, negated_conjecture, (~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).# trainpos
cnf(c_0_593, negated_conjecture, (v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_600, negated_conjecture, (v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_616, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_614, negated_conjecture, (v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_615, negated_conjecture, (v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_606, negated_conjecture, (m1_subset_1(esk4_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_619, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|esk4_1(o_1_14_waybel30(esk1_0))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).# trainpos
cnf(c_0_626, negated_conjecture, (v2_struct_0(esk1_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_34, plain, (~v2_lattice3(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_46, plain, (l1_pre_topc(X1)|~l1_waybel_9(X1))).#trainneg
cnf(i_0_61, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk2_2(X3,X2)))).#trainneg
cnf(i_0_69, plain, (v3_waybel_3(k2_yellow_1(k5_waybel11(X1)))|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~v2_waybel_2(X1)|~l1_waybel_9(X1)|~v3_waybel_3(X1))).#trainneg
cnf(i_0_53, plain, (~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~v2_waybel19(X1)|~l1_waybel_9(X1)|~v7_struct_0(k2_yellow_1(k5_waybel11(X1))))).#trainneg
cnf(i_0_44, plain, (v1_orders_2(k2_yellow_1(X1)))).#trainneg
cnf(i_0_62, plain, (r2_hidden(esk2_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_117, plain, (v1_tops_2(esk14_2(X1,X2),X1)|m1_subset_1(esk12_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_99, plain, (v3_pre_topc(X1,X2)|esk12_1(X2)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),X4)|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v6_waybel_6(esk13_2(X2,X4),k2_yellow_1(k5_waybel11(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X2))))))).#trainneg
cnf(i_0_76, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(X2)))|X4!=a_1_0_waybel14(X2)|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~l1_waybel_9(X2)|~v4_waybel11(X2)|~v1_tops_2(X4,X2)|~v1_cantor_1(X4,X2)|~r2_hidden(X1,esk5_2(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).#trainneg
cnf(i_0_98, plain, (v2_waybel_0(X1,X2)|esk12_1(X2)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),X4)|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v6_waybel_6(esk13_2(X2,X4),k2_yellow_1(k5_waybel11(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X2))))))).#trainneg
cnf(i_0_116, plain, (v1_yellow_8(esk14_2(X1,X2),X1,X2)|m1_subset_1(esk12_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_115, plain, (m1_subset_1(esk12_1(X1),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|m1_subset_1(esk14_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_102, plain, (v1_tops_2(esk14_2(X1,X2),X1)|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v6_waybel_6(esk13_2(X1,X3),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_114, plain, (v3_pre_topc(X1,X2)|m1_subset_1(esk12_1(X2),u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_113, plain, (v2_waybel_0(X1,X2)|m1_subset_1(esk12_1(X2),u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_104, plain, (v3_pre_topc(X1,X2)|r2_hidden(esk13_2(X2,X4),X4)|esk12_1(X2)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),X4)|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X2))))))).#trainneg
cnf(i_0_103, plain, (v2_waybel_0(X1,X2)|r2_hidden(esk13_2(X2,X4),X4)|esk12_1(X2)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),X4)|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X2))))))).#trainneg
cnf(i_0_101, plain, (v1_yellow_8(esk14_2(X1,X2),X1,X2)|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v6_waybel_6(esk13_2(X1,X3),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_100, plain, (m1_subset_1(esk14_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v6_waybel_6(esk13_2(X1,X3),k2_yellow_1(k5_waybel11(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_77, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),esk5_2(X2,X1))=X1|X3!=a_1_0_waybel14(X2)|~v2_pre_topc(X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v2_lattice3(X2)|~v3_lattice3(X2)|~l1_waybel_9(X2)|~v4_waybel11(X2)|~v1_tops_2(X3,X2)|~v1_cantor_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(X2)))))).#trainneg
cnf(i_0_109, plain, (v3_pre_topc(X1,X2)|m1_subset_1(esk13_2(X2,X4),u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|esk12_1(X2)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),X4)|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X2))))))).#trainneg
cnf(i_0_108, plain, (v2_waybel_0(X1,X2)|m1_subset_1(esk13_2(X2,X4),u1_struct_0(k2_yellow_1(k5_waybel11(X2))))|esk12_1(X2)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X2)),X4)|~epred1_1(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,esk14_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X2))))))).#trainneg
cnf(i_0_107, plain, (r2_hidden(esk13_2(X1,X3),X3)|v1_tops_2(esk14_2(X1,X2),X1)|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_78, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))|X3!=a_1_0_waybel14(X1)|~v2_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_lattice3(X1)|~l1_waybel_9(X1)|~v4_waybel11(X1)|~v1_tops_2(X3,X1)|~v1_cantor_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(X1)))))).#trainneg
cnf(i_0_106, plain, (r2_hidden(esk13_2(X1,X3),X3)|v1_yellow_8(esk14_2(X1,X2),X1,X2)|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_105, plain, (r2_hidden(esk13_2(X1,X3),X3)|m1_subset_1(esk14_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_112, plain, (v1_tops_2(esk14_2(X1,X2),X1)|m1_subset_1(esk13_2(X1,X3),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_111, plain, (v1_yellow_8(esk14_2(X1,X2),X1,X2)|m1_subset_1(esk13_2(X1,X3),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_110, plain, (m1_subset_1(esk13_2(X1,X3),u1_struct_0(k2_yellow_1(k5_waybel11(X1))))|m1_subset_1(esk14_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|esk12_1(X1)!=k1_yellow_0(k2_yellow_1(k5_waybel11(X1)),X3)|~epred1_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(X1))))))).#trainneg
cnf(i_0_133, negated_conjecture, (l1_pre_topc(esk1_0))).#trainneg
cnf(i_0_142, negated_conjecture, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk6_1(esk1_0),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(esk1_0))).#trainneg
cnf(i_0_143, negated_conjecture, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk6_1(esk1_0),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(esk1_0))).#trainneg
cnf(i_0_144, negated_conjecture, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk6_1(esk1_0),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(esk1_0))).#trainneg
cnf(i_0_145, negated_conjecture, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk6_1(esk1_0),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(esk1_0))).#trainneg
cnf(i_0_173, negated_conjecture, (v3_waybel_3(o_1_14_waybel30(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_167, negated_conjecture, (l1_pre_topc(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_168, negated_conjecture, (v4_waybel11(o_1_14_waybel30(o_1_14_waybel30(esk1_0)))|~v2_waybel_2(o_1_14_waybel30(esk1_0))|~v2_waybel19(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_184, negated_conjecture, (v1_lattice3(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_lattice3(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_169, negated_conjecture, (v3_lattice3(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|~v2_waybel19(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_191, negated_conjecture, (v3_waybel30(k2_yellow_1(k5_waybel11(esk1_0)))|v2_waybel_8(esk1_0)|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_waybel19(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_195, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_196, negated_conjecture, (v3_waybel30(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_waybel19(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_197, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_waybel_2(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_waybel19(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_170, negated_conjecture, (m1_yellow_9(o_1_14_waybel30(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v2_waybel_2(o_1_14_waybel30(esk1_0))|~v2_waybel19(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v5_orders_2(o_1_14_waybel30(esk1_0))|~v4_orders_2(o_1_14_waybel30(esk1_0))|~v3_orders_2(o_1_14_waybel30(esk1_0))|~v2_pre_topc(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_199, negated_conjecture, (v3_waybel30(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|~v2_waybel19(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_200, negated_conjecture, (v3_waybel_3(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|~v2_waybel_2(o_1_14_waybel30(esk1_0))|~v2_waybel19(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_149, negated_conjecture, (m1_yellow_9(X1,esk1_0)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_waybel_9(X1))).#trainneg
cnf(i_0_201, negated_conjecture, (m1_yellow_9(esk1_0,esk1_0))).#trainneg
cnf(i_0_212, negated_conjecture, (m1_yellow_9(o_1_14_waybel30(esk1_0),o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_176, negated_conjecture, (k5_waybel11(X1)=k5_waybel11(o_1_14_waybel30(esk1_0))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(o_1_14_waybel30(esk1_0)),u1_orders_2(o_1_14_waybel30(esk1_0)))|~l1_orders_2(X1)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v3_lattice3(X1)|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(X1)|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_234, negated_conjecture, (v2_struct_0(esk1_0)|~v1_xboole_0(k5_waybel11(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_235, negated_conjecture, (v2_struct_0(esk1_0)|~v7_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_waybel19(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_230, negated_conjecture, (k5_waybel11(X1)=k5_waybel11(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~l1_orders_2(X1)|~v3_lattice3(X1)|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(X1)|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_250, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~v6_waybel_6(esk13_2(o_1_14_waybel30(esk1_0),X2),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_269, plain, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~v6_waybel_6(esk13_2(o_1_14_waybel30(esk1_0),X2),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X3,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X3)))).#trainneg
cnf(i_0_284, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_289, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_286, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_288, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_290, plain, (m1_subset_1(esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_283, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_291, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_285, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_270, plain, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~v6_waybel_6(esk13_2(o_1_14_waybel30(esk1_0),X2),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X3,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X3)))).#trainneg
cnf(i_0_287, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_261, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~v6_waybel_6(esk13_2(o_1_14_waybel30(esk1_0),X2),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_258, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~v6_waybel_6(esk13_2(o_1_14_waybel30(esk1_0),X2),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_141, plain, (k5_waybel11(X1)=k5_waybel11(k2_yellow_1(X2))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(k2_yellow_1(X2)),u1_orders_2(k2_yellow_1(X2)))|~l1_orders_2(X1)|~v3_lattice3(k2_yellow_1(X2))|~v3_lattice3(X1)|~v2_lattice3(k2_yellow_1(X2))|~v2_lattice3(X1)|~v1_lattice3(k2_yellow_1(X2))|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_299, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(k2_yellow_1(X1))|g1_orders_2(u1_struct_0(k2_yellow_1(X1)),u1_orders_2(k2_yellow_1(X1)))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~v3_lattice3(k2_yellow_1(X1))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(k2_yellow_1(X1))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(k2_yellow_1(X1))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_300, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(k2_yellow_1(X1))|v2_struct_0(esk1_0)|g1_orders_2(u1_struct_0(k2_yellow_1(X1)),u1_orders_2(k2_yellow_1(X1)))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~v3_lattice3(k2_yellow_1(X1))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(k2_yellow_1(X1))|~v1_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(k2_yellow_1(X1)))).#trainneg
cnf(i_0_301, negated_conjecture, (k5_waybel11(o_1_14_waybel30(esk1_0))=k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|g1_orders_2(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))),u1_orders_2(k2_yellow_1(k5_waybel11(esk1_0))))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0)))).#trainneg
cnf(i_0_249, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),X2),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_243, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),X2),X2)|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_140, plain, (m1_yellow_9(X1,k2_yellow_1(X2))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(k2_yellow_1(X2)),u1_orders_2(k2_yellow_1(X2)))|~l1_waybel_9(X1))).#trainneg
cnf(i_0_306, plain, (m1_yellow_9(k2_yellow_1(X1),k2_yellow_1(X1))|~l1_waybel_9(k2_yellow_1(X1)))).#trainneg
cnf(i_0_267, plain, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),X2),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X3,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X3)))).#trainneg
cnf(i_0_307, negated_conjecture, (m1_yellow_9(o_1_14_waybel30(esk1_0),k2_yellow_1(X1))|g1_orders_2(u1_struct_0(k2_yellow_1(X1)),u1_orders_2(k2_yellow_1(X1)))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0)))).#trainneg
cnf(i_0_253, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),X2),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_257, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),X2),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_245, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),X2),X2)|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_248, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),X2),X2)|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_268, plain, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),X2),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X3,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X3)))).#trainneg
cnf(i_0_265, plain, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),X2),X2)|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X3,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X3)))).#trainneg
cnf(i_0_266, plain, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),X2),X2)|v2_struct_0(esk1_0)|esk12_1(o_1_14_waybel30(esk1_0))!=k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),X2)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X3,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X3)))).#trainneg
cnf(i_0_139, plain, (v2_struct_0(X1)|v2_struct_0(k2_yellow_1(X2))|v2_waybel_8(k2_yellow_1(X2))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(k2_yellow_1(X2)),u1_orders_2(k2_yellow_1(X2)))|~l1_orders_2(X1)|~v2_waybel_8(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_311, negated_conjecture, (v2_struct_0(k2_yellow_1(X1))|v2_struct_0(o_1_14_waybel30(esk1_0))|v2_waybel_8(k2_yellow_1(X1))|g1_orders_2(u1_struct_0(k2_yellow_1(X1)),u1_orders_2(k2_yellow_1(X1)))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~v2_waybel_8(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_327, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_332, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_329, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_331, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_333, plain, (m1_subset_1(esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_326, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_334, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_328, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_321, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_346, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_339, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_344, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_341, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_343, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_345, plain, (m1_subset_1(esk11_2(esk1_0,esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_338, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_330, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_340, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_342, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_312, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk5_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_350, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk5_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_322, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk5_2(o_1_14_waybel30(esk1_0),X2))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_325, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_354, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk5_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_349, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_347, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0))).#trainneg
cnf(i_0_317, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_315, negated_conjecture, (m1_subset_1(esk5_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_319, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_358, negated_conjecture, (m1_subset_1(esk5_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_313, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),esk8_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_314, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_316, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),esk8_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_318, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_323, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),esk8_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_379, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_383, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_384, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_391, plain, (v3_pre_topc(X1,esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_373, plain, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_386, plain, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_378, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_366, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_412, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_417, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_380, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_381, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_375, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_416, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_374, plain, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_392, plain, (v2_waybel_0(X1,esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_414, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_369, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_370, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_419, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_411, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_385, plain, (m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_376, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_377, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_365, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_389, plain, (v3_pre_topc(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_448, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_449, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_434, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_436, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_438, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_440, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk5_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_462, plain, (v2_waybel_8(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_450, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk5_2(o_1_14_waybel30(esk1_0),X2))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_463, plain, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_443, negated_conjecture, (m1_subset_1(esk5_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_454, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk5_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_464, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk5_2(o_1_14_waybel30(esk1_0),X2))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_480, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_481, plain, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_457, plain, (m1_subset_1(esk5_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_459, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_324, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk8_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk6_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_467, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_513, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_517, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_518, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_525, plain, (v3_pre_topc(X1,esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_507, plain, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_520, plain, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_512, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_500, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_514, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_515, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_509, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_508, plain, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_503, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_504, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_526, plain, (v2_waybel_0(X1,esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_519, plain, (m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_510, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_511, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_499, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_550, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_555, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_571, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_576, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_557, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_505, plain, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_506, plain, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_516, plain, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_554, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_575, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_501, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_556, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_577, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_502, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_496, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_578, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_606, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_611, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_549, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_570, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_545, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_523, plain, (v3_pre_topc(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_546, plain, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_610, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_524, plain, (v2_waybel_0(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_552, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_573, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_497, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_612, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_498, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_530, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_636, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_461, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_613, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_608, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_553, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_574, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_605, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_521, plain, (v3_pre_topc(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_655, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_658, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_522, plain, (v2_waybel_0(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_685, negated_conjecture, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_697, plain, (v3_pre_topc(X1,esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_689, negated_conjecture, (v1_tops_2(esk8_1(esk1_0),esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_690, negated_conjecture, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_692, negated_conjecture, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_684, negated_conjecture, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_672, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_718, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_723, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_686, negated_conjecture, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_687, negated_conjecture, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_681, negated_conjecture, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_679, negated_conjecture, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_698, plain, (v2_waybel_0(X1,esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_680, negated_conjecture, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_722, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_675, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_676, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_724, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_725, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_717, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_691, negated_conjecture, (m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_682, negated_conjecture, (v1_tops_2(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_683, negated_conjecture, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_720, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_695, plain, (v3_pre_topc(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_677, plain, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_671, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_721, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_678, plain, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_688, negated_conjecture, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_673, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_674, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_668, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_696, plain, (v2_waybel_0(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_719, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_669, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_670, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_702, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_745, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_693, plain, (v3_pre_topc(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_694, plain, (v2_waybel_0(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_703, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_748, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_609, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_704, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_755, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_699, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_758, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_707, plain, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_761, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_708, plain, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_764, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_607, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_749, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_700, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_767, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_701, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_770, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_466, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_775, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_778, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_705, plain, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_786, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_706, plain, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_789, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_774, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_465, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_792, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_795, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_531, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_803, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_456, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_797, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_806, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_776, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_816, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_736, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_532, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_822, plain, (m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_809, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_810, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_527, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_825, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_800, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_737, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_807, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_862, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_848, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_850, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_852, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_876, plain, (v2_waybel_8(o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_808, plain, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_738, negated_conjecture, (r2_hidden(esk2_2(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0))).#trainneg
cnf(i_0_829, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_873, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_897, negated_conjecture, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk2_2(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk2_2(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_881, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_902, plain, (a_1_0_waybel14(o_1_14_waybel30(esk1_0))=esk8_1(o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_348, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0))).#trainneg
cnf(i_0_945, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_949, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_957, plain, (v3_pre_topc(X1,esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_939, plain, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_950, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_952, plain, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_944, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_932, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_946, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_947, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_941, plain, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_940, plain, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_935, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_936, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_958, plain, (v2_waybel_0(X1,esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_955, plain, (v3_pre_topc(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_951, plain, (m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_942, plain, (v1_tops_2(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_943, plain, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_931, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_970, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_977, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_982, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_971, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_997, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1002, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_984, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_937, plain, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_938, plain, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_948, plain, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_981, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1001, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_933, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_983, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_1016, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_1004, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1017, plain, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))=esk6_1(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_1003, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_934, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_928, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_973, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1025, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1030, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_976, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_996, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_969, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_972, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1046, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_956, plain, (v2_waybel_0(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_979, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1053, negated_conjecture, (v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).#trainneg
cnf(i_0_1029, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_999, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_929, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1031, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_930, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_962, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1056, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1032, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1000, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1050, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_953, plain, (v3_pre_topc(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_1027, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1058, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).#trainneg
cnf(i_0_1024, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_974, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_954, plain, (v2_waybel_0(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_963, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1064, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_964, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1065, plain, (m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_959, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1028, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1066, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1061, plain, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1083, negated_conjecture, (m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))).#trainneg
cnf(i_0_1093, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(o_1_14_waybel30(esk1_0)),u1_struct_0(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|a_1_0_waybel14(o_1_14_waybel30(esk1_0))!=esk8_1(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1099, negated_conjecture, (v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1101, negated_conjecture, (v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1103, negated_conjecture, (m1_subset_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1113, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1098, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1100, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1102, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1127, plain, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1117, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1109, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1111, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1106, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1105, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1110, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1131, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1123, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1108, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_875, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1116, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_870, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1137, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1115, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_869, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1141, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1142, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_1143, plain, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_1125, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1120, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1119, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_874, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1150, negated_conjecture, (m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_880, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1157, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1163, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1156, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1166, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_1188, negated_conjecture, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1200, plain, (v3_pre_topc(X1,esk1_0)|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_1192, negated_conjecture, (v1_tops_2(esk8_1(esk1_0),esk1_0)|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1193, negated_conjecture, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1195, negated_conjecture, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v6_waybel_6(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1187, negated_conjecture, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1175, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1189, negated_conjecture, (v1_tops_2(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1190, negated_conjecture, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1184, negated_conjecture, (a_1_0_waybel14(esk1_0)=esk8_1(esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1182, negated_conjecture, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1183, negated_conjecture, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),X1))|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1178, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1201, plain, (v2_waybel_0(X1,esk1_0)|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_1179, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~v6_waybel_6(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1180, plain, (v3_pre_topc(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_1194, negated_conjecture, (m1_subset_1(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1185, negated_conjecture, (v1_tops_2(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1186, negated_conjecture, (v1_cantor_1(esk8_1(esk1_0),esk1_0)|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1174, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1219, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1224, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_1244, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1249, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_1226, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1181, plain, (v2_waybel_0(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))),esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),X1)))).#trainneg
cnf(i_0_1191, negated_conjecture, (m1_subset_1(esk8_1(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk7_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk6_1(esk1_0)|~v4_waybel11(esk1_0)|~v2_waybel_8(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1223, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1248, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1176, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1225, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_1250, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_1177, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1171, plain, (v1_tops_2(esk14_2(esk1_0,X1),esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1251, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1198, plain, (v3_pre_topc(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_1273, negated_conjecture, (v1_tops_2(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1278, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))))=esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_1218, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1243, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1277, negated_conjecture, (m1_subset_1(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1199, plain, (v2_waybel_0(X1,esk1_0)|m1_subset_1(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_1221, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1246, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1172, plain, (v1_yellow_8(esk14_2(esk1_0,X1),esk1_0,X1)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1279, plain, (m1_subset_1(esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0))).#trainneg
cnf(i_0_1173, plain, (m1_subset_1(esk14_2(esk1_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1280, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(esk1_0,esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1222, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1205, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1305, negated_conjecture, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1247, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1272, plain, (v1_tops_2(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1196, plain, (v3_pre_topc(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_1197, plain, (v2_waybel_0(X1,esk1_0)|r2_hidden(esk13_2(esk1_0,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(esk1_0)|~epred1_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~r2_hidden(X1,esk14_2(esk1_0,X2)))).#trainneg
cnf(i_0_1275, negated_conjecture, (v1_yellow_8(esk3_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~v4_waybel11(k2_yellow_1(k5_waybel11(esk1_0)))|~l1_waybel_9(k2_yellow_1(k5_waybel11(esk1_0)))|~v2_pre_topc(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1206, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1306, negated_conjecture, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1207, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1208, plain, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1220, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1245, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1307, negated_conjecture, (m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1308, negated_conjecture, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1202, plain, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1309, negated_conjecture, (v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1276, plain, (m1_subset_1(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1210, plain, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1310, negated_conjecture, (v3_pre_topc(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1211, plain, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1311, negated_conjecture, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|m1_subset_1(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1203, plain, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1209, plain, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1274, plain, (v1_yellow_8(esk14_2(k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))),k2_yellow_1(k5_waybel11(esk1_0)),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(k2_yellow_1(k5_waybel11(esk1_0))),u1_struct_0(k2_yellow_1(k5_waybel11(k2_yellow_1(k5_waybel11(esk1_0))))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(k2_yellow_1(k5_waybel11(esk1_0))))).#trainneg
cnf(i_0_1312, negated_conjecture, (v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),X1),o_1_14_waybel30(esk1_0),X1)|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1313, negated_conjecture, (v2_waybel_0(X1,o_1_14_waybel30(esk1_0))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|~m1_subset_1(X2,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(X1,esk14_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1204, plain, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1318, negated_conjecture, (m1_subset_1(esk14_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|r2_hidden(esk13_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|esk6_1(o_1_14_waybel30(esk1_0))!=esk12_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0))))).#trainneg
cnf(i_0_1328, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1334, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1130, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1344, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1350, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_879, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1367, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1377, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))).#trainneg
cnf(i_0_1368, plain, (k1_yellow_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1382, negated_conjecture, (k1_yellow_0(k2_yellow_1(k5_waybel11(esk1_0)),esk11_2(o_1_14_waybel30(esk1_0),X1))=X1|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1329, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1388, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1345, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1394, negated_conjecture, (v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1129, plain, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),X2)))).#trainneg
cnf(i_0_1241, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1266, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1242, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1267, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1124, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1271, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1404, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1405, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|v1_tops_2(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1300, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1408, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1409, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1369, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_872, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1410, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1411, negated_conjecture, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1295, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1122, plain, (m1_subset_1(esk11_2(o_1_14_waybel30(esk1_0),X1),k1_zfmisc_1(u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~v1_yellow_8(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))|~v1_tops_2(esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))),o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))))).#trainneg
cnf(i_0_1406, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v1_cantor_1(esk8_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_1296, negated_conjecture, (v6_waybel_6(X1,k2_yellow_1(k5_waybel11(esk1_0)))|r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|~r2_hidden(X1,esk11_2(o_1_14_waybel30(esk1_0),esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))))))).#trainneg
cnf(i_0_1304, negated_conjecture, (v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|~m1_subset_1(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0)))))).#trainneg
cnf(i_0_1447, negated_conjecture, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1448, negated_conjecture, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk3_2(o_1_14_waybel30(esk1_0),X1))|~v3_waybel_3(o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1445, plain, (v3_pre_topc(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_1446, plain, (v2_waybel_0(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),o_1_14_waybel30(esk1_0))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0)))))|v6_waybel_6(esk7_2(o_1_14_waybel30(esk1_0),esk11_2(o_1_14_waybel30(esk1_0),esk6_1(o_1_14_waybel30(esk1_0)))),k2_yellow_1(k5_waybel11(esk1_0)))|m1_subset_1(esk8_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|m1_subset_1(esk12_1(o_1_14_waybel30(esk1_0)),u1_struct_0(k2_yellow_1(k5_waybel11(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(o_1_14_waybel30(esk1_0))|v2_waybel_8(esk1_0)|~epred1_1(o_1_14_waybel30(esk1_0))|~m1_subset_1(X1,u1_struct_0(o_1_14_waybel30(esk1_0)))|~r2_hidden(esk10_2(o_1_14_waybel30(esk1_0),esk14_2(o_1_14_waybel30(esk1_0),esk9_1(o_1_14_waybel30(esk1_0)))),esk14_2(o_1_14_waybel30(esk1_0),X1)))).#trainneg
cnf(i_0_1474, negated_conjecture, (esk4_1(o_1_14_waybel30(esk1_0))=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1471, negated_conjecture, (v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1472, negated_conjecture, (v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1473, negated_conjecture, (m1_subset_1(esk4_1(o_1_14_waybel30(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(o_1_14_waybel30(esk1_0)))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1489, negated_conjecture, (v2_waybel_8(k2_yellow_1(k5_waybel11(o_1_14_waybel30(esk1_0))))|v2_struct_0(esk1_0)|v2_waybel_8(esk1_0)|esk4_1(o_1_14_waybel30(esk1_0))!=a_1_0_waybel14(o_1_14_waybel30(esk1_0))|~v1_cantor_1(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v1_tops_2(esk4_1(o_1_14_waybel30(esk1_0)),o_1_14_waybel30(esk1_0))|~v3_lattice3(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1501, negated_conjecture, (v2_waybel_8(esk1_0)|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1502, negated_conjecture, (v2_waybel_8(esk1_0)|~v2_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1509, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|~v1_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
cnf(i_0_1510, negated_conjecture, (v2_waybel_8(o_1_14_waybel30(esk1_0))|~v2_lattice3(o_1_14_waybel30(esk1_0)))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 32
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 129
# Removed in clause preprocessing      : 14
# Initial clauses in saturation        : 115
# Processed clauses                    : 1038
# ...of these trivial                  : 3
# ...subsumed                          : 134
# ...remaining for further processing  : 901
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 661
# Backward-rewritten                   : 77
# Generated clauses                    : 1415
# ...of the previous two non-trivial   : 1386
# Contextual simplify-reflections      : 721
# Paramodulations                      : 1406
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 163
#    Positive orientable unit clauses  : 38
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 124
# Current number of unprocessed clauses: 231
# ...number of literals in the above   : 1719
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1639
# Proof object given clauses           : 334
# Proof search given clauses           : 901
# Clause-clause subsumption calls (NU) : 258277
# Rec. Clause-clause subsumption calls : 3083
# Non-unit clause-clause subsumptions  : 1514
# Unit Clause-clause subsumption calls : 1964
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 107146

# -------------------------------------------------
# User time                : 0.325 s
# System time              : 0.000 s
# Total time               : 0.325 s
# Maximum resident set size: 7476 pages
