2020-05-31 10:12:32.457887: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55
2020-05-31 10:12:32.541534: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 10:12:32.663025: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 10:12:33.200796: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 10:12:35.274948: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2817084 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.814 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_series_1, conjecture, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>~(((![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k3_funct_2(k5_numbers,k1_numbers,X1,X3)=k2_power(X3,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,X2),X3)))&?[X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>(r1_xxreal_0(X3,X4)=>r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,X1,X4))))))&v1_series_1(X2))))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t41_series_1)).
fof(dt_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', dt_k2_nat_1)).
fof(redefinition_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k2_nat_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t6_boole)).
fof(redefinition_k2_power, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k1_numbers))=>k2_power(X1,X2)=k1_power(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k2_power)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', commutativity_k2_xcmplx_0)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k6_numbers)).
fof(rc4_xreal_0, axiom, ?[X1]:(((v1_xboole_0(X1)&v1_xcmplx_0(X1))&v1_xxreal_0(X1))&v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', rc4_xreal_0)).
fof(t4_power, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v7_ordinal1(X2)=>(~((~((r1_xxreal_0(np__1,X2)&r1_xxreal_0(k6_numbers,X1)))&v1_abian(X2)))=>(k1_newton(k1_power(X2,X1),X2)=X1&k1_power(X2,k1_newton(X1,X2))=X1)))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t4_power)).
fof(cc2_xcmplx_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc2_xcmplx_0)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc1_xreal_0)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', fc6_ordinal1)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t1_arithm)).
fof(t1_numerals, axiom, m1_subset_1(k1_xboole_0,k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t1_numerals)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k5_numbers)).
fof(t12_seq_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))<=>![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_seq_1(X1,X3)=k18_complex1(k1_seq_1(X2,X3)))))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t12_seq_1)).
fof(redefinition_k56_valued_1, axiom, ![X1, X2, X3]:(((v1_membered(X2)&v1_funct_1(X3))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k56_valued_1(X1,X2,X3)=k54_valued_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k56_valued_1)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc5_membered)).
fof(projectivity_k56_valued_1, axiom, ![X1, X2, X3]:(((v1_membered(X2)&v1_funct_1(X3))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k56_valued_1(X1,X2,k56_valued_1(X1,X2,X3))=k56_valued_1(X1,X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', projectivity_k56_valued_1)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:(((((~(v1_xboole_0(X1))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&m1_subset_1(X4,X1))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k3_funct_2)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc8_ordinal1)).
fof(redefinition_k18_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>k18_complex1(X1)=k16_complex1(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k18_complex1)).
fof(redefinition_k17_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>k17_complex1(X1)=k16_complex1(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k17_complex1)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', fc3_membered)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:((((((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X4))&v1_funct_2(X4,X1,X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_r2_funct_2)).
fof(dt_k56_valued_1, axiom, ![X1, X2, X3]:(((v1_membered(X2)&v1_funct_1(X3))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(v1_funct_1(k56_valued_1(X1,X2,X3))&m1_subset_1(k56_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers))))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', dt_k56_valued_1)).
fof(fc99_valued_1, axiom, ![X1, X2, X3]:(((((~(v1_xboole_0(X2))&v1_membered(X2))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(v1_funct_1(k54_valued_1(X3))&v1_partfun1(k54_valued_1(X3),X1))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', fc99_valued_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', fc1_numbers)).
fof(t46_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>r1_xxreal_0(k6_numbers,k17_complex1(X1))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t46_complex1)).
fof(t4_series_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v1_series_1(X1)=>(v2_comseq_2(X1)&k2_seq_2(X1)=k6_numbers))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t4_series_1)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_k1_seq_1)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_partfun1(X3,X1)=>v1_funct_2(X3,X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc1_funct_2)).
fof(fc1_seq_4, axiom, ![X1]:((((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&v2_comseq_2(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>((v1_funct_1(k54_valued_1(X1))&v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers))&v2_comseq_2(k54_valued_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', fc1_seq_4)).
fof(fc61_valued_0, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_valued_0(X1))=>v1_xcmplx_0(k1_funct_1(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', fc61_valued_0)).
fof(t38_series_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>~((((~(r1_xxreal_0(X1,k6_numbers))&?[X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>(r1_xxreal_0(X3,X4)=>r1_xxreal_0(X1,k18_complex1(k3_funct_2(k5_numbers,k1_numbers,X2,X4)))))))&v2_comseq_2(X2))&k2_seq_2(X2)=k6_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t38_series_1)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', dt_k1_seq_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc1_relset_1)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc18_valued_0)).
fof(cc6_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v3_valued_0(X1))=>(v1_relat_1(X1)&v1_valued_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc6_valued_0)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', redefinition_m2_subset_1)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', dt_k5_numbers)).
fof(t11_prepower, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v7_ordinal1(X2)=>(r1_xxreal_0(np__1,X1)=>r1_xxreal_0(np__1,k1_newton(X1,X2))))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t11_prepower)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc14_membered)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r0, axiom, ~(r1_xxreal_0(np__1,np__0)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', rqLessOrEqual__r1_xxreal_0__r1_r0)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', spc0_boole)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', spc1_numerals)).
fof(dt_k3_funct_2, axiom, ![X1, X2, X3, X4]:(((((~(v1_xboole_0(X1))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&m1_subset_1(X4,X1))=>m1_subset_1(k3_funct_2(X1,X2,X3,X4),X2)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', dt_k3_funct_2)).
fof(t2_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X3))=>r1_xxreal_0(X1,X3))))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', t2_xxreal_0)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc2_xxreal_0)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', fc6_membered)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', connectedness_r1_xxreal_0)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc4_xreal_0)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t41_series_1', cc21_membered)).
fof(c_0_54, negated_conjecture, ~(![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>~(((![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k3_funct_2(k5_numbers,k1_numbers,X1,X3)=k2_power(X3,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,X2),X3)))&?[X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>(r1_xxreal_0(X3,X4)=>r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,X1,X4))))))&v1_series_1(X2)))))), inference(assume_negation,[status(cth)],[t41_series_1])).
fof(c_0_55, plain, ![X103, X104]:(~m1_subset_1(X103,k5_numbers)|~v7_ordinal1(X104)|m2_subset_1(k2_nat_1(X103,X104),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_nat_1])])).
fof(c_0_56, plain, ![X86, X87]:(~m1_subset_1(X86,k5_numbers)|~v7_ordinal1(X87)|k2_nat_1(X86,X87)=k2_xcmplx_0(X86,X87)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_1])])).
fof(c_0_57, plain, ![X64]:(~v1_xboole_0(X64)|X64=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_58, negated_conjecture, ![X7, X9]:(((v1_funct_1(esk1_0)&v1_funct_2(esk1_0,k5_numbers,k1_numbers))&m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))&(((v1_funct_1(esk2_0)&v1_funct_2(esk2_0,k5_numbers,k1_numbers))&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))&(((~m2_subset_1(X7,k1_numbers,k5_numbers)|k3_funct_2(k5_numbers,k1_numbers,esk1_0,X7)=k2_power(X7,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X7)))&(m2_subset_1(esk3_0,k1_numbers,k5_numbers)&(~m2_subset_1(X9,k1_numbers,k5_numbers)|(~r1_xxreal_0(esk3_0,X9)|r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X9))))))&v1_series_1(esk2_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_54])])])])).
fof(c_0_59, plain, ![X98, X99]:(~v7_ordinal1(X98)|~m1_subset_1(X99,k1_numbers)|k2_power(X98,X99)=k1_power(X98,X99)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_power])])).
cnf(c_0_60, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_62, plain, ![X78, X79]:(~v1_xcmplx_0(X78)|~v1_xcmplx_0(X79)|k2_xcmplx_0(X78,X79)=k2_xcmplx_0(X79,X78)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_63, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_65, plain, (((v1_xboole_0(esk7_0)&v1_xcmplx_0(esk7_0))&v1_xxreal_0(esk7_0))&v1_xreal_0(esk7_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
fof(c_0_66, plain, ![X18, X19]:(((k1_newton(k1_power(X19,X18),X19)=X18|(~r1_xxreal_0(np__1,X19)|~r1_xxreal_0(k6_numbers,X18))|~v7_ordinal1(X19)|~v1_xreal_0(X18))&(k1_power(X19,k1_newton(X18,X19))=X18|(~r1_xxreal_0(np__1,X19)|~r1_xxreal_0(k6_numbers,X18))|~v7_ordinal1(X19)|~v1_xreal_0(X18)))&((k1_newton(k1_power(X19,X18),X19)=X18|v1_abian(X19)|~v7_ordinal1(X19)|~v1_xreal_0(X18))&(k1_power(X19,k1_newton(X18,X19))=X18|v1_abian(X19)|~v7_ordinal1(X19)|~v1_xreal_0(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_power])])])])).
cnf(c_0_67, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)=k2_power(X1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (k2_power(X1,X2)=k1_power(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_70, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_71, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_72, plain, ![X141]:(~v7_ordinal1(X141)|v1_xcmplx_0(X141)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xcmplx_0])])).
cnf(c_0_73, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_74, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_75, plain, (v1_xboole_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_76, plain, (k1_newton(k1_power(X1,X2),X1)=X2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77, negated_conjecture, (k2_power(X1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_78, plain, (k2_power(X1,X2)=k1_power(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_68])).
fof(c_0_79, plain, ![X138]:(~m1_subset_1(X138,k1_numbers)|v1_xreal_0(X138)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
fof(c_0_80, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_81, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70])])).
cnf(c_0_82, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_83, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_84, plain, ![X73]:(~v1_xcmplx_0(X73)|k2_xcmplx_0(X73,k6_numbers)=X73), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_85, plain, (v1_xcmplx_0(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_86, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_87, plain, (v1_xboole_0(esk7_0)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_88, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(split_conjunct,[status(thm)],[t1_numerals])).
cnf(c_0_89, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_90, plain, ![X10, X11, X12]:((~r2_funct_2(k5_numbers,k1_numbers,X10,k56_valued_1(k5_numbers,k1_numbers,X11))|(~m2_subset_1(X12,k1_numbers,k5_numbers)|k1_seq_1(X10,X12)=k18_complex1(k1_seq_1(X11,X12)))|(~v1_funct_1(X11)|~v1_funct_2(X11,k5_numbers,k1_numbers)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((m2_subset_1(esk4_2(X10,X11),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,X10,k56_valued_1(k5_numbers,k1_numbers,X11))|(~v1_funct_1(X11)|~v1_funct_2(X11,k5_numbers,k1_numbers)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(k1_seq_1(X10,esk4_2(X10,X11))!=k18_complex1(k1_seq_1(X11,esk4_2(X10,X11)))|r2_funct_2(k5_numbers,k1_numbers,X10,k56_valued_1(k5_numbers,k1_numbers,X11))|(~v1_funct_1(X11)|~v1_funct_2(X11,k5_numbers,k1_numbers)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_seq_1])])])])])).
fof(c_0_91, plain, ![X34, X35, X36]:(~v1_membered(X35)|~v1_funct_1(X36)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X34,X35)))|k56_valued_1(X34,X35,X36)=k54_valued_1(X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k56_valued_1])])).
fof(c_0_92, plain, ![X148]:(~v3_membered(X148)|v1_membered(X148)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
fof(c_0_93, plain, ![X67, X68, X69]:(~v1_membered(X68)|~v1_funct_1(X69)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68)))|k56_valued_1(X67,X68,k56_valued_1(X67,X68,X69))=k56_valued_1(X67,X68,X69)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[projectivity_k56_valued_1])])).
cnf(c_0_94, plain, (k1_newton(k1_power(X1,X2),X1)=X2|~v1_xreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_95, negated_conjecture, (k1_power(X1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78])])).
cnf(c_0_96, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_79])).
fof(c_0_97, plain, ![X24, X25, X26, X27]:(v1_xboole_0(X24)|~v1_funct_1(X26)|~v1_funct_2(X26,X24,X25)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|~m1_subset_1(X27,X24)|k3_funct_2(X24,X25,X26,X27)=k1_funct_1(X26,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_98, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_99, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82])])).
cnf(c_0_100, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_101, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_102, plain, (v1_xcmplx_0(esk7_0)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_103, plain, (esk7_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_104, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_105, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_89])).
fof(c_0_106, plain, ![X54]:(~m1_subset_1(X54,k4_ordinal1)|v7_ordinal1(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_107, plain, ![X51]:(~v1_xcmplx_0(X51)|k18_complex1(X51)=k16_complex1(X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k18_complex1])])).
fof(c_0_108, plain, ![X37]:(~v1_xcmplx_0(X37)|k17_complex1(X37)=k16_complex1(X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k17_complex1])])).
cnf(c_0_109, plain, (k1_seq_1(X1,X3)=k18_complex1(k1_seq_1(X2,X3))|~r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_110, plain, (k56_valued_1(X3,X1,X2)=k54_valued_1(X2)|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_111, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_112, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_113, plain, ![X74, X75, X76, X77]:((~r2_funct_2(X74,X75,X76,X77)|X76=X77|(~v1_funct_1(X76)|~v1_funct_2(X76,X74,X75)|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(X74,X75)))|~v1_funct_1(X77)|~v1_funct_2(X77,X74,X75)|~m1_subset_1(X77,k1_zfmisc_1(k2_zfmisc_1(X74,X75)))))&(X76!=X77|r2_funct_2(X74,X75,X76,X77)|(~v1_funct_1(X76)|~v1_funct_2(X76,X74,X75)|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(X74,X75)))|~v1_funct_1(X77)|~v1_funct_2(X77,X74,X75)|~m1_subset_1(X77,k1_zfmisc_1(k2_zfmisc_1(X74,X75)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_114, plain, (k56_valued_1(X3,X1,k56_valued_1(X3,X1,X2))=k56_valued_1(X3,X1,X2)|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_93])).
fof(c_0_115, plain, ![X46, X47, X48]:((v1_funct_1(k56_valued_1(X46,X47,X48))|(~v1_membered(X47)|~v1_funct_1(X48)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))))&(m1_subset_1(k56_valued_1(X46,X47,X48),k1_zfmisc_1(k2_zfmisc_1(X46,k1_numbers)))|(~v1_membered(X47)|~v1_funct_1(X48)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k56_valued_1])])])).
fof(c_0_116, plain, ![X31, X32, X33]:((v1_funct_1(k54_valued_1(X33))|(v1_xboole_0(X32)|~v1_membered(X32)|~v1_funct_1(X33)|~v1_funct_2(X33,X31,X32)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))))&(v1_partfun1(k54_valued_1(X33),X31)|(v1_xboole_0(X32)|~v1_membered(X32)|~v1_funct_1(X33)|~v1_funct_2(X33,X31,X32)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc99_valued_1])])])])).
fof(c_0_117, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
cnf(c_0_118, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_119, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_120, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_121, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_122, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v1_xcmplx_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(csr,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_123, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_124, plain, (v1_xcmplx_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_125, plain, (m1_subset_1(k6_numbers,k5_numbers)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_74]), c_0_105])).
cnf(c_0_126, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_106])).
fof(c_0_127, plain, ![X45]:(~v1_xcmplx_0(X45)|r1_xxreal_0(k6_numbers,k17_complex1(X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t46_complex1])])).
cnf(c_0_128, plain, (k18_complex1(X1)=k16_complex1(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_129, plain, (k17_complex1(X1)=k16_complex1(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_130, plain, (k1_seq_1(X1,X3)=k18_complex1(k1_seq_1(X2,X3))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_131, plain, (k56_valued_1(X3,X1,X2)=k54_valued_1(X2)|~v1_funct_1(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_132, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_133, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_134, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_135, plain, (k56_valued_1(X3,X1,k56_valued_1(X3,X1,X2))=k56_valued_1(X3,X1,X2)|~v1_funct_1(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_136, plain, (v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_137, plain, (m1_subset_1(k56_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_138, plain, (v1_partfun1(k54_valued_1(X1),X2)|v1_xboole_0(X3)|~v1_membered(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_139, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_140, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_141, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_142, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_117])).
fof(c_0_143, plain, ![X20]:((v2_comseq_2(X20)|~v1_series_1(X20)|(~v1_funct_1(X20)|~v1_funct_2(X20,k5_numbers,k1_numbers)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(k2_seq_2(X20)=k6_numbers|~v1_series_1(X20)|(~v1_funct_1(X20)|~v1_funct_2(X20,k5_numbers,k1_numbers)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_series_1])])])).
cnf(c_0_144, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)), inference(csr,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_145, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_146, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_121, c_0_105])).
cnf(c_0_147, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_124]), c_0_125])])])).
cnf(c_0_148, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_126])).
fof(c_0_149, plain, ![X49, X50]:(~v1_relat_1(X49)|~v1_funct_1(X49)|~v3_valued_0(X49)|k1_seq_1(X49,X50)=k1_funct_1(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
cnf(c_0_150, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_151, plain, (k16_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_152, plain, (k16_complex1(X1)=k17_complex1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_153, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(X3,X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_154, plain, (v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_155, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X4)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_156, plain, (k54_valued_1(k56_valued_1(X1,X2,X3))=k56_valued_1(X1,X2,X3)|~v1_membered(X2)|~m1_subset_1(k56_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k56_valued_1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_135])])).
cnf(c_0_157, plain, (v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_funct_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_158, plain, (m1_subset_1(k56_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_137])).
fof(c_0_159, plain, ![X88, X89, X90]:(~m1_subset_1(X90,k1_zfmisc_1(k2_zfmisc_1(X88,X89)))|(~v1_partfun1(X90,X88)|v1_funct_2(X90,X88,X89))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])).
cnf(c_0_160, plain, (v1_xboole_0(X3)|v1_partfun1(k54_valued_1(X1),X2)|~v1_funct_1(X1)|~v1_membered(X3)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_161, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_162, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_163, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_164, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_142])).
fof(c_0_165, plain, ![X105]:(((v1_funct_1(k54_valued_1(X105))|(~v1_funct_1(X105)|~v1_funct_2(X105,k5_numbers,k1_numbers)|~v2_comseq_2(X105)|~m1_subset_1(X105,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(v1_funct_2(k54_valued_1(X105),k5_numbers,k1_numbers)|(~v1_funct_1(X105)|~v1_funct_2(X105,k5_numbers,k1_numbers)|~v2_comseq_2(X105)|~m1_subset_1(X105,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(v2_comseq_2(k54_valued_1(X105))|(~v1_funct_1(X105)|~v1_funct_2(X105,k5_numbers,k1_numbers)|~v2_comseq_2(X105)|~m1_subset_1(X105,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_seq_4])])])).
cnf(c_0_166, plain, (v2_comseq_2(X1)|~v1_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_167, negated_conjecture, (v1_series_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_168, plain, ![X81, X82]:(~v1_relat_1(X81)|~v1_funct_1(X81)|~v1_valued_0(X81)|v1_xcmplx_0(k1_funct_1(X81,X82))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc61_valued_0])])).
fof(c_0_169, plain, ![X14, X15, X16]:((m2_subset_1(esk5_3(X14,X15,X16),k1_numbers,k5_numbers)|~m2_subset_1(X16,k1_numbers,k5_numbers)|r1_xxreal_0(X14,k6_numbers)|~v2_comseq_2(X15)|k2_seq_2(X15)!=k6_numbers|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k1_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~v1_xreal_0(X14))&((r1_xxreal_0(X16,esk5_3(X14,X15,X16))|~m2_subset_1(X16,k1_numbers,k5_numbers)|r1_xxreal_0(X14,k6_numbers)|~v2_comseq_2(X15)|k2_seq_2(X15)!=k6_numbers|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k1_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~v1_xreal_0(X14))&(~r1_xxreal_0(X14,k18_complex1(k3_funct_2(k5_numbers,k1_numbers,X15,esk5_3(X14,X15,X16))))|~m2_subset_1(X16,k1_numbers,k5_numbers)|r1_xxreal_0(X14,k6_numbers)|~v2_comseq_2(X15)|k2_seq_2(X15)!=k6_numbers|(~v1_funct_1(X15)|~v1_funct_2(X15,k5_numbers,k1_numbers)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~v1_xreal_0(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_series_1])])])])])])).
cnf(c_0_170, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_146])])).
cnf(c_0_171, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_147, c_0_100])).
cnf(c_0_172, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_148, c_0_105])).
cnf(c_0_173, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
fof(c_0_174, plain, ![X65, X66]:(~v1_relat_1(X65)|~v1_funct_1(X65)|~v3_valued_0(X65)|m1_subset_1(k1_seq_1(X65,X66),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
fof(c_0_175, plain, ![X38, X39, X40]:(~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))|v1_relat_1(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_176, plain, ![X70, X71, X72]:(~v3_membered(X71)|(~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X70,X71)))|v3_valued_0(X72))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
cnf(c_0_177, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_178, plain, (k17_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152])])).
cnf(c_0_179, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(X3,X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_180, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_155])])).
cnf(c_0_181, plain, (k54_valued_1(k56_valued_1(X1,X2,X3))=k56_valued_1(X1,X2,X3)|~v1_membered(X2)|~m1_subset_1(k56_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_182, plain, (m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_131])])).
cnf(c_0_183, plain, (v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_partfun1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_184, negated_conjecture, (v1_partfun1(k54_valued_1(esk2_0),k5_numbers)|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161]), c_0_162]), c_0_163])]), c_0_164])])).
cnf(c_0_185, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_186, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_187, negated_conjecture, (v1_series_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_188, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_168])).
fof(c_0_189, plain, ![X124]:((v1_relat_1(X124)|(~v1_relat_1(X124)|~v3_valued_0(X124)))&(v1_valued_0(X124)|(~v1_relat_1(X124)|~v3_valued_0(X124)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_valued_0])])])).
cnf(c_0_190, plain, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k18_complex1(k3_funct_2(k5_numbers,k1_numbers,X2,esk5_3(X1,X2,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|k2_seq_2(X2)!=k6_numbers|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_169])).
fof(c_0_191, plain, ![X21, X22, X23]:((~m2_subset_1(X23,X21,X22)|m1_subset_1(X23,X22)|(v1_xboole_0(X21)|v1_xboole_0(X22)|~m1_subset_1(X22,k1_zfmisc_1(X21))))&(~m1_subset_1(X23,X22)|m2_subset_1(X23,X21,X22)|(v1_xboole_0(X21)|v1_xboole_0(X22)|~m1_subset_1(X22,k1_zfmisc_1(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_192, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_170, c_0_171]), c_0_172])).
cnf(c_0_193, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_194, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_195, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_196, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_197, plain, (r1_xxreal_0(k6_numbers,k18_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178])])).
cnf(c_0_198, plain, (k18_complex1(k1_seq_1(X1,X2))=k1_seq_1(k54_valued_1(X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(k54_valued_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180])])).
cnf(c_0_199, plain, (k54_valued_1(k54_valued_1(X3))=k54_valued_1(X3)|~v1_membered(X2)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_131])])).
cnf(c_0_200, negated_conjecture, (m1_subset_1(k54_valued_1(esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_161]), c_0_154]), c_0_163])])])).
cnf(c_0_201, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_202, negated_conjecture, (v1_partfun1(k54_valued_1(esk2_0),k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_184, c_0_154])])).
cnf(c_0_203, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_204, negated_conjecture, (v2_comseq_2(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_161]), c_0_162]), c_0_163])])])).
cnf(c_0_205, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_188])).
cnf(c_0_206, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_189])).
cnf(c_0_207, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v1_funct_1(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X1,k18_complex1(k3_funct_2(k5_numbers,k1_numbers,X2,esk5_3(X1,X2,X3))))), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_208, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_191])).
cnf(c_0_209, plain, (m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|k2_seq_2(X2)!=k6_numbers|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_210, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_211, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~v3_valued_0(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~v1_relat_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193])])).
cnf(c_0_212, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_213, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_214, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_215, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_216, negated_conjecture, (k54_valued_1(k54_valued_1(esk2_0))=k54_valued_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_154]), c_0_161]), c_0_163])])])).
cnf(c_0_217, negated_conjecture, (v1_funct_2(k54_valued_1(esk2_0),k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_200]), c_0_202])])])).
cnf(c_0_218, negated_conjecture, (v1_funct_1(k54_valued_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_204]), c_0_161]), c_0_162]), c_0_163])])])).
cnf(c_0_219, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_193])])).
cnf(c_0_220, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_221, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(X2,esk5_3(X1,X2,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_145]), c_0_146])])).
cnf(c_0_222, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_223, plain, (r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|k2_seq_2(X2)!=k6_numbers|~v1_funct_1(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_224, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_210])).
fof(c_0_225, plain, ![X52, X53]:(~v1_xreal_0(X52)|(~v7_ordinal1(X53)|(~r1_xxreal_0(np__1,X52)|r1_xxreal_0(np__1,k1_newton(X52,X53))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_prepower])])])).
cnf(c_0_226, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v3_valued_0(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~r1_xxreal_0(k6_numbers,k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_211, c_0_212]), c_0_213])).
cnf(c_0_227, negated_conjecture, (v3_valued_0(k54_valued_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_200]), c_0_133])])])).
cnf(c_0_228, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_200]), c_0_217]), c_0_218])])])).
cnf(c_0_229, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_219, c_0_220])).
cnf(c_0_230, negated_conjecture, (v1_relat_1(k54_valued_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_200])])).
cnf(c_0_231, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k18_complex1(k1_seq_1(X2,esk5_3(X1,X2,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_193])])).
cnf(c_0_232, plain, (m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)|r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223]), c_0_224])]), c_0_164]), c_0_146])])).
cnf(c_0_233, plain, (r1_xxreal_0(np__1,k1_newton(X1,X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_234, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_131]), c_0_227]), c_0_200]), c_0_217]), c_0_218]), c_0_154]), c_0_161]), c_0_163])])])).
cnf(c_0_235, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_229]), c_0_227]), c_0_230]), c_0_218])])])).
fof(c_0_236, plain, ![X125, X126]:(~v3_membered(X125)|(~m1_subset_1(X126,X125)|v1_xreal_0(X126))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_237, plain, ~r1_xxreal_0(np__1,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r0])).
cnf(c_0_238, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_239, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v3_valued_0(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k18_complex1(k1_seq_1(X2,esk5_3(X1,X2,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_231, c_0_232]), c_0_213])).
cnf(c_0_240, plain, (v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_241, plain, (r1_xxreal_0(np__1,k1_newton(X1,X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X1)), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_242, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235])])).
cnf(c_0_243, plain, (k2_seq_2(X1)=k6_numbers|~v1_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_244, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_245, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_246, plain, (~r1_xxreal_0(np__1,np__0)), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_247, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_238])).
fof(c_0_248, plain, ![X112, X113, X114, X115]:(v1_xboole_0(X112)|~v1_funct_1(X114)|~v1_funct_2(X114,X112,X113)|~m1_subset_1(X114,k1_zfmisc_1(k2_zfmisc_1(X112,X113)))|~m1_subset_1(X115,X112)|m1_subset_1(k3_funct_2(X112,X113,X114,X115),X113)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_funct_2])])])).
cnf(c_0_249, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v3_valued_0(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(X2),esk5_3(X1,X2,X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_198])])).
cnf(c_0_250, plain, (v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_251, negated_conjecture, (r1_xxreal_0(np__1,k1_seq_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_242])])).
cnf(c_0_252, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_171]), c_0_224])]), c_0_164]), c_0_146])])).
cnf(c_0_253, plain, (k2_seq_2(X1)=k6_numbers|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_254, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_255, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_245])).
cnf(c_0_256, plain, (~r1_xxreal_0(np__1,np__0)), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_257, plain, (np__0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_247])])).
cnf(c_0_258, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_248])).
cnf(c_0_259, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v3_valued_0(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(X2),esk5_3(X1,X2,X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_249, c_0_203]), c_0_250]), c_0_223])).
cnf(c_0_260, negated_conjecture, (r1_xxreal_0(np__1,k1_seq_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_251, c_0_252]), c_0_171])).
cnf(c_0_261, negated_conjecture, (k2_seq_2(esk2_0)=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_187]), c_0_161]), c_0_162]), c_0_163])])])).
cnf(c_0_262, negated_conjecture, (v3_valued_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_161]), c_0_133])])])).
cnf(c_0_263, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255]), c_0_133])])])).
cnf(c_0_264, plain, (~r1_xxreal_0(np__1,k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_256, c_0_257])])).
cnf(c_0_265, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_258])).
cnf(c_0_266, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_267, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_268, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_269, negated_conjecture, (~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260]), c_0_261]), c_0_262]), c_0_204]), c_0_263]), c_0_200]), c_0_161]), c_0_162]), c_0_163])]), c_0_264])])).
cnf(c_0_270, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_265])])).
cnf(c_0_271, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_266])).
cnf(c_0_272, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_273, negated_conjecture, (v1_funct_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_268])).
fof(c_0_274, plain, ![X28, X29, X30]:(~v1_xxreal_0(X28)|(~v1_xxreal_0(X29)|(~v1_xxreal_0(X30)|(~r1_xxreal_0(X28,X29)|~r1_xxreal_0(X29,X30)|r1_xxreal_0(X28,X30))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_xxreal_0])])])).
cnf(c_0_275, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,esk2_0,X1),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270]), c_0_133]), c_0_271]), c_0_272]), c_0_273])]), c_0_146])])).
cnf(c_0_276, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_0,X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_277, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_278, plain, (r1_xxreal_0(X1,esk5_3(X2,X3,X1))|r1_xxreal_0(X2,k6_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v2_comseq_2(X3)|k2_seq_2(X3)!=k6_numbers|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_169])).
fof(c_0_279, plain, ![X154]:(~v7_ordinal1(X154)|v1_xxreal_0(X154)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
fof(c_0_280, plain, ![X41, X42]:(~v6_membered(X41)|(~m1_subset_1(X42,X41)|v7_ordinal1(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_281, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_282, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_275, c_0_252])).
cnf(c_0_283, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(esk3_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_284, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_285, plain, (r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,esk5_3(X2,X3,X1))|k2_seq_2(X3)!=k6_numbers|~v1_funct_1(X3)|~v1_xreal_0(X2)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_286, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_287, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_280])).
cnf(c_0_288, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_281])).
fof(c_0_289, plain, ![X116, X117]:(~v1_xxreal_0(X116)|~v1_xxreal_0(X117)|(r1_xxreal_0(X116,X117)|r1_xxreal_0(X117,X116))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
fof(c_0_290, plain, ![X155]:(~v1_xreal_0(X155)|v1_xxreal_0(X155)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
cnf(c_0_291, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_0,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(esk5_3(np__1,esk2_0,X1),k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283])])).
cnf(c_0_292, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_293, plain, (r1_xxreal_0(X1,esk5_3(X2,X3,X4))|r1_xxreal_0(X2,k6_numbers)|~v1_xxreal_0(esk5_3(X2,X3,X4))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|k2_seq_2(X3)!=k6_numbers|~v2_comseq_2(X3)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_285])])).
cnf(c_0_294, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_286])).
cnf(c_0_295, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_287])).
cnf(c_0_296, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_288, c_0_105])).
fof(c_0_297, plain, ![X170, X171]:(~v3_membered(X170)|(~m1_subset_1(X171,k1_zfmisc_1(X170))|v3_membered(X171))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_298, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_289])).
cnf(c_0_299, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_290])).
cnf(c_0_300, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(esk3_0,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_291, c_0_171])).
cnf(c_0_301, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_292])).
cnf(c_0_302, plain, (r1_xxreal_0(X1,esk5_3(X2,X3,X4))|r1_xxreal_0(X2,k6_numbers)|k2_seq_2(X3)!=k6_numbers|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v2_comseq_2(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v7_ordinal1(esk5_3(X2,X3,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_294])])).
cnf(c_0_303, plain, (v7_ordinal1(esk5_3(X1,X2,X3))|r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_232]), c_0_296])])])).
cnf(c_0_304, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_297])).
cnf(c_0_305, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_298])).
cnf(c_0_306, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_299])).
cnf(c_0_307, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,esk3_0))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_285]), c_0_301]), c_0_261]), c_0_204]), c_0_263]), c_0_161]), c_0_162]), c_0_163])]), c_0_264])])).
cnf(c_0_308, plain, (r1_xxreal_0(X1,esk5_3(X2,X3,X4))|r1_xxreal_0(X2,k6_numbers)|k2_seq_2(X3)!=k6_numbers|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v2_comseq_2(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_302, c_0_303])).
cnf(c_0_309, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_304])).
cnf(c_0_310, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_311, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306])])).
cnf(c_0_312, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,esk3_0))|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_308]), c_0_261]), c_0_204]), c_0_263]), c_0_301]), c_0_161]), c_0_162]), c_0_163])]), c_0_264])])).
cnf(c_0_313, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_301]), c_0_224])]), c_0_164]), c_0_146])])).
cnf(c_0_314, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_224]), c_0_133])])])).
cnf(c_0_315, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(X1)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_308]), c_0_261]), c_0_204]), c_0_263]), c_0_161]), c_0_162]), c_0_163])]), c_0_264])])).
cnf(c_0_316, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_317, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_306])])).
cnf(c_0_318, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_303]), c_0_261]), c_0_204]), c_0_263]), c_0_301]), c_0_161]), c_0_162]), c_0_163])]), c_0_264])])).
cnf(c_0_319, negated_conjecture, (v1_xreal_0(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_313]), c_0_314])])])).
cnf(c_0_320, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~v7_ordinal1(esk5_3(np__1,esk2_0,np__1))|~r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_285]), c_0_316]), c_0_261]), c_0_204]), c_0_263]), c_0_161]), c_0_162]), c_0_163])]), c_0_264])])).
cnf(c_0_321, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_263])])).
cnf(c_0_322, negated_conjecture, (~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_306]), c_0_319])])])).
cnf(c_0_323, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_303]), c_0_261]), c_0_204]), c_0_263]), c_0_316]), c_0_161]), c_0_162]), c_0_163])]), c_0_264])])).
cnf(c_0_324, negated_conjecture, (r1_xxreal_0(np__1,esk3_0)|r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_319])])).
cnf(c_0_325, negated_conjecture, (~r1_xxreal_0(np__1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_322, c_0_306]), c_0_263])])])).
cnf(c_0_326, negated_conjecture, (~v1_xxreal_0(np__1)|~r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_306]), c_0_319])])])).
cnf(c_0_327, negated_conjecture, (r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_324, c_0_325])])).
cnf(c_0_328, negated_conjecture, (~r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_306]), c_0_263])])])).
cnf(c_0_329, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_327, c_0_328]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 330
# Proof object clause steps            : 229
# Proof object formula steps           : 101
# Proof object conjectures             : 67
# Proof object clause conjectures      : 64
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 70
# Proof object initial formulas used   : 54
# Proof object generating inferences   : 64
# Proof object simplifying inferences  : 181
# Training examples: 138 positive, 990 negative
# Training: Positive examples begin
cnf(c_0_105, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_74, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_273, negated_conjecture, (v1_funct_1(esk1_0))).# trainpos
cnf(c_0_163, negated_conjecture, (v1_funct_1(esk2_0))).# trainpos
cnf(c_0_187, negated_conjecture, (v1_series_1(esk2_0))).# trainpos
cnf(c_0_247, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_87, plain, (v1_xboole_0(esk7_0))).# trainpos
cnf(c_0_272, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers))).# trainpos
cnf(c_0_162, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers))).# trainpos
cnf(c_0_256, plain, (~r1_xxreal_0(np__1,np__0))).# trainpos
cnf(c_0_102, plain, (v1_xcmplx_0(esk7_0))).# trainpos
cnf(c_0_316, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_164, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_121, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_133, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_288, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_73, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_255, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_104, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).# trainpos
cnf(c_0_224, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_301, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_271, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_161, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_253, plain, (k2_seq_2(X1)=k6_numbers|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_306, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_132, plain, (v1_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_294, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_100, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_119, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_151, plain, (k16_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_152, plain, (k16_complex1(X1)=k17_complex1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_123, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_148, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_177, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_220, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_283, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(esk3_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_77, negated_conjecture, (k2_power(X1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_207, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v1_funct_1(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X1,k18_complex1(k3_funct_2(k5_numbers,k1_numbers,X2,esk5_3(X1,X2,X3)))))).# trainpos
cnf(c_0_254, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_295, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_213, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_309, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_82, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_201, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_305, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_78, plain, (k2_power(X1,X2)=k1_power(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_70, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_214, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_186, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_284, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_203, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_193, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_69, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_157, plain, (v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_funct_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_222, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_205, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_212, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_241, plain, (r1_xxreal_0(np__1,k1_newton(X1,X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X1))).# trainpos
cnf(c_0_131, plain, (k56_valued_1(X3,X1,X2)=k54_valued_1(X2)|~v1_funct_1(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_94, plain, (k1_newton(k1_power(X1,X2),X1)=X2|~v1_xreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2))).# trainpos
cnf(c_0_250, plain, (v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_160, plain, (v1_xboole_0(X3)|v1_partfun1(k54_valued_1(X1),X2)|~v1_funct_1(X1)|~v1_membered(X3)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_158, plain, (m1_subset_1(k56_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_135, plain, (k56_valued_1(X3,X1,k56_valued_1(X3,X1,X2))=k56_valued_1(X3,X1,X2)|~v1_funct_1(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_285, plain, (r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,esk5_3(X2,X3,X1))|k2_seq_2(X3)!=k6_numbers|~v1_funct_1(X3)|~v1_xreal_0(X2)|~v2_comseq_2(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_223, plain, (r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|k2_seq_2(X2)!=k6_numbers|~v1_funct_1(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_145, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_265, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_155, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X4)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).# trainpos
cnf(c_0_130, plain, (k1_seq_1(X1,X3)=k18_complex1(k1_seq_1(X2,X3))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2)))).# trainpos
cnf(c_0_314, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_154, plain, (v1_membered(k1_numbers))).# trainpos
cnf(c_0_257, plain, (np__0=k6_numbers)).# trainpos
cnf(c_0_261, negated_conjecture, (k2_seq_2(esk2_0)=k6_numbers)).# trainpos
cnf(c_0_263, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_103, plain, (esk7_0=k6_numbers)).# trainpos
cnf(c_0_180, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).# trainpos
cnf(c_0_81, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_178, plain, (k17_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_313, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers))).# trainpos
cnf(c_0_95, negated_conjecture, (k1_power(X1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers))).# trainpos
cnf(c_0_204, negated_conjecture, (v2_comseq_2(esk2_0))).# trainpos
cnf(c_0_262, negated_conjecture, (v3_valued_0(esk2_0))).# trainpos
cnf(c_0_182, plain, (m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).# trainpos
cnf(c_0_311, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_221, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k18_complex1(k1_funct_1(X2,esk5_3(X1,X2,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk5_3(X1,X2,X3),k5_numbers))).# trainpos
cnf(c_0_156, plain, (k54_valued_1(k56_valued_1(X1,X2,X3))=k56_valued_1(X1,X2,X3)|~v1_membered(X2)|~m1_subset_1(k56_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k56_valued_1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).# trainpos
cnf(c_0_293, plain, (r1_xxreal_0(X1,esk5_3(X2,X3,X4))|r1_xxreal_0(X2,k6_numbers)|~v1_xxreal_0(esk5_3(X2,X3,X4))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|k2_seq_2(X3)!=k6_numbers|~v2_comseq_2(X3)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).# trainpos
cnf(c_0_232, plain, (m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)|r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).# trainpos
cnf(c_0_184, negated_conjecture, (v1_partfun1(k54_valued_1(esk2_0),k5_numbers)|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_270, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).# trainpos
cnf(c_0_153, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(X3,X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_219, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_197, plain, (r1_xxreal_0(k6_numbers,k18_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_218, negated_conjecture, (v1_funct_1(k54_valued_1(esk2_0)))).# trainpos
cnf(c_0_264, plain, (~r1_xxreal_0(np__1,k6_numbers))).# trainpos
cnf(c_0_118, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers))).# trainpos
cnf(c_0_124, plain, (v1_xcmplx_0(k6_numbers))).# trainpos
cnf(c_0_319, negated_conjecture, (v1_xreal_0(esk3_0))).# trainpos
cnf(c_0_99, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_199, plain, (k54_valued_1(k54_valued_1(X3))=k54_valued_1(X3)|~v1_membered(X2)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).# trainpos
cnf(c_0_200, negated_conjecture, (m1_subset_1(k54_valued_1(esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_317, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_198, plain, (k18_complex1(k1_seq_1(X1,X2))=k1_seq_1(k54_valued_1(X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(k54_valued_1(X1)))).# trainpos
cnf(c_0_231, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k18_complex1(k1_seq_1(X2,esk5_3(X1,X2,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_303, plain, (v7_ordinal1(esk5_3(X1,X2,X3))|r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).# trainpos
cnf(c_0_302, plain, (r1_xxreal_0(X1,esk5_3(X2,X3,X4))|r1_xxreal_0(X2,k6_numbers)|k2_seq_2(X3)!=k6_numbers|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v2_comseq_2(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v7_ordinal1(esk5_3(X2,X3,X4)))).# trainpos
cnf(c_0_217, negated_conjecture, (v1_funct_2(k54_valued_1(esk2_0),k5_numbers,k1_numbers))).# trainpos
cnf(c_0_147, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_230, negated_conjecture, (v1_relat_1(k54_valued_1(esk2_0)))).# trainpos
cnf(c_0_227, negated_conjecture, (v3_valued_0(k54_valued_1(esk2_0)))).# trainpos
cnf(c_0_321, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_249, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v3_valued_0(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(X2),esk5_3(X1,X2,X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2)))).# trainpos
cnf(c_0_216, negated_conjecture, (k54_valued_1(k54_valued_1(esk2_0))=k54_valued_1(esk2_0))).# trainpos
cnf(c_0_215, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1))).# trainpos
cnf(c_0_170, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0)))).# trainpos
cnf(c_0_252, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_324, negated_conjecture, (r1_xxreal_0(np__1,esk3_0)|r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_228, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_211, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~v3_valued_0(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~v1_relat_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0)))).# trainpos
cnf(c_0_235, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_234, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_242, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_251, negated_conjecture, (r1_xxreal_0(np__1,k1_seq_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_269, negated_conjecture, (~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1)))).# trainpos
cnf(c_0_275, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,esk2_0,X1),k5_numbers))).# trainpos
cnf(c_0_291, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_0,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(esk5_3(np__1,esk2_0,X1),k1_numbers,k5_numbers))).# trainpos
cnf(c_0_307, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,esk3_0))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,esk3_0)))).# trainpos
cnf(c_0_315, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(X1)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X1))).# trainpos
cnf(c_0_312, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,esk3_0))|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk3_0))).# trainpos
cnf(c_0_320, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~v7_ordinal1(esk5_3(np__1,esk2_0,np__1))|~r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_318, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk3_0))).# trainpos
cnf(c_0_323, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_322, negated_conjecture, (~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk3_0))).# trainpos
cnf(c_0_326, negated_conjecture, (~v1_xxreal_0(np__1)|~r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_325, negated_conjecture, (~r1_xxreal_0(np__1,esk3_0))).# trainpos
cnf(c_0_328, negated_conjecture, (~r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_327, negated_conjecture, (r1_xxreal_0(esk3_0,np__1))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_128, plain, (v1_xreal_0(esk7_0))).#trainneg
cnf(i_0_113, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_123, plain, (v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_26, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_122, plain, (v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_129, plain, (v1_xxreal_0(esk7_0))).#trainneg
cnf(i_0_68, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_139, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_115, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_66, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_57, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_56, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_108, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_138, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_58, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_137, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_149, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_133, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_107, plain, (k18_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_89, plain, (k17_complex1(k17_complex1(X1))=k17_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_142, plain, (m1_subset_1(k18_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(k17_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_119, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_87, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_88, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_82, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_114, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_100, plain, (v1_seq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_99, plain, (v2_seq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_158, plain, (m1_subset_1(k2_power(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_54, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_55, plain, (v1_relat_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_150, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_148, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_81, plain, (v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_157, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_152, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_62, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_61, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_40, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_143, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_53, plain, (v3_valued_0(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_74, plain, (X3=X4|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,X1,X2)|~v1_funct_2(X3,X1,X2)|~r2_funct_2(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_120, plain, (k54_valued_1(k54_valued_1(X1))=k54_valued_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_146, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_147, plain, (v1_relat_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_145, plain, (v3_valued_0(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_24, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_94, plain, (v1_comseq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_144, plain, (v4_valued_0(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_60, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_98, plain, (k2_newton(X1,X2)=k1_newton(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_151, plain, (m1_subset_1(k2_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_154, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_155, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_127, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_97, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_126, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_109, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_85, plain, (v1_xreal_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_63, plain, (m1_subset_1(k2_seq_2(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_91, plain, (v2_comseq_2(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_118, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_17, plain, (k1_power(X1,k1_newton(X2,X1))=X2|v1_abian(X1)|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18, plain, (k1_newton(k1_power(X1,X2),X1)=X2|v1_abian(X1)|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_76, plain, (k1_seq_2(X1)=k2_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_78, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_19, plain, (k1_power(X1,k1_newton(X2,X1))=X2|~v1_xreal_0(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2))).#trainneg
cnf(i_0_84, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_117, plain, (r2_funct_2(X2,X3,X1,X1)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_funct_2(X4,X2,X3)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_50, plain, (v1_partfun1(k54_valued_1(X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,X2)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_11, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))|k1_seq_1(X1,esk4_2(X1,X2))!=k18_complex1(k1_seq_1(X2,esk4_2(X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_135, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_34, plain, (v1_xboole_0(X2)|v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_membered(X2)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_12, plain, (m2_subset_1(esk4_2(X1,X2),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_295, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_197, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_173, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_160, plain, (esk6_0=k6_numbers)).#trainneg
cnf(i_0_235, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_344, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_345, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_224, negated_conjecture, (v1_relat_1(esk2_0))).#trainneg
cnf(i_0_225, negated_conjecture, (v1_relat_1(esk1_0))).#trainneg
cnf(i_0_337, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_333, plain, (r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v1_xreal_0(k1_power(X1,X2))|~r1_xxreal_0(np__1,k1_power(X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2))).#trainneg
cnf(i_0_381, negated_conjecture, (r1_xxreal_0(np__1,k1_funct_1(esk1_0,X1))|~r1_xxreal_0(esk3_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_256, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_194, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_201, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_236, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers))).#trainneg
cnf(i_0_283, negated_conjecture, (v3_valued_0(esk1_0))).#trainneg
cnf(i_0_184, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_237, negated_conjecture, (v5_relat_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_238, negated_conjecture, (v5_relat_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_296, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_326, negated_conjecture, (k2_power(X1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_370, plain, (v1_xboole_0(X1)|~v1_xboole_0(k56_valued_1(X1,X2,X3))|~v1_funct_2(k56_valued_1(X1,X2,X3),X1,k1_numbers)|~v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_250, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_240, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_177, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_297, plain, (v1_xcmplx_0(k1_seq_2(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_329, negated_conjecture, (v1_partfun1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_327, plain, (v1_funct_1(k54_valued_1(X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_207, plain, (v1_xxreal_0(k6_numbers)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_209, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_210, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_220, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_178, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_366, plain, (v3_valued_0(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_267, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_324, plain, (r1_xxreal_0(X1,k1_newton(X2,X3))|~v1_xxreal_0(k1_newton(X2,X3))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,X2))).#trainneg
cnf(i_0_185, plain, (v1_xreal_0(k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_183, plain, (v1_xreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_272, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_334, negated_conjecture, (m1_subset_1(esk3_0,k1_numbers))).#trainneg
cnf(i_0_270, plain, (m1_subset_1(k1_power(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_331, plain, (v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(k1_power(X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2))).#trainneg
cnf(i_0_362, plain, (v1_funct_2(k56_valued_1(X1,X2,X3),X1,k1_numbers)|~v1_partfun1(k56_valued_1(X1,X2,X3),X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_309, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_218, plain, (v1_xcmplx_0(k6_numbers)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_302, plain, (v1_xreal_0(k2_seq_2(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_330, negated_conjecture, (v1_partfun1(esk1_0,k5_numbers))).#trainneg
cnf(i_0_179, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_284, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_285, negated_conjecture, (v1_valued_0(esk1_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_288, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_291, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_180, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_188, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_359, plain, (v1_relat_1(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_363, plain, (v1_partfun1(k56_valued_1(X1,X2,X3),X1)|~v1_funct_2(k56_valued_1(X1,X2,X3),X1,k1_numbers)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_382, negated_conjecture, (k2_power(X1,k1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0)))).#trainneg
cnf(i_0_300, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_294, plain, (r1_xxreal_0(X1,k17_complex1(X2))|~v1_xxreal_0(k17_complex1(X2))|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_335, plain, (m1_subset_1(k2_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_249, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(k6_numbers)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_189, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_226, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_198, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_212, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_223, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_199, plain, (v1_xreal_0(k2_power(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_190, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_193, plain, (v7_ordinal1(k17_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_401, plain, (v7_ordinal1(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_360, plain, (v1_xboole_0(k56_valued_1(X1,X2,X3))|~v1_xboole_0(X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_217, plain, (v1_xcmplx_0(k17_complex1(X1))|~v1_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_219, plain, (v1_xcmplx_0(k18_complex1(X1))|~v1_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_339, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_298, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_314, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_303, plain, (v7_ordinal1(k2_seq_2(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_349, negated_conjecture, (v1_partfun1(k54_valued_1(esk1_0),k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_411, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(X3,X2))|m2_subset_1(esk4_2(X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_195, plain, (v7_ordinal1(k18_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_227, negated_conjecture, (v3_membered(esk1_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_202, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_203, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_254, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_255, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_200, plain, (v7_ordinal1(k2_power(X1,X2))|~v6_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_206, plain, (v1_xxreal_0(k17_complex1(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_403, plain, (v1_xcmplx_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_277, plain, (m1_subset_1(k1_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_379, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_numbers)|r1_xxreal_0(X1,k6_numbers)|k2_seq_2(X2)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_368, plain, (v5_relat_1(k56_valued_1(X1,X2,X3),k1_numbers)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_361, plain, (v4_relat_1(k56_valued_1(X1,X2,X3),X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_404, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X3,X4),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_410, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,X3),X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X3))|~v1_funct_1(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_305, plain, (v1_xcmplx_0(k2_seq_2(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_376, plain, (k56_valued_1(X1,X2,k54_valued_1(X3))=k54_valued_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_292, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_213, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_304, plain, (v1_xxreal_0(k2_seq_2(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_214, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_357, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k1_numbers))|~r2_hidden(X1,k56_valued_1(X2,X3,X4))|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_208, plain, (v1_xxreal_0(k18_complex1(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_257, plain, (v1_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_211, plain, (v1_xxreal_0(k2_power(X1,X2))|~v2_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_222, plain, (v1_xcmplx_0(k2_power(X1,X2))|~v1_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_262, plain, (v1_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_273, plain, (v1_xreal_0(k2_newton(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_395, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xboole_0(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X4,X5))|~v1_funct_2(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X4,X5),X1,X2)|~v1_funct_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))|~m1_subset_1(X5,X3)|~v1_funct_2(X4,X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_367, plain, (v1_valued_0(k56_valued_1(X1,X2,X3))|~v1_membered(k1_numbers)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_354, plain, (v3_membered(k56_valued_1(X1,X2,X3))|~v3_membered(k2_zfmisc_1(X1,k1_numbers))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_406, plain, (X1=k56_valued_1(k5_numbers,k1_numbers,X2)|m2_subset_1(esk4_2(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_321, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_306, plain, (v1_funct_1(k54_valued_1(k54_valued_1(X1)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_276, plain, (v1_xcmplx_0(k2_newton(X1,X2))|~v1_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_396, plain, (v3_membered(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_390, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_274, plain, (v7_ordinal1(k2_newton(X1,X2))|~v6_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_275, plain, (v1_xxreal_0(k2_newton(X1,X2))|~v2_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_322, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_268, plain, (v4_valued_0(k54_valued_1(X1))|~v5_relat_1(k54_valued_1(X1),k4_numbers)|~v1_relat_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_318, plain, (v7_ordinal1(X2)|v1_abian(X1)|~v7_ordinal1(X1)|~v7_ordinal1(k1_power(X1,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_316, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_353, plain, (v1_xcmplx_0(k56_valued_1(X1,X2,X3))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_315, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_364, plain, (v1_partfun1(k54_valued_1(k56_valued_1(X1,X2,X3)),X1)|~v1_membered(k1_numbers)|~v1_funct_2(k56_valued_1(X1,X2,X3),X1,k1_numbers)|~v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_408, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(X1,k56_valued_1(k5_numbers,k1_numbers,X2)),k1_numbers,k5_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X2))|~v1_funct_1(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_351, plain, (v7_ordinal1(k56_valued_1(X1,X2,X3))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_385, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_407, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X2))|m2_subset_1(esk4_2(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_325, plain, (r1_xxreal_0(np__1,X2)|v1_abian(X1)|~v7_ordinal1(X1)|~v1_xreal_0(k1_power(X1,X2))|~r1_xxreal_0(np__1,k1_power(X1,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_387, plain, (v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2,X3)|v1_xboole_0(X1)|~v1_partfun1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_397, plain, (v3_valued_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_391, plain, (v3_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_388, plain, (v1_partfun1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|v1_xboole_0(X3)|v1_xboole_0(X1)|~v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_384, plain, (v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_350, plain, (v1_xreal_0(k56_valued_1(X1,X2,X3))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_399, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,k3_funct_2(X3,k1_zfmisc_1(X2),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(X2))))|~m1_subset_1(X5,X3)|~v1_funct_2(X4,X3,k1_zfmisc_1(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_402, plain, (v1_xxreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v2_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_352, plain, (v1_xxreal_0(k56_valued_1(X1,X2,X3))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_389, plain, (v1_partfun1(k54_valued_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5)),X2)|v1_xboole_0(X3)|v1_xboole_0(X1)|~v1_membered(X3)|~v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2,X3)|~v1_funct_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_386, plain, (v4_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_398, plain, (v1_valued_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v1_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_392, plain, (v1_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_393, plain, (v5_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X3)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_412, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_115, plain, (k2_xcmplx_0(k6_numbers,np__1)=np__1)).#trainneg
cnf(i_0_129, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_139, plain, (k2_xcmplx_0(np__1,k6_numbers)=np__1)).#trainneg
cnf(i_0_415, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2)))).#trainneg
cnf(i_0_450, negated_conjecture, (v7_ordinal1(esk3_0))).#trainneg
cnf(i_0_422, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_423, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_416, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2)))).#trainneg
cnf(i_0_417, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_464, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_468, negated_conjecture, (r1_xxreal_0(np__1,k1_seq_1(esk1_0,X1))|~r1_xxreal_0(esk3_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v3_valued_0(esk1_0))).#trainneg
cnf(i_0_466, plain, (r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,k1_power(X2,X1))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(k1_power(X2,X1),k1_numbers))).#trainneg
cnf(i_0_493, plain, (r1_xxreal_0(k6_numbers,X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_442, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(k6_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_512, negated_conjecture, (v1_funct_1(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_418, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_542, negated_conjecture, (v1_xcmplx_0(esk3_0))).#trainneg
cnf(i_0_428, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_429, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_419, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_420, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_480, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_527, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v1_xreal_0(k1_power(X2,X3))|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,k1_power(X2,X3))|~v1_xreal_0(X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k6_numbers,X3))).#trainneg
cnf(i_0_440, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_506, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_524, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_481, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_465, plain, (r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,k1_power(X2,X1))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k6_numbers,X1)|~v7_ordinal1(k1_power(X2,X1)))).#trainneg
cnf(i_0_508, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_451, negated_conjecture, (v1_xxreal_0(esk3_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_424, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_569, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_430, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_425, plain, (v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_partfun1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_426, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X1)|v1_xboole_0(X2)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2))).#trainneg
cnf(i_0_492, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_441, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_622, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_535, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_446, plain, (k17_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_517, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_427, plain, (v1_partfun1(k54_valued_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(X2)|~v1_membered(X2)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_640, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_541, negated_conjecture, (v1_xxreal_0(esk3_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_431, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_433, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_439, plain, (k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,X1),X2)=k18_complex1(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X1)))).#trainneg
cnf(i_0_445, plain, (k18_complex1(k17_complex1(X1))=k17_complex1(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k17_complex1(X1)))).#trainneg
cnf(i_0_528, plain, (r1_xxreal_0(X1,k1_newton(X2,X3))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,X2)|~v1_xreal_0(k1_newton(X2,X3)))).#trainneg
cnf(i_0_414, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk1_0,X2))|~v1_xxreal_0(k1_funct_1(esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_554, plain, (v1_funct_2(k54_valued_1(X3),X1,k1_numbers)|~v1_partfun1(k54_valued_1(X3),X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_485, plain, (v1_xboole_0(X1)|~v1_membered(X2)|~v1_xboole_0(k54_valued_1(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(k54_valued_1(X3),X1,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_557, plain, (v7_ordinal1(k2_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_488, negated_conjecture, (m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_472, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),k1_numbers)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0)))).#trainneg
cnf(i_0_453, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_455, negated_conjecture, (k1_power(X1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_529, plain, (r1_xxreal_0(X1,k1_newton(X2,X3))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k1_newton(X2,X3)))).#trainneg
cnf(i_0_484, negated_conjecture, (k2_power(X1,k1_funct_1(k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k54_valued_1(esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(esk2_0)))).#trainneg
cnf(i_0_570, plain, (r1_xxreal_0(X1,esk5_3(X2,X3,X4))|r1_xxreal_0(X2,k6_numbers)|k2_seq_2(X3)!=k6_numbers|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v2_comseq_2(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_xreal_0(esk5_3(X2,X3,X4)))).#trainneg
cnf(i_0_553, plain, (v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(k1_power(X2,X1),k5_numbers))).#trainneg
cnf(i_0_568, plain, (k54_valued_1(k56_valued_1(X1,k1_numbers,X2))=k56_valued_1(X1,k1_numbers,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_543, plain, (v1_xreal_0(k1_power(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_755, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_648, plain, (v1_xreal_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_470, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),k1_numbers)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_665, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_711, plain, (r1_xxreal_0(k2_power(X1,X2),k2_power(X1,X2))|~v1_xreal_0(X3)|~v2_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_712, plain, (r1_xxreal_0(k2_power(X1,X2),k2_power(X1,X2))|~v7_ordinal1(X3)|~v2_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_489, plain, (m1_subset_1(k54_valued_1(k56_valued_1(X1,X2,X3)),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_663, plain, (v4_relat_1(k54_valued_1(X3),X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_556, plain, (v1_xreal_0(k2_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_606, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_611, plain, (v1_xboole_0(X1)|v7_ordinal1(k1_funct_1(X3,X4))|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_625, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_572, plain, (v1_relat_1(k54_valued_1(X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_651, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_666, negated_conjecture, (m1_subset_1(k1_funct_1(esk1_0,X1),k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_490, plain, (m1_subset_1(k54_valued_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|v1_xboole_0(X1)|~v1_membered(X3)|~v1_funct_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_817, plain, (r1_xxreal_0(k6_numbers,np__1))).#trainneg
cnf(i_0_820, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_491, plain, (m1_subset_1(k54_valued_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_504, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1279, plain, (m1_subset_1(k56_valued_1(X1,k1_numbers,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_294, plain, (r1_xxreal_0(X1,k17_complex1(X2))|~v1_xcmplx_0(X2)|~v1_xxreal_0(k17_complex1(X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k6_numbers))).#trainneg
cnf(i_0_507, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k1_funct_1(X1,X2))|~v1_xreal_0(X3)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_828, plain, (r1_xxreal_0(k6_numbers,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_900, negated_conjecture, (~v1_xboole_0(k54_valued_1(esk2_0))|~v1_funct_2(k54_valued_1(esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1052, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_830, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_513, plain, (v1_funct_1(k54_valued_1(k56_valued_1(X1,X2,X3)))|~v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_514, plain, (v1_funct_1(k54_valued_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5)))|v1_xboole_0(X1)|~v1_membered(X3)|~v1_funct_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_240, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1011, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_983, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k6_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_869, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1015, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X3,X2))|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_973, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_897, negated_conjecture, (v1_valued_0(k54_valued_1(esk2_0)))).#trainneg
cnf(i_0_1191, plain, (k54_valued_1(k54_valued_1(X2))=k54_valued_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_515, plain, (v1_funct_1(k54_valued_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_982, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k6_numbers))).#trainneg
cnf(i_0_1104, negated_conjecture, (~v1_xboole_0(k54_valued_1(esk1_0))|~v1_funct_2(k54_valued_1(esk1_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_940, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_520, plain, (v3_valued_0(k54_valued_1(X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_986, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_1148, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_987, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_890, negated_conjecture, (v4_relat_1(k54_valued_1(esk2_0),k5_numbers))).#trainneg
cnf(i_0_1125, negated_conjecture, (v1_relat_1(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_1134, negated_conjecture, (v3_valued_0(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_1105, negated_conjecture, (~v1_xboole_0(k54_valued_1(k54_valued_1(esk2_0)))|~v1_funct_2(k54_valued_1(k54_valued_1(esk2_0)),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1139, negated_conjecture, (~v1_xboole_0(k54_valued_1(k54_valued_1(esk1_0)))|~v1_funct_2(k54_valued_1(k54_valued_1(esk1_0)),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_870, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_522, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_526, plain, (r1_xxreal_0(X1,X3)|v1_abian(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v1_xreal_0(k1_power(X2,X3))|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,k1_power(X2,X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1094, negated_conjecture, (v1_funct_2(k54_valued_1(esk1_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_845, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_959, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k6_numbers))).#trainneg
cnf(i_0_1135, negated_conjecture, (v1_valued_0(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_882, negated_conjecture, (v3_membered(k54_valued_1(esk2_0))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_921, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_544, plain, (v7_ordinal1(k1_power(X1,X2))|~v6_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_545, plain, (v1_xxreal_0(k1_power(X1,X2))|~v2_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_958, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X4,X2))).#trainneg
cnf(i_0_865, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)))).#trainneg
cnf(i_0_1132, negated_conjecture, (m1_subset_1(k54_valued_1(k54_valued_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_848, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_1003, plain, (m1_subset_1(k3_funct_2(X1,k5_numbers,X2,X3),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k5_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_850, negated_conjecture, (r1_xxreal_0(X1,k1_seq_1(esk1_0,X2))|~v1_xxreal_0(k1_seq_1(esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1137, negated_conjecture, (v1_funct_1(k54_valued_1(k54_valued_1(esk1_0))))).#trainneg
cnf(i_0_1141, negated_conjecture, (k54_valued_1(k54_valued_1(esk1_0))=k54_valued_1(esk1_0))).#trainneg
cnf(i_0_546, plain, (v1_xcmplx_0(k1_power(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_831, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_898, negated_conjecture, (v5_relat_1(k54_valued_1(esk2_0),k1_numbers))).#trainneg
cnf(i_0_944, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_957, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_547, plain, (m1_subset_1(X2,k1_numbers)|v1_abian(X1)|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(X2,X1),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_548, plain, (m1_subset_1(X2,k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2))).#trainneg
cnf(i_0_822, plain, (r1_xxreal_0(X1,k18_complex1(X2))|~v1_xxreal_0(k18_complex1(X2))|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1053, plain, (r2_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k18_complex1(k1_seq_1(X1,esk4_2(k56_valued_1(k5_numbers,k1_numbers,X1),X2)))!=k18_complex1(k1_seq_1(X2,esk4_2(k56_valued_1(k5_numbers,k1_numbers,X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X1),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X1))|~m2_subset_1(esk4_2(k56_valued_1(k5_numbers,k1_numbers,X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_864, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)))).#trainneg
cnf(i_0_843, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_842, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2),k1_numbers))).#trainneg
cnf(i_0_854, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v1_xcmplx_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1200, plain, (r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,k1_power(X2,X1))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1127, negated_conjecture, (v4_relat_1(k54_valued_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1136, negated_conjecture, (v5_relat_1(k54_valued_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_1143, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0)))).#trainneg
cnf(i_0_832, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1119, negated_conjecture, (v3_membered(k54_valued_1(esk1_0))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1222, plain, (r1_xxreal_0(k2_power(X1,X2),k2_power(X1,X2))|~v2_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_552, negated_conjecture, (v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_1493, plain, (v1_funct_1(k56_valued_1(X1,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1906, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1469, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_1104, negated_conjecture, (~v1_xboole_0(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_1546, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0))).#trainneg
cnf(i_0_1562, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_607, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1432, plain, (m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1436, plain, (m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1736, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,k6_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1631, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)))).#trainneg
cnf(i_0_484, negated_conjecture, (k2_power(X1,k1_funct_1(k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1477, plain, (v1_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_612, plain, (k56_valued_1(X1,X2,X3)=k6_numbers|~v1_membered(X2)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_613, plain, (v1_xboole_0(k54_valued_1(X3))|~v1_membered(X2)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1509, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1870, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k18_complex1(k1_seq_1(X1,esk4_2(k54_valued_1(X1),X2)))!=k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1543, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X3,X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1529, plain, (r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1747, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_620, plain, (v1_xboole_0(k56_valued_1(X1,X2,X3))|v1_xboole_0(k2_zfmisc_1(X1,k1_numbers))|m1_subset_1(X4,k2_zfmisc_1(X1,k1_numbers))|~m1_subset_1(X4,k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_621, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X2)|m1_subset_1(X5,X2)|v1_xboole_0(X1)|~m1_subset_1(X5,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1785, plain, (v1_xreal_0(k3_funct_2(X1,k5_numbers,X2,X3))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k5_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1510, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1883, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(np__1)|~r1_xxreal_0(X3,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1675, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1784, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)))).#trainneg
cnf(i_0_1885, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(np__1)|~r1_xxreal_0(X3,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2),k1_numbers))).#trainneg
cnf(i_0_1530, plain, (r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1748, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0)))).#trainneg
cnf(i_0_624, plain, (v1_xcmplx_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_628, negated_conjecture, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(esk2_0,X2))|m2_subset_1(esk4_2(X1,esk2_0),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1788, plain, (v1_xcmplx_0(k3_funct_2(X1,k5_numbers,X2,X3))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k5_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1810, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk1_0),X1))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(esk1_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1632, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(k54_valued_1(X2))!=k6_numbers|~v3_valued_0(k54_valued_1(X2))|~v2_comseq_2(k54_valued_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(X2),esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1790, negated_conjecture, (r1_xxreal_0(X1,k1_seq_1(esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(k1_seq_1(esk1_0,X2)))).#trainneg
cnf(i_0_1904, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1791, negated_conjecture, (r1_xxreal_0(X1,k1_seq_1(esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(k1_seq_1(esk1_0,X2)))).#trainneg
cnf(i_0_1541, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X2,k5_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1884, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk1_0,X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1743, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_1891, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1593, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1548, plain, (r1_xxreal_0(k1_seq_2(X1),k1_seq_2(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1907, negated_conjecture, (r1_xxreal_0(k6_numbers,esk3_0)|r1_xxreal_0(esk3_0,k6_numbers))).#trainneg
cnf(i_0_629, negated_conjecture, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(esk1_0,X2))|m2_subset_1(esk4_2(X1,esk1_0),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1694, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1660, negated_conjecture, (v4_valued_0(k54_valued_1(esk2_0))|~v5_relat_1(k54_valued_1(esk2_0),k4_numbers))).#trainneg
cnf(i_0_1608, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_630, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k56_valued_1(k5_numbers,X3,X4),X2))|m2_subset_1(esk4_2(X1,k56_valued_1(k5_numbers,X3,X4)),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,X3,X4),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,X3,X4))|~v1_funct_1(X1)|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_631, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X4,X5),X2))|m2_subset_1(esk4_2(X1,k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X4,X5)),k1_numbers,k5_numbers)|v1_xboole_0(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X4,X5),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),X4,X5))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|~m1_subset_1(X5,X3)|~v1_funct_2(X4,X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_1739, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X3,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X2))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1903, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1654, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v2_comseq_2(k54_valued_1(esk2_0))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(esk2_0),esk5_3(X1,k54_valued_1(esk2_0),X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1872, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k18_complex1(k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2)))!=k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1425, plain, (v1_xboole_0(X1)|~v1_xboole_0(k56_valued_1(X1,k1_numbers,X2))|~v1_funct_2(k56_valued_1(X1,k1_numbers,X2),X1,k1_numbers)|~v1_funct_1(k56_valued_1(X1,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1554, plain, (r1_xxreal_0(k2_seq_2(X1),k2_seq_2(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1698, negated_conjecture, (v1_membered(k54_valued_1(esk2_0))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1695, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_funct_1(k54_valued_1(esk2_0),X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1801, negated_conjecture, (v4_valued_0(k54_valued_1(esk1_0))|~v5_relat_1(k54_valued_1(esk1_0),k4_numbers))).#trainneg
cnf(i_0_1715, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_632, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k2_zfmisc_1(k5_numbers,k1_numbers),X2))|m2_subset_1(esk4_2(X1,k2_zfmisc_1(k5_numbers,k1_numbers)),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_635, negated_conjecture, (v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v6_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1414, plain, (v1_partfun1(k56_valued_1(X1,k1_numbers,X2),X1)|~v1_funct_2(k56_valued_1(X1,k1_numbers,X2),X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1199, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,k1_power(X3,X2))|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k6_numbers,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1693, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1924, negated_conjecture, (v1_membered(k54_valued_1(esk1_0))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_636, negated_conjecture, (v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v6_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1708, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_878, negated_conjecture, (v1_xreal_0(k54_valued_1(esk2_0))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_879, negated_conjecture, (v7_ordinal1(k54_valued_1(esk2_0))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1590, plain, (r1_xxreal_0(np__1,k2_power(X1,X2))|r1_xxreal_0(k2_power(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_641, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2305, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1990, plain, (v4_relat_1(k54_valued_1(X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2086, plain, (k54_valued_1(k6_numbers)=k6_numbers|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_642, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2464, negated_conjecture, (r1_xxreal_0(esk3_0,np__1)|~r1_xxreal_0(esk3_0,k6_numbers))).#trainneg
cnf(i_0_880, negated_conjecture, (v1_xxreal_0(k54_valued_1(esk2_0))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_643, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v7_ordinal1(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1961, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1971, plain, (m1_subset_1(k3_funct_2(X1,X2,X3,X4),k5_numbers)|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2058, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_2175, plain, (r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2369, negated_conjecture, (r1_xxreal_0(esk3_0,np__1)|r1_xxreal_0(k6_numbers,esk3_0))).#trainneg
cnf(i_0_2105, plain, (k54_valued_1(k6_numbers)=k6_numbers|~v1_membered(X2)|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_647, plain, (v1_xcmplx_0(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_881, negated_conjecture, (v1_xcmplx_0(k54_valued_1(esk2_0))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1115, negated_conjecture, (v1_xreal_0(k54_valued_1(esk1_0))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2304, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_649, plain, (v7_ordinal1(k1_newton(X1,X2))|~v6_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_650, plain, (v1_xxreal_0(k1_newton(X1,X2))|~v2_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1963, plain, (m1_subset_1(k1_newton(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2259, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk1_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2170, plain, (r1_xxreal_0(k1_seq_1(X1,X2),X3)|~v7_ordinal1(X3)|~r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2019, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2159, negated_conjecture, (r1_xxreal_0(esk3_0,X1)|~v7_ordinal1(X1)|~r1_xxreal_0(esk3_0,k6_numbers))).#trainneg
cnf(i_0_652, plain, (m1_subset_1(X2,k1_numbers)|v1_abian(X1)|~v7_ordinal1(X1)|~m1_subset_1(k1_power(X1,X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2481, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1116, negated_conjecture, (v7_ordinal1(k54_valued_1(esk1_0))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1117, negated_conjecture, (v1_xxreal_0(k54_valued_1(esk1_0))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_653, plain, (m1_subset_1(X2,k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k1_power(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2))).#trainneg
cnf(i_0_658, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X3))|~v1_funct_1(X1)|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2173, plain, (r1_xxreal_0(k3_funct_2(X1,X2,X3,X4),X5)|v1_xboole_0(X1)|~v7_ordinal1(X5)|~r1_xxreal_0(k3_funct_2(X1,X2,X3,X4),k6_numbers)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2021, plain, (m1_subset_1(k1_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1991, plain, (v1_funct_2(k54_valued_1(X1),X2,k1_numbers)|~v1_partfun1(k54_valued_1(X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2066, negated_conjecture, (k2_power(X1,k1_seq_1(k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2298, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2006, plain, (k54_valued_1(k54_valued_1(k54_valued_1(X1)))=k54_valued_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1118, negated_conjecture, (v1_xcmplx_0(k54_valued_1(esk1_0))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_661, plain, (v5_relat_1(k54_valued_1(X3),k1_numbers)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_2401, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1)))).#trainneg
cnf(i_0_2057, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_funct_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(k1_funct_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1948, plain, (v1_funct_1(k54_valued_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2169, plain, (r1_xxreal_0(k18_complex1(X1),X2)|~v7_ordinal1(X2)|~r1_xxreal_0(k18_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1531, plain, (k2_zfmisc_1(X1,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_667, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k56_valued_1(X1,X2,X3),X4),k1_numbers)|~m1_subset_1(X4,X1)|~v1_funct_2(k56_valued_1(X1,X2,X3),X1,k1_numbers)|~v1_funct_1(k56_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_668, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X6),X3)|v1_xboole_0(X2)|~m1_subset_1(X6,X1)|~v1_funct_2(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X1,X3)|~v1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2420, plain, (v1_xboole_0(X1)|~v1_xboole_0(k54_valued_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),X1,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2396, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_2225, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_1083, plain, (k18_complex1(k1_seq_1(k54_valued_1(X1),X2))=k1_seq_1(k54_valued_1(X1),X2)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2302, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,k6_numbers))).#trainneg
cnf(i_0_1550, plain, (r1_xxreal_0(k17_complex1(X1),k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2355, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k18_complex1(k1_seq_1(esk2_0,X1))|~v3_valued_0(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~r1_xxreal_0(k6_numbers,k18_complex1(k1_seq_1(esk2_0,X1)))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_669, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k2_zfmisc_1(X1,X2),X3),X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2085, plain, (v1_funct_1(k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1677, plain, (k2_zfmisc_1(k6_numbers,X1)=k6_numbers)).#trainneg
cnf(i_0_1559, plain, (r1_xxreal_0(k2_power(X1,X2),k2_power(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1410, plain, (v1_relat_1(k56_valued_1(X1,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_670, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X3))|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2138, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k18_complex1(k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),X1)))!=k18_complex1(k1_seq_1(X1,esk4_2(k54_valued_1(esk1_0),X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2397, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_2762, plain, (v1_funct_2(k54_valued_1(X1),X2,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_partfun1(X1,X2))).#trainneg
cnf(i_0_2776, plain, (r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3021, plain, (m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k6_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3011, plain, (v1_funct_1(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2868, plain, (v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_3018, plain, (v3_valued_0(k6_numbers))).#trainneg
cnf(i_0_2562, plain, (k54_valued_1(k6_numbers)=k6_numbers|~v1_xboole_0(X1)|~v1_funct_1(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_2881, plain, (v5_relat_1(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_1420, plain, (v3_valued_0(k56_valued_1(X1,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1421, plain, (v1_valued_0(k56_valued_1(X1,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_672, plain, (k18_complex1(k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,X1),X2))=k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X1),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_673, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k56_valued_1(k5_numbers,k1_numbers,X3),X2))|m2_subset_1(esk4_2(X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X3))|~v1_funct_1(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2898, plain, (~v1_xboole_0(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2834, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2931, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(k54_valued_1(X2))!=k6_numbers|~v3_valued_0(k54_valued_1(X2))|~v2_comseq_2(k54_valued_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(X2),esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_xcmplx_0(k1_seq_1(X2,esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(esk5_3(X1,k54_valued_1(X2),X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2736, negated_conjecture, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k54_valued_1(esk2_0),X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(esk2_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3019, plain, (v1_valued_0(k6_numbers))).#trainneg
cnf(i_0_3012, plain, (k54_valued_1(k6_numbers)=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3066, plain, (v4_relat_1(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_918, plain, (r1_xxreal_0(X1,k2_power(X2,X3))|r1_xxreal_0(k2_power(X2,X3),X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_941, plain, (r1_xxreal_0(X1,k2_power(X2,X3))|r1_xxreal_0(k2_power(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_675, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X2))|m2_subset_1(esk4_2(X1,k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3022, plain, (v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3051, plain, (v3_valued_0(X1)|~v3_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3112, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k18_complex1(k1_seq_1(X1,esk4_2(k54_valued_1(esk1_0),X1)))!=k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),X1))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3052, plain, (v1_valued_0(X1)|~v1_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2737, negated_conjecture, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k54_valued_1(esk1_0),X2))|~r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(esk1_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2870, plain, (v4_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_2782, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2831, plain, (r1_xxreal_0(k1_seq_1(k54_valued_1(X1),X2),X3)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v7_ordinal1(X3)|~r1_xxreal_0(k1_seq_1(k54_valued_1(X1),X2),k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2815, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(k1_seq_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_2925, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2875, plain, (v1_funct_2(k6_numbers,X1,k6_numbers)|~v1_partfun1(k6_numbers,X1))).#trainneg
cnf(i_0_2837, plain, (v1_funct_2(X1,X2,k6_numbers)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3077, plain, (v5_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_3071, plain, (v1_funct_2(k6_numbers,k6_numbers,X1)|~v1_partfun1(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_2873, plain, (v1_funct_1(k54_valued_1(k6_numbers))|~v1_membered(k6_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_714, plain, (r1_xxreal_0(X1,k2_power(X2,X3))|r1_xxreal_0(k2_power(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_691, plain, (k56_valued_1(X1,X2,k54_valued_1(X3))=k54_valued_1(X3)|~v1_membered(X2)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k54_valued_1(X3))|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3129, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2942, plain, (k18_complex1(k1_seq_1(k54_valued_1(X1),X2))=k1_seq_1(k54_valued_1(X1),X2)|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2940, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(esk2_0),X1))=k1_seq_1(k54_valued_1(esk2_0),X1)|~v1_xcmplx_0(k1_seq_1(k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2584, plain, (m1_subset_1(k3_funct_2(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2534, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X3,X1))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2855, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_696, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|r1_xxreal_0(k2_seq_2(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_697, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|r1_xxreal_0(k2_seq_2(X2),X1)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2941, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(esk1_0),X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~v1_xcmplx_0(k1_seq_1(k54_valued_1(esk1_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2939, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2))!=k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_seq_1(X1,esk4_2(k54_valued_1(X1),X2)))|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2894, plain, (v1_xboole_0(X1)|~v1_xboole_0(k54_valued_1(X2))|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),X1,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2533, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X1,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3107, plain, (k18_complex1(k1_seq_1(k54_valued_1(X1),X2))=k1_seq_1(k54_valued_1(X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2684, plain, (v1_xreal_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3237, plain, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3155, plain, (v1_xboole_0(X1)|~v1_xboole_0(k54_valued_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2)|~v1_valued_0(X2)|~v1_partfun1(X2,X1))).#trainneg
cnf(i_0_3203, plain, (k54_valued_1(k6_numbers)=k6_numbers|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3588, negated_conjecture, (~m1_subset_1(esk2_0,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_698, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|r1_xxreal_0(k2_seq_2(X2),X1)|~v7_ordinal1(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3281, plain, (r1_xxreal_0(k2_power(X1,np__1),X2)|r1_xxreal_0(X2,k2_power(X1,np__1))|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3304, plain, (r1_xxreal_0(k2_power(X1,np__1),X2)|r1_xxreal_0(X2,k2_power(X1,np__1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_700, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k1_numbers))|~r2_hidden(X1,k54_valued_1(X4))|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3400, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_3678, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3353, negated_conjecture, (r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(esk3_0,k6_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3376, negated_conjecture, (k1_power(X1,k1_seq_1(k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(k1_seq_1(k54_valued_1(esk2_0),X1)))).#trainneg
cnf(i_0_3684, plain, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(k54_valued_1(X2))!=k6_numbers|~v3_valued_0(k54_valued_1(X2))|~v2_comseq_2(k54_valued_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(X2),esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~m2_subset_1(esk5_3(X1,k54_valued_1(X2),X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3596, plain, (X1=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3409, plain, (v1_funct_1(k54_valued_1(k6_numbers))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3589, negated_conjecture, (~m1_subset_1(esk1_0,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3593, plain, (~m1_subset_1(k1_numbers,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3594, plain, (~m1_subset_1(k5_numbers,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1712, plain, (r1_xxreal_0(X1,k2_power(X2,X3))|r1_xxreal_0(k2_power(X2,X3),X1)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_706, plain, (r1_xxreal_0(X1,k18_complex1(X2))|r1_xxreal_0(k18_complex1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_707, plain, (r1_xxreal_0(X1,k18_complex1(X2))|r1_xxreal_0(k18_complex1(X2),X1)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3326, negated_conjecture, (k1_power(X1,k1_seq_1(k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3694, plain, (k1_seq_1(k54_valued_1(k54_valued_1(X1)),X2)=k1_seq_1(k54_valued_1(X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(k54_valued_1(X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(k54_valued_1(X1)),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(k54_valued_1(X1)))|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3695, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2))!=k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3342, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3370, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3419, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k18_complex1(k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2)))!=k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3323, plain, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3410, plain, (v1_relat_1(k54_valued_1(k6_numbers))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3590, negated_conjecture, (~m1_subset_1(k54_valued_1(esk2_0),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_708, plain, (r1_xxreal_0(X1,k18_complex1(X2))|r1_xxreal_0(k18_complex1(X2),X1)|~v7_ordinal1(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3072, plain, (v1_partfun1(k6_numbers,k6_numbers)|v1_xboole_0(X1)|~v1_funct_2(k6_numbers,k6_numbers,X1))).#trainneg
cnf(i_0_3591, negated_conjecture, (~m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1411, plain, (v1_xboole_0(k56_valued_1(X1,k1_numbers,X2))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_718, negated_conjecture, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v2_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3499, plain, (m1_subset_1(k1_seq_1(k54_valued_1(X1),X2),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3181, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_3503, plain, (v1_xreal_0(k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3243, plain, (~v1_xboole_0(k54_valued_1(X1))|~v2_comseq_2(k54_valued_1(X1))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3380, negated_conjecture, (r1_xxreal_0(np__1,k1_seq_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(k1_seq_1(k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_3675, plain, (m1_subset_1(k1_seq_1(k54_valued_1(X1),X2),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3313, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k54_valued_1(X2))|m2_subset_1(esk4_2(X1,k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(k54_valued_1(X2))|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3350, plain, (v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2836, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3592, plain, (~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_719, negated_conjecture, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v2_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1412, plain, (v4_relat_1(k56_valued_1(X1,k1_numbers,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1422, plain, (v5_relat_1(k56_valued_1(X1,k1_numbers,X2),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_722, negated_conjecture, (v1_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2717, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k1_seq_1(k54_valued_1(esk2_0),X1)))).#trainneg
cnf(i_0_2812, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_2816, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_funct_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_3347, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(esk1_0),X1))!=k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),X1))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_3679, plain, (v1_xreal_0(k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3685, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X2)))|k1_seq_1(X1,esk4_2(X1,k54_valued_1(X2)))!=k1_seq_1(k54_valued_1(X2),esk4_2(X1,k54_valued_1(X2)))|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m2_subset_1(esk4_2(X1,k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3183, negated_conjecture, (r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_3009, plain, (v1_funct_1(k6_numbers)|~v1_xboole_0(X1)|~v1_funct_1(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_2858, plain, (v5_relat_1(X1,k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_723, negated_conjecture, (v1_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1404, plain, (v3_membered(k56_valued_1(X1,k1_numbers,X2))|~v3_membered(k2_zfmisc_1(X1,k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3431, plain, (m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k54_valued_1(X3))|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2783, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3752, plain, (~v2_comseq_2(k6_numbers)|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3718, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(X2),X1))=k1_seq_1(esk2_0,X1)|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3720, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(X2),X1))=k1_seq_1(k54_valued_1(esk2_0),X1)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X2),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3719, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(X2),X1))=k1_seq_1(esk1_0,X1)|m2_subset_1(esk4_2(esk1_0,X2),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3747, plain, (v4_valued_0(k6_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3790, plain, (r1_xxreal_0(X1,k2_power(np__1,np__1))|r1_xxreal_0(k2_power(np__1,np__1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3807, plain, (r1_xxreal_0(np__1,k2_power(X1,np__1))|r1_xxreal_0(k2_power(X1,np__1),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_728, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xboole_0(k1_funct_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))|~m1_subset_1(X5,X1)|~v1_funct_2(k1_funct_1(X4,X5),X3,X2)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(k1_funct_1(X4,X5))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3721, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(X2),X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X2),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_729, plain, (v1_valued_0(k54_valued_1(X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3793, plain, (r1_xxreal_0(X1,k2_power(X2,np__1))|r1_xxreal_0(k2_power(X2,np__1),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3810, plain, (r1_xxreal_0(X1,k2_power(X2,np__1))|r1_xxreal_0(k2_power(X2,np__1),X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4271, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v2_comseq_2(k54_valued_1(esk2_0))|~m2_subset_1(esk5_3(np__1,k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1)))).#trainneg
cnf(i_0_3971, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k18_complex1(k1_seq_1(X1,esk4_2(k54_valued_1(esk2_0),X1)))!=k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(esk2_0),X1))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3849, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_3053, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2557, plain, (k54_valued_1(k6_numbers)=k6_numbers|~v1_xboole_0(X1)|~v1_funct_1(k54_valued_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_732, plain, (v1_membered(k56_valued_1(X1,X2,X3))|~v3_membered(k2_zfmisc_1(X1,k1_numbers))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3791, plain, (r1_xxreal_0(X1,k2_power(k2_power(X2,X3),np__1))|r1_xxreal_0(k2_power(k2_power(X2,X3),np__1),X1)|~v1_xreal_0(X1)|~v6_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_733, plain, (v3_membered(k54_valued_1(X3))|~v3_membered(k2_zfmisc_1(X1,k1_numbers))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_4002, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k18_complex1(k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(esk2_0),X1)))!=k18_complex1(k1_seq_1(X1,esk4_2(k54_valued_1(esk2_0),X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3725, plain, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(X1),X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4196, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)=k1_funct_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1),k5_numbers)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_4179, negated_conjecture, (k1_power(X1,k1_funct_1(k54_valued_1(esk2_0),X1))=k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_3597, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3808, plain, (r1_xxreal_0(k2_power(X1,X2),k2_power(X3,np__1))|r1_xxreal_0(k2_power(X3,np__1),k2_power(X1,X2))|~v7_ordinal1(X3)|~v6_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3740, plain, (v1_xboole_0(X1)|~v1_partfun1(k6_numbers,X1)|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3735, plain, (v1_xboole_0(X1)|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_2(k6_numbers,X1,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4070, plain, (r1_xxreal_0(k1_seq_1(k54_valued_1(X1),X2),X3)|~v7_ordinal1(X3)|~r1_xxreal_0(k1_seq_1(k54_valued_1(X1),X2),k6_numbers)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3872, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(k54_valued_1(esk1_0))!=k6_numbers|~v2_comseq_2(k54_valued_1(esk1_0))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(esk1_0),esk5_3(X1,k54_valued_1(esk1_0),X2)))|~m2_subset_1(esk5_3(X1,k54_valued_1(esk1_0),X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3916, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k1_seq_1(k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)))).#trainneg
cnf(i_0_3732, plain, (v1_xboole_0(X1)|~v1_valued_0(k54_valued_1(X2))|~v1_partfun1(k54_valued_1(X2),X1)|~v1_xboole_0(k54_valued_1(X2))|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(k54_valued_1(X2))|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1556, plain, (r1_xxreal_0(k18_complex1(X1),k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4201, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,esk1_0))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),esk1_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_741, plain, (v1_membered(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3108, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(esk1_0)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4257, plain, (v1_membered(k56_valued_1(X1,k1_numbers,X2))|~v3_membered(k2_zfmisc_1(X1,k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3282, plain, (r1_xxreal_0(k2_power(X1,k2_power(X2,X3)),X4)|r1_xxreal_0(X4,k2_power(X1,k2_power(X2,X3)))|~v7_ordinal1(X1)|~v1_xreal_0(X4)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_742, plain, (v3_membered(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_743, plain, (k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5)=k6_numbers|v1_xboole_0(X1)|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4509, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~v1_xreal_0(k1_funct_1(esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,k1_funct_1(esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,esk2_0,X1),k5_numbers))).#trainneg
cnf(i_0_4799, negated_conjecture, (k18_complex1(k1_seq_1(esk1_0,X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),esk1_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4766, plain, (v1_xboole_0(X1)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_xboole_0(k54_valued_1(X2))|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2)|~v4_relat_1(X2,X1)|~v1_partfun1(X2,X1))).#trainneg
cnf(i_0_4321, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k1_seq_1(esk2_0,X2)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4391, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k1_seq_1(esk1_0,X2)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4316, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))!=k1_seq_1(esk2_0,esk4_2(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4373, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk4_2(esk1_0,X2),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X2))!=k6_numbers|~v3_valued_0(k54_valued_1(X2))|~v2_comseq_2(k54_valued_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(esk1_0,esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~m2_subset_1(esk5_3(X1,k54_valued_1(X2),X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4680, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_744, plain, (v1_xboole_0(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4415, plain, (r1_xxreal_0(k2_power(np__1,np__1),np__1)|r1_xxreal_0(np__1,k2_power(np__1,np__1)))).#trainneg
cnf(i_0_4419, plain, (r1_xxreal_0(k2_power(np__1,np__1),X1)|r1_xxreal_0(X1,k2_power(np__1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_749, plain, (r1_xxreal_0(X1,k2_newton(X2,X3))|r1_xxreal_0(k2_newton(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4755, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(k6_numbers,k1_seq_1(k54_valued_1(esk2_0),X1))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_4386, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4315, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))!=k1_seq_1(esk2_0,esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4303, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X2))!=k6_numbers|~v3_valued_0(k54_valued_1(X2))|~v2_comseq_2(k54_valued_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(esk2_0,esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~m2_subset_1(esk5_3(X1,k54_valued_1(X2),X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4392, negated_conjecture, (k1_seq_1(esk2_0,X2)=k1_seq_1(esk1_0,X2)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4513, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~v1_xreal_0(k1_funct_1(esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,k1_funct_1(esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_4385, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4338, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X2),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X2))!=k6_numbers|~v3_valued_0(k54_valued_1(X2))|~v2_comseq_2(k54_valued_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(esk2_0),esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~m2_subset_1(esk5_3(X1,k54_valued_1(X2),X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4463, negated_conjecture, (k1_seq_1(k54_valued_1(esk1_0),X2)=k1_seq_1(esk2_0,X2)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4356, negated_conjecture, (k1_seq_1(k54_valued_1(esk2_0),X2)=k1_seq_1(esk2_0,X2)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_750, plain, (r1_xxreal_0(X1,k2_newton(X2,X3))|r1_xxreal_0(k2_newton(X2,X3),X1)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_751, plain, (r1_xxreal_0(X1,k2_newton(X2,X3))|r1_xxreal_0(k2_newton(X2,X3),X1)|~v7_ordinal1(X1)|~v2_membered(k1_numbers)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4393, negated_conjecture, (k1_seq_1(k54_valued_1(esk2_0),X2)=k1_seq_1(esk1_0,X2)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4514, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1))))).#trainneg
cnf(i_0_4065, negated_conjecture, (r1_xxreal_0(np__1,k1_funct_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4462, negated_conjecture, (k1_seq_1(k54_valued_1(esk1_0),X2)=k1_seq_1(esk1_0,X2)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4464, negated_conjecture, (k1_seq_1(k54_valued_1(esk2_0),X2)=k1_seq_1(k54_valued_1(esk1_0),X2)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4798, negated_conjecture, (k56_valued_1(k5_numbers,k1_numbers,esk1_0)=k54_valued_1(esk1_0)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk1_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk1_0))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),esk1_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_756, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4420, plain, (r1_xxreal_0(k2_power(np__1,np__1),X1)|r1_xxreal_0(X1,k2_power(np__1,np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4834, plain, (r1_xxreal_0(X1,k2_power(X2,k2_power(X3,np__1)))|r1_xxreal_0(k2_power(X2,k2_power(X3,np__1)),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_757, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_758, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v7_ordinal1(X1)|~v2_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4014, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k18_complex1(k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2)))!=k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2)))).#trainneg
cnf(i_0_4310, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X1)))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1))))!=k1_seq_1(esk2_0,esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4444, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X2),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X2))!=k6_numbers|~v3_valued_0(k54_valued_1(X2))|~v2_comseq_2(k54_valued_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(esk1_0),esk5_3(X1,k54_valued_1(X2),X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~m2_subset_1(esk5_3(X1,k54_valued_1(X2),X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4350, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))!=k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4516, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_4804, negated_conjecture, (k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0))=k54_valued_1(esk1_0)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(esk1_0)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4416, plain, (r1_xxreal_0(k2_power(np__1,np__1),k2_power(X1,X2))|r1_xxreal_0(k2_power(X1,X2),k2_power(np__1,np__1))|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3305, plain, (r1_xxreal_0(k2_power(X1,k2_power(X2,X3)),X4)|r1_xxreal_0(X4,k2_power(X1,k2_power(X2,X3)))|~v7_ordinal1(X4)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_3111, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k18_complex1(k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),X1)))!=k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(esk1_0),X1))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_2536, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3))).#trainneg
cnf(i_0_4355, negated_conjecture, (k1_seq_1(k54_valued_1(esk2_0),X2)=k1_seq_1(k54_valued_1(X1),X2)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_3464, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,k6_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5172, negated_conjecture, (k56_valued_1(k5_numbers,k1_numbers,esk1_0)=k54_valued_1(esk1_0)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk1_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk1_0)))).#trainneg
cnf(i_0_763, plain, (v1_abian(X1)|v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_power(X1,X2),k5_numbers))).#trainneg
cnf(i_0_762, negated_conjecture, (v1_abian(X1)|v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_767, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5287, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_0,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(esk5_3(np__1,k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4882, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,k1_funct_1(esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,esk2_0,X1),k5_numbers))).#trainneg
cnf(i_0_4955, negated_conjecture, (m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v2_comseq_2(k54_valued_1(X1))|~m2_subset_1(esk5_3(np__1,k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(esk3_0,esk5_3(np__1,k54_valued_1(X1),X2))|~m1_subset_1(esk5_3(np__1,k54_valued_1(X1),X2),k5_numbers))).#trainneg
cnf(i_0_5001, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~v1_xcmplx_0(k1_seq_1(esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3470, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5292, negated_conjecture, (k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0))=k54_valued_1(esk1_0)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0))))).#trainneg
cnf(i_0_768, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xreal_0(X1)|~v2_valued_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4996, plain, (r1_xxreal_0(X1,k2_power(np__1,np__1))|r1_xxreal_0(k2_power(np__1,np__1),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5313, plain, (r1_xxreal_0(k2_power(X1,np__1),k2_power(np__1,np__1))|r1_xxreal_0(k2_power(np__1,np__1),k2_power(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_769, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v7_ordinal1(X1)|~v2_valued_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_770, plain, (v1_xcmplx_0(k54_valued_1(X3))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_4999, negated_conjecture, (k1_newton(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5077, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2)))).#trainneg
cnf(i_0_5036, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))!=k1_seq_1(esk2_0,esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2)))).#trainneg
cnf(i_0_5282, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))!=k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2)))).#trainneg
cnf(i_0_4889, negated_conjecture, (k18_complex1(k1_seq_1(esk1_0,X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(esk4_2(k54_valued_1(esk1_0),esk1_0)))).#trainneg
cnf(i_0_5072, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X2)))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X2),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5210, plain, (r1_xxreal_0(k2_power(np__1,k2_power(X1,np__1)),X2)|r1_xxreal_0(X2,k2_power(np__1,k2_power(X1,np__1)))|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5337, plain, (r1_xxreal_0(X1,k2_power(X2,k2_power(X3,np__1)))|r1_xxreal_0(k2_power(X2,k2_power(X3,np__1)),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5279, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X2)))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X2),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5032, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X2)))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5057, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~v1_xreal_0(k1_seq_1(esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,k1_seq_1(esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5089, negated_conjecture, (m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v2_comseq_2(k54_valued_1(X1))|~m2_subset_1(esk5_3(np__1,k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),X2))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(X1),X2)))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(X1),X2)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(X1),X2)))).#trainneg
cnf(i_0_5283, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(np__1,k1_funct_1(esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_5214, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k18_complex1(k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(esk2_0),X1)))!=k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(esk2_0),X1))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5247, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X1)))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)))!=k1_seq_1(esk2_0,esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5237, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2))!=k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1983, plain, (v3_valued_0(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_781, plain, (k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5)=k6_numbers|v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4887, negated_conjecture, (k18_complex1(k1_seq_1(esk1_0,X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk4_2(k54_valued_1(esk1_0),esk1_0),k5_numbers))).#trainneg
cnf(i_0_4197, negated_conjecture, (k1_funct_1(k54_valued_1(esk2_0),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(k1_seq_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers))).#trainneg
cnf(i_0_3545, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_4461, negated_conjecture, (k1_seq_1(k54_valued_1(esk1_0),X2)=k1_seq_1(k54_valued_1(X1),X2)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4177, negated_conjecture, (r1_xxreal_0(np__1,k1_funct_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(k1_funct_1(k54_valued_1(esk2_0),X1))|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4578, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X1)))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(esk2_0),k54_valued_1(X1))))!=k1_seq_1(esk2_0,esk4_2(k54_valued_1(esk2_0),k54_valued_1(X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),k54_valued_1(X1)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4380, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X1)))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1))))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1988, plain, (v1_relat_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2860, plain, (v1_relat_1(k54_valued_1(X1))|~v1_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4805, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(esk1_0),X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(esk1_0)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_782, plain, (v1_xboole_0(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3031, plain, (v1_funct_2(X1,k6_numbers,X2)|~v1_partfun1(X1,k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_5213, plain, (r1_xxreal_0(k2_power(X1,k2_power(X2,np__1)),X3)|r1_xxreal_0(X3,k2_power(X1,k2_power(X2,np__1)))|~v7_ordinal1(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4456, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(X2,esk4_2(k54_valued_1(X1),X2)))!=k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4577, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X1)))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k18_complex1(k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(esk2_0),k54_valued_1(X1))))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(esk2_0),k54_valued_1(X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),k54_valued_1(X1)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5073, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X2)))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|k1_seq_1(esk2_0,esk4_2(k54_valued_1(X1),k54_valued_1(X2)))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(X1),k54_valued_1(X2)))|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3406, plain, (v1_funct_2(k54_valued_1(X1),X2,k6_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k6_numbers))|~v1_valued_0(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_partfun1(X1,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2861, plain, (v3_valued_0(k54_valued_1(X1))|~v1_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4400, negated_conjecture, (r1_xxreal_0(k2_power(np__1,np__1),esk3_0)|r1_xxreal_0(esk3_0,k2_power(np__1,np__1)))).#trainneg
cnf(i_0_788, plain, (v1_abian(X1)|r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,k1_power(X1,X2))|~v7_ordinal1(k1_power(X1,X2)))).#trainneg
cnf(i_0_4724, plain, (v1_funct_2(k6_numbers,X1,X2)|~v1_partfun1(k6_numbers,X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_5428, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),esk3_0))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),esk3_0)))).#trainneg
cnf(i_0_5506, plain, (r1_xxreal_0(k2_power(np__1,np__1),k2_power(np__1,np__1)))).#trainneg
cnf(i_0_787, negated_conjecture, (v1_abian(X1)|r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~v7_ordinal1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X1),k1_numbers))).#trainneg
cnf(i_0_789, plain, (v1_abian(X1)|r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,k1_power(X1,X2))|~m1_subset_1(k1_power(X1,X2),k1_numbers))).#trainneg
cnf(i_0_5430, negated_conjecture, (~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(np__1,k1_seq_1(esk1_0,esk5_3(np__1,esk2_0,X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5574, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk1_0,X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X3))|~v1_funct_1(k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_5627, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X3))|~v1_funct_1(k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_5644, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk2_0,X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X3))|~v1_funct_1(k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1999, plain, (v1_valued_0(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_5714, plain, (v3_valued_0(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4410, plain, (r1_xxreal_0(k2_power(np__1,np__1),k6_numbers)|r1_xxreal_0(k6_numbers,k2_power(np__1,np__1)))).#trainneg
cnf(i_0_5406, negated_conjecture, (k56_valued_1(k5_numbers,k1_numbers,esk1_0)=k54_valued_1(esk1_0)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk1_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk1_0)))).#trainneg
cnf(i_0_790, plain, (v1_xboole_0(X1)|v1_funct_2(k1_funct_1(X4,X5),X2,X3)|~v1_partfun1(k1_funct_1(X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_5603, plain, (r1_xxreal_0(X1,k2_power(np__1,k2_power(np__1,np__1)))|r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_791, plain, (v1_valued_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_792, plain, (v3_valued_0(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_5532, negated_conjecture, (k1_funct_1(k54_valued_1(esk2_0),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1)))).#trainneg
cnf(i_0_5429, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(X1)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X1))).#trainneg
cnf(i_0_5691, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X1))|~v1_valued_0(X1)|~m2_subset_1(esk4_2(k54_valued_1(X1),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5916, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))!=k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2)))).#trainneg
cnf(i_0_5662, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(np__1,k1_seq_1(esk1_0,esk5_3(np__1,k54_valued_1(esk2_0),X1)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(esk2_0),X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5694, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2))!=k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))|~v1_valued_0(k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5671, negated_conjecture, (m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),X2))|~v2_comseq_2(k54_valued_1(X1))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,esk5_3(np__1,k54_valued_1(X1),X2)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk5_3(np__1,k54_valued_1(X1),X2),k5_numbers))).#trainneg
cnf(i_0_5914, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X2)))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X2),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5678, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k56_valued_1(k5_numbers,k1_numbers,X1))|k1_seq_1(k54_valued_1(esk2_0),esk4_2(k54_valued_1(esk2_0),X1))!=k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(esk2_0),X1))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5850, plain, (v1_relat_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4979, plain, (r1_xxreal_0(np__1,k1_power(np__1,np__1))|r1_xxreal_0(k1_power(np__1,np__1),np__1))).#trainneg
cnf(i_0_793, plain, (v1_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_5510, plain, (r1_xxreal_0(k2_power(np__1,np__1),k2_power(X1,np__1))|r1_xxreal_0(k2_power(X1,np__1),k2_power(np__1,np__1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5607, plain, (r1_xxreal_0(X1,k2_power(np__1,k2_power(X2,np__1)))|r1_xxreal_0(k2_power(np__1,k2_power(X2,np__1)),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_794, plain, (v3_valued_0(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_795, plain, (v1_partfun1(k1_funct_1(X4,X5),X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(k1_funct_1(X4,X5),X2,X3)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_5653, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5871, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(esk1_0),X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(esk1_0))))).#trainneg
cnf(i_0_5583, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(esk1_0,X2),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5440, negated_conjecture, (m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v2_comseq_2(k54_valued_1(X1))|~r1_xxreal_0(esk3_0,esk5_3(np__1,k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk5_3(np__1,k54_valued_1(X1),X2),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),X2)))).#trainneg
cnf(i_0_5533, negated_conjecture, (k1_newton(k1_funct_1(esk1_0,X1),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~v1_xreal_0(k1_funct_1(esk1_0,X1))|~r1_xxreal_0(k6_numbers,k1_funct_1(esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3194, plain, (m1_subset_1(X1,k6_numbers)|~r2_hidden(X1,k54_valued_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3190, plain, (v3_membered(k54_valued_1(X1))|~v3_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_5479, negated_conjecture, (k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0))=k54_valued_1(esk1_0)|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0))))).#trainneg
cnf(i_0_796, plain, (v1_relat_1(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4801, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k54_valued_1(esk1_0))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),esk1_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4807, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k54_valued_1(esk1_0))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(esk1_0)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5625, plain, (r1_xxreal_0(k2_power(X1,k2_power(X2,np__1)),X3)|r1_xxreal_0(X3,k2_power(X1,k2_power(X2,np__1)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5621, plain, (r1_xxreal_0(k2_power(X1,k2_power(X2,np__1)),np__1)|r1_xxreal_0(np__1,k2_power(X1,k2_power(X2,np__1)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_797, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~r2_hidden(X2,k1_funct_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X3))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_801, plain, (r1_xxreal_0(X1,k3_funct_2(X2,X3,X4,X5))|r1_xxreal_0(k3_funct_2(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|~v1_xxreal_0(X1)|~v2_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_5869, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(esk1_0),X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1084, plain, (m1_subset_1(k1_seq_1(k54_valued_1(X1),X2),k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3687, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X2)))|k18_complex1(k1_seq_1(X1,esk4_2(k54_valued_1(X1),k54_valued_1(X2))))!=k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),k54_valued_1(X2)))|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1795, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(k54_valued_1(esk1_0))!=k6_numbers|~v2_comseq_2(k54_valued_1(esk1_0))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(esk1_0),esk5_3(X1,k54_valued_1(esk1_0),X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4472, plain, (v1_valued_0(k54_valued_1(X1))|~v1_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4673, plain, (k54_valued_1(k54_valued_1(X1))=k54_valued_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_802, plain, (r1_xxreal_0(X1,k3_funct_2(X2,X3,X4,X5))|r1_xxreal_0(k3_funct_2(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v2_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_6115, plain, (r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),X1)|r1_xxreal_0(X1,k2_power(np__1,k2_power(np__1,np__1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_803, plain, (r1_xxreal_0(X1,k3_funct_2(X2,X3,X4,X5))|r1_xxreal_0(k3_funct_2(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|~v7_ordinal1(X1)|~v2_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_805, plain, (v1_xxreal_0(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v2_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6052, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_3(np__1,esk2_0,X1))|~r1_xxreal_0(esk3_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_6196, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X3))|~v1_funct_1(k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6194, negated_conjecture, (m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),X2))|~v2_comseq_2(k54_valued_1(X1))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(esk3_0,esk5_3(np__1,k54_valued_1(X1),X2))|~m2_subset_1(esk5_3(np__1,k54_valued_1(X1),X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6314, negated_conjecture, (m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),esk3_0))|~v2_comseq_2(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6190, negated_conjecture, (m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),X2))|~v2_comseq_2(k54_valued_1(X1))|~r1_xxreal_0(np__1,k1_funct_1(esk1_0,esk5_3(np__1,k54_valued_1(X1),X2)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk5_3(np__1,k54_valued_1(X1),X2),k5_numbers))).#trainneg
cnf(i_0_6053, plain, (v1_valued_0(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_806, plain, (v1_partfun1(k54_valued_1(k1_funct_1(X4,X5)),X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(k1_funct_1(X4,X5),X2,X3)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(k1_funct_1(X4,X5))|~v1_funct_1(X4))).#trainneg
cnf(i_0_6086, plain, (r1_xxreal_0(k6_numbers,k1_power(np__1,np__1))|r1_xxreal_0(k1_power(np__1,np__1),k6_numbers))).#trainneg
cnf(i_0_6116, plain, (r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),X1)|r1_xxreal_0(X1,k2_power(np__1,k2_power(np__1,np__1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6110, plain, (r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),np__1)|r1_xxreal_0(np__1,k2_power(np__1,k2_power(np__1,np__1))))).#trainneg
cnf(i_0_807, plain, (v4_relat_1(k1_funct_1(X4,X5),X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_808, plain, (v1_valued_0(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v1_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_6205, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X2),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6049, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X3),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6134, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X1))|~v1_valued_0(X1)|~m2_subset_1(esk4_2(k54_valued_1(X1),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6445, plain, (v1_xreal_0(k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6015, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X3),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6129, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),np__1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(esk3_0,np__1))).#trainneg
cnf(i_0_6130, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),X1))|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(esk3_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_5980, plain, (r1_xxreal_0(k1_power(np__1,np__1),k1_power(np__1,np__1)))).#trainneg
cnf(i_0_809, plain, (v1_valued_0(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_5972, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(esk2_0),esk3_0))|~v2_comseq_2(k54_valued_1(esk2_0))|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk3_0))).#trainneg
cnf(i_0_6111, plain, (r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),k2_power(X1,X2))|r1_xxreal_0(k2_power(X1,X2),k2_power(np__1,k2_power(np__1,np__1)))|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_810, plain, (v1_funct_1(k54_valued_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)),X3,X4)))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)),X3,X4))|~v1_funct_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)),X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_811, plain, (v1_relat_1(k54_valued_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)),X3,X4)))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)),X3,X4))|~v1_funct_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)),X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,k4_numbers)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_6290, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6027, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k5_numbers))).#trainneg
cnf(i_0_6044, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k5_numbers))).#trainneg
cnf(i_0_6029, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v7_ordinal1(esk4_2(k54_valued_1(X1),k54_valued_1(X3))))).#trainneg
cnf(i_0_6315, negated_conjecture, (m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),X2))|~v2_comseq_2(k54_valued_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X2)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6144, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,esk1_0))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),esk1_0),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6046, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v7_ordinal1(esk4_2(k54_valued_1(X1),k54_valued_1(X3))))).#trainneg
cnf(i_0_6010, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k5_numbers))).#trainneg
cnf(i_0_6289, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_4671, plain, (k54_valued_1(k54_valued_1(k54_valued_1(X1)))=k54_valued_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_812, plain, (v3_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),X3,X4))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_6358, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_3910, plain, (r1_xxreal_0(k2_power(X1,X2),k2_power(X3,X4))|r1_xxreal_0(k2_power(X3,X4),k2_power(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers)|~v7_ordinal1(X3)|~m1_subset_1(X4,k1_numbers))).#trainneg
cnf(i_0_5604, plain, (r1_xxreal_0(X1,k2_power(np__1,k2_power(k2_power(X2,X3),np__1)))|r1_xxreal_0(k2_power(np__1,k2_power(k2_power(X2,X3),np__1)),X1)|~v1_xreal_0(X1)|~v6_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_813, plain, (v5_relat_1(k1_funct_1(X4,X5),X3)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_6308, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6448, plain, (v1_xcmplx_0(k1_seq_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6307, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6012, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v7_ordinal1(esk4_2(k54_valued_1(X1),k54_valued_1(X3))))).#trainneg
cnf(i_0_6203, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(k54_valued_1(X2)))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X2),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X2)))).#trainneg
cnf(i_0_6166, plain, (r2_funct_2(k5_numbers,k1_numbers,k6_numbers,k56_valued_1(k5_numbers,k1_numbers,X1))|k1_seq_1(k54_valued_1(X1),esk4_2(k6_numbers,X1))!=k1_seq_1(k6_numbers,esk4_2(k6_numbers,X1))|~m2_subset_1(esk4_2(k6_numbers,X1),k1_numbers,k5_numbers)|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_6340, plain, (v1_membered(k54_valued_1(X1))|~v3_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_6772, negated_conjecture, (k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~v2_comseq_2(k54_valued_1(esk2_0))|~r1_xxreal_0(esk3_0,np__1))).#trainneg
cnf(i_0_6554, plain, (r1_xxreal_0(X1,k2_power(np__1,k2_power(np__1,np__1)))|r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_846, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(np__1)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~r1_xxreal_0(X3,np__1)|~r1_xxreal_0(esk3_0,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6881, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(esk1_0,X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),esk1_0),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6586, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6911, negated_conjecture, (k54_valued_1(X1)=k54_valued_1(X2)|m2_subset_1(esk4_2(esk2_0,X2),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6998, negated_conjecture, (k54_valued_1(X1)=k54_valued_1(X2)|m2_subset_1(esk4_2(esk1_0,X2),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2))|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6914, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(X3,X2))|m2_subset_1(esk4_2(esk2_0,X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7001, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(X3,X2))|m2_subset_1(esk4_2(esk1_0,X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6912, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk2_0,X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk2_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6938, negated_conjecture, (k18_complex1(k1_seq_1(esk1_0,X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6999, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(esk1_0,X3),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6616, negated_conjecture, (m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v2_comseq_2(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk5_3(np__1,k54_valued_1(X1),esk3_0),k5_numbers))).#trainneg
cnf(i_0_6617, negated_conjecture, (m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v2_comseq_2(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_3020, plain, (m1_subset_1(k56_valued_1(k6_numbers,k1_numbers,X1),k1_zfmisc_1(k6_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_7045, plain, (r2_funct_2(k5_numbers,k1_numbers,k6_numbers,k56_valued_1(k5_numbers,k1_numbers,k6_numbers))|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_6810, plain, (r1_xxreal_0(k2_power(X1,np__1),k2_power(np__1,k2_power(np__1,np__1)))|r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),k2_power(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_885, negated_conjecture, (v1_xboole_0(k54_valued_1(esk2_0))|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k54_valued_1(esk2_0)))).#trainneg
cnf(i_0_6597, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X3)),k5_numbers))).#trainneg
cnf(i_0_6731, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X1))|~v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk4_2(k54_valued_1(X1),X1),k5_numbers))).#trainneg
cnf(i_0_6609, negated_conjecture, (m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),esk3_0))|~v2_comseq_2(k54_valued_1(X1))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(X1),esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6623, negated_conjecture, (m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v7_ordinal1(esk5_3(np__1,k54_valued_1(X1),X2))|~v2_comseq_2(k54_valued_1(X1))|~r1_xxreal_0(np__1,k1_seq_1(esk1_0,esk5_3(np__1,k54_valued_1(X1),X2)))|~r1_xxreal_0(np__1,esk5_3(np__1,k54_valued_1(X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7043, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k6_numbers,k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk1_0)))|k1_seq_1(k54_valued_1(esk1_0),esk4_2(k6_numbers,k54_valued_1(esk1_0)))!=k1_seq_1(k6_numbers,esk4_2(k6_numbers,k54_valued_1(esk1_0)))|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_7042, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k6_numbers,k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(esk2_0)))|k1_seq_1(k54_valued_1(esk2_0),esk4_2(k6_numbers,k54_valued_1(esk2_0)))!=k1_seq_1(k6_numbers,esk4_2(k6_numbers,k54_valued_1(esk2_0)))|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1989, plain, (v1_xboole_0(k54_valued_1(X1))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1995, plain, (v4_relat_1(k54_valued_1(k54_valued_1(X1)),X2)|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4603, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X1)=k1_funct_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_886, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,k54_valued_1(esk2_0)))).#trainneg
cnf(i_0_6959, plain, (r1_xxreal_0(k2_power(X1,X2),k2_power(X3,np__1))|r1_xxreal_0(k2_power(X3,np__1),k2_power(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_887, negated_conjecture, (k1_seq_1(X1,X2)=k18_complex1(k1_seq_1(k54_valued_1(esk2_0),X2))|m2_subset_1(esk4_2(X1,k54_valued_1(esk2_0)),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(esk2_0),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6599, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X3),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v7_ordinal1(esk4_2(k54_valued_1(X1),k54_valued_1(X3))))).#trainneg
cnf(i_0_6032, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(X3),X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X3),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X3),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(k54_valued_1(X1)))).#trainneg
cnf(i_0_6023, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(X2),X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),k54_valued_1(esk1_0)),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X2),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6127, negated_conjecture, (k1_funct_1(k54_valued_1(esk2_0),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1050, plain, (v1_funct_1(k54_valued_1(k2_zfmisc_1(X1,k4_numbers)))|~v1_funct_1(k2_zfmisc_1(X1,k4_numbers)))).#trainneg
cnf(i_0_5964, negated_conjecture, (r1_xxreal_0(esk3_0,k1_power(np__1,np__1))|r1_xxreal_0(k1_power(np__1,np__1),esk3_0))).#trainneg
cnf(i_0_3993, negated_conjecture, (k1_funct_1(k54_valued_1(esk2_0),X1)=k1_seq_1(k54_valued_1(esk2_0),X1)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(k1_funct_1(k54_valued_1(esk2_0),X1),k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6726, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k54_valued_1(esk2_0))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),k54_valued_1(esk2_0)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1403, plain, (v1_xcmplx_0(k56_valued_1(X1,k1_numbers,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_905, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X2))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k56_valued_1(k5_numbers,k1_numbers,esk2_0),X2),k1_numbers))).#trainneg
cnf(i_0_906, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,k1_power(X3,X2))|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k6_numbers,X2)|~v7_ordinal1(k1_power(X3,X2)))).#trainneg
cnf(i_0_5636, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X2),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(X2)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2740, plain, (k18_complex1(k1_seq_1(k54_valued_1(X1),X2))=k1_seq_1(k54_valued_1(X1),X2)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1156, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X2))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X2))|~r1_xxreal_0(np__1,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X2))|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),X2),k1_numbers))).#trainneg
cnf(i_0_3983, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X2))|k1_seq_1(k54_valued_1(X1),esk4_2(k54_valued_1(X1),X2))!=k1_seq_1(k54_valued_1(X2),esk4_2(k54_valued_1(X1),X2))|~m2_subset_1(esk4_2(k54_valued_1(X1),X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k54_valued_1(X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X2),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X2)))).#trainneg
cnf(i_0_2564, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7353, negated_conjecture, (v4_relat_1(k54_valued_1(esk2_0),X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers))))).#trainneg
cnf(i_0_7269, plain, (k56_valued_1(k5_numbers,k1_numbers,k6_numbers)=k6_numbers|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,k6_numbers),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,k6_numbers),k5_numbers,k1_numbers)|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,k6_numbers))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_907, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,k1_power(X3,X2))|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k6_numbers,X2)|~m1_subset_1(k1_power(X3,X2),k1_numbers))).#trainneg
cnf(i_0_7376, plain, (r1_xxreal_0(k2_power(X1,np__1),k2_power(X2,np__1))|r1_xxreal_0(k2_power(X2,np__1),k2_power(X1,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_909, plain, (r1_xxreal_0(X1,k2_newton(X2,X3))|r1_xxreal_0(k2_newton(X2,X3),X1)|~v1_xreal_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_910, plain, (r1_xxreal_0(X1,k1_seq_2(X2))|r1_xxreal_0(k1_seq_2(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7141, negated_conjecture, (~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7248, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,X1),k5_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v2_comseq_2(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7354, negated_conjecture, (v4_relat_1(k54_valued_1(esk1_0),X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers))))).#trainneg
cnf(i_0_7444, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k54_valued_1(esk2_0)))).#trainneg
cnf(i_0_911, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3073, plain, (v1_partfun1(k54_valued_1(k6_numbers),k6_numbers)|v1_xboole_0(X1)|~v1_membered(X1)|~v1_funct_2(k6_numbers,k6_numbers,X1)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_7291, plain, (r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),k2_power(X1,np__1))|r1_xxreal_0(k2_power(X1,np__1),k2_power(np__1,k2_power(np__1,np__1)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7287, plain, (r1_xxreal_0(k2_power(np__1,k2_power(np__1,np__1)),k2_power(np__1,np__1))|r1_xxreal_0(k2_power(np__1,np__1),k2_power(np__1,k2_power(np__1,np__1))))).#trainneg
cnf(i_0_912, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_913, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|r1_xxreal_0(k2_seq_2(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7537, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7219, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|k2_seq_2(esk1_0)!=k6_numbers|~v2_comseq_2(esk1_0)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(esk1_0),esk5_3(X1,esk1_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(esk5_3(X1,esk1_0,X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7136, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(esk1_0,X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk4_2(k54_valued_1(X1),esk1_0),k5_numbers))).#trainneg
cnf(i_0_7333, plain, (v1_xboole_0(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_7249, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,X1),k1_numbers)|k2_seq_2(k54_valued_1(X1))!=k6_numbers|~v2_comseq_2(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_915, plain, (r1_xxreal_0(X1,k18_complex1(X2))|r1_xxreal_0(k18_complex1(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7272, plain, (r2_funct_2(k5_numbers,k1_numbers,k6_numbers,k54_valued_1(k6_numbers))|~m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_916, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_919, plain, (r1_xxreal_0(X1,k3_funct_2(X2,X3,X4,X5))|r1_xxreal_0(k3_funct_2(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_7138, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(esk1_0,X2))|m2_subset_1(esk4_2(k54_valued_1(esk1_0),X1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk4_2(k54_valued_1(X1),esk1_0)))).#trainneg
cnf(i_0_7072, negated_conjecture, (r1_xxreal_0(esk3_0,k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|r1_xxreal_0(np__1,esk3_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk3_0)|~v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,esk1_0,X1))|~r1_xxreal_0(esk3_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7573, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X2))).#trainneg
cnf(i_0_7255, plain, (v3_valued_0(k56_valued_1(k6_numbers,k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_932, plain, (r1_xxreal_0(X1,k2_newton(X2,X3))|r1_xxreal_0(k2_newton(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3014, plain, (m1_subset_1(k54_valued_1(k6_numbers),k1_zfmisc_1(k6_numbers))|~v1_membered(X1)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_6212, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk2_0),k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),esk2_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7377, plain, (r1_xxreal_0(k2_power(X1,np__1),k2_power(X2,k2_power(X3,X4)))|r1_xxreal_0(k2_power(X2,k2_power(X3,X4)),k2_power(X1,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X4,k1_numbers))).#trainneg
cnf(i_0_933, plain, (r1_xxreal_0(X1,k1_seq_2(X2))|r1_xxreal_0(k1_seq_2(X2),X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_934, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7220, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,esk1_0))|k1_seq_1(X1,esk4_2(X1,esk1_0))!=k1_seq_1(k54_valued_1(esk1_0),esk4_2(X1,esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(esk4_2(X1,esk1_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7512, plain, (k1_seq_1(k54_valued_1(k54_valued_1(X1)),X2)=k1_seq_1(k54_valued_1(X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(k54_valued_1(X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(k54_valued_1(X1)),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(k54_valued_1(X1)))|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7650, negated_conjecture, (~v1_xxreal_0(X1)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7660, negated_conjecture, (k18_complex1(k1_seq_1(esk2_0,X1))=k1_seq_1(k54_valued_1(esk2_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7983, plain, (m1_subset_1(k54_valued_1(k6_numbers),k1_zfmisc_1(k6_numbers))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_935, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7611, plain, (r1_xxreal_0(k2_power(X1,np__1),k2_power(X2,np__1))|r1_xxreal_0(k2_power(X2,np__1),k2_power(X1,np__1))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_936, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|r1_xxreal_0(k2_seq_2(X2),X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_938, plain, (r1_xxreal_0(X1,k18_complex1(X2))|r1_xxreal_0(k18_complex1(X2),X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7747, plain, (k56_valued_1(k5_numbers,k1_numbers,X1)=k54_valued_1(X1)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,X1),k5_numbers,k1_numbers)|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,X1))|~v1_funct_1(k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7902, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7750, plain, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k54_valued_1(X1))|~m2_subset_1(esk4_2(k54_valued_1(X1),X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7903, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X3,X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_7523, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(esk1_0),k56_valued_1(k5_numbers,k1_numbers,k54_valued_1(X1)))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)))|~m2_subset_1(esk4_2(k54_valued_1(esk1_0),k54_valued_1(X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7258, plain, (m1_subset_1(X1,k6_numbers)|~r2_hidden(X1,k56_valued_1(k6_numbers,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7256, plain, (v1_valued_0(k56_valued_1(k6_numbers,k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_funct_1(X1))).#trainneg
cnf(i_0_7692, plain, (v1_partfun1(k54_valued_1(k6_numbers),k6_numbers)|~v1_funct_2(k6_numbers,k6_numbers,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_7985, negated_conjecture, (k56_valued_1(k5_numbers,k1_numbers,esk2_0)=k54_valued_1(esk2_0)|~m1_subset_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k56_valued_1(k5_numbers,k1_numbers,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k56_valued_1(k5_numbers,k1_numbers,esk2_0))|~m2_subset_1(esk4_2(k54_valued_1(esk2_0),esk2_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_939, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v7_ordinal1(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7693, plain, (v1_partfun1(k54_valued_1(k6_numbers),k6_numbers)|~v1_funct_2(k6_numbers,k6_numbers,k5_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_7686, negated_conjecture, (v1_partfun1(k54_valued_1(k6_numbers),k6_numbers)|~v1_funct_2(k6_numbers,k6_numbers,k54_valued_1(esk2_0))|~v1_funct_1(k6_numbers)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_8009, plain, (r1_xxreal_0(k2_power(X1,k2_power(X2,np__1)),k2_power(X3,np__1))|r1_xxreal_0(k2_power(X3,np__1),k2_power(X1,k2_power(X2,np__1)))|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_942, plain, (r1_xxreal_0(X1,k3_funct_2(X2,X3,X4,X5))|r1_xxreal_0(k3_funct_2(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|~v7_ordinal1(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_7224, negated_conjecture, (r2_funct_2(k5_numbers,k1_numbers,k54_valued_1(X1),k56_valued_1(k5_numbers,k1_numbers,esk1_0))|m2_subset_1(esk4_2(esk1_0,X1),k1_numbers,k5_numbers)|k1_seq_1(k54_valued_1(esk1_0),esk4_2(k54_valued_1(X1),esk1_0))!=k1_seq_1(esk1_0,esk4_2(k54_valued_1(X1),esk1_0))|~m2_subset_1(esk4_2(k54_valued_1(X1),esk1_0),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7529, negated_conjecture, (k18_complex1(k1_seq_1(k54_valued_1(esk1_0),X1))=k1_seq_1(k54_valued_1(esk1_0),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2692, plain, (m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(k1_power(X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k6_numbers,X2))).#trainneg
cnf(i_0_6018, negated_conjecture, (k1_seq_1(k54_valued_1(X1),X2)=k18_complex1(k1_seq_1(k54_valued_1(esk1_0),X2))|m2_subset_1(esk4_2(k54_valued_1(esk2_0),X1),k1_numbers,k5_numbers)|m2_subset_1(esk4_2(k54_valued_1(esk2_0),k54_valued_1(esk1_0)),k1_numbers,k5_numbers)|~m2_subset_1(esk4_2(k54_valued_1(X1),k54_valued_1(esk1_0)),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2003, plain, (v1_funct_1(k54_valued_1(k54_valued_1(X1)))|~v1_funct_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2828, plain, (v1_funct_1(k54_valued_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_7687, negated_conjecture, (v1_partfun1(k54_valued_1(k6_numbers),k6_numbers)|~v1_funct_2(k6_numbers,k6_numbers,k54_valued_1(esk1_0))|~v1_funct_1(k6_numbers)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_947, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7652, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,k54_valued_1(esk2_0)),k5_numbers)|k2_seq_2(k54_valued_1(esk2_0))!=k6_numbers|~v2_comseq_2(k54_valued_1(esk2_0)))).#trainneg
cnf(i_0_7653, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,k54_valued_1(esk1_0)),k5_numbers)|k2_seq_2(k54_valued_1(esk1_0))!=k6_numbers|~v2_comseq_2(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_7608, plain, (r1_xxreal_0(k2_power(X1,np__1),k2_power(k2_power(X2,X3),np__1))|r1_xxreal_0(k2_power(k2_power(X2,X3),np__1),k2_power(X1,np__1))|~v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_5211, plain, (r1_xxreal_0(k2_power(k2_power(X1,X2),k2_power(X3,np__1)),X4)|r1_xxreal_0(X4,k2_power(k2_power(X1,X2),k2_power(X3,np__1)))|~v7_ordinal1(X3)|~v1_xreal_0(X4)|~v6_membered(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_949, plain, (r1_xxreal_0(X1,k1_newton(X2,X3))|r1_xxreal_0(k1_newton(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2537, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X4),X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~r1_xxreal_0(X1,X4))).#trainneg
cnf(i_0_8145, negated_conjecture, (~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8144, negated_conjecture, (~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8157, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k1_seq_1(k54_valued_1(esk2_0),esk5_3(X1,esk2_0,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(esk5_3(X1,esk2_0,X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8177, plain, (v3_valued_0(k54_valued_1(k6_numbers))|~v1_funct_1(k6_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 105
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 158
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 149
# Processed clauses                    : 2031
# ...of these trivial                  : 80
# ...subsumed                          : 823
# ...remaining for further processing  : 1127
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 128
# Backward-rewritten                   : 29
# Generated clauses                    : 8574
# ...of the previous two non-trivial   : 7925
# Contextual simplify-reflections      : 314
# Paramodulations                      : 8556
# Factorizations                       : 0
# Equation resolutions                 : 16
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 967
#    Positive orientable unit clauses  : 98
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 854
# Current number of unprocessed clauses: 5708
# ...number of literals in the above   : 43798
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1287
# Proof object given clauses           : 138
# Proof search given clauses           : 1128
# Clause-clause subsumption calls (NU) : 703060
# Rec. Clause-clause subsumption calls : 30761
# Non-unit clause-clause subsumptions  : 1150
# Unit Clause-clause subsumption calls : 8998
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 81
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 297945

# -------------------------------------------------
# User time                : 22.182 s
# System time              : 0.516 s
# Total time               : 22.699 s
# Maximum resident set size: 410908 pages
22.18user 0.60system 0:23.01elapsed 99%CPU (0avgtext+0avgdata 410908maxresident)k
0inputs+688outputs (0major+212325minor)pagefaults 0swaps
