# ENIGMA: LightGBM model 'Enigma/mizar40-all-T10/loop02-VHSLCAXPh/lgb-d30-l1800-e0.15/model.lgb' loaded. (hash_base: 32768; conj_feats: 83; version: 991; iters: 150)
# Preprocessing time       : 0.488 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t42_fdiff_4, conjecture, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))&![X2]:(m1_subset_1(X2,k1_numbers)=>~((r2_hidden(X2,X1)&r1_xxreal_0(k1_seq_1(k19_sin_cos,X2),k6_numbers)))))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1),X2)=k1_real_1(k1_sin_cos4(X2))))))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t42_fdiff_4)).
fof(t9_fdiff_1, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r2_fdiff_1(X2,X1)<=>(r1_tarski(X1,k1_relset_1(k1_numbers,X2))&![X3]:(m1_subset_1(X3,k1_numbers)=>(r2_hidden(X3,X1)=>r1_fdiff_1(X2,X3))))))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t9_fdiff_1)).
fof(t20_taylor_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r1_fdiff_1(X2,X1)=>(r1_xxreal_0(k1_seq_1(X2,X1),k6_numbers)|(r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k3_taylor_1),X1)&k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k3_taylor_1),X1)=k10_real_1(k1_fdiff_1(X2,X1),k1_seq_1(X2,X1))))))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t20_taylor_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t4_subset)).
fof(dt_k19_sin_cos, axiom, ((v1_funct_1(k19_sin_cos)&v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k19_sin_cos)).
fof(t63_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>(r1_fdiff_1(k19_sin_cos,X1)&k1_fdiff_1(k19_sin_cos,X1)=k1_real_1(k1_seq_1(k16_sin_cos,X1)))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t63_sin_cos)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', cc12_membered)).
fof(dt_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k1_real_1(X1),k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k1_real_1)).
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', redefinition_k1_real_1)).
fof(involutiveness_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(k1_real_1(X1))=X1), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', involutiveness_k1_real_1)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', cc14_membered)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', fc3_membered)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', cc2_relset_1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k1_partfun1)).
fof(redefinition_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', redefinition_k10_real_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t6_boole)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', cc5_membered)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t3_subset)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', redefinition_k1_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', cc1_relset_1)).
fof(dt_k1_fdiff_1, axiom, ![X1, X2]:(((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))&v1_xreal_0(X2))=>m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k1_fdiff_1)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', redefinition_k6_numbers)).
fof(d7_fdiff_1, axiom, ![X1]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X2]:(r2_fdiff_1(X1,X2)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k2_fdiff_1(X1,X2)<=>(k1_relset_1(k1_numbers,X3)=X2&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(X3,X4)=k1_fdiff_1(X1,X4)))))))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', d7_fdiff_1)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', spc0_boole)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', redefinition_k1_partfun1)).
fof(dt_k2_fdiff_1, axiom, ![X1, X2]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v1_funct_1(k2_fdiff_1(X1,X2))&m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k2_fdiff_1)).
fof(t187_xcmplx_1, axiom, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>k4_xcmplx_0(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t187_xcmplx_1)).
fof(dt_k3_taylor_1, axiom, (v1_funct_1(k3_taylor_1)&m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k3_taylor_1)).
fof(d1_sin_cos4, axiom, ![X1]:(v1_xreal_0(X1)=>k1_sin_cos4(X1)=k7_xcmplx_0(k17_sin_cos(X1),k20_sin_cos(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', d1_sin_cos4)).
fof(d19_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', d19_sin_cos)).
fof(d17_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', d17_sin_cos)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', t1_subset)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', cc21_membered)).
fof(dt_k1_sin_cos4, axiom, ![X1]:(v1_xreal_0(X1)=>m1_subset_1(k1_sin_cos4(X1),k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k1_sin_cos4)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', cc18_valued_0)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k1_seq_1)).
fof(dt_k16_sin_cos, axiom, ((v1_funct_1(k16_sin_cos)&v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/yan/atp/benchmarks/mizar40/all/t42_fdiff_4', dt_k16_sin_cos)).
fof(c_0_37, negated_conjecture, ~(![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))&![X2]:(m1_subset_1(X2,k1_numbers)=>~((r2_hidden(X2,X1)&r1_xxreal_0(k1_seq_1(k19_sin_cos,X2),k6_numbers)))))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1),X2)=k1_real_1(k1_sin_cos4(X2)))))))), inference(assume_negation,[status(cth)],[t42_fdiff_4])).
fof(c_0_38, plain, ![X276, X277, X278]:(((r1_tarski(X276,k1_relset_1(k1_numbers,X277))|~r2_fdiff_1(X277,X276)|(~v1_funct_1(X277)|~m1_subset_1(X277,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X276)|~m1_subset_1(X276,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X278,k1_numbers)|(~r2_hidden(X278,X276)|r1_fdiff_1(X277,X278))|~r2_fdiff_1(X277,X276)|(~v1_funct_1(X277)|~m1_subset_1(X277,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X276)|~m1_subset_1(X276,k1_zfmisc_1(k1_numbers)))))&((m1_subset_1(esk19_2(X276,X277),k1_numbers)|~r1_tarski(X276,k1_relset_1(k1_numbers,X277))|r2_fdiff_1(X277,X276)|(~v1_funct_1(X277)|~m1_subset_1(X277,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X276)|~m1_subset_1(X276,k1_zfmisc_1(k1_numbers))))&((r2_hidden(esk19_2(X276,X277),X276)|~r1_tarski(X276,k1_relset_1(k1_numbers,X277))|r2_fdiff_1(X277,X276)|(~v1_funct_1(X277)|~m1_subset_1(X277,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X276)|~m1_subset_1(X276,k1_zfmisc_1(k1_numbers))))&(~r1_fdiff_1(X277,esk19_2(X276,X277))|~r1_tarski(X276,k1_relset_1(k1_numbers,X277))|r2_fdiff_1(X277,X276)|(~v1_funct_1(X277)|~m1_subset_1(X277,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X276)|~m1_subset_1(X276,k1_zfmisc_1(k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_fdiff_1])])])])])])])).
fof(c_0_39, plain, ![X249, X250]:((r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X250,k3_taylor_1),X249)|r1_xxreal_0(k1_seq_1(X250,X249),k6_numbers)|~r1_fdiff_1(X250,X249)|(~v1_funct_1(X250)|~m1_subset_1(X250,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~v1_xreal_0(X249))&(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X250,k3_taylor_1),X249)=k10_real_1(k1_fdiff_1(X250,X249),k1_seq_1(X250,X249))|r1_xxreal_0(k1_seq_1(X250,X249),k6_numbers)|~r1_fdiff_1(X250,X249)|(~v1_funct_1(X250)|~m1_subset_1(X250,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~v1_xreal_0(X249))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_taylor_1])])])])])])).
fof(c_0_40, negated_conjecture, ![X8]:((v3_rcomp_1(esk1_0)&m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))&((r1_tarski(esk1_0,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))&(~m1_subset_1(X8,k1_numbers)|(~r2_hidden(X8,esk1_0)|~r1_xxreal_0(k1_seq_1(k19_sin_cos,X8),k6_numbers))))&((m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0))&((r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0))&(k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_37])])])])])])])).
cnf(c_0_41, plain, (r2_fdiff_1(X1,X2)|~r1_fdiff_1(X1,esk19_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_43, plain, ![X261, X262, X263]:(~r2_hidden(X261,X262)|~m1_subset_1(X262,k1_zfmisc_1(X263))|m1_subset_1(X261,X263)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_44, negated_conjecture, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk19_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_47, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
fof(c_0_48, plain, ![X268]:((r1_fdiff_1(k19_sin_cos,X268)|~v1_xreal_0(X268))&(k1_fdiff_1(k19_sin_cos,X268)=k1_real_1(k1_seq_1(k16_sin_cos,X268))|~v1_xreal_0(X268))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_sin_cos])])])).
cnf(c_0_49, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_51, plain, ![X20, X21]:(~v1_membered(X20)|(~m1_subset_1(X21,X20)|v1_xcmplx_0(X21))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])])])).
fof(c_0_52, plain, ![X137]:(~m1_subset_1(X137,k1_numbers)|m1_subset_1(k1_real_1(X137),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_real_1])])).
fof(c_0_53, plain, ![X228]:(~m1_subset_1(X228,k1_numbers)|k1_real_1(X228)=k4_xcmplx_0(X228)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
fof(c_0_54, plain, ![X201]:(~m1_subset_1(X201,k1_numbers)|k1_real_1(k1_real_1(X201))=X201), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k1_real_1])])).
cnf(c_0_55, negated_conjecture, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k6_numbers)), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_56, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))),k6_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~r1_fdiff_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~v1_funct_1(X1)|~v1_xreal_0(esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_57, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_50])).
fof(c_0_61, plain, ![X28, X29]:(~v3_membered(X28)|(~m1_subset_1(X29,X28)|v1_xreal_0(X29))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])])])).
cnf(c_0_62, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_63, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~r2_hidden(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),esk1_0)|~m1_subset_1(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k1_numbers)|~r1_fdiff_1(k19_sin_cos,esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_funct_1(k19_sin_cos)|~v1_xreal_0(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])])])).
cnf(c_0_67, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60])])).
cnf(c_0_69, plain, (v1_funct_1(k19_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
cnf(c_0_70, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_72, plain, ![X83, X84, X85]:((v4_relat_1(X85,X83)|~m1_subset_1(X85,k1_zfmisc_1(k2_zfmisc_1(X83,X84))))&(v5_relat_1(X85,X84)|~m1_subset_1(X85,k1_zfmisc_1(k2_zfmisc_1(X83,X84))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_73, plain, ![X131, X132, X133, X134, X135, X136]:((v1_funct_1(k1_partfun1(X131,X132,X133,X134,X135,X136))|(~v1_funct_1(X135)|~m1_subset_1(X135,k1_zfmisc_1(k2_zfmisc_1(X131,X132)))|~v1_funct_1(X136)|~m1_subset_1(X136,k1_zfmisc_1(k2_zfmisc_1(X133,X134)))))&(m1_subset_1(k1_partfun1(X131,X132,X133,X134,X135,X136),k1_zfmisc_1(k2_zfmisc_1(X131,X134)))|(~v1_funct_1(X135)|~m1_subset_1(X135,k1_zfmisc_1(k2_zfmisc_1(X131,X132)))|~v1_funct_1(X136)|~m1_subset_1(X136,k1_zfmisc_1(k2_zfmisc_1(X133,X134)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
fof(c_0_74, plain, ![X219, X220]:(~m1_subset_1(X219,k1_numbers)|~v1_xreal_0(X220)|k10_real_1(X219,X220)=k7_xcmplx_0(X219,X220)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k10_real_1])])).
fof(c_0_75, plain, ![X271]:(~v1_xboole_0(X271)|X271=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_76, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_77, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_63])).
fof(c_0_78, plain, ![X102]:(~v3_membered(X102)|v1_membered(X102)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_79, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_80, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_81, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_funct_1(k19_sin_cos)|~v1_xreal_0(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),esk1_0)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_82, plain, (v1_funct_1(k19_sin_cos)), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_83, plain, (r2_hidden(esk19_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_84, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_85, plain, ![X256, X257, X256, X257]:((~m1_subset_1(X256,k1_zfmisc_1(X257))|r1_tarski(X256,X257))&(~r1_tarski(X256,X257)|m1_subset_1(X256,k1_zfmisc_1(X257)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_86, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_87, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_71])).
fof(c_0_88, plain, ![X229, X230]:(~v1_relat_1(X230)|~v4_relat_1(X230,X229)|k1_relset_1(X229,X230)=k9_xtuple_0(X230)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_89, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_90, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_73])).
fof(c_0_91, plain, ![X58, X59, X60]:(~m1_subset_1(X60,k1_zfmisc_1(k2_zfmisc_1(X58,X59)))|v1_relat_1(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_92, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_93, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)=k10_real_1(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_94, plain, ![X129, X130]:(~v1_funct_1(X129)|~m1_subset_1(X129,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X130)|m1_subset_1(k1_fdiff_1(X129,X130),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_fdiff_1])])).
cnf(c_0_95, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_96, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_97, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77])])).
cnf(c_0_98, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_99, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80])])).
cnf(c_0_100, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_xreal_0(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),esk1_0)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82])])])).
cnf(c_0_101, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk19_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_102, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_103, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_104, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_77]), c_0_87])])])).
cnf(c_0_105, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_79])])).
cnf(c_0_106, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_107, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_108, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_109, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_110, plain, ![X121, X122, X123, X124]:(((k1_relset_1(k1_numbers,X123)=X122|X123!=k2_fdiff_1(X121,X122)|(~v1_funct_1(X123)|~m1_subset_1(X123,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X121,X122)|(~v1_funct_1(X121)|~m1_subset_1(X121,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(~m1_subset_1(X124,k1_numbers)|(~r2_hidden(X124,X122)|k1_seq_1(X123,X124)=k1_fdiff_1(X121,X124))|X123!=k2_fdiff_1(X121,X122)|(~v1_funct_1(X123)|~m1_subset_1(X123,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X121,X122)|(~v1_funct_1(X121)|~m1_subset_1(X121,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&((m1_subset_1(esk3_3(X121,X122,X123),k1_numbers)|k1_relset_1(k1_numbers,X123)!=X122|X123=k2_fdiff_1(X121,X122)|(~v1_funct_1(X123)|~m1_subset_1(X123,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X121,X122)|(~v1_funct_1(X121)|~m1_subset_1(X121,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk3_3(X121,X122,X123),X122)|k1_relset_1(k1_numbers,X123)!=X122|X123=k2_fdiff_1(X121,X122)|(~v1_funct_1(X123)|~m1_subset_1(X123,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X121,X122)|(~v1_funct_1(X121)|~m1_subset_1(X121,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(k1_seq_1(X123,esk3_3(X121,X122,X123))!=k1_fdiff_1(X121,esk3_3(X121,X122,X123))|k1_relset_1(k1_numbers,X123)!=X122|X123=k2_fdiff_1(X121,X122)|(~v1_funct_1(X123)|~m1_subset_1(X123,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X121,X122)|(~v1_funct_1(X121)|~m1_subset_1(X121,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_fdiff_1])])])])])])])).
cnf(c_0_111, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_112, plain, (k10_real_1(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_113, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_114, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_115, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_116, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_117, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_79])])).
cnf(c_0_118, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_119, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_79])])).
cnf(c_0_120, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_79])])).
cnf(c_0_121, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_xreal_0(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_60]), c_0_102])])])).
cnf(c_0_122, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_123, plain, (v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105])])).
cnf(c_0_124, plain, (m1_subset_1(esk19_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_125, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_126, negated_conjecture, (r1_tarski(esk1_0,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_127, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_128, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108])])).
cnf(c_0_129, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_109])).
fof(c_0_130, plain, ![X222, X223, X224, X225, X226, X227]:(~v1_funct_1(X226)|~m1_subset_1(X226,k1_zfmisc_1(k2_zfmisc_1(X222,X223)))|~v1_funct_1(X227)|~m1_subset_1(X227,k1_zfmisc_1(k2_zfmisc_1(X224,X225)))|k1_partfun1(X222,X223,X224,X225,X226,X227)=k3_relat_1(X226,X227)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_131, plain, (k1_seq_1(X3,X1)=k1_fdiff_1(X4,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_110])).
fof(c_0_132, plain, ![X144, X145, X145]:((v1_funct_1(k2_fdiff_1(X144,X145))|(~v1_funct_1(X144)|~m1_subset_1(X144,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(m1_subset_1(k2_fdiff_1(X144,X145),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_funct_1(X144)|~m1_subset_1(X144,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_fdiff_1])])])])])).
cnf(c_0_133, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_134, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_135, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_136, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_137, plain, (k1_fdiff_1(k19_sin_cos,X1)=k1_real_1(k1_seq_1(k16_sin_cos,X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_138, plain, ![X243, X244]:(~v1_xcmplx_0(X243)|(~v1_xcmplx_0(X244)|k4_xcmplx_0(k7_xcmplx_0(X243,X244))=k7_xcmplx_0(k4_xcmplx_0(X243),X244))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t187_xcmplx_1])])])])])).
cnf(c_0_139, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_87])])])).
cnf(c_0_140, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_141, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_xreal_0(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122])])).
cnf(c_0_142, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_123, c_0_120])).
cnf(c_0_143, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk19_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_144, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_145, negated_conjecture, (r1_tarski(esk1_0,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_146, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_147, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_108])])).
cnf(c_0_148, plain, (v1_funct_1(k3_taylor_1)), inference(split_conjunct,[status(thm)],[dt_k3_taylor_1])).
cnf(c_0_149, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k3_taylor_1])).
cnf(c_0_150, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_151, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_152, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_hidden(X1,X2)|~r2_fdiff_1(X4,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_153, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_154, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_155, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_156, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136])])).
cnf(c_0_157, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_158, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_159, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140])])).
cnf(c_0_160, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))|~m1_subset_1(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_161, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk19_2(X2,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_122])])).
cnf(c_0_162, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_163, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_164, plain, (v1_funct_1(k3_taylor_1)), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_165, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_166, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_167, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_168, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_fdiff_1(X1,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_152])])).
cnf(c_0_169, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_170, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_171, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_172, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_155, c_0_156])])).
cnf(c_0_173, plain, (k1_fdiff_1(k19_sin_cos,X1)=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_157])])).
cnf(c_0_174, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_175, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_159, c_0_120])).
fof(c_0_176, plain, ![X120]:(~v1_xreal_0(X120)|k1_sin_cos4(X120)=k7_xcmplx_0(k17_sin_cos(X120),k20_sin_cos(X120))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_sin_cos4])])).
fof(c_0_177, plain, ![X119]:(~v1_xreal_0(X119)|k20_sin_cos(X119)=k1_seq_1(k19_sin_cos,X119)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_sin_cos])])).
cnf(c_0_178, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161]), c_0_60]), c_0_102])])])).
cnf(c_0_179, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163]), c_0_164]), c_0_82]), c_0_165]), c_0_57])])])).
cnf(c_0_180, negated_conjecture, (k1_seq_1(k2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_165]), c_0_57])])])).
cnf(c_0_181, plain, (k1_seq_1(k2_fdiff_1(X1,X2),X3)=k1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_168, c_0_169]), c_0_170])).
cnf(c_0_182, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_183, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1)),k1_seq_1(k19_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_82]), c_0_57])])])).
cnf(c_0_184, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175])])).
cnf(c_0_185, plain, (k1_sin_cos4(X1)=k7_xcmplx_0(k17_sin_cos(X1),k20_sin_cos(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_186, plain, (k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
fof(c_0_187, plain, ![X118]:(~v1_xreal_0(X118)|k17_sin_cos(X118)=k1_seq_1(k16_sin_cos,X118)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d17_sin_cos])])).
cnf(c_0_188, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_178, c_0_179])])])).
cnf(c_0_189, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_hidden(esk2_0,esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181])])).
cnf(c_0_190, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_167]), c_0_165]), c_0_57])])])).
cnf(c_0_191, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1)),k1_seq_1(k19_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)), inference(csr,[status(thm)],[c_0_183, c_0_67])).
cnf(c_0_192, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_175])])).
cnf(c_0_193, plain, (k7_xcmplx_0(k17_sin_cos(X1),k20_sin_cos(X1))=k1_sin_cos4(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_194, plain, (k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_195, plain, (k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_187])).
fof(c_0_196, plain, ![X247, X248]:(~r2_hidden(X247,X248)|m1_subset_1(X247,X248)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_197, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_188])])).
cnf(c_0_198, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_73])).
fof(c_0_199, plain, ![X67, X68]:(~v3_membered(X67)|(~m1_subset_1(X68,k1_zfmisc_1(X67))|v3_membered(X68))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])])])).
cnf(c_0_200, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_189, c_0_68]), c_0_190])).
cnf(c_0_201, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192])])).
cnf(c_0_202, plain, (k7_xcmplx_0(k17_sin_cos(X1),k1_seq_1(k19_sin_cos,X1))=k1_sin_cos4(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_203, plain, (k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_204, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_196])).
cnf(c_0_205, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_108]), c_0_164]), c_0_82]), c_0_165]), c_0_57])])])).
cnf(c_0_206, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_207, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_199])).
cnf(c_0_208, plain, (m1_subset_1(k3_relat_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_167])])).
cnf(c_0_209, negated_conjecture, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_79])])).
cnf(c_0_210, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)), inference(csr,[status(thm)],[c_0_201, c_0_142])).
cnf(c_0_211, plain, (k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_seq_1(k19_sin_cos,X1))=k1_sin_cos4(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203])])).
cnf(c_0_212, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_213, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_164]), c_0_82]), c_0_165]), c_0_57])])])).
cnf(c_0_214, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_207])).
cnf(c_0_215, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_167]), c_0_164]), c_0_82]), c_0_165]), c_0_57])])])).
cnf(c_0_216, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_167]), c_0_164]), c_0_82]), c_0_165]), c_0_57])])])).
cnf(c_0_217, plain, (m1_subset_1(k3_relat_1(X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_165]), c_0_164])])])).
cnf(c_0_218, plain, (v1_funct_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_167])])).
cnf(c_0_219, negated_conjecture, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(k3_taylor_1)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_209, c_0_82])])])).
cnf(c_0_220, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)=k4_xcmplx_0(k1_sin_cos4(X1))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211])])).
cnf(c_0_221, negated_conjecture, (m1_subset_1(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213])])).
cnf(c_0_222, negated_conjecture, (v3_membered(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_60]), c_0_87])])])).
cnf(c_0_223, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_216])])])).
cnf(c_0_224, plain, (m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_57]), c_0_82])])])).
cnf(c_0_225, plain, (v1_funct_1(k3_relat_1(X1,k3_taylor_1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_165]), c_0_164])])])).
cnf(c_0_226, negated_conjecture, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_219, c_0_164])])).
cnf(c_0_227, plain, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),X1)=k4_xcmplx_0(k1_sin_cos4(X1))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_167]), c_0_164]), c_0_82]), c_0_165]), c_0_57])])])).
cnf(c_0_228, negated_conjecture, (v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_221]), c_0_222])])])).
cnf(c_0_229, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_223, c_0_224])])])).
cnf(c_0_230, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_57]), c_0_82])])])).
cnf(c_0_231, negated_conjecture, (r1_xxreal_0(k1_seq_1(k19_sin_cos,esk2_0),np__0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227]), c_0_228])])])).
cnf(c_0_232, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_229, c_0_230])])])).
cnf(c_0_233, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_156])])).
cnf(c_0_234, negated_conjecture, (r1_xxreal_0(k1_seq_1(k19_sin_cos,esk2_0),np__0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_231, c_0_224])])])).
cnf(c_0_235, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_161]), c_0_230]), c_0_224]), c_0_216]), c_0_60]), c_0_102])])])).
fof(c_0_236, plain, ![X142]:(~v1_xreal_0(X142)|m1_subset_1(k1_sin_cos4(X142),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_sin_cos4])])).
fof(c_0_237, plain, ![X47, X48, X49]:(~v3_membered(X48)|(~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|v3_valued_0(X49))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])])])).
cnf(c_0_238, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[c_0_233, c_0_68])).
cnf(c_0_239, negated_conjecture, (r1_xxreal_0(k1_seq_1(k19_sin_cos,esk2_0),np__0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_234, c_0_230]), c_0_235])])).
cnf(c_0_240, plain, (m1_subset_1(k1_sin_cos4(X1),k1_numbers)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_236])).
fof(c_0_241, plain, ![X140, X141]:(~v1_relat_1(X140)|~v1_funct_1(X140)|~v3_valued_0(X140)|m1_subset_1(k1_seq_1(X140,X141),k1_numbers)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])])])).
cnf(c_0_242, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_243, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
cnf(c_0_244, negated_conjecture, (~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_213])])])).
cnf(c_0_245, plain, (m1_subset_1(k1_sin_cos4(X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_246, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_241])).
cnf(c_0_247, plain, (v1_funct_1(k16_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
cnf(c_0_248, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_249, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_250, negated_conjecture, (~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245]), c_0_228])])])).
cnf(c_0_251, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_252, plain, (v1_funct_1(k16_sin_cos)), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_253, plain, (v3_valued_0(k16_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_249]), c_0_87])])])).
cnf(c_0_254, plain, (v1_relat_1(k16_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_249])])).
cnf(c_0_255, negated_conjecture, (~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251]), c_0_252]), c_0_253]), c_0_254])])])).
cnf(c_0_256, plain, (v3_valued_0(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_57]), c_0_87])])])).
cnf(c_0_257, plain, (v1_relat_1(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_57])])).
cnf(c_0_258, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_251]), c_0_82]), c_0_256]), c_0_257])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 259
# Proof object clause steps            : 190
# Proof object formula steps           : 69
# Proof object conjectures             : 54
# Proof object clause conjectures      : 51
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 52
# Proof object initial formulas used   : 37
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 123
# Training examples: 123 positive, 1945 negative
# Training: Positive examples begin
cnf(c_0_115, plain, (k1_xboole_0=k6_numbers)).#trainpos
cnf(c_0_114, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainpos
cnf(c_0_60, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))).#trainpos
cnf(c_0_67, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_203, plain, (k17_sin_cos(X1)=k1_seq_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_194, plain, (k20_sin_cos(X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_55, negated_conjecture, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),k6_numbers))).#trainpos
cnf(c_0_102, negated_conjecture, (v3_rcomp_1(esk1_0))).#trainpos
cnf(c_0_166, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0))).#trainpos
cnf(c_0_87, plain, (v3_membered(k1_numbers))).#trainpos
cnf(c_0_79, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_80, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_76, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainpos
cnf(c_0_86, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainpos
cnf(c_0_99, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_105, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_212, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainpos
cnf(c_0_193, plain, (k7_xcmplx_0(k17_sin_cos(X1),k20_sin_cos(X1))=k1_sin_cos4(X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_77, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_144, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainpos
cnf(c_0_157, plain, (k1_real_1(k1_seq_1(k16_sin_cos,X1))=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_127, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainpos
cnf(c_0_111, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_104, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_97, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_57, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_165, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_249, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_251, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainpos
cnf(c_0_59, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainpos
cnf(c_0_129, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_248, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainpos
cnf(c_0_68, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0))).#trainpos
cnf(c_0_254, plain, (v1_relat_1(k16_sin_cos))).#trainpos
cnf(c_0_257, plain, (v1_relat_1(k19_sin_cos))).#trainpos
cnf(c_0_107, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_170, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_134, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_169, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_143, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk19_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_101, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk19_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_45, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk19_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_152, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_hidden(X1,X2)|~r2_fdiff_1(X4,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_253, plain, (v3_valued_0(k16_sin_cos))).#trainpos
cnf(c_0_256, plain, (v3_valued_0(k19_sin_cos))).#trainpos
cnf(c_0_167, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_46, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_112, plain, (k10_real_1(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_206, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainpos
cnf(c_0_108, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainpos
cnf(c_0_182, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0))).#trainpos
cnf(c_0_145, negated_conjecture, (r1_tarski(esk1_0,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainpos
cnf(c_0_202, plain, (k7_xcmplx_0(k17_sin_cos(X1),k1_seq_1(k19_sin_cos,X1))=k1_sin_cos4(X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_180, negated_conjecture, (k1_seq_1(k2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos))).#trainpos
cnf(c_0_120, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_168, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_fdiff_1(X1,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainpos
cnf(c_0_190, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos))).#trainpos
cnf(c_0_189, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_hidden(esk2_0,esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_numbers))).#trainpos
cnf(c_0_119, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_56, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))),k6_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~r1_fdiff_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~v1_funct_1(X1)|~v1_xreal_0(esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_133, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_128, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainpos
cnf(c_0_162, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainpos
cnf(c_0_211, plain, (k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_seq_1(k19_sin_cos,X1))=k1_sin_cos4(X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_66, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~r2_hidden(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),esk1_0)|~m1_subset_1(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k1_numbers)|~r1_fdiff_1(k19_sin_cos,esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_funct_1(k19_sin_cos)|~v1_xreal_0(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainpos
cnf(c_0_209, negated_conjecture, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainpos
cnf(c_0_136, plain, (v1_xboole_0(np__0))).#trainpos
cnf(c_0_156, plain, (k6_numbers=np__0)).#trainpos
cnf(c_0_172, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_233, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_82, plain, (v1_funct_1(k19_sin_cos))).#trainpos
cnf(c_0_100, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_xreal_0(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),esk1_0)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainpos
cnf(c_0_121, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_xreal_0(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_164, plain, (v1_funct_1(k3_taylor_1))).#trainpos
cnf(c_0_252, plain, (v1_funct_1(k16_sin_cos))).#trainpos
cnf(c_0_219, negated_conjecture, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(k3_taylor_1)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainpos
cnf(c_0_173, plain, (k1_fdiff_1(k19_sin_cos,X1)=k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_123, plain, (v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_118, plain, (v1_membered(X1)|~v3_membered(X1))).#trainpos
cnf(c_0_117, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_139, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_245, plain, (m1_subset_1(k1_sin_cos4(X1),k1_numbers)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_214, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainpos
cnf(c_0_122, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainpos
cnf(c_0_141, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_xreal_0(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainpos
cnf(c_0_161, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk19_2(X2,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainpos
cnf(c_0_222, negated_conjecture, (v3_membered(esk1_0))).#trainpos
cnf(c_0_160, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))|~m1_subset_1(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k1_numbers))).#trainpos
cnf(c_0_215, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers))).#trainpos
cnf(c_0_178, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainpos
cnf(c_0_174, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainpos
cnf(c_0_218, plain, (v1_funct_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainpos
cnf(c_0_147, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainpos
cnf(c_0_208, plain, (m1_subset_1(k3_relat_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainpos
cnf(c_0_146, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainpos
cnf(c_0_179, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainpos
cnf(c_0_188, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_216, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainpos
cnf(c_0_197, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_223, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers))).#trainpos
cnf(c_0_205, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))).#trainpos
cnf(c_0_213, negated_conjecture, (r2_hidden(esk2_0,esk1_0))).#trainpos
cnf(c_0_221, negated_conjecture, (m1_subset_1(esk2_0,esk1_0))).#trainpos
cnf(c_0_228, negated_conjecture, (v1_xreal_0(esk2_0))).#trainpos
cnf(c_0_183, plain, (k7_xcmplx_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1)),k1_seq_1(k19_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainpos
cnf(c_0_159, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_184, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainpos
cnf(c_0_192, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainpos
cnf(c_0_201, plain, (k4_xcmplx_0(k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers))).#trainpos
cnf(c_0_220, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)=k4_xcmplx_0(k1_sin_cos4(X1))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers))).#trainpos
cnf(c_0_227, plain, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),X1)=k4_xcmplx_0(k1_sin_cos4(X1))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers))).#trainpos
cnf(c_0_231, negated_conjecture, (r1_xxreal_0(k1_seq_1(k19_sin_cos,esk2_0),np__0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers))).#trainpos
cnf(c_0_225, plain, (v1_funct_1(k3_relat_1(X1,k3_taylor_1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_217, plain, (m1_subset_1(k3_relat_1(X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_224, plain, (m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainpos
cnf(c_0_229, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers))).#trainpos
cnf(c_0_230, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainpos
cnf(c_0_232, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(esk19_2(esk1_0,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers))).#trainpos
cnf(c_0_235, negated_conjecture, (r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainpos
cnf(c_0_234, negated_conjecture, (r1_xxreal_0(k1_seq_1(k19_sin_cos,esk2_0),np__0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainpos
cnf(c_0_244, negated_conjecture, (~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainpos
cnf(c_0_250, negated_conjecture, (~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers))).#trainpos
cnf(c_0_255, negated_conjecture, (~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers))).#trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_228, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_187, plain, (m1_subset_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_195, plain, (m1_subset_1(esk9_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_182, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_339, plain, (r2_fdiff_1(k19_sin_cos,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_44, plain, (v6_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_191, plain, (v3_rcomp_1(esk9_0))).#trainneg
cnf(i_0_299, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_102, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_136, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_143, plain, (~v1_xboole_0(k4_numbers))).#trainneg
cnf(i_0_147, plain, (v7_membered(k1_numbers))).#trainneg
cnf(i_0_159, plain, (v5_membered(k4_numbers))).#trainneg
cnf(i_0_203, plain, (v1_relat_1(esk11_0))).#trainneg
cnf(i_0_30, plain, (v1_int_1(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_79, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_222, plain, (k18_sin_cos(X1)=k17_sin_cos(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_227, plain, (k21_sin_cos(X1)=k20_sin_cos(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_20, plain, (v1_rat_1(X2)|~v4_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_22, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_24, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_358, plain, (v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_369, plain, (r2_fdiff_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_168, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_244, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_352, plain, (v1_xcmplx_0(esk6_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_140, plain, (~v1_xboole_0(k3_numbers))).#trainneg
cnf(i_0_171, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_362, negated_conjecture, (v1_rat_1(esk1_0)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_368, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_351, plain, (v1_xcmplx_0(esk9_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_100, plain, (m1_subset_1(k18_sin_cos(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_116, plain, (m1_subset_1(k21_sin_cos(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_86, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_226, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_355, plain, (v1_xxreal_0(esk6_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_188, plain, (v6_membered(esk7_0))).#trainneg
cnf(i_0_239, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_353, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_364, plain, (v1_int_1(esk6_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_310, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_315, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_378, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_354, plain, (v1_xxreal_0(esk9_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_356, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_95, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_110, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_23, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_25, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_40, plain, (v4_valued_0(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_450, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk9_0))).#trainneg
cnf(i_0_205, plain, (v6_membered(esk13_0))).#trainneg
cnf(i_0_245, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_357, plain, (v1_xreal_0(esk9_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_189, plain, (~v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_206, plain, (~v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_464, plain, (v1_relat_1(k3_taylor_1))).#trainneg
cnf(i_0_359, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_360, plain, (v1_rat_1(esk9_0)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_49, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_327, plain, (r2_fdiff_1(X2,X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r1_tarski(X1,X3)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_348, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_347, plain, (r1_fdiff_1(X3,X1)|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r2_hidden(X1,X2)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_476, plain, (v3_valued_0(k3_taylor_1))).#trainneg
cnf(i_0_212, plain, (v6_membered(esk15_0))).#trainneg
cnf(i_0_237, plain, (~r1_xxreal_0(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_361, plain, (v1_rat_1(esk6_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_483, plain, (v5_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_210, plain, (~v1_xboole_0(esk14_0))).#trainneg
cnf(i_0_213, plain, (~v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_484, plain, (v5_relat_1(k3_taylor_1,k1_numbers))).#trainneg
cnf(i_0_363, plain, (v1_int_1(esk9_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_365, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_92, plain, (X3=k2_fdiff_1(X1,X2)|m1_subset_1(esk3_3(X1,X2,X3),k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_91, plain, (X3=k2_fdiff_1(X1,X2)|r2_hidden(esk3_3(X1,X2,X3),X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_90, plain, (X1=k2_fdiff_1(X2,X3)|k1_relset_1(k1_numbers,X1)!=X3|k1_fdiff_1(X2,esk3_3(X2,X3,X1))!=k1_seq_1(X1,esk3_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_549, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_485, plain, (v5_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_194, plain, (v1_membered(esk9_0))).#trainneg
cnf(i_0_366, plain, (v7_ordinal1(esk9_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_487, plain, (v4_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_220, plain, (~v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_178, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_488, plain, (v4_relat_1(k3_taylor_1,k1_numbers))).#trainneg
cnf(i_0_489, plain, (v4_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_367, plain, (v7_ordinal1(esk6_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_388, plain, (v1_xreal_0(k18_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_392, plain, (m1_subset_1(k17_sin_cos(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_395, plain, (v1_xreal_0(k21_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_399, plain, (m1_subset_1(k20_sin_cos(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_402, plain, (m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_408, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_449, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_461, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_510, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_516, plain, (m1_subset_1(esk19_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_185, plain, (v1_xcmplx_0(esk6_0))).#trainneg
cnf(i_0_588, plain, (v1_xreal_0(np__0)|~m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_169, plain, (~v1_finset_1(k4_numbers))).#trainneg
cnf(i_0_172, plain, (~v1_finset_1(k3_numbers))).#trainneg
cnf(i_0_468, plain, (v1_valued_0(k3_taylor_1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_469, plain, (v1_valued_0(k19_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_527, plain, (m1_subset_1(esk3_3(X1,X2,X3),X2)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_547, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_552, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_553, negated_conjecture, (r1_tarski(esk1_0,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_566, plain, (v1_xreal_0(k17_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_574, plain, (v1_xreal_0(k20_sin_cos(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_467, plain, (v1_valued_0(k16_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_193, plain, (v2_membered(esk9_0))).#trainneg
cnf(i_0_609, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~r2_hidden(esk2_0,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_374, plain, (k1_real_1(np__0)=np__0|~m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_242, plain, (~r1_xxreal_0(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_247, plain, (~r1_xxreal_0(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_472, plain, (v2_valued_0(k3_taylor_1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_473, plain, (v2_valued_0(k19_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_376, plain, (k7_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k20_sin_cos(X1))=k1_sin_cos4(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_380, plain, (v1_rat_1(k1_real_1(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_603, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,X2))|~r2_hidden(X1,X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3))).#trainneg
cnf(i_0_607, plain, (k10_real_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3)|r1_xxreal_0(k1_fdiff_1(X1,X3),k6_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xreal_0(X3)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_635, plain, (k7_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3)|r1_xxreal_0(k1_fdiff_1(X1,X3),k6_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X1,X3))|~v1_xreal_0(X3)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_645, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_647, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_679, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0)|~v1_funct_1(k3_taylor_1)|~v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_186, plain, (v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_192, plain, (v3_membered(esk9_0))).#trainneg
cnf(i_0_471, plain, (v2_valued_0(k16_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_590, plain, (m1_subset_1(k7_xcmplx_0(X1,esk6_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_142, plain, (v4_membered(k3_numbers))).#trainneg
cnf(i_0_381, plain, (v1_int_1(k1_real_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_479, plain, (v4_valued_0(k16_sin_cos)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_243, plain, (~r1_xxreal_0(np__1,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_248, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_480, plain, (v4_valued_0(k3_taylor_1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_481, plain, (v4_valued_0(k19_sin_cos)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_382, plain, (v7_ordinal1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_386, plain, (v1_xcmplx_0(k18_sin_cos(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_201, plain, (v4_valued_0(esk11_0))).#trainneg
cnf(i_0_183, plain, (v1_int_1(esk6_0))).#trainneg
cnf(i_0_204, plain, (v1_int_1(esk12_0))).#trainneg
cnf(i_0_218, plain, (v7_ordinal1(esk17_0))).#trainneg
cnf(i_0_219, plain, (v7_ordinal1(esk18_0))).#trainneg
cnf(i_0_635, plain, (k7_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3)|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_607, plain, (k10_real_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3)|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_529, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))),np__0)|~r1_fdiff_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))|~v1_funct_1(X1)|~v1_xreal_0(esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_325, plain, (k10_real_1(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_341, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_689, plain, (v1_xreal_0(k7_xcmplx_0(X1,esk6_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_229, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_387, plain, (v1_xxreal_0(k18_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_698, plain, (v4_valued_0(k16_sin_cos)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_253, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__0))).#trainneg
cnf(i_0_170, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_653, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_389, plain, (v1_rat_1(k18_sin_cos(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_390, plain, (v1_int_1(k18_sin_cos(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_326, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_190, plain, (v3_ordinal1(esk8_0))).#trainneg
cnf(i_0_207, plain, (v3_ordinal1(esk14_0))).#trainneg
cnf(i_0_209, plain, (v1_ordinal1(esk14_0))).#trainneg
cnf(i_0_208, plain, (v2_ordinal1(esk14_0))).#trainneg
cnf(i_0_153, plain, (v7_membered(k4_numbers))).#trainneg
cnf(i_0_150, plain, (v7_membered(k3_numbers))).#trainneg
cnf(i_0_699, plain, (v4_valued_0(k3_taylor_1)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_700, plain, (v4_valued_0(k19_sin_cos)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_156, plain, (v7_membered(k4_ordinal1))).#trainneg
cnf(i_0_654, negated_conjecture, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainneg
cnf(i_0_391, plain, (v7_ordinal1(k18_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_211, plain, (v7_membered(esk15_0))).#trainneg
cnf(i_0_439, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_581, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_393, plain, (v1_xcmplx_0(k21_sin_cos(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_394, plain, (v1_xxreal_0(k21_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_679, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_funct_1(k3_taylor_1)|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_553, negated_conjecture, (r1_tarski(esk1_0,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_funct_1(k3_taylor_1))).#trainneg
cnf(i_0_550, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k3_taylor_1)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_528, negated_conjecture, (k1_seq_1(k2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~v1_funct_1(k3_taylor_1)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_528, negated_conjecture, (k1_seq_1(k2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_550, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_679, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_553, negated_conjecture, (r1_tarski(esk1_0,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_753, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_hidden(esk2_0,esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_549, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k3_taylor_1)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_760, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_396, plain, (v1_rat_1(k21_sin_cos(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_202, plain, (v1_funct_1(esk11_0))).#trainneg
cnf(i_0_314, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_730, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,X1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_397, plain, (v1_int_1(k21_sin_cos(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_398, plain, (v7_ordinal1(k21_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_761, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_319, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_302, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_303, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_217, plain, (v1_relat_1(esk16_1(X1)))).#trainneg
cnf(i_0_214, plain, (v5_ordinal1(esk16_1(X1)))).#trainneg
cnf(i_0_215, plain, (v1_funct_1(esk16_1(X1)))).#trainneg
cnf(i_0_292, plain, (k7_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_298, plain, (k7_xcmplx_0(np__2,np__2)=np__1)).#trainneg
cnf(i_0_269, plain, (k6_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_270, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_776, plain, (v1_xreal_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_418, plain, (v1_xxreal_0(k4_xcmplx_0(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_322, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_792, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_257, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_787, plain, (v1_xcmplx_0(np__0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_731, plain, (v1_xxreal_0(k1_seq_1(k19_sin_cos,X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_733, plain, (v1_rat_1(k1_seq_1(k19_sin_cos,X1))|~v4_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_401, plain, (k1_real_1(k1_fdiff_1(k19_sin_cos,X1))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_403, plain, (k4_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))=k1_seq_1(k16_sin_cos,X1)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_274, plain, (k6_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_275, plain, (k6_xcmplx_0(np__2,np__1)=np__1)).#trainneg
cnf(i_0_276, plain, (k6_xcmplx_0(np__2,np__2)=np__0)).#trainneg
cnf(i_0_307, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_311, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_312, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_316, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_317, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_234, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_235, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_236, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_588, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_826, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_835, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_838, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_847, plain, (v7_ordinal1(np__2))).#trainneg
cnf(i_0_734, plain, (v1_int_1(k1_seq_1(k19_sin_cos,X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_735, plain, (v7_ordinal1(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_407, plain, (v1_xreal_0(k1_fdiff_1(k19_sin_cos,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_850, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_240, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_241, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_246, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_31, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_83, plain, (v6_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_26, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_63, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_58, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_57, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_891, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_876, plain, (v1_valued_0(k16_sin_cos))).#trainneg
cnf(i_0_877, plain, (v1_valued_0(k3_taylor_1))).#trainneg
cnf(i_0_878, plain, (v1_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_374, plain, (k1_real_1(np__0)=np__0)).#trainneg
cnf(i_0_409, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_real_1(X1)))=k1_real_1(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_881, plain, (v2_valued_0(k16_sin_cos))).#trainneg
cnf(i_0_882, plain, (v2_valued_0(k3_taylor_1))).#trainneg
cnf(i_0_883, plain, (v2_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_892, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_47, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_46, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_941, plain, (v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_942, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_943, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_75, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_72, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_972, plain, (v3_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_59, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_33, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_32, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_884, plain, (v3_membered(k3_numbers))).#trainneg
cnf(i_0_885, plain, (v1_xreal_0(esk12_0))).#trainneg
cnf(i_0_929, plain, (v1_int_1(esk18_0))).#trainneg
cnf(i_0_412, plain, (v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_930, plain, (v1_int_1(esk17_0))).#trainneg
cnf(i_0_934, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_935, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_958, plain, (v1_int_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_414, plain, (r1_xxreal_0(X1,k1_real_1(X2))|r1_xxreal_0(k1_real_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_45, plain, (v7_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_64, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_263, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_264, plain, (k6_xcmplx_0(np__0,np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_271, plain, (k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_232, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_265, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_266, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_272, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2)).#trainneg
cnf(i_0_127, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_249, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0))).#trainneg
cnf(i_0_1001, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_993, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_936, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_959, plain, (v3_ordinal1(esk18_0))).#trainneg
cnf(i_0_960, plain, (v3_ordinal1(esk17_0))).#trainneg
cnf(i_0_964, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_980, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1007, plain, (v1_rat_1(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_417, plain, (v1_xxreal_0(k1_fdiff_1(k19_sin_cos,X1))|~v2_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1019, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1031, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2),k1_numbers)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_xboole_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_250, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_251, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_343, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_126, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_138, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_137, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_294, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_277, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_965, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_966, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_974, plain, (v1_xreal_0(esk18_0))).#trainneg
cnf(i_0_342, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_419, plain, (v1_xcmplx_0(k19_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_977, plain, (v1_xreal_0(esk17_0))).#trainneg
cnf(i_0_984, plain, (v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_1009, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_928, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_420, plain, (v1_xxreal_0(k19_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_421, plain, (v1_xreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1073, plain, (r2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_278, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_73, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_70, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_80, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_84, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_48, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_334, plain, (k7_xcmplx_0(k6_numbers,X1)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_134, plain, (m1_subset_1(esk4_1(X1),X1))).#trainneg
cnf(i_0_216, plain, (v5_relat_1(esk16_1(X1),X1))).#trainneg
cnf(i_0_279, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_280, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=np__1)).#trainneg
cnf(i_0_1109, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk4_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1080, plain, (v3_membered(k4_numbers))).#trainneg
cnf(i_0_1085, plain, (v3_valued_0(esk11_0))).#trainneg
cnf(i_0_1074, plain, (X3=k2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=X1|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_422, plain, (v1_rat_1(k19_sin_cos)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_788, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_790, plain, (v1_rat_1(np__0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_1081, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1089, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_423, plain, (v1_int_1(k19_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_424, plain, (v7_ordinal1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_283, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))=np__0)).#trainneg
cnf(i_0_252, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_340, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_331, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_282, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_35, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_1153, plain, (v5_relat_1(esk11_0,k4_numbers))).#trainneg
cnf(i_0_76, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_1158, plain, (v5_relat_1(esk11_0,k3_numbers))).#trainneg
cnf(i_0_233, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_121, plain, (v1_relat_1(k3_relat_1(X1,X2)))).#trainneg
cnf(i_0_1163, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1171, plain, (r1_xxreal_0(esk6_0,esk6_0))).#trainneg
cnf(i_0_819, plain, (v1_xxreal_0(np__0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_425, plain, (v1_xcmplx_0(k3_taylor_1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1154, plain, (v5_relat_1(k16_sin_cos,k4_numbers)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1149, plain, (v1_xreal_0(k1_sin_cos4(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1147, plain, (v1_xcmplx_0(k1_sin_cos4(X1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_426, plain, (v1_xxreal_0(k3_taylor_1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_427, plain, (v1_xreal_0(k3_taylor_1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_198, plain, (v1_relat_1(esk10_2(X1,X2)))).#trainneg
cnf(i_0_199, plain, (v1_xboole_0(esk10_2(X1,X2)))).#trainneg
cnf(i_0_295, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_9, plain, (v5_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_12, plain, (v6_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_78, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_82, plain, (v4_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_53, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_61, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_66, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_55, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_144, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v5_ordinal1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1198, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1208, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_821, plain, (v1_rat_1(np__0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_257, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_428, plain, (v1_rat_1(k3_taylor_1)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1155, plain, (v5_relat_1(k3_taylor_1,k4_numbers)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_824, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1206, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1148, plain, (v1_xxreal_0(k1_sin_cos4(X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_429, plain, (v1_int_1(k3_taylor_1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_430, plain, (v7_ordinal1(k3_taylor_1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_296, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=np__2)).#trainneg
cnf(i_0_238, plain, (r1_xxreal_0(np__0,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_254, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__1))).#trainneg
cnf(i_0_255, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__2))).#trainneg
cnf(i_0_69, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_173, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_174, plain, (v2_valued_0(k2_zfmisc_1(X2,X1))|~v2_membered(X1))).#trainneg
cnf(i_0_175, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_180, plain, (v4_valued_0(k2_zfmisc_1(X2,X1))|~v6_membered(X1))).#trainneg
cnf(i_0_304, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_196, plain, (v5_relat_1(esk10_2(X1,X2),X2))).#trainneg
cnf(i_0_197, plain, (v4_relat_1(esk10_2(X1,X2),X1))).#trainneg
cnf(i_0_1222, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1010, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_825, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_257, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_431, plain, (v1_xcmplx_0(k16_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1156, plain, (v5_relat_1(k19_sin_cos,k4_numbers)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_827, plain, (v1_rat_1(np__1)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1150, plain, (v1_rat_1(k1_sin_cos4(X1))|~v4_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1151, plain, (v1_int_1(k1_sin_cos4(X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_432, plain, (v1_xxreal_0(k16_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_433, plain, (v1_xreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_51, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_328, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_258, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__0))).#trainneg
cnf(i_0_259, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__1))).#trainneg
cnf(i_0_260, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2))).#trainneg
cnf(i_0_273, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_284, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__0)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_42, plain, (v5_membered(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10, plain, (v6_ordinal1(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_43, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_28, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_37, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_848, plain, (m1_subset_1(k7_xcmplx_0(X1,np__0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_831, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_295, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_434, plain, (v1_rat_1(k16_sin_cos)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1159, plain, (v5_relat_1(k16_sin_cos,k3_numbers)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1312, negated_conjecture, (v5_membered(esk1_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1327, negated_conjecture, (v6_ordinal1(esk1_0)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_435, plain, (v1_int_1(k16_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_436, plain, (v7_ordinal1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_41, plain, (v4_membered(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_305, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_297, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_141, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_8, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1465, plain, (r2_fdiff_1(X2,X1)|~r2_hidden(X1,esk19_2(X1,X2))|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_177, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k4_numbers)|~v5_membered(X1))).#trainneg
cnf(i_0_176, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k3_numbers)|~v4_membered(X1))).#trainneg
cnf(i_0_1411, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_833, plain, (v1_rat_1(np__1)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_437, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1160, plain, (v5_relat_1(k3_taylor_1,k3_numbers)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_836, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1342, negated_conjecture, (v6_membered(esk1_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1357, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_438, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_440, plain, (v1_rat_1(k1_seq_1(X1,X2))|~v4_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1471, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk19_2(X2,X1),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1474, plain, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1477, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))|~m1_subset_1(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k4_numbers))).#trainneg
cnf(i_0_1478, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))|~m1_subset_1(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k5_numbers))).#trainneg
cnf(i_0_1479, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))|~v1_xboole_0(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1520, plain, (r2_fdiff_1(X1,esk9_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1525, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1492, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1449, plain, (v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_837, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_441, plain, (v1_int_1(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1161, plain, (v5_relat_1(k19_sin_cos,k3_numbers)|~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_839, plain, (v1_rat_1(np__2)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1372, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1423, negated_conjecture, (v4_membered(esk1_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_442, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_443, plain, (v1_xcmplx_0(k10_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1540, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X2,k1_numbers))).#trainneg
cnf(i_0_1541, plain, (r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X2,esk9_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,k1_numbers))).#trainneg
cnf(i_0_1543, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1546, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X2,k1_numbers))).#trainneg
cnf(i_0_1547, negated_conjecture, (r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X2,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,k1_numbers))).#trainneg
cnf(i_0_1578, plain, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,k1_numbers)|~r2_hidden(esk19_2(X2,X1),esk9_0))).#trainneg
cnf(i_0_1586, negated_conjecture, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,k1_numbers)|~r2_hidden(esk19_2(X2,X1),esk1_0))).#trainneg
cnf(i_0_1587, plain, (r2_fdiff_1(X1,X2)|v1_xboole_0(esk9_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(esk19_2(X2,X1),esk9_0))).#trainneg
cnf(i_0_1589, negated_conjecture, (r2_fdiff_1(X1,X2)|v1_xboole_0(esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(esk19_2(X2,X1),esk1_0))).#trainneg
cnf(i_0_267, plain, (k6_xcmplx_0(np__0,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1454, plain, (v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_843, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_444, plain, (v1_xxreal_0(k10_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_845, plain, (v1_rat_1(np__2)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_1152, plain, (v7_ordinal1(k1_sin_cos4(X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_445, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_446, plain, (v1_rat_1(k10_real_1(X1,X2))|~v4_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_285, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_323, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_97, plain, (v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_128, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_313, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_318, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13, plain, (v2_valued_0(X2)|~v1_relat_1(X1)|~v2_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_15, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_21, plain, (v4_valued_0(X2)|~v1_relat_1(X1)|~v4_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_447, plain, (v1_int_1(k10_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1640, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1656, negated_conjecture, (v2_valued_0(esk1_0)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_448, plain, (v7_ordinal1(k10_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_452, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_145, plain, (v1_valued_0(k3_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2))).#trainneg
cnf(i_0_148, plain, (v2_valued_0(k3_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v2_valued_0(X2))).#trainneg
cnf(i_0_162, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_151, plain, (v3_valued_0(k3_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X2))).#trainneg
cnf(i_0_163, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_164, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_160, plain, (v4_valued_0(k3_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v4_valued_0(X2))).#trainneg
cnf(i_0_167, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v4_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1555, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_962, plain, (v3_ordinal1(esk6_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_453, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_1236, plain, (v3_ordinal1(esk6_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1307, plain, (v5_membered(esk9_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1315, plain, (v5_membered(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1672, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1688, negated_conjecture, (v4_valued_0(esk1_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_454, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_455, plain, (v1_xcmplx_0(k1_relset_1(X1,X2))|~v1_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_268, plain, (k6_xcmplx_0(np__0,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_286, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__2))=np__0)).#trainneg
cnf(i_0_256, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_320, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_288, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__0)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_289, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k4_xcmplx_0(np__1))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_281, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_335, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_287, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__1)).#trainneg
cnf(i_0_261, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_1322, plain, (v6_ordinal1(esk9_0)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1598, plain, (v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1330, plain, (v6_ordinal1(k5_numbers)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_456, plain, (v1_xxreal_0(k1_relset_1(X1,X2))|~v2_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1337, plain, (v6_membered(esk9_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1360, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1375, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1267, plain, (v3_ordinal1(k1_sin_cos4(X1))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_457, plain, (v1_xreal_0(k1_relset_1(X1,X2))|~v3_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_458, plain, (v1_rat_1(k1_relset_1(X1,X2))|~v4_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_165, plain, (v1_rat_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_166, plain, (v1_int_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_290, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__1,np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_291, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__0)).#trainneg
cnf(i_0_262, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_19, plain, (v5_relat_1(X2,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1824, plain, (v5_relat_1(k1_relset_1(X1,X2),k4_numbers)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_17, plain, (v5_relat_1(X2,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1840, plain, (v5_relat_1(k1_relset_1(X1,X2),k3_numbers)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_157, plain, (v5_relat_1(k3_relat_1(X1,X2),k4_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v5_relat_1(X2,k4_numbers))).#trainneg
cnf(i_0_154, plain, (v5_relat_1(k3_relat_1(X1,X2),k3_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v5_relat_1(X2,k3_numbers))).#trainneg
cnf(i_0_1418, plain, (v4_membered(esk9_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1384, plain, (v1_xreal_0(k7_xcmplx_0(X1,np__0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1426, plain, (v4_membered(k5_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_459, plain, (v1_int_1(k1_relset_1(X1,X2))|~v5_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1818, plain, (v5_relat_1(esk9_0,k4_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_868, plain, (v1_int_1(esk9_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_869, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1856, negated_conjecture, (v1_int_1(esk1_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1859, negated_conjecture, (v1_xreal_0(esk1_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_460, plain, (v7_ordinal1(k1_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_462, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_200, plain, (m1_subset_1(esk10_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1883, plain, (r2_fdiff_1(X1,esk10_2(X2,X3))|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_zfmisc_1(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk10_2(X2,X3),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(esk10_2(X2,X3)))).#trainneg
cnf(i_0_230, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_135, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk5_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_65, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_60, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_38, plain, (v5_relat_1(X2,k4_numbers)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_29, plain, (v5_relat_1(X2,k3_numbers)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_231, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_133, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_94, plain, (k1_relset_1(k1_numbers,X1)=X2|X1!=k2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1947, plain, (k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))=X2|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_849, plain, (m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_466, plain, (v1_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1823, negated_conjecture, (v5_relat_1(esk1_0,k4_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_470, plain, (v2_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v2_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_474, plain, (v3_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1948, plain, (m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1957, plain, (m1_subset_1(X1,k1_numbers)|~v4_relat_1(k2_fdiff_1(X2,X3),k1_numbers)|~v1_relat_1(k2_fdiff_1(X2,X3))|~r2_hidden(X1,X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_179, plain, (v1_relat_1(k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))))).#trainneg
cnf(i_0_478, plain, (v4_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v6_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_482, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_486, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_490, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_491, plain, (v1_xxreal_0(k1_fdiff_1(X1,X2))|~v2_membered(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_492, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_493, plain, (v1_rat_1(k1_fdiff_1(X1,X2))|~v4_membered(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_494, plain, (v1_int_1(k1_fdiff_1(X1,X2))|~v5_membered(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_495, plain, (v7_ordinal1(k1_fdiff_1(X1,X2))|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1300, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(esk9_0)|~m1_subset_1(X1,esk9_0))).#trainneg
cnf(i_0_1971, plain, (v1_xreal_0(k7_xcmplx_0(X1,np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_555, plain, (k9_xtuple_0(k16_sin_cos)=k1_relset_1(k1_numbers,k16_sin_cos))).#trainneg
cnf(i_0_1783, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_496, plain, (v1_xcmplx_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1826, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_874, plain, (v1_xcmplx_0(esk9_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_875, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2023, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1606, plain, (v1_int_1(k1_sin_cos4(X1))|~v1_xreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_497, plain, (v1_xxreal_0(k2_fdiff_1(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_498, plain, (v1_xreal_0(k2_fdiff_1(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_499, plain, (v1_rat_1(k2_fdiff_1(X1,X2))|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_500, plain, (v1_int_1(k2_fdiff_1(X1,X2))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_501, plain, (v7_ordinal1(k2_fdiff_1(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_502, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_503, plain, (v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_504, plain, (v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_505, plain, (v1_valued_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_506, plain, (v2_valued_0(k2_fdiff_1(X1,X2))|~v2_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_507, plain, (v3_valued_0(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_508, plain, (v4_valued_0(k2_fdiff_1(X1,X2))|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_509, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_511, plain, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1302, negated_conjecture, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_851, plain, (m1_subset_1(k7_xcmplx_0(X1,np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_556, plain, (k9_xtuple_0(k3_taylor_1)=k1_relset_1(k1_numbers,k3_taylor_1))).#trainneg
cnf(i_0_2012, plain, (X2=k2_fdiff_1(X1,k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k3_taylor_1|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_512, plain, (r1_fdiff_1(k19_sin_cos,X1)|~v1_funct_1(k19_sin_cos)|~r2_hidden(X1,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_numbers))).#trainneg
cnf(i_0_1834, plain, (v5_relat_1(esk9_0,k3_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_879, plain, (v1_xxreal_0(esk9_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_880, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2069, negated_conjecture, (v1_xxreal_0(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2068, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_514, plain, (m1_subset_1(esk19_2(esk9_0,X1),k1_numbers)|r2_fdiff_1(X1,esk9_0)|~v1_funct_1(X1)|~r1_tarski(esk9_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_515, negated_conjecture, (m1_subset_1(esk19_2(esk1_0,X1),k1_numbers)|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_517, plain, (r2_fdiff_1(k19_sin_cos,X1)|~v1_funct_1(k19_sin_cos)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk19_2(X1,k19_sin_cos)))).#trainneg
cnf(i_0_519, plain, (v1_xcmplx_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v1_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_520, plain, (v1_xxreal_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v2_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_521, plain, (v1_xreal_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2092, plain, (m1_subset_1(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)),k1_numbers)|X4=k2_fdiff_1(X2,X3)|~m1_subset_1(X1,k1_numbers)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_522, plain, (v1_rat_1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v4_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_523, plain, (v1_int_1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v5_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_524, plain, (v7_ordinal1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v6_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_534, plain, (v1_xcmplx_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2054, plain, (v1_xreal_0(k7_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_557, plain, (k9_xtuple_0(k19_sin_cos)=k1_relset_1(k1_numbers,k19_sin_cos))).#trainneg
cnf(i_0_2028, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(esk1_0))=esk1_0|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_535, plain, (v1_xxreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1839, negated_conjecture, (v5_relat_1(esk1_0,k3_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_961, plain, (v3_ordinal1(esk9_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1607, plain, (v3_ordinal1(k1_sin_cos4(X1))|~v1_xreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_536, plain, (v1_xreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_537, plain, (v1_rat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_538, plain, (v1_int_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_539, plain, (v7_ordinal1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_540, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k1_partfun1(X2,X4,X5,X3,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_543, plain, (v1_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_544, plain, (v2_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_545, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_546, plain, (v4_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v6_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_561, plain, (v1_rat_1(k4_xcmplx_0(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_973, plain, (m1_subset_1(k7_xcmplx_0(X1,esk12_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_998, plain, (v1_rat_1(k5_numbers)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2029, negated_conjecture, (k7_xcmplx_0(np__0,esk1_0)=np__0|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_562, plain, (v1_int_1(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1842, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_963, negated_conjecture, (v3_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_995, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2185, negated_conjecture, (v3_ordinal1(esk1_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2075, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_563, plain, (v7_ordinal1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_567, plain, (v1_xcmplx_0(k17_sin_cos(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_568, plain, (v1_xxreal_0(k17_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_570, plain, (v1_rat_1(k17_sin_cos(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_571, plain, (v1_int_1(k17_sin_cos(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_572, plain, (v7_ordinal1(k17_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_573, plain, (m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_575, plain, (v1_xcmplx_0(k20_sin_cos(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_576, plain, (v1_xxreal_0(k20_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_578, plain, (v1_rat_1(k20_sin_cos(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_579, plain, (v1_int_1(k20_sin_cos(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_580, plain, (v7_ordinal1(k20_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_996, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2171, plain, (v1_xreal_0(k7_xcmplx_0(X1,esk12_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1000, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2030, negated_conjecture, (k7_xcmplx_0(esk1_0,np__1)=esk1_0|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1920, plain, (v5_relat_1(k16_sin_cos,k4_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_997, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_999, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_939, plain, (v1_int_1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_585, plain, (v1_rat_1(k1_fdiff_1(k19_sin_cos,X1))|~v4_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_586, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_587, plain, (v7_ordinal1(k1_fdiff_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_591, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_592, plain, (m1_subset_1(k7_xcmplx_0(X1,k21_sin_cos(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_593, plain, (m1_subset_1(k7_xcmplx_0(X1,k18_sin_cos(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_594, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_595, plain, (m1_subset_1(esk19_2(k1_relset_1(X1,X2),X3),X1)|r2_fdiff_1(X3,k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X3)|~r1_tarski(k1_relset_1(X1,X2),k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(X1,X2)))).#trainneg
cnf(i_0_596, plain, (m1_subset_1(esk3_3(X1,k1_relset_1(X2,X3),X4),X2)|X4=k2_fdiff_1(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|k1_relset_1(k1_numbers,X4)!=k1_relset_1(X2,X3)|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_relset_1(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_597, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_598, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_599, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_600, plain, (v1_rat_1(X1)|~v4_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1004, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2278, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_1853, plain, (v1_int_1(esk9_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1779, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_601, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1923, plain, (v5_relat_1(k3_taylor_1,k4_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_969, plain, (v3_ordinal1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1266, plain, (v3_ordinal1(k1_real_1(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_602, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_606, plain, (m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_2400, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~r2_hidden(X2,esk9_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~r2_fdiff_1(X1,k1_numbers))).#trainneg
cnf(i_0_2401, negated_conjecture, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~r2_hidden(X2,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~r2_fdiff_1(X1,k1_numbers))).#trainneg
cnf(i_0_605, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X4)|k1_fdiff_1(X3,esk3_3(X3,X4,k2_fdiff_1(X1,X2)))!=k1_fdiff_1(X1,esk3_3(X3,X4,k2_fdiff_1(X1,X2)))|k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))!=X4|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(esk3_3(X3,X4,k2_fdiff_1(X1,X2)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_3(X3,X4,k2_fdiff_1(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_610, plain, (k4_xcmplx_0(k4_xcmplx_0(esk6_0))=esk6_0)).#trainneg
cnf(i_0_618, plain, (v1_relat_1(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_619, plain, (v4_relat_1(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4),X2)|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_620, plain, (v1_valued_0(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_membered(X3)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_621, plain, (v2_valued_0(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v2_membered(X3)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_622, plain, (v3_valued_0(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_membered(X3)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_623, plain, (v4_valued_0(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v6_membered(X3)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1005, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2217, plain, (v1_xreal_0(k1_seq_1(k16_sin_cos,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2065, plain, (v1_xxreal_0(esk9_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2376, plain, (X2=k2_fdiff_1(X1,k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k16_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k16_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_626, plain, (v1_xcmplx_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v1_membered(X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1924, plain, (v5_relat_1(k19_sin_cos,k4_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1732, negated_conjecture, (v3_valued_0(esk1_0)|~v1_relat_1(esk1_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_624, plain, (v5_relat_1(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4),X3)|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_625, plain, (m1_subset_1(X1,X2)|X4=k2_fdiff_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,esk3_3(X3,k1_zfmisc_1(X2),X4))|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2446, plain, (k7_xcmplx_0(esk3_3(X1,X2,X3),np__1)=esk3_3(X1,X2,X3)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2447, plain, (k6_xcmplx_0(esk3_3(X1,X2,X3),np__0)=esk3_3(X1,X2,X3)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_627, plain, (v1_xxreal_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v2_membered(X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_629, plain, (v1_rat_1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v4_membered(X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_630, plain, (v1_int_1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v5_membered(X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_631, plain, (v7_ordinal1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v6_membered(X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2468, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_2470, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_747, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~v1_xreal_0(esk19_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2526, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1072, plain, (m1_subset_1(k7_xcmplx_0(X1,esk18_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2127, plain, (v3_ordinal1(esk9_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_637, plain, (v4_relat_1(k3_relat_1(X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1931, plain, (v5_relat_1(k16_sin_cos,k3_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1122, plain, (v1_xreal_0(esk9_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1123, negated_conjecture, (v1_xreal_0(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_638, plain, (v5_relat_1(k3_relat_1(X5,X6),X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_639, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_2501, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_2539, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2540, negated_conjecture, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2541, negated_conjecture, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2619, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk19_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),esk1_0))).#trainneg
cnf(i_0_2628, negated_conjecture, (m1_subset_1(k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),X1),k1_numbers)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2534, plain, (v1_xreal_0(k7_xcmplx_0(X1,esk18_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2374, plain, (v1_xreal_0(esk9_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_640, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_1934, plain, (v5_relat_1(k3_taylor_1,k3_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1235, plain, (v3_ordinal1(esk9_0)|~v3_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1242, negated_conjecture, (v3_ordinal1(esk1_0)|~v3_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1006, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_641, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_642, negated_conjecture, (v1_rat_1(esk1_0)|~v4_membered(k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_2629, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2671, negated_conjecture, (~r2_hidden(esk1_0,esk2_0))).#trainneg
cnf(i_0_2672, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_643, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_644, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_650, plain, (m1_subset_1(k7_xcmplx_0(X1,k17_sin_cos(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1076, plain, (m1_subset_1(k7_xcmplx_0(X1,esk17_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1784, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_652, plain, (m1_subset_1(k7_xcmplx_0(X1,k20_sin_cos(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1935, plain, (v5_relat_1(k19_sin_cos,k3_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_2669, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_2653, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1008, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_668, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~v1_funct_1(k19_sin_cos)|~r2_hidden(X1,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_numbers))).#trainneg
cnf(i_0_673, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_674, negated_conjecture, (v1_xxreal_0(X1)|~v2_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_675, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_676, negated_conjecture, (v1_rat_1(X1)|~v4_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_677, negated_conjecture, (v1_int_1(X1)|~v5_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_678, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_686, plain, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_687, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,esk6_0))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2793, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,esk6_0)))=k7_xcmplx_0(X1,esk6_0)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_688, plain, (v1_xxreal_0(k7_xcmplx_0(X1,esk6_0))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1078, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2708, plain, (v1_xreal_0(k7_xcmplx_0(X1,esk17_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2674, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_2715, plain, (X2=k2_fdiff_1(X1,k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k19_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k19_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_690, plain, (v1_rat_1(k7_xcmplx_0(X1,esk6_0))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2241, plain, (v5_relat_1(k16_sin_cos,k4_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2673, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(esk1_0))).#trainneg
cnf(i_0_2812, negated_conjecture, (v1_xcmplx_0(esk2_0))).#trainneg
cnf(i_0_2801, plain, (v1_int_1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2819, plain, (v1_xreal_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_691, plain, (v1_int_1(k7_xcmplx_0(X1,esk6_0))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_692, plain, (v7_ordinal1(k7_xcmplx_0(X1,esk6_0))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_706, plain, (k4_xcmplx_0(k4_xcmplx_0(k18_sin_cos(X1)))=k18_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_709, plain, (k7_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3)|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_710, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3),k1_numbers)|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1079, plain, (v3_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2676, negated_conjecture, (v1_rat_1(esk2_0)|~v4_membered(esk1_0))).#trainneg
cnf(i_0_717, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(X2,esk6_0)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2677, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_2678, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_2852, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2868, negated_conjecture, (v3_ordinal1(esk1_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_712, plain, (r2_fdiff_1(k3_relat_1(X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,esk19_2(X2,k3_relat_1(X1,k3_taylor_1))),np__0)|~r1_fdiff_1(X1,esk19_2(X2,k3_relat_1(X1,k3_taylor_1)))|~v1_funct_1(k3_relat_1(X1,k3_taylor_1))|~v1_funct_1(X1)|~v1_xreal_0(esk19_2(X2,k3_relat_1(X1,k3_taylor_1)))|~r1_tarski(X2,k1_relset_1(k1_numbers,k3_relat_1(X1,k3_taylor_1)))|~m1_subset_1(k3_relat_1(X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(k3_taylor_1))).#trainneg
cnf(i_0_715, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2),k1_numbers)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_713, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3)|r1_xxreal_0(k1_fdiff_1(X1,esk19_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1))),np__0)|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk19_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xreal_0(esk19_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1)))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(esk19_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_719, plain, (r1_xxreal_0(X1,k18_sin_cos(X2))|r1_xxreal_0(k18_sin_cos(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_721, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_seq_1(k19_sin_cos,X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2891, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1),k1_numbers)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~m1_subset_1(k1_fdiff_1(k19_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_725, plain, (r1_fdiff_1(k3_relat_1(X1,k3_taylor_1),X2)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k3_taylor_1))).#trainneg
cnf(i_0_728, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_seq_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_729, plain, (v1_xreal_0(k1_fdiff_1(X1,X3))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_736, plain, (k4_xcmplx_0(k4_xcmplx_0(k21_sin_cos(X1)))=k21_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1277, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2679, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(esk1_0))).#trainneg
cnf(i_0_1773, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk4_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_2683, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_2935, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_738, plain, (r1_xxreal_0(X1,k21_sin_cos(X2))|r1_xxreal_0(k21_sin_cos(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_754, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_755, negated_conjecture, (v1_xxreal_0(X1)|~v2_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_756, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_757, negated_conjecture, (v1_rat_1(X1)|~v4_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_758, negated_conjecture, (v1_int_1(X1)|~v5_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_759, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_763, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_764, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_765, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_766, negated_conjecture, (v1_rat_1(esk1_0)|~v4_membered(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_767, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_768, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_2020, plain, (v1_xcmplx_0(esk9_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1011, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2685, negated_conjecture, (v1_rat_1(esk2_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_2937, plain, (X3=k2_fdiff_1(X2,esk4_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=esk4_1(k1_zfmisc_1(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,esk4_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2381, plain, (v5_relat_1(k3_taylor_1,k4_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2686, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_2687, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2024, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2984, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_771, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_seq_1(k19_sin_cos,X1)))=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2975, plain, (k1_relset_1(X1,k2_zfmisc_1(X1,X2))=k9_xtuple_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2988, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(esk1_0))=esk1_0|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3023, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_xtuple_0(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_774, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_3030, plain, (k1_real_1(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_777, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2)|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_781, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3026, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2688, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1781, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2862, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_2366, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3111, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_793, plain, (X1=k2_fdiff_1(k19_sin_cos,X2)|k4_xcmplx_0(k1_seq_1(k16_sin_cos,esk3_3(k19_sin_cos,X2,X1)))!=k1_seq_1(X1,esk3_3(k19_sin_cos,X2,X1))|k1_relset_1(k1_numbers,X1)!=X2|~v1_funct_1(X1)|~r2_fdiff_1(k19_sin_cos,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(esk3_3(k19_sin_cos,X2,X1))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk3_3(k19_sin_cos,X2,X1)),k1_numbers))).#trainneg
cnf(i_0_795, plain, (k10_real_1(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X1)),k1_seq_1(k19_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),np__0)|~r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_797, plain, (k10_real_1(k1_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X2),k4_xcmplx_0(k1_seq_1(k16_sin_cos,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k3_taylor_1),X2)|r1_xxreal_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X2)),np__0)|~r1_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X2)|~v1_xreal_0(X2)|~r2_fdiff_1(k19_sin_cos,X1)|~r2_hidden(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_798, plain, (k7_xcmplx_0(k1_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X2),k4_xcmplx_0(k1_seq_1(k16_sin_cos,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k19_sin_cos,X1),k3_taylor_1),X2)|r1_xxreal_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X2)),np__0)|~r1_fdiff_1(k2_fdiff_1(k19_sin_cos,X1),X2)|~v1_xreal_0(k4_xcmplx_0(k1_seq_1(k16_sin_cos,X2)))|~v1_xreal_0(X2)|~r2_fdiff_1(k19_sin_cos,X1)|~r2_hidden(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_800, plain, (r1_xxreal_0(X1,k1_seq_1(k19_sin_cos,X2))|r1_xxreal_0(k1_seq_1(k19_sin_cos,X2),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3117, plain, (k7_xcmplx_0(np__0,X1)=np__0|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2866, negated_conjecture, (v3_ordinal1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_2980, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2438, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3151, plain, (v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_852, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_fdiff_1(k19_sin_cos,X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_854, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_855, plain, (v1_xxreal_0(k7_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_856, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_857, plain, (v1_rat_1(k7_xcmplx_0(X1,X2))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3157, plain, (r1_xxreal_0(X1,X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3116, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_858, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2448, plain, (v5_relat_1(k19_sin_cos,k4_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2983, negated_conjecture, (v3_ordinal1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3150, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3273, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_859, plain, (v7_ordinal1(k7_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_862, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3),k1_numbers)|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_863, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2),k1_numbers)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3309, plain, (r1_xxreal_0(k1_seq_1(X1,X2),np__0)|m1_subset_1(k1_fdiff_1(k3_relat_1(X1,k3_taylor_1),X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_890, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_903, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k4_xcmplx_0(k1_real_1(X1)),k1_numbers)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3246, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2188, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2654, plain, (v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_911, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1)))=k1_fdiff_1(k19_sin_cos,X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_912, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_914, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_915, plain, (v1_xxreal_0(k7_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_3392, plain, (r1_xxreal_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_916, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_917, plain, (v1_rat_1(k7_xcmplx_0(X1,X2))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_918, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_919, plain, (v7_ordinal1(k7_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_923, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2),k1_numbers)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),k4_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3247, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1016, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_3333, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)|~v1_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(k1_real_1(X1)),k1_numbers))).#trainneg
cnf(i_0_2237, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2239, plain, (v1_int_1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3112, plain, (v1_xcmplx_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_937, plain, (v1_int_1(k21_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_938, plain, (v1_int_1(k18_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_940, plain, (v1_int_1(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_944, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_945, plain, (v1_xxreal_0(k7_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3575, plain, (r1_xxreal_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_946, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_947, plain, (v1_rat_1(k7_xcmplx_0(X1,X2))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_948, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_949, plain, (v7_ordinal1(k7_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_953, plain, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2),k1_numbers)|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),k5_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3248, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2557, plain, (v5_relat_1(k16_sin_cos,k3_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3542, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2073, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_967, plain, (v3_ordinal1(k21_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_968, plain, (v3_ordinal1(k18_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_970, plain, (v3_ordinal1(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_982, plain, (r1_xxreal_0(k1_real_1(X1),X2)|r1_xxreal_0(X2,k1_real_1(X1))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3749, plain, (r1_xxreal_0(X1,k1_real_1(k1_fdiff_1(X2,X3)))|r1_xxreal_0(k1_real_1(k1_fdiff_1(X2,X3)),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X3,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3750, negated_conjecture, (r1_xxreal_0(X1,k1_real_1(k1_fdiff_1(X2,X3)))|r1_xxreal_0(k1_real_1(k1_fdiff_1(X2,X3)),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X3,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_983, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2244, plain, (v1_xreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3219, plain, (v1_xreal_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_2647, plain, (v5_relat_1(k3_taylor_1,k3_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3716, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3843, plain, (k7_xcmplx_0(k4_xcmplx_0(esk6_0),X1)=k4_xcmplx_0(k7_xcmplx_0(esk6_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3846, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(esk2_0),X1)=k4_xcmplx_0(k7_xcmplx_0(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3848, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(esk1_0),X1)=k4_xcmplx_0(k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3851, plain, (k4_xcmplx_0(k7_xcmplx_0(np__0,X1))=k7_xcmplx_0(np__0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4035, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2972, plain, (v1_xcmplx_0(esk9_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4089, plain, (m1_subset_1(k7_xcmplx_0(np__1,np__2),k1_numbers))).#trainneg
cnf(i_0_2681, negated_conjecture, (v1_xboole_0(esk2_0)|~m1_subset_1(esk1_0,esk2_0))).#trainneg
cnf(i_0_4178, negated_conjecture, (esk2_0=np__0|~m1_subset_1(esk1_0,esk2_0))).#trainneg
cnf(i_0_2727, plain, (v5_relat_1(k19_sin_cos,k3_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3344, plain, (v1_xcmplx_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3536, plain, (v1_xxreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3549, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2072, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_991, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk19_2(k1_relset_1(k1_numbers,X1),X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_992, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk19_2(k1_relset_1(k1_numbers,X1),X1),k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_4165, plain, (k4_xcmplx_0(k4_xcmplx_0(esk9_0))=esk9_0|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4169, plain, (k7_xcmplx_0(k4_xcmplx_0(esk9_0),X1)=k4_xcmplx_0(k7_xcmplx_0(esk9_0,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1002, plain, (m1_subset_1(esk19_2(k5_numbers,X1),k1_numbers)|r2_fdiff_1(X1,k5_numbers)|~v1_funct_1(X1)|~r1_tarski(k5_numbers,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k5_numbers))).#trainneg
cnf(i_0_1020, plain, (r1_xxreal_0(X1,k1_fdiff_1(k19_sin_cos,X2))|r1_xxreal_0(k1_fdiff_1(k19_sin_cos,X2),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v2_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_1022, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1023, plain, (v1_xxreal_0(k7_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3539, plain, (v1_int_1(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4172, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_3832, plain, (v1_xreal_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1024, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3709, plain, (v5_relat_1(k16_sin_cos,k3_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4182, plain, (v1_xcmplx_0(k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1225, plain, (r1_xxreal_0(k1_sin_cos4(X1),k1_sin_cos4(X1))|~v1_xreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1025, plain, (v1_rat_1(k7_xcmplx_0(X1,X2))|~v4_membered(k1_numbers)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1026, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1027, plain, (v7_ordinal1(k7_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1036, plain, (v1_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_membered(k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1037, plain, (v1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v2_membered(k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1038, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1039, plain, (v1_rat_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v4_membered(k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1040, plain, (v1_int_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v5_membered(k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1041, plain, (v7_ordinal1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v6_membered(k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1200, plain, (v4_valued_0(k16_sin_cos)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4102, plain, (m1_subset_1(k7_xcmplx_0(np__0,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1201, plain, (v4_valued_0(k3_taylor_1)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3118, plain, (k7_xcmplx_0(X1,np__1)=X1|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1042, plain, (r1_xxreal_0(k1_seq_1(X1,X2),np__0)|m1_subset_1(k1_fdiff_1(k3_relat_1(X1,k3_taylor_1),X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1202, plain, (v4_valued_0(k19_sin_cos)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1314, plain, (v5_membered(k16_sin_cos)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1177, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1044, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1052, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1060, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1064, plain, (v1_int_1(k1_real_1(X1))|~v1_int_1(k4_xcmplx_0(k1_real_1(X1)))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1065, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1066, plain, (v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_4457, plain, (r1_xxreal_0(X1,X1)|~v2_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4625, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1635, plain, (v1_valued_0(esk9_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3108, plain, (X4=k2_fdiff_1(X2,k2_fdiff_1(X1,X3))|~v1_funct_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|k1_relset_1(k1_numbers,X4)!=k2_fdiff_1(X1,X3)|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,k2_fdiff_1(X1,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1067, plain, (v1_int_1(X1)|~v1_int_1(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1318, plain, (v5_membered(k3_taylor_1)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1319, plain, (v5_membered(k19_sin_cos)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4643, plain, (v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1075, plain, (v1_xcmplx_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1082, plain, (v1_xxreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1083, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~v1_xboole_0(k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_1092, plain, (k7_xcmplx_0(np__0,esk6_0)=np__0)).#trainneg
cnf(i_0_1093, plain, (k7_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_1094, plain, (k7_xcmplx_0(np__0,k4_xcmplx_0(X1))=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1095, plain, (k7_xcmplx_0(np__0,k4_xcmplx_0(X1))=np__0|~v1_int_1(X1))).#trainneg
cnf(i_0_1096, plain, (k7_xcmplx_0(np__0,k4_xcmplx_0(X1))=np__0|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1097, plain, (k7_xcmplx_0(np__0,k21_sin_cos(X1))=np__0|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1098, plain, (k7_xcmplx_0(np__0,k18_sin_cos(X1))=np__0|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1099, plain, (k7_xcmplx_0(np__0,k1_fdiff_1(k19_sin_cos,X1))=np__0|~v1_xreal_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_4429, plain, (v1_xreal_0(k7_xcmplx_0(np__0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4876, plain, (k7_xcmplx_0(np__0,X1)=np__0|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1100, plain, (k7_xcmplx_0(np__0,k1_real_1(X1))=np__0|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3836, plain, (v5_relat_1(k3_taylor_1,k3_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1101, plain, (k7_xcmplx_0(np__0,k1_seq_1(k19_sin_cos,X1))=np__0|~v1_xreal_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1102, plain, (v1_relat_1(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1103, plain, (v4_relat_1(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_1104, plain, (v1_valued_0(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_1105, plain, (v2_valued_0(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v2_membered(X2))).#trainneg
cnf(i_0_1106, plain, (v3_valued_0(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_1107, plain, (v4_valued_0(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v6_membered(X2))).#trainneg
cnf(i_0_1108, plain, (v5_relat_1(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_1110, plain, (v1_xcmplx_0(esk4_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1111, plain, (v1_xxreal_0(esk4_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1112, plain, (v1_xreal_0(esk4_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1113, plain, (v1_rat_1(esk4_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_1114, plain, (v1_int_1(esk4_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_1115, plain, (v7_ordinal1(esk4_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1119, plain, (v1_rat_1(k19_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3776, plain, (r1_xxreal_0(X1,k1_real_1(k1_sin_cos4(X2)))|r1_xxreal_0(k1_real_1(k1_sin_cos4(X2)),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5217, plain, (r1_xxreal_0(k1_real_1(k1_sin_cos4(X1)),X2)|r1_xxreal_0(X2,k1_real_1(k1_sin_cos4(X1)))|~v1_xxreal_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1116, plain, (m1_subset_1(esk19_2(esk4_1(k1_zfmisc_1(X1)),X2),X1)|r2_fdiff_1(X2,esk4_1(k1_zfmisc_1(X1)))|~v1_funct_1(X2)|~r1_tarski(esk4_1(k1_zfmisc_1(X1)),k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_1(k1_zfmisc_1(X1)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(esk4_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1117, plain, (m1_subset_1(esk3_3(X1,esk4_1(k1_zfmisc_1(X2)),X3),X2)|X3=k2_fdiff_1(X1,esk4_1(k1_zfmisc_1(X2)))|k1_relset_1(k1_numbers,X3)!=esk4_1(k1_zfmisc_1(X2))|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk4_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1118, plain, (k2_fdiff_1(X2,k1_relset_1(k1_numbers,X1))=X1|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_xboole_0(k1_relset_1(k1_numbers,X1))|~r2_fdiff_1(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1120, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1121, plain, (v1_xreal_0(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1124, plain, (v1_int_1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1125, plain, (v7_ordinal1(k19_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1126, plain, (v7_ordinal1(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_1127, plain, (k7_xcmplx_0(esk6_0,np__1)=esk6_0)).#trainneg
cnf(i_0_1128, plain, (k7_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_5382, plain, (m1_subset_1(k4_xcmplx_0(esk6_0),k1_numbers))).#trainneg
cnf(i_0_1129, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1269, plain, (v3_ordinal1(k1_seq_1(k19_sin_cos,X1))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2227, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1130, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k4_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1131, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5563, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,np__1))=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1132, plain, (k7_xcmplx_0(k21_sin_cos(X1),np__1)=k21_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1133, plain, (k7_xcmplx_0(k18_sin_cos(X1),np__1)=k18_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1134, plain, (k7_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1),np__1)=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_1135, plain, (k7_xcmplx_0(k1_real_1(X1),np__1)=k1_real_1(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1137, plain, (k6_xcmplx_0(esk6_0,np__0)=esk6_0)).#trainneg
cnf(i_0_1138, plain, (k6_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_1139, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1140, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1141, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1142, plain, (k6_xcmplx_0(k21_sin_cos(X1),np__0)=k21_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5423, plain, (v1_xreal_0(k4_xcmplx_0(esk6_0)))).#trainneg
cnf(i_0_5806, plain, (k6_xcmplx_0(X1,np__0)=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1143, plain, (k6_xcmplx_0(k18_sin_cos(X1),np__0)=k18_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4180, plain, (v5_relat_1(k19_sin_cos,k3_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2230, plain, (v1_xxreal_0(k1_seq_1(k19_sin_cos,X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2231, plain, (v1_rat_1(k1_seq_1(k19_sin_cos,X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1144, plain, (k6_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1),np__0)=k1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_1145, plain, (k6_xcmplx_0(k1_real_1(X1),np__0)=k1_real_1(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1157, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1162, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1173, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1174, plain, (r1_xxreal_0(k21_sin_cos(X1),k21_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1175, plain, (r1_xxreal_0(k18_sin_cos(X1),k18_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1176, plain, (r1_xxreal_0(k1_fdiff_1(k19_sin_cos,X1),k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_1179, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1181, plain, (v1_xcmplx_0(k3_taylor_1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1182, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_sin_cos4(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1183, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_sin_cos4(X1)))=k1_sin_cos4(X1)|~v1_xreal_0(X1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1184, plain, (k7_xcmplx_0(np__0,k1_sin_cos4(X1))=np__0|~v1_xreal_0(X1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1185, plain, (k7_xcmplx_0(k1_sin_cos4(X1),np__1)=k1_sin_cos4(X1)|~v1_xreal_0(X1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1186, plain, (k6_xcmplx_0(k1_sin_cos4(X1),np__0)=k1_sin_cos4(X1)|~v1_xreal_0(X1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1187, plain, (v1_xxreal_0(k3_taylor_1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1188, plain, (v1_xreal_0(k3_taylor_1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1189, plain, (esk10_2(X1,X2)=np__0)).#trainneg
cnf(i_0_1195, plain, (v1_int_1(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1196, plain, (v1_int_1(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1197, negated_conjecture, (v1_int_1(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1199, plain, (v7_ordinal1(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1203, plain, (v1_xreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1204, plain, (v1_xreal_0(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1205, plain, (v1_xreal_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1207, plain, (v1_rat_1(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1209, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_1210, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_1211, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_198, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_197, plain, (v4_relat_1(np__0,X1))).#trainneg
cnf(i_0_1883, plain, (r2_fdiff_1(X1,np__0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_zfmisc_1(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(np__0,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(np__0))).#trainneg
cnf(i_0_1212, plain, (v1_xreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1213, plain, (v1_xreal_0(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1214, plain, (v1_xreal_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1215, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1216, plain, (v1_xreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1217, plain, (v1_xreal_0(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1218, plain, (v1_xreal_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1219, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1220, plain, (v1_rat_1(k3_taylor_1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1221, plain, (v1_rat_1(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1223, negated_conjecture, (m1_subset_1(k7_xcmplx_0(X1,esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1224, plain, (r1_xxreal_0(X1,k1_sin_cos4(X2))|r1_xxreal_0(k1_sin_cos4(X2),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1226, plain, (v1_int_1(k3_taylor_1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1227, plain, (v1_int_1(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1228, plain, (v7_ordinal1(k3_taylor_1)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_196, plain, (v5_relat_1(np__0,X2))).#trainneg
cnf(i_0_1643, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_200, plain, (m1_subset_1(np__0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1229, plain, (v7_ordinal1(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_2770, plain, (r1_xxreal_0(esk6_0,np__2)|r1_xxreal_0(np__2,esk6_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1329, plain, (v6_ordinal1(k16_sin_cos)|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1333, plain, (v6_ordinal1(k3_taylor_1)|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6036, plain, (r1_xxreal_0(k1_sin_cos4(X1),X2)|r1_xxreal_0(X2,k1_sin_cos4(X1))|~v1_xreal_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_6148, plain, (r1_xxreal_0(X1,k1_sin_cos4(X2))|r1_xxreal_0(k1_sin_cos4(X2),X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1230, plain, (v7_ordinal1(k3_taylor_1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1237, plain, (v3_ordinal1(esk4_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1238, plain, (v3_ordinal1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v3_ordinal1(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1239, plain, (v3_ordinal1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v3_ordinal1(X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1240, negated_conjecture, (v3_ordinal1(esk1_0)|~v3_ordinal1(k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_1241, negated_conjecture, (v3_ordinal1(esk1_0)|~v3_ordinal1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_1249, plain, (v3_ordinal1(k4_xcmplx_0(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1250, plain, (v3_ordinal1(k21_sin_cos(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1251, plain, (v3_ordinal1(k18_sin_cos(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1252, plain, (v3_ordinal1(k10_real_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1253, plain, (v3_ordinal1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v3_ordinal1(k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(k1_seq_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1254, plain, (v3_ordinal1(k1_fdiff_1(k19_sin_cos,X1))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_1255, plain, (v3_ordinal1(k1_fdiff_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1256, plain, (v3_ordinal1(k1_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6071, plain, (v1_valued_0(np__0)|~v1_membered(X1))).#trainneg
cnf(i_0_6199, plain, (v1_valued_0(np__0))).#trainneg
cnf(i_0_5579, plain, (k7_xcmplx_0(X1,np__1)=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1257, plain, (v3_ordinal1(k7_xcmplx_0(X1,esk6_0))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6085, plain, (r1_xxreal_0(np__2,esk6_0)|r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_1334, plain, (v6_ordinal1(k19_sin_cos)|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1344, plain, (v6_membered(k16_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1258, plain, (v3_ordinal1(k7_xcmplx_0(X1,X2))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1259, plain, (v3_ordinal1(k7_xcmplx_0(X1,X2))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_numbers))).#trainneg
cnf(i_0_1260, plain, (v3_ordinal1(k7_xcmplx_0(X1,X2))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1261, plain, (v3_ordinal1(k7_xcmplx_0(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1262, plain, (v3_ordinal1(k20_sin_cos(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1263, plain, (v3_ordinal1(k16_sin_cos)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1264, plain, (v3_ordinal1(k17_sin_cos(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1268, plain, (v3_ordinal1(k2_fdiff_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1270, plain, (v3_ordinal1(k1_seq_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1271, plain, (v3_ordinal1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1272, plain, (v3_ordinal1(k3_taylor_1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1273, plain, (v3_ordinal1(k19_sin_cos)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6067, plain, (v1_relat_1(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_1274, negated_conjecture, (v3_ordinal1(X1)|~v3_ordinal1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1348, plain, (v6_membered(k3_taylor_1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1275, negated_conjecture, (v3_ordinal1(X1)|~v3_ordinal1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1276, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1278, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1279, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k4_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1280, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1281, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_1282, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(np__2),k4_numbers))).#trainneg
cnf(i_0_1283, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_1284, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers)|~v1_xboole_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1290, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1292, plain, (v1_xcmplx_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1294, plain, (v1_xxreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1295, plain, (v1_xreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6025, plain, (k1_relset_1(X1,np__0)=k9_xtuple_0(np__0))).#trainneg
cnf(i_0_1651, plain, (v2_valued_0(esk9_0)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1778, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_relset_1(X1,X3))|~v4_relat_1(X3,X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1296, plain, (v1_xreal_0(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2772, plain, (r1_xxreal_0(esk6_0,np__1)|r1_xxreal_0(np__1,esk6_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1349, plain, (v6_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1359, plain, (v1_membered(k16_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1297, plain, (v1_xreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1298, plain, (v1_xreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_6553, plain, (X4=k2_fdiff_1(X3,k1_relset_1(X2,X1))|~v4_relat_1(X1,X2)|~v1_xboole_0(X2)|~v1_relat_1(X1)|k1_relset_1(k1_numbers,X4)!=k1_relset_1(X2,X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X3,k1_relset_1(X2,X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1301, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk4_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk4_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1303, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k1_relset_1(X2,X3)))).#trainneg
cnf(i_0_1304, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1308, plain, (v5_membered(esk4_1(k1_zfmisc_1(X1)))|~v5_membered(X1))).#trainneg
cnf(i_0_1309, plain, (v5_membered(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v5_membered(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1310, negated_conjecture, (v5_membered(esk1_0)|~v5_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1311, negated_conjecture, (v5_membered(esk1_0)|~v5_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1313, plain, (v5_membered(k1_relset_1(X1,X2))|~v5_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1316, plain, (v5_membered(k2_fdiff_1(X1,X2))|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1317, plain, (v5_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v5_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1320, plain, (v5_membered(X1)|~v5_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6545, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_6617, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X3,np__0)))).#trainneg
cnf(i_0_1659, plain, (v2_valued_0(k5_numbers)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_6565, plain, (X1=k2_fdiff_1(X2,k9_xtuple_0(np__0))|k1_relset_1(k1_numbers,X1)!=k9_xtuple_0(np__0)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_xboole_0(X3)|~r2_fdiff_1(X2,k9_xtuple_0(np__0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1323, plain, (v6_ordinal1(esk4_1(k1_zfmisc_1(X1)))|~v6_ordinal1(X1))).#trainneg
cnf(i_0_6558, plain, (r1_xxreal_0(np__1,esk6_0)|r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_1363, plain, (v1_membered(k3_taylor_1)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1364, plain, (v1_membered(k19_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1324, plain, (v6_ordinal1(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v6_ordinal1(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1325, negated_conjecture, (v6_ordinal1(esk1_0)|~v6_ordinal1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1326, negated_conjecture, (v6_ordinal1(esk1_0)|~v6_ordinal1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1328, plain, (v6_ordinal1(k1_relset_1(X1,X2))|~v6_ordinal1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1331, plain, (v6_ordinal1(k2_fdiff_1(X1,X2))|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1332, plain, (v6_ordinal1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v6_ordinal1(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1335, plain, (v6_ordinal1(X1)|~v6_ordinal1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1338, plain, (v6_membered(esk4_1(k1_zfmisc_1(X1)))|~v6_membered(X1))).#trainneg
cnf(i_0_1339, plain, (v6_membered(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v6_membered(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1340, negated_conjecture, (v6_membered(esk1_0)|~v6_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1341, negated_conjecture, (v6_membered(esk1_0)|~v6_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1343, plain, (v6_membered(k1_relset_1(X1,X2))|~v6_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1346, plain, (v6_membered(k2_fdiff_1(X1,X2))|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1347, plain, (v6_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v6_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6013, plain, (v5_relat_1(esk16_1(k5_numbers),k4_numbers))).#trainneg
cnf(i_0_1667, plain, (v3_valued_0(esk9_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_6549, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_1350, plain, (v6_membered(X1)|~v6_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1374, plain, (v2_membered(k16_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1378, plain, (v2_membered(k3_taylor_1)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1353, plain, (v1_membered(esk4_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_1354, plain, (v1_membered(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v1_membered(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1355, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1356, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1358, plain, (v1_membered(k1_relset_1(X1,X2))|~v1_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1361, plain, (v1_membered(k2_fdiff_1(X1,X2))|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1362, plain, (v1_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1365, plain, (v1_membered(X1)|~v1_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1368, plain, (v2_membered(esk4_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1369, plain, (v2_membered(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v2_membered(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1370, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1371, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1373, plain, (v2_membered(k1_relset_1(X1,X2))|~v2_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1376, plain, (v2_membered(k2_fdiff_1(X1,X2))|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6072, plain, (v2_valued_0(np__0)|~v2_membered(X1))).#trainneg
cnf(i_0_6690, plain, (v2_valued_0(np__0))).#trainneg
cnf(i_0_1675, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1377, plain, (v2_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2774, plain, (r1_xxreal_0(esk6_0,np__0)|r1_xxreal_0(np__0,esk6_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1379, plain, (v2_membered(k19_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1410, plain, (v3_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1380, plain, (v2_membered(X1)|~v2_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1382, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,np__0))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1383, plain, (v1_xxreal_0(k7_xcmplx_0(X1,np__0))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1385, plain, (v1_rat_1(k7_xcmplx_0(X1,np__0))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1386, plain, (v1_int_1(k7_xcmplx_0(X1,np__0))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1387, plain, (v7_ordinal1(k7_xcmplx_0(X1,np__0))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1388, plain, (v3_ordinal1(k7_xcmplx_0(X1,np__0))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1389, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1391, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_1395, plain, (v1_rat_1(k16_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1396, plain, (v1_rat_1(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_6073, plain, (v3_valued_0(np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_6812, plain, (v3_valued_0(np__0))).#trainneg
cnf(i_0_6657, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_relset_1(X3,np__0)))).#trainneg
cnf(i_0_1398, plain, (v1_int_1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6699, plain, (r1_xxreal_0(np__0,esk6_0)|r1_xxreal_0(esk6_0,np__0))).#trainneg
cnf(i_0_1399, plain, (v1_int_1(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1400, plain, (v7_ordinal1(k16_sin_cos)|~v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1401, plain, (v7_ordinal1(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_1402, plain, (v7_ordinal1(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1404, plain, (v3_membered(esk4_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1405, plain, (v3_membered(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v3_membered(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1409, plain, (v3_membered(k1_relset_1(X1,X2))|~v3_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1412, plain, (v3_membered(k2_fdiff_1(X1,X2))|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1413, plain, (v3_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1414, plain, (v3_membered(k3_taylor_1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1415, plain, (v3_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6074, plain, (v4_valued_0(np__0)|~v6_membered(X1))).#trainneg
cnf(i_0_6848, plain, (v4_valued_0(np__0))).#trainneg
cnf(i_0_1419, plain, (v4_membered(esk4_1(k1_zfmisc_1(X1)))|~v4_membered(X1))).#trainneg
cnf(i_0_1425, plain, (v4_membered(k16_sin_cos)|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1429, plain, (v4_membered(k3_taylor_1)|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6137, plain, (r1_xxreal_0(X1,k1_sin_cos4(k1_sin_cos4(X2)))|r1_xxreal_0(k1_sin_cos4(k1_sin_cos4(X2)),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1420, plain, (v4_membered(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v4_membered(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1421, negated_conjecture, (v4_membered(esk1_0)|~v4_membered(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1422, negated_conjecture, (v4_membered(esk1_0)|~v4_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1424, plain, (v4_membered(k1_relset_1(X1,X2))|~v4_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1427, plain, (v4_membered(k2_fdiff_1(X1,X2))|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1428, plain, (v4_membered(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v4_membered(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1430, plain, (v4_membered(k19_sin_cos)|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1431, plain, (v4_membered(X1)|~v4_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1433, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__2))|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_1434, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__2))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_1435, plain, (v1_int_1(k7_xcmplx_0(np__1,np__2))|~v1_int_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_1436, plain, (k1_real_1(k7_xcmplx_0(np__1,np__2))=k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_6019, plain, (v5_relat_1(esk16_1(k5_numbers),k3_numbers))).#trainneg
cnf(i_0_1683, plain, (v4_valued_0(esk9_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_6566, plain, (X1=k2_fdiff_1(X2,k9_xtuple_0(k2_zfmisc_1(X3,X4)))|k1_relset_1(k1_numbers,X1)!=k9_xtuple_0(k2_zfmisc_1(X3,X4))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_xboole_0(X3)|~r2_fdiff_1(X2,k9_xtuple_0(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1439, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__2))|~v2_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_1690, plain, (v4_valued_0(k16_sin_cos)|~v4_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6930, plain, (r1_xxreal_0(k1_sin_cos4(k1_sin_cos4(X1)),X2)|r1_xxreal_0(X2,k1_sin_cos4(k1_sin_cos4(X1)))|~v1_xxreal_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1440, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__2))|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_1441, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_7024, plain, (k7_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),X1)=k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(np__1,np__2),X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7055, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__0)),k1_numbers))).#trainneg
cnf(i_0_1445, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1466, plain, (X3=k2_fdiff_1(X1,X2)|~r2_hidden(X2,esk3_3(X1,X2,X3))|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1467, plain, (v1_xboole_0(X2)|~r2_hidden(X2,X1)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_1468, plain, (r2_fdiff_1(X1,X2)|v1_xboole_0(esk19_2(X2,X1))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,esk19_2(X2,X1)))).#trainneg
cnf(i_0_1472, plain, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1480, plain, (v1_xcmplx_0(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6527, plain, (v1_relat_1(k1_relset_1(X1,np__0)))).#trainneg
cnf(i_0_777, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_6550, plain, (~v1_xboole_0(X1)|~r2_hidden(X3,k9_xtuple_0(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1481, plain, (v1_xxreal_0(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v2_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1691, plain, (v4_valued_0(k5_numbers)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1482, plain, (v1_xreal_0(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1483, plain, (v1_rat_1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v4_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_7196, plain, (k1_real_1(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_1484, plain, (v1_int_1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v5_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1485, plain, (v7_ordinal1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v6_membered(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1486, plain, (v3_ordinal1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v3_ordinal1(k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1487, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_seq_1(X1,X2)))=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_membered(k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1488, plain, (k7_xcmplx_0(np__0,k1_seq_1(X1,X2))=np__0|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_membered(k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1489, plain, (k7_xcmplx_0(k1_seq_1(X1,X2),np__1)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_membered(k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1490, plain, (k6_xcmplx_0(k1_seq_1(X1,X2),np__0)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_membered(k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1491, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X3))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1493, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~v2_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1494, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1495, plain, (v1_xxreal_0(k1_fdiff_1(X1,X3))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_6533, plain, (v5_relat_1(k9_xtuple_0(np__0),X2))).#trainneg
cnf(i_0_1512, plain, (v1_xcmplx_0(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1694, plain, (v4_valued_0(k3_taylor_1)|~v4_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4640, plain, (v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xboole_0(k1_real_1(X1)))).#trainneg
cnf(i_0_1496, plain, (v1_rat_1(k1_fdiff_1(X1,X3))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v4_membered(k1_numbers)|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1497, plain, (v1_relat_1(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1498, plain, (v4_relat_1(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X1)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1499, plain, (v1_valued_0(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1500, plain, (v2_valued_0(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v2_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1501, plain, (v3_valued_0(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1502, plain, (v4_valued_0(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v6_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1503, plain, (v5_relat_1(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X2)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1504, plain, (v5_membered(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v5_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1505, plain, (v6_ordinal1(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v6_ordinal1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1506, plain, (v6_membered(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v6_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1507, plain, (v1_membered(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1508, plain, (v2_membered(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v2_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1509, plain, (v3_membered(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v3_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6534, plain, (v4_relat_1(k9_xtuple_0(np__0),X1))).#trainneg
cnf(i_0_1854, plain, (v1_int_1(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1513, plain, (v1_xxreal_0(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v2_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1695, plain, (v4_valued_0(k19_sin_cos)|~v4_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4609, plain, (v1_int_1(k1_real_1(X1))|~v1_int_1(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1510, plain, (v4_membered(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v4_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1511, plain, (m1_subset_1(X1,X2)|r2_fdiff_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,esk19_2(k1_zfmisc_1(X2),X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1515, plain, (v1_rat_1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v4_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1516, plain, (v1_int_1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v5_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1517, plain, (v7_ordinal1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v6_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1518, plain, (v3_ordinal1(esk19_2(X1,X2))|r2_fdiff_1(X2,X1)|~v3_ordinal1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1519, plain, (r2_fdiff_1(X1,esk19_2(k1_zfmisc_1(X2),X3))|r2_fdiff_1(X3,k1_zfmisc_1(X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(k1_zfmisc_1(X2),X3),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(esk19_2(k1_zfmisc_1(X2),X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1521, plain, (r2_fdiff_1(X1,esk4_1(k1_zfmisc_1(X2)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_1(k1_zfmisc_1(X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(esk4_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1522, plain, (r2_fdiff_1(X1,esk3_3(X2,k1_zfmisc_1(X3),X4))|X4=k2_fdiff_1(X2,k1_zfmisc_1(X3))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_3(X2,k1_zfmisc_1(X3),X4),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(esk3_3(X2,k1_zfmisc_1(X3),X4))|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(X3)|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1526, plain, (r2_fdiff_1(X1,k1_relset_1(X2,X3))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(X2,X3),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1527, plain, (r2_fdiff_1(X1,k16_sin_cos)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k16_sin_cos))).#trainneg
cnf(i_0_1528, plain, (r2_fdiff_1(X1,k5_numbers)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k5_numbers))).#trainneg
cnf(i_0_1529, plain, (r2_fdiff_1(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1857, negated_conjecture, (v1_int_1(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_6535, plain, (m1_subset_1(k9_xtuple_0(np__0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2128, plain, (v3_ordinal1(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1531, plain, (r2_fdiff_1(X1,k3_taylor_1)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k3_taylor_1))).#trainneg
cnf(i_0_6829, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk6_0)|~v2_xxreal_0(esk6_0))).#trainneg
cnf(i_0_2186, negated_conjecture, (v3_ordinal1(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2560, plain, (v1_xreal_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1530, plain, (r2_fdiff_1(X1,k1_partfun1(X2,X3,X4,X5,X6,X7))|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_zfmisc_1(X2,X5))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,X6,X7),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_partfun1(X2,X3,X4,X5,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1532, plain, (r2_fdiff_1(X1,k19_sin_cos)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k19_sin_cos))).#trainneg
cnf(i_0_1552, plain, (k7_xcmplx_0(np__0,np__2)=np__0)).#trainneg
cnf(i_0_1553, plain, (k7_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_1556, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1559, plain, (v1_int_1(k1_fdiff_1(X1,X3))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v5_membered(k1_numbers)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1561, plain, (v1_int_1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v6_membered(k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1562, plain, (v3_ordinal1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v6_membered(k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1563, plain, (v7_ordinal1(k1_fdiff_1(X1,X3))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v6_membered(k1_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1564, plain, (k4_xcmplx_0(k4_xcmplx_0(k10_real_1(X1,X2)))=k10_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1565, plain, (k7_xcmplx_0(np__0,k10_real_1(X1,X2))=np__0|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1566, plain, (k7_xcmplx_0(k10_real_1(X1,X2),np__1)=k10_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7486, plain, (v1_relat_1(k9_xtuple_0(k9_xtuple_0(np__0))))).#trainneg
cnf(i_0_1567, plain, (k6_xcmplx_0(k10_real_1(X1,X2),np__0)=k10_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1568, plain, (v1_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3))|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~v1_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1569, plain, (v1_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1570, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1571, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1572, plain, (v1_xxreal_0(X1)|~v2_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1573, plain, (v1_xreal_0(X1)|~v3_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1574, plain, (v1_rat_1(X1)|~v4_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1575, plain, (v1_int_1(X1)|~v5_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1576, plain, (v7_ordinal1(X1)|~v6_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1577, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1579, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1580, negated_conjecture, (v1_xxreal_0(X1)|~v2_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1581, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1582, negated_conjecture, (v1_rat_1(X1)|~v4_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5836, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7218, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_relset_1(X1,k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_1583, negated_conjecture, (v1_int_1(X1)|~v5_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1584, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1585, negated_conjecture, (v3_ordinal1(X1)|~v3_ordinal1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1591, plain, (v1_xboole_0(esk9_0)|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(k1_relset_1(k1_numbers,X1),X1),esk9_0)|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_1592, plain, (v1_xboole_0(esk9_0)|r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(X2,X1),esk9_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1594, negated_conjecture, (v1_xboole_0(esk1_0)|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(k1_relset_1(k1_numbers,X1),X1),esk1_0)|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_7967, negated_conjecture, (r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(esk19_2(X2,k2_fdiff_1(X1,X2)),esk1_0)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1595, negated_conjecture, (v1_xboole_0(esk1_0)|r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(X2,X1),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1599, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xxreal_0(X1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1601, plain, (r1_xxreal_0(X1,k10_real_1(X2,X3))|r1_xxreal_0(k10_real_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1602, plain, (r1_xxreal_0(k10_real_1(X1,X2),k10_real_1(X1,X2))|~v1_xreal_0(X2)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1603, plain, (v1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3))|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~v2_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1604, plain, (v1_xxreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_xreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1605, plain, (v1_xxreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1608, plain, (m1_subset_1(k7_xcmplx_0(X1,k10_real_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1609, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3))|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1610, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1611, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1612, plain, (v1_rat_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3))|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v4_membered(k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1613, plain, (v1_rat_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v4_membered(k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1614, plain, (v1_rat_1(k7_xcmplx_0(X1,X2))|~v4_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1623, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__2))|~v2_xxreal_0(np__0)|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_1628, plain, (v2_xxreal_0(np__0)|~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1629, plain, (v2_xxreal_0(np__0)|~v2_xxreal_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))|~v1_xreal_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_1633, plain, (v2_xxreal_0(np__0)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v2_xxreal_0(k1_seq_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))))|~v1_xreal_0(k1_seq_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))))|~r1_fdiff_1(X1,esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))|~v1_funct_1(X1)|~v1_xreal_0(esk19_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_1634, plain, (v1_valued_0(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1636, plain, (v1_valued_0(esk4_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1638, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1637, plain, (v1_valued_0(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v1_valued_0(X2)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1639, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1641, plain, (v1_valued_0(k1_relset_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1644, plain, (v1_valued_0(k2_fdiff_1(X1,X2))|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1645, plain, (v1_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_valued_0(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1648, plain, (v1_valued_0(X1)|~v1_valued_0(k1_relset_1(k1_numbers,X2))|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1650, plain, (v2_valued_0(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1652, plain, (v2_valued_0(esk4_1(k1_zfmisc_1(X1)))|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1653, plain, (v2_valued_0(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v2_valued_0(X2)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1654, negated_conjecture, (v2_valued_0(esk1_0)|~v2_valued_0(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1655, negated_conjecture, (v2_valued_0(esk1_0)|~v2_valued_0(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1657, plain, (v2_valued_0(k1_relset_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1660, plain, (v2_valued_0(k2_fdiff_1(X1,X2))|~v2_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1661, plain, (v2_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v2_valued_0(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1664, plain, (v2_valued_0(X1)|~v2_valued_0(k1_relset_1(k1_numbers,X2))|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_7407, plain, (v1_int_1(k1_real_1(X1))|~v1_xboole_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1666, plain, (v3_valued_0(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1668, plain, (v3_valued_0(esk4_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1669, plain, (v3_valued_0(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v3_valued_0(X2)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1670, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1671, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1673, plain, (v3_valued_0(k1_relset_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1677, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_valued_0(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1680, plain, (v3_valued_0(X1)|~v3_valued_0(k1_relset_1(k1_numbers,X2))|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1682, plain, (v4_valued_0(esk19_2(k1_zfmisc_1(X1),X2))|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1684, plain, (v4_valued_0(esk4_1(k1_zfmisc_1(X1)))|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1685, plain, (v4_valued_0(esk3_3(X1,k1_zfmisc_1(X2),X3))|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v4_valued_0(X2)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1686, negated_conjecture, (v4_valued_0(esk1_0)|~v4_valued_0(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1687, negated_conjecture, (v4_valued_0(esk1_0)|~v4_valued_0(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_7000, plain, (X1=k2_fdiff_1(X2,k1_relset_1(X3,k2_zfmisc_1(X3,X4)))|k1_relset_1(k1_numbers,X1)!=k1_relset_1(X3,k2_zfmisc_1(X3,X4))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_xboole_0(X3)|~r2_fdiff_1(X2,k1_relset_1(X3,k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1689, plain, (v4_valued_0(k1_relset_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2563, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1692, plain, (v4_valued_0(k2_fdiff_1(X1,X2))|~v4_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1693, plain, (v4_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v4_valued_0(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1696, plain, (v4_valued_0(X1)|~v4_valued_0(k1_relset_1(k1_numbers,X2))|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1699, plain, (v1_int_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3))|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~v5_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1700, plain, (v1_int_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_xreal_0(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1701, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1702, plain, (v1_int_1(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1703, plain, (v3_ordinal1(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1704, plain, (v7_ordinal1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k3_taylor_1),X3))|r1_xxreal_0(k1_fdiff_1(X1,X3),np__0)|~v1_xreal_0(k1_fdiff_1(X1,X3))|~v6_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1705, plain, (v7_ordinal1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|r1_xxreal_0(k1_seq_1(X1,X2),np__0)|~v1_xreal_0(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1706, plain, (v7_ordinal1(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1707, plain, (m1_subset_1(esk19_2(k16_sin_cos,X1),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X1,k16_sin_cos)|~v1_funct_1(X1)|~r1_tarski(k16_sin_cos,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k16_sin_cos))).#trainneg
cnf(i_0_1708, plain, (m1_subset_1(esk3_3(X1,k16_sin_cos,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|X2=k2_fdiff_1(X1,k16_sin_cos)|k1_relset_1(k1_numbers,X2)!=k16_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k16_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1709, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xboole_0(k16_sin_cos)|~m1_subset_1(X1,k16_sin_cos))).#trainneg
cnf(i_0_2863, negated_conjecture, (v1_int_1(esk2_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1710, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_funct_1(X1,X2)))=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1711, plain, (k7_xcmplx_0(np__0,k1_funct_1(X1,X2))=np__0|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1712, plain, (k7_xcmplx_0(k1_funct_1(X1,X2),np__1)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1713, plain, (k6_xcmplx_0(k1_funct_1(X1,X2),np__0)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1714, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1715, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~v2_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1716, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1717, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1718, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_funct_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1720, plain, (v3_valued_0(k3_relat_1(X1,X2))|~v4_valued_0(X2)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1721, plain, (v5_relat_1(k3_relat_1(X1,X2),k4_numbers)|~v4_valued_0(X2)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1722, plain, (v5_relat_1(k3_relat_1(X1,X2),k3_numbers)|~v4_valued_0(X2)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1723, plain, (v1_int_1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1724, plain, (v3_ordinal1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2802, plain, (v1_int_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_2867, plain, (v3_ordinal1(esk9_0)|~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_1725, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_2973, plain, (v1_xcmplx_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2981, negated_conjecture, (v1_int_1(esk2_0)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_8684, plain, (v1_int_1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8685, plain, (v3_ordinal1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1726, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(np__1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1727, plain, (m1_subset_1(esk19_2(k3_taylor_1,X1),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X1,k3_taylor_1)|~v1_funct_1(X1)|~r1_tarski(k3_taylor_1,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k3_taylor_1))).#trainneg
cnf(i_0_8695, plain, (k7_xcmplx_0(k4_xcmplx_0(esk9_0),X1)=k4_xcmplx_0(k7_xcmplx_0(esk9_0,X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1728, plain, (m1_subset_1(esk3_3(X1,k3_taylor_1,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|X2=k2_fdiff_1(X1,k3_taylor_1)|k1_relset_1(k1_numbers,X2)!=k3_taylor_1|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1729, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xboole_0(k3_taylor_1)|~m1_subset_1(X1,k3_taylor_1))).#trainneg
cnf(i_0_1733, negated_conjecture, (v5_relat_1(esk1_0,k4_numbers)|~v1_relat_1(esk1_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1734, negated_conjecture, (v5_relat_1(esk1_0,k3_numbers)|~v1_relat_1(esk1_0)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1735, plain, (m1_subset_1(esk19_2(k19_sin_cos,X1),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X1,k19_sin_cos)|~v1_funct_1(X1)|~r1_tarski(k19_sin_cos,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k19_sin_cos))).#trainneg
cnf(i_0_1736, plain, (m1_subset_1(esk3_3(X1,k19_sin_cos,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|X2=k2_fdiff_1(X1,k19_sin_cos)|k1_relset_1(k1_numbers,X2)!=k19_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k19_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1737, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xboole_0(k19_sin_cos)|~m1_subset_1(X1,k19_sin_cos))).#trainneg
cnf(i_0_1738, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_relset_1(X1,X2)))=k1_relset_1(X1,X2)|~v4_relat_1(X2,X1)|~v1_membered(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1739, plain, (k7_xcmplx_0(np__0,k1_relset_1(X1,X2))=np__0|~v4_relat_1(X2,X1)|~v1_membered(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1740, plain, (k7_xcmplx_0(k1_relset_1(X1,X2),np__1)=k1_relset_1(X1,X2)|~v4_relat_1(X2,X1)|~v1_membered(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1741, plain, (k6_xcmplx_0(k1_relset_1(X1,X2),np__0)=k1_relset_1(X1,X2)|~v4_relat_1(X2,X1)|~v1_membered(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_2853, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_6017, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k4_numbers))).#trainneg
cnf(i_0_2985, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_5300, plain, (X1=k2_fdiff_1(X2,X4)|k1_relset_1(k1_numbers,X1)!=X4|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_fdiff_1(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X4)|~v1_xboole_0(k1_relset_1(k1_numbers,X4))|~r2_fdiff_1(X3,k1_relset_1(k1_numbers,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8825, negated_conjecture, (v5_relat_1(esk1_0,k4_numbers)|~v1_relat_1(esk1_0)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3540, plain, (v1_int_1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3833, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_8700, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1745, plain, (v1_int_1(k6_xcmplx_0(X2,X1))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1750, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8996, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(esk1_0),X1)=k4_xcmplx_0(k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_9067, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9069, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9181, plain, (k6_xcmplx_0(np__0,esk6_0)=k4_xcmplx_0(esk6_0))).#trainneg
cnf(i_0_9184, plain, (k6_xcmplx_0(np__0,esk3_3(X1,X2,X3))=k4_xcmplx_0(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9185, negated_conjecture, (k6_xcmplx_0(np__0,esk2_0)=k4_xcmplx_0(esk2_0))).#trainneg
cnf(i_0_1751, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__2)=k6_xcmplx_0(k4_xcmplx_0(np__2),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1752, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1753, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k7_xcmplx_0(np__1,np__2))=k6_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),X1)|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1754, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_real_1(X2))=k6_xcmplx_0(k4_xcmplx_0(k1_real_1(X2)),X1)|~v1_xcmplx_0(k4_xcmplx_0(k1_real_1(X2)))|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9233, plain, (v1_int_1(k4_xcmplx_0(esk6_0)))).#trainneg
cnf(i_0_4183, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1755, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8827, negated_conjecture, (v5_relat_1(esk1_0,k3_numbers)|~v1_relat_1(esk1_0)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2074, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1756, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k1_seq_1(k16_sin_cos,X2))=k6_xcmplx_0(k1_fdiff_1(k19_sin_cos,X2),X1)|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_1757, plain, (k6_xcmplx_0(k1_real_1(X2),X1)=k6_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k1_real_1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1758, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1759, plain, (k6_xcmplx_0(np__2,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1760, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1761, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_1762, plain, (k6_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(k1_real_1(X1)))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1763, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3348, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_6023, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k3_numbers))).#trainneg
cnf(i_0_9294, negated_conjecture, (r1_xxreal_0(k1_seq_1(k19_sin_cos,esk2_0),np__0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_1764, plain, (k6_xcmplx_0(k1_seq_1(k16_sin_cos,X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k1_fdiff_1(k19_sin_cos,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_4450, plain, (v5_membered(k16_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_3270, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6145, plain, (r1_xxreal_0(X1,k1_sin_cos4(X2))|r1_xxreal_0(k1_sin_cos4(X2),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1765, plain, (k6_xcmplx_0(X2,k1_real_1(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9675, plain, (r1_xxreal_0(k1_sin_cos4(k1_fdiff_1(X1,X2)),X3)|r1_xxreal_0(X3,k1_sin_cos4(k1_fdiff_1(X1,X2)))|~v1_xxreal_0(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~r2_hidden(X2,esk9_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9676, negated_conjecture, (r1_xxreal_0(k1_sin_cos4(k1_fdiff_1(X1,X2)),X3)|r1_xxreal_0(X3,k1_sin_cos4(k1_fdiff_1(X1,X2)))|~v1_xxreal_0(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_numbers)|~r2_hidden(X2,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1771, plain, (r2_fdiff_1(X3,k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~r2_hidden(X2,esk19_2(k1_zfmisc_1(X1),X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1774, plain, (X4=k2_fdiff_1(X3,k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~r2_hidden(X2,esk3_3(X3,k1_zfmisc_1(X1),X4))|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1775, negated_conjecture, (~v1_xboole_0(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1776, negated_conjecture, (~v1_xboole_0(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1782, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,k1_partfun1(X1,X4,X5,X2,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X4))))).#trainneg
cnf(i_0_1785, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,X1))|~r2_hidden(X2,X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3))).#trainneg
cnf(i_0_1788, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__2))|~v2_xxreal_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))|~v1_xreal_0(k7_xcmplx_0(np__1,np__2))|~v1_xreal_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))))).#trainneg
cnf(i_0_1789, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),esk6_0)=k4_xcmplx_0(k7_xcmplx_0(X1,esk6_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1790, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__2)=k4_xcmplx_0(k7_xcmplx_0(X1,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1791, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k4_xcmplx_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3543, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_7386, plain, (v5_relat_1(k1_relset_1(X2,np__0),X1))).#trainneg
cnf(i_0_9768, plain, (~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xboole_0(X2)|~r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~r2_hidden(X4,X3)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1792, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(k7_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4644, plain, (v5_membered(k3_taylor_1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4646, plain, (v5_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_3154, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1793, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k1_funct_1(X2,X3))=k4_xcmplx_0(k7_xcmplx_0(X1,k1_funct_1(X2,X3)))|~v1_xcmplx_0(X1)|~v1_funct_1(X2)|~v1_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1794, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1795, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_1796, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1797, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k21_sin_cos(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k21_sin_cos(X2)))|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1798, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k18_sin_cos(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k18_sin_cos(X2)))|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1799, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k10_real_1(X2,X3))=k4_xcmplx_0(k7_xcmplx_0(X1,k10_real_1(X2,X3)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1800, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k1_fdiff_1(k19_sin_cos,X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k1_fdiff_1(k19_sin_cos,X2)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X2),k1_numbers))).#trainneg
cnf(i_0_1801, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k1_relset_1(X2,X3))=k4_xcmplx_0(k7_xcmplx_0(X1,k1_relset_1(X2,X3)))|~v1_xcmplx_0(X1)|~v4_relat_1(X3,X2)|~v1_membered(k1_zfmisc_1(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1802, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k1_real_1(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k1_real_1(X2)))|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1803, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k1_sin_cos4(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k1_sin_cos4(X2)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1806, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(np__2)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1807, plain, (r1_xxreal_0(X1,k1_relset_1(X2,X3))|r1_xxreal_0(k1_relset_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v4_relat_1(X3,X2)|~v2_membered(k1_zfmisc_1(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_7423, plain, (v4_relat_1(k1_relset_1(X2,np__0),X1))).#trainneg
cnf(i_0_9946, negated_conjecture, (~v1_funct_1(X1)|~v1_xboole_0(X2)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1808, plain, (r1_xxreal_0(k1_relset_1(X1,X2),k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v2_membered(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5206, plain, (r1_xxreal_0(k1_real_1(k1_sin_cos4(k1_sin_cos4(X1))),X2)|r1_xxreal_0(X2,k1_real_1(k1_sin_cos4(k1_sin_cos4(X1))))|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6919, plain, (r1_xxreal_0(k1_sin_cos4(k1_sin_cos4(k1_sin_cos4(X1))),X2)|r1_xxreal_0(X2,k1_sin_cos4(k1_sin_cos4(k1_sin_cos4(X1))))|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1809, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_relset_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v4_relat_1(X3,X2)|~v3_membered(k1_zfmisc_1(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1810, plain, (v1_rat_1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1811, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1812, plain, (v1_int_1(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1817, plain, (v5_relat_1(esk19_2(k1_zfmisc_1(X1),X2),k4_numbers)|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1819, plain, (v5_relat_1(esk4_1(k1_zfmisc_1(X1)),k4_numbers)|~v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1820, plain, (v5_relat_1(esk3_3(X1,k1_zfmisc_1(X2),X3),k4_numbers)|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v5_relat_1(X2,k4_numbers)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1821, negated_conjecture, (v5_relat_1(esk1_0,k4_numbers)|~v5_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k4_numbers)|~v1_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1822, negated_conjecture, (v5_relat_1(esk1_0,k4_numbers)|~v5_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)),k4_numbers)|~v1_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1825, plain, (v5_relat_1(k16_sin_cos,k4_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k4_numbers))).#trainneg
cnf(i_0_1827, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k4_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k4_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1828, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k4_numbers)|~v5_relat_1(k2_zfmisc_1(X1,X4),k4_numbers)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1829, plain, (v5_relat_1(k3_taylor_1,k4_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k4_numbers))).#trainneg
cnf(i_0_1830, plain, (v5_relat_1(k19_sin_cos,k4_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k4_numbers))).#trainneg
cnf(i_0_7773, plain, (v1_relat_1(k9_xtuple_0(k1_relset_1(X1,np__0))))).#trainneg
cnf(i_0_9945, plain, (~v1_funct_1(X1)|~v1_xboole_0(X2)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,esk9_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1831, plain, (v5_relat_1(X1,k4_numbers)|~v5_relat_1(k1_relset_1(k1_numbers,X2),k4_numbers)|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6086, plain, (v6_ordinal1(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_10003, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10906, plain, (r1_xxreal_0(X1,k1_real_1(k1_sin_cos4(k1_sin_cos4(X2))))|r1_xxreal_0(k1_real_1(k1_sin_cos4(k1_sin_cos4(X2))),X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1833, plain, (v5_relat_1(esk19_2(k1_zfmisc_1(X1),X2),k3_numbers)|r2_fdiff_1(X2,k1_zfmisc_1(X1))|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1835, plain, (v5_relat_1(esk4_1(k1_zfmisc_1(X1)),k3_numbers)|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1836, plain, (v5_relat_1(esk3_3(X1,k1_zfmisc_1(X2),X3),k3_numbers)|X3=k2_fdiff_1(X1,k1_zfmisc_1(X2))|~v5_relat_1(X2,k3_numbers)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X3)!=k1_zfmisc_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1837, negated_conjecture, (v5_relat_1(esk1_0,k3_numbers)|~v5_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1)),k3_numbers)|~v1_relat_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1838, negated_conjecture, (v5_relat_1(esk1_0,k3_numbers)|~v5_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1)),k3_numbers)|~v1_relat_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1841, plain, (v5_relat_1(k16_sin_cos,k3_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k3_numbers))).#trainneg
cnf(i_0_1843, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k3_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k3_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1844, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k3_numbers)|~v5_relat_1(k2_zfmisc_1(X1,X4),k3_numbers)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1845, plain, (v5_relat_1(k3_taylor_1,k3_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k3_numbers))).#trainneg
cnf(i_0_1846, plain, (v5_relat_1(k19_sin_cos,k3_numbers)|~v5_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers),k3_numbers))).#trainneg
cnf(i_0_1847, plain, (v5_relat_1(X1,k3_numbers)|~v5_relat_1(k1_relset_1(k1_numbers,X2),k3_numbers)|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_11062, plain, (v5_relat_1(X1,k3_numbers)|~v1_funct_1(X2)|~v1_relat_1(k1_relset_1(k1_numbers,X2))|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v5_relat_1(k1_relset_1(k1_numbers,X2),k4_numbers))).#trainneg
cnf(i_0_1850, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(X2,np__0)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1852, plain, (v1_xreal_0(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6087, plain, (v6_ordinal1(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_10827, negated_conjecture, (~v1_funct_1(X1)|~v1_xboole_0(k1_relset_1(k1_numbers,X2))|~r2_fdiff_1(X2,k1_numbers)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,X2))|~r2_hidden(X3,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1855, plain, (v1_int_1(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_6248, plain, (v6_ordinal1(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6249, plain, (v6_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_10979, plain, (r1_xxreal_0(X1,k1_sin_cos4(k1_sin_cos4(k1_sin_cos4(X2))))|r1_xxreal_0(k1_sin_cos4(k1_sin_cos4(k1_sin_cos4(X2))),X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1858, negated_conjecture, (v1_int_1(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1860, negated_conjecture, (m1_subset_1(k7_xcmplx_0(X1,esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1861, plain, (v1_int_1(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v6_membered(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1862, plain, (v3_ordinal1(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v6_membered(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1893, plain, (v1_xboole_0(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1895, plain, (v1_xboole_0(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1896, plain, (v1_xboole_0(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X3)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1897, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1900, plain, (v1_xboole_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_xboole_0(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1903, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1904, plain, (v1_xboole_0(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1906, plain, (v1_xboole_0(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1907, plain, (v1_xboole_0(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X2)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1908, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_6430, plain, (v6_membered(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1911, plain, (v1_xboole_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_xboole_0(X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1914, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1915, plain, (v5_relat_1(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),k4_numbers)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v5_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1917, plain, (v5_relat_1(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),k4_numbers)|~v5_membered(X2))).#trainneg
cnf(i_0_1918, plain, (v5_relat_1(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4),k4_numbers)|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v5_membered(X3)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1919, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),k4_numbers)|~v5_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1921, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k4_numbers)|~v5_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1922, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k4_numbers)|~v5_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1926, plain, (v5_relat_1(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),k3_numbers)|r2_fdiff_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v4_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1928, plain, (v5_relat_1(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),k3_numbers)|~v4_membered(X2))).#trainneg
cnf(i_0_1929, plain, (v5_relat_1(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4),k3_numbers)|X4=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v4_membered(X3)|k1_relset_1(k1_numbers,X4)!=k1_zfmisc_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1930, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),k3_numbers)|~v4_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_11034, plain, (~v1_funct_1(X1)|~v1_xboole_0(k1_relset_1(k1_numbers,X2))|~r2_fdiff_1(X2,k1_numbers)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,X2))|~r2_hidden(X3,esk9_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1932, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k3_numbers)|~v4_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6559, plain, (v6_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1933, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k3_numbers)|~v4_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1937, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(esk5_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1942, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(esk5_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1946, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1949, plain, (v1_xcmplx_0(X2)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1950, plain, (v1_xxreal_0(X2)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1951, plain, (v1_xreal_0(X2)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1952, plain, (v1_rat_1(X2)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v4_membered(k1_zfmisc_1(k1_numbers))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1953, plain, (v1_int_1(X2)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v5_membered(k1_zfmisc_1(k1_numbers))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1954, plain, (v7_ordinal1(X2)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1955, plain, (v5_relat_1(X2,k4_numbers)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1956, plain, (v5_relat_1(X2,k3_numbers)|~v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1958, plain, (r1_tarski(X1,X3)|~v1_funct_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1969, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6562, plain, (v1_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6628, plain, (v1_membered(k3_taylor_1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1959, plain, (m1_subset_1(X1,k1_zfmisc_1(X3))|~v1_funct_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1960, plain, (m1_subset_1(X1,X3)|~v1_funct_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(k2_fdiff_1(X2,X3),k1_numbers)|~r2_hidden(X1,esk9_0)|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1961, negated_conjecture, (m1_subset_1(X1,X3)|~v1_funct_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(k2_fdiff_1(X2,X3),k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1962, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk19_2(X3,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1963, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk19_2(X3,k2_fdiff_1(X1,X2)),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_11505, negated_conjecture, (r2_fdiff_1(k2_fdiff_1(X1,esk1_0),esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(k2_fdiff_1(X1,esk1_0),k1_numbers)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1964, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xboole_0(X3)|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1966, plain, (v1_xboole_0(esk9_0)|r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),k1_numbers)|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(X3,k2_fdiff_1(X1,X2)),esk9_0)|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1967, negated_conjecture, (v1_xboole_0(esk1_0)|r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),k1_numbers)|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk19_2(X3,k2_fdiff_1(X1,X2)),esk1_0)|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1965, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X3)|~v1_xboole_0(X2)|~r2_fdiff_1(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_7518, plain, (m1_subset_1(k1_relset_1(X2,np__0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1970, plain, (v1_xxreal_0(k7_xcmplx_0(X1,np__1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1972, plain, (v1_rat_1(k7_xcmplx_0(X1,np__1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1973, plain, (v1_int_1(k7_xcmplx_0(X1,np__1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1974, plain, (v7_ordinal1(k7_xcmplx_0(X1,np__1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1975, plain, (v3_ordinal1(k7_xcmplx_0(X1,np__1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1977, plain, (v1_relat_1(k9_xtuple_0(esk19_2(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),X4)))|r2_fdiff_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),k1_zfmisc_1(k1_relset_1(k1_numbers,X4)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_1979, plain, (v1_relat_1(k9_xtuple_0(esk4_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))))).#trainneg
cnf(i_0_1980, plain, (v1_relat_1(k9_xtuple_0(esk3_3(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4)),X5)))|X5=k2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4)))|k1_relset_1(k1_numbers,X5)!=k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))|~v1_funct_1(X5)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1981, plain, (v1_relat_1(k9_xtuple_0(k1_relset_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)))|~v4_relat_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))|~v1_relat_1(X4))).#trainneg
cnf(i_0_1982, plain, (v1_relat_1(k9_xtuple_0(k1_partfun1(k2_zfmisc_1(X1,X2),X3,X4,X5,X6,X7)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_1983, plain, (v1_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_1984, plain, (v3_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v6_membered(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1985, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),k4_numbers)|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v6_membered(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1986, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),k3_numbers)|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v6_membered(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1987, plain, (k1_relset_1(X1,k1_relset_1(k2_zfmisc_1(X1,X2),X3))=k9_xtuple_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1988, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k1_fdiff_1(X2,X3))=k4_xcmplx_0(k7_xcmplx_0(X1,k1_fdiff_1(X2,X3)))|~v1_xcmplx_0(X1)|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1989, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_fdiff_1(X1,X2)))=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1990, plain, (k7_xcmplx_0(np__0,k1_fdiff_1(X1,X2))=np__0|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1991, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),np__1)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1992, plain, (k6_xcmplx_0(k1_fdiff_1(X1,X2),np__0)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1994, plain, (r1_xxreal_0(X1,k1_fdiff_1(X2,X3))|r1_xxreal_0(k1_fdiff_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1995, plain, (r1_xxreal_0(k1_fdiff_1(X1,X2),k1_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1997, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_fdiff_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2002, plain, (v1_int_1(k1_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2003, plain, (v3_ordinal1(k1_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2005, plain, (v1_xboole_0(esk9_0)|m1_subset_1(esk19_2(esk9_0,X1),k1_numbers)|r2_fdiff_1(X1,esk9_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk9_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_2006, plain, (v1_xboole_0(esk9_0)|m1_subset_1(esk4_1(esk9_0),k1_numbers))).#trainneg
cnf(i_0_2008, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(X2,np__1)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2010, plain, (v3_ordinal1(k1_relset_1(k1_numbers,k16_sin_cos))|~v5_ordinal1(k16_sin_cos))).#trainneg
cnf(i_0_2441, plain, (r1_xxreal_0(esk9_0,esk9_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2014, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k2_fdiff_1(X2,X3))=k4_xcmplx_0(k7_xcmplx_0(X1,k2_fdiff_1(X2,X3)))|~v1_xcmplx_0(X1)|~v1_funct_1(X2)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6629, plain, (v1_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6666, plain, (v2_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2015, plain, (k4_xcmplx_0(k4_xcmplx_0(k2_fdiff_1(X1,X2)))=k2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2016, plain, (k7_xcmplx_0(np__0,k2_fdiff_1(X1,X2))=np__0|~v1_funct_1(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2011, plain, (r2_fdiff_1(X1,k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~r1_tarski(k3_taylor_1,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k3_taylor_1))).#trainneg
cnf(i_0_2013, plain, (v1_xboole_0(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k3_taylor_1))).#trainneg
cnf(i_0_2017, plain, (k7_xcmplx_0(k2_fdiff_1(X1,X2),np__1)=k2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2018, plain, (k6_xcmplx_0(k2_fdiff_1(X1,X2),np__0)=k2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2019, plain, (v1_xcmplx_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2021, plain, (v1_xcmplx_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2022, plain, (v1_xcmplx_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2025, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2026, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2027, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0)=k4_xcmplx_0(k7_xcmplx_0(X1,esk1_0))|~v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2031, negated_conjecture, (k6_xcmplx_0(esk1_0,np__0)=esk1_0|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2033, plain, (r1_xxreal_0(X1,k2_fdiff_1(X2,X3))|r1_xxreal_0(k2_fdiff_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X2)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4188, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2034, plain, (r1_xxreal_0(k2_fdiff_1(X1,X2),k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6667, plain, (v2_membered(k3_taylor_1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6700, plain, (v2_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2035, plain, (m1_subset_1(k7_xcmplx_0(X1,k2_fdiff_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_1(X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2036, plain, (v1_xreal_0(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2037, plain, (v1_int_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2038, plain, (v3_ordinal1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2039, plain, (m1_subset_1(esk19_2(k2_fdiff_1(X1,X2),X3),k2_zfmisc_1(k1_numbers,k1_numbers))|r2_fdiff_1(X3,k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X3)|~r1_tarski(k2_fdiff_1(X1,X2),k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_fdiff_1(X1,X2)))).#trainneg
cnf(i_0_2040, plain, (m1_subset_1(esk3_3(X1,k2_fdiff_1(X2,X3),X4),k2_zfmisc_1(k1_numbers,k1_numbers))|X4=k2_fdiff_1(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|k1_relset_1(k1_numbers,X4)!=k2_fdiff_1(X2,X3)|~v1_funct_1(X4)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k2_fdiff_1(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2041, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|v1_xboole_0(k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k2_fdiff_1(X2,X3)))).#trainneg
cnf(i_0_2042, plain, (k9_xtuple_0(k2_fdiff_1(X1,X2))=k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2043, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X2)|~v1_relat_1(k2_fdiff_1(X2,X1))|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2044, plain, (m1_subset_1(X1,k1_numbers)|~v1_funct_1(X2)|~v1_relat_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(X2,X3)|~r2_hidden(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2046, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k4_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9589, plain, (v5_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2047, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k3_numbers)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6702, plain, (v3_membered(k16_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6841, plain, (v3_membered(k3_taylor_1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2048, plain, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X3,X4))|~r2_fdiff_1(X1,X4)|~r2_fdiff_1(k2_fdiff_1(X3,X4),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r2_fdiff_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2049, negated_conjecture, (v1_xboole_0(esk1_0)|m1_subset_1(esk19_2(esk1_0,X1),k1_numbers)|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_2050, negated_conjecture, (v1_xboole_0(esk1_0)|m1_subset_1(esk4_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_2052, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,np__2))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2053, plain, (v1_xxreal_0(k7_xcmplx_0(X1,np__2))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2055, plain, (v1_rat_1(k7_xcmplx_0(X1,np__2))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2056, plain, (v1_int_1(k7_xcmplx_0(X1,np__2))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2057, plain, (v7_ordinal1(k7_xcmplx_0(X1,np__2))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2058, plain, (v3_ordinal1(k7_xcmplx_0(X1,np__2))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2061, plain, (v3_ordinal1(k1_relset_1(k1_numbers,k3_taylor_1))|~v5_ordinal1(k3_taylor_1))).#trainneg
cnf(i_0_2062, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,k3_taylor_1)|X2!=k3_taylor_1|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X3)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_fdiff_1(X3,k3_taylor_1)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2063, plain, (r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk19_2(X1,k19_sin_cos),k1_numbers)|~v3_rcomp_1(k1_numbers))).#trainneg
cnf(i_0_12756, negated_conjecture, (k1_seq_1(k2_fdiff_1(X1,k3_taylor_1),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|k3_taylor_1!=esk1_0|~v1_funct_1(X1)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2064, plain, (v1_xxreal_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_12674, negated_conjecture, (v3_ordinal1(X1)|~r2_fdiff_1(k3_taylor_1,k1_numbers)|~r2_hidden(X1,esk1_0)|~v5_ordinal1(k3_taylor_1))).#trainneg
cnf(i_0_12862, negated_conjecture, (v3_ordinal1(X1)|~r2_fdiff_1(k3_taylor_1,k1_numbers)|~r2_hidden(X1,esk1_0)|~v1_xboole_0(k3_taylor_1))).#trainneg
cnf(i_0_2066, plain, (v1_xxreal_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2067, plain, (v1_xxreal_0(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2070, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2071, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2076, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),esk9_0)|m1_subset_1(esk19_2(esk9_0,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(esk9_0,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2077, plain, (r2_fdiff_1(X1,esk9_0)|m1_subset_1(esk19_2(esk9_0,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk9_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_2079, negated_conjecture, (r2_fdiff_1(k2_fdiff_1(X1,X2),esk1_0)|m1_subset_1(esk19_2(esk1_0,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(esk1_0,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2082, plain, (r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k19_sin_cos))|~v1_xreal_0(esk19_2(k1_relset_1(k1_numbers,k19_sin_cos),k19_sin_cos))|~m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_2083, plain, (r2_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(esk19_2(X1,k19_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_12973, plain, (r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk19_2(X1,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_12974, plain, (r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk19_2(X1,k19_sin_cos),k4_numbers))).#trainneg
cnf(i_0_12975, plain, (r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk19_2(X1,k19_sin_cos),k5_numbers))).#trainneg
cnf(i_0_12980, plain, (r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6845, plain, (v3_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6866, plain, (v4_membered(k16_sin_cos)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2085, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),esk3_3(X2,X3,X4))=k4_xcmplx_0(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|~v1_xcmplx_0(X1)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~v1_membered(k1_numbers)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2086, plain, (k4_xcmplx_0(k4_xcmplx_0(esk3_3(X1,X2,X3)))=esk3_3(X1,X2,X3)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(k1_numbers)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2087, plain, (k7_xcmplx_0(np__0,esk3_3(X1,X2,X3))=np__0|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(k1_numbers)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2088, plain, (k7_xcmplx_0(esk3_3(X1,X2,X3),np__1)=esk3_3(X1,X2,X3)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(k1_numbers)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2089, plain, (k6_xcmplx_0(esk3_3(X1,X2,X3),np__0)=esk3_3(X1,X2,X3)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_membered(k1_numbers)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2090, plain, (r1_xxreal_0(X1,esk3_3(X2,X3,X4))|r1_xxreal_0(esk3_3(X2,X3,X4),X1)|X4=k2_fdiff_1(X2,X3)|~v1_xxreal_0(X1)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~v2_membered(k1_numbers)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2091, plain, (r1_xxreal_0(esk3_3(X1,X2,X3),esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v2_membered(k1_numbers)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2093, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|~v1_membered(k1_numbers)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2094, plain, (v1_xxreal_0(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|~v2_membered(k1_numbers)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2095, plain, (v1_xreal_0(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2096, plain, (v1_rat_1(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|~v4_membered(k1_numbers)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2097, plain, (v1_int_1(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|~v5_membered(k1_numbers)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2098, plain, (v7_ordinal1(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|~v6_membered(k1_numbers)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2099, plain, (v3_ordinal1(k7_xcmplx_0(X1,esk3_3(X2,X3,X4)))|X4=k2_fdiff_1(X2,X3)|~v3_ordinal1(k1_numbers)|k1_relset_1(k1_numbers,X4)!=X3|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12290, plain, (r1_xxreal_0(esk9_0,esk9_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8708, plain, (v1_xreal_0(k7_xcmplx_0(X1,k4_xcmplx_0(np__1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12572, negated_conjecture, (v1_xreal_0(esk4_1(esk1_0)))).#trainneg
cnf(i_0_12843, negated_conjecture, (k1_seq_1(k2_fdiff_1(X1,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|k3_taylor_1!=esk1_0|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2101, plain, (v1_int_1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v6_membered(k1_numbers)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6868, plain, (v4_membered(k3_taylor_1)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6949, plain, (v4_membered(k19_sin_cos)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2102, plain, (v3_ordinal1(esk3_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~v6_membered(k1_numbers)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2103, plain, (v1_funct_1(k3_relat_1(X1,esk19_2(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4)))|r2_fdiff_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),k1_zfmisc_1(k1_relset_1(k1_numbers,X4)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2105, plain, (v1_funct_1(k3_relat_1(X1,esk4_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~v1_funct_1(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_2106, plain, (v1_funct_1(k3_relat_1(X1,esk3_3(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)),X5)))|X5=k2_fdiff_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(esk3_3(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)),X5))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|k1_relset_1(k1_numbers,X5)!=k1_zfmisc_1(k2_zfmisc_1(X3,X4))|~v1_funct_1(X5)|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2107, plain, (v1_funct_1(k3_relat_1(X1,k1_relset_1(k2_zfmisc_1(X2,X3),X4)))|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X3),X4))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~v4_relat_1(X4,k2_zfmisc_1(X2,X3))|~v1_relat_1(X4))).#trainneg
cnf(i_0_2108, plain, (v1_funct_1(k3_relat_1(X1,k16_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2109, plain, (v1_funct_1(k3_relat_1(X1,k2_fdiff_1(X2,X3)))|~v1_funct_1(k2_fdiff_1(X2,X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2110, plain, (v1_funct_1(k3_relat_1(X1,k1_partfun1(X2,X3,X4,X5,X6,X7)))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,X6,X7))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2112, plain, (v1_funct_1(k3_relat_1(X1,k19_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2113, plain, (v1_funct_1(k3_relat_1(X1,k2_zfmisc_1(X2,X3)))|~v1_funct_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_2114, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k1_partfun1(X2,X3,X4,X5,X6,X7))=k4_xcmplx_0(k7_xcmplx_0(X1,k1_partfun1(X2,X3,X4,X5,X6,X7)))|~v1_xcmplx_0(X1)|~v1_funct_1(X7)|~v1_funct_1(X6)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2115, plain, (k4_xcmplx_0(k4_xcmplx_0(k1_partfun1(X1,X2,X3,X4,X5,X6)))=k1_partfun1(X1,X2,X3,X4,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2116, plain, (k7_xcmplx_0(np__0,k1_partfun1(X1,X2,X3,X4,X5,X6))=np__0|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12441, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_13441, plain, (v1_funct_1(k3_relat_1(k16_sin_cos,k16_sin_cos)))).#trainneg
cnf(i_0_9997, plain, (v5_membered(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2117, plain, (k7_xcmplx_0(k1_partfun1(X1,X2,X3,X4,X5,X6),np__1)=k1_partfun1(X1,X2,X3,X4,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7011, plain, (v4_valued_0(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2118, plain, (k6_xcmplx_0(k1_partfun1(X1,X2,X3,X4,X5,X6),np__0)=k1_partfun1(X1,X2,X3,X4,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2119, plain, (v1_xcmplx_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2120, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(X2,np__2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2123, plain, (v3_ordinal1(k1_relset_1(k1_numbers,k19_sin_cos))|~v5_ordinal1(k19_sin_cos))).#trainneg
cnf(i_0_2124, plain, (r1_xxreal_0(X1,k1_partfun1(X2,X3,X4,X5,X6,X7))|r1_xxreal_0(k1_partfun1(X2,X3,X4,X5,X6,X7),X1)|~v1_xxreal_0(X1)|~v1_funct_1(X7)|~v1_funct_1(X6)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2125, plain, (r1_xxreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2126, plain, (v1_xxreal_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2129, plain, (v3_ordinal1(esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2130, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_partfun1(X2,X3,X4,X5,X6,X7)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_1(X7)|~v1_funct_1(X6)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2131, plain, (v1_xreal_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2132, plain, (v1_rat_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2133, plain, (v1_xreal_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2134, plain, (v1_int_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_13445, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k16_sin_cos)))).#trainneg
cnf(i_0_10000, plain, (v5_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2135, plain, (v1_int_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7396, plain, (v4_valued_0(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_13661, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)|~v5_ordinal1(k19_sin_cos))).#trainneg
cnf(i_0_2136, plain, (v3_ordinal1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2137, plain, (v7_ordinal1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2138, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~v1_funct_1(X4)|~v1_funct_1(X5)|~r2_hidden(X1,k3_relat_1(X5,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X7,X3)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X6))))).#trainneg
cnf(i_0_2139, plain, (m1_subset_1(esk19_2(k1_partfun1(X1,X2,X3,X4,X5,X6),X7),k2_zfmisc_1(X1,X4))|r2_fdiff_1(X7,k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X7)|~r1_tarski(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_relset_1(k1_numbers,X7))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_partfun1(X1,X2,X3,X4,X5,X6)))).#trainneg
cnf(i_0_2140, plain, (m1_subset_1(esk3_3(X1,k1_partfun1(X2,X3,X4,X5,X6,X7),X8),k2_zfmisc_1(X2,X5))|X8=k2_fdiff_1(X1,k1_partfun1(X2,X3,X4,X5,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|k1_relset_1(k1_numbers,X8)!=k1_partfun1(X2,X3,X4,X5,X6,X7)|~v1_funct_1(X8)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_partfun1(X2,X3,X4,X5,X6,X7))|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2141, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|v1_xboole_0(k1_partfun1(X2,X6,X7,X3,X5,X4))|~v1_funct_1(X4)|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X7,X3)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X6)))|~m1_subset_1(X1,k1_partfun1(X2,X6,X7,X3,X5,X4)))).#trainneg
cnf(i_0_2143, plain, (v1_valued_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2144, plain, (v2_valued_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v2_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2145, plain, (v3_valued_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v3_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2146, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v6_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2147, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k4_numbers)|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v6_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2148, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k3_numbers)|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v6_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2149, plain, (v4_valued_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v6_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2150, plain, (m1_subset_1(k3_relat_1(X1,esk19_2(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4)),k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|r2_fdiff_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(esk19_2(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),k1_zfmisc_1(k1_relset_1(k1_numbers,X4)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_13772, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)|~v1_xboole_0(k19_sin_cos))).#trainneg
cnf(i_0_13446, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k16_sin_cos)))).#trainneg
cnf(i_0_12543, plain, (v3_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2152, plain, (m1_subset_1(k3_relat_1(X1,esk4_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))),k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(esk4_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3))))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_7437, plain, (v4_valued_0(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2153, plain, (m1_subset_1(k3_relat_1(X1,esk3_3(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)),X5)),k1_zfmisc_1(k2_zfmisc_1(X6,X4)))|X5=k2_fdiff_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(esk3_3(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)),X5))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|k1_relset_1(k1_numbers,X5)!=k1_zfmisc_1(k2_zfmisc_1(X3,X4))|~v1_funct_1(X5)|~v1_funct_1(X2)|~r2_fdiff_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2154, plain, (m1_subset_1(k3_relat_1(X1,k1_relset_1(k2_zfmisc_1(X2,X3),X4)),k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X3),X4))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~v4_relat_1(X4,k2_zfmisc_1(X2,X3))|~v1_relat_1(X4))).#trainneg
cnf(i_0_2155, plain, (m1_subset_1(k3_relat_1(X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_14009, plain, (m1_subset_1(k3_relat_1(k3_taylor_1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14010, plain, (m1_subset_1(k3_relat_1(k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2156, plain, (m1_subset_1(k3_relat_1(X1,k2_fdiff_1(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(X4,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X2,X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14159, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k3_taylor_1,k16_sin_cos),k2_fdiff_1(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14160, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k19_sin_cos,k16_sin_cos),k2_fdiff_1(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2157, plain, (m1_subset_1(k3_relat_1(X1,k1_partfun1(X2,X3,X4,X5,X6,X7)),k1_zfmisc_1(k2_zfmisc_1(X8,X5)))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,X6,X7))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X8,X9)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_14357, plain, (m1_subset_1(k3_relat_1(k16_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_753, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_3717, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_13495, plain, (v1_funct_1(k3_relat_1(k16_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_611, negated_conjecture, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2159, plain, (m1_subset_1(k3_relat_1(X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11429, plain, (v1_membered(k16_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_11433, plain, (v1_membered(k3_taylor_1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2160, plain, (m1_subset_1(k3_relat_1(X1,k2_zfmisc_1(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~v1_funct_1(k2_zfmisc_1(X2,X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_14531, plain, (m1_subset_1(k3_relat_1(k16_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14535, plain, (m1_subset_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2164, plain, (v1_rat_1(k7_xcmplx_0(np__1,np__2))|~v4_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_2169, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,esk12_0))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14717, plain, (k4_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(X1,esk12_0)))=k7_xcmplx_0(X1,esk12_0)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2170, plain, (v1_xxreal_0(k7_xcmplx_0(X1,esk12_0))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2172, plain, (v1_rat_1(k7_xcmplx_0(X1,esk12_0))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2173, plain, (v1_int_1(k7_xcmplx_0(X1,esk12_0))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12291, plain, (r1_xxreal_0(esk9_0,esk9_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_13499, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_12550, plain, (v3_membered(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_12843, negated_conjecture, (k1_seq_1(k2_fdiff_1(X1,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|k3_taylor_1!=esk1_0|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2174, plain, (v7_ordinal1(k7_xcmplx_0(X1,esk12_0))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12295, plain, (v1_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_12299, plain, (v2_membered(k16_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2175, plain, (v3_ordinal1(k7_xcmplx_0(X1,esk12_0))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2180, plain, (v1_int_1(k7_xcmplx_0(np__1,np__2))|~v5_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_2187, negated_conjecture, (v3_ordinal1(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2190, plain, (v1_int_1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2191, plain, (v3_ordinal1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2195, plain, (v7_ordinal1(k7_xcmplx_0(np__1,np__2))|~v6_membered(k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_13001, plain, (v3_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_753, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0))).#trainneg
cnf(i_0_12452, plain, (v2_membered(k3_taylor_1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_12456, plain, (v2_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2200, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k17_sin_cos(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k17_sin_cos(X2)))|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2201, plain, (k4_xcmplx_0(k4_xcmplx_0(k17_sin_cos(X1)))=k17_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_753, negated_conjecture, (k1_real_1(k1_sin_cos4(esk2_0))!=k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0))).#trainneg
cnf(i_0_528, negated_conjecture, (k1_seq_1(k2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0)))).#trainneg
cnf(i_0_15004, negated_conjecture, (m1_subset_1(k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),X1),k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15011, negated_conjecture, (k1_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2202, plain, (k7_xcmplx_0(np__0,k17_sin_cos(X1))=np__0|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2203, plain, (k7_xcmplx_0(k17_sin_cos(X1),np__1)=k17_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12442, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_13513, plain, (v1_funct_1(k3_relat_1(k16_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_2204, plain, (k6_xcmplx_0(k17_sin_cos(X1),np__0)=k17_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13004, plain, (v4_membered(k16_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_13935, negated_conjecture, (v3_ordinal1(esk1_0)|~v1_xboole_0(k19_sin_cos)|~r2_hidden(k1_zfmisc_1(k1_numbers),esk1_0))).#trainneg
cnf(i_0_2205, plain, (v1_xcmplx_0(k1_seq_1(k16_sin_cos,X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2206, plain, (r1_xxreal_0(X1,k17_sin_cos(X2))|r1_xxreal_0(k17_sin_cos(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2207, plain, (r1_xxreal_0(k17_sin_cos(X1),k17_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2208, plain, (v1_xxreal_0(k1_seq_1(k16_sin_cos,X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2209, plain, (v1_rat_1(k1_seq_1(k16_sin_cos,X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2211, plain, (v1_int_1(k1_seq_1(k16_sin_cos,X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2212, plain, (v1_int_1(k17_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2213, plain, (v3_ordinal1(k17_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2214, plain, (v7_ordinal1(k1_seq_1(k16_sin_cos,X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13660, plain, (v3_ordinal1(X1)|~r2_hidden(X1,esk9_0)|~v5_ordinal1(k19_sin_cos))).#trainneg
cnf(i_0_13517, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_12843, negated_conjecture, (k1_seq_1(k2_fdiff_1(X1,k3_taylor_1),esk2_0)!=k4_xcmplx_0(k1_sin_cos4(esk2_0))|k3_taylor_1!=esk1_0|~v1_funct_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainneg
cnf(i_0_13354, plain, (v4_membered(k3_taylor_1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_13356, plain, (v4_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_13936, negated_conjecture, (v3_ordinal1(X1)|~v1_xboole_0(k19_sin_cos)|~r2_hidden(k1_zfmisc_1(X1),esk1_0))).#trainneg
cnf(i_0_4220, negated_conjecture, (r1_xxreal_0(esk1_0,k1_sin_cos4(X1))|r1_xxreal_0(k1_sin_cos4(X1),esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_2221, plain, (v3_ordinal1(k1_seq_1(k16_sin_cos,X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2222, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),k20_sin_cos(X2))=k4_xcmplx_0(k7_xcmplx_0(X1,k20_sin_cos(X2)))|~v1_xcmplx_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2223, plain, (k4_xcmplx_0(k4_xcmplx_0(k20_sin_cos(X1)))=k20_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2224, plain, (k7_xcmplx_0(np__0,k20_sin_cos(X1))=np__0|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2225, plain, (k7_xcmplx_0(k20_sin_cos(X1),np__1)=k20_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2226, plain, (k6_xcmplx_0(k20_sin_cos(X1),np__0)=k20_sin_cos(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2228, plain, (r1_xxreal_0(X1,k20_sin_cos(X2))|r1_xxreal_0(k20_sin_cos(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2229, plain, (r1_xxreal_0(k20_sin_cos(X1),k20_sin_cos(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2234, plain, (v1_int_1(k20_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2235, plain, (v3_ordinal1(k20_sin_cos(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2238, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(X2,esk12_0)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2242, plain, (v5_relat_1(k16_sin_cos,k4_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2243, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15178, plain, (v3_ordinal1(X1)|~r2_hidden(X1,esk9_0)|~v1_xboole_0(k19_sin_cos))).#trainneg
cnf(i_0_13518, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_12756, negated_conjecture, (k1_seq_1(k2_fdiff_1(X1,k3_taylor_1),esk2_0)!=k1_real_1(k1_sin_cos4(esk2_0))|k3_taylor_1!=esk1_0|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_funct_1(X1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),esk1_0)|~r2_fdiff_1(X1,k3_taylor_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2245, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_14538, plain, (v1_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_14541, plain, (v1_membered(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2246, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2248, plain, (v1_int_1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2254, plain, (v1_int_1(k1_fdiff_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2260, plain, (v3_ordinal1(k1_fdiff_1(k19_sin_cos,X1))|~v1_xreal_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_2262, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14049, plain, (v3_valued_0(k3_relat_1(k3_taylor_1,k16_sin_cos)))).#trainneg
cnf(i_0_14793, plain, (v1_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_9294, negated_conjecture, (r1_xxreal_0(k1_seq_1(k19_sin_cos,esk2_0),np__0)|~v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk2_0),k1_numbers)|~m1_subset_1(k1_sin_cos4(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2263, plain, (v1_xxreal_0(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14796, plain, (v2_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_14869, plain, (v2_membered(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_3148, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,X1),X2)|r1_xxreal_0(X2,k1_seq_1(k19_sin_cos,X1))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2264, plain, (v1_xreal_0(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2265, plain, (v1_rat_1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2266, plain, (v1_int_1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2267, plain, (v7_ordinal1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2268, plain, (v3_ordinal1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2272, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(np__1,np__2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_2280, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,k21_sin_cos(X2)))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2281, plain, (v1_xxreal_0(k7_xcmplx_0(X1,k21_sin_cos(X2)))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14118, plain, (v3_valued_0(k3_relat_1(k19_sin_cos,k16_sin_cos)))).#trainneg
cnf(i_0_14872, plain, (v2_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2282, plain, (v1_xreal_0(k7_xcmplx_0(X1,k21_sin_cos(X2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_15130, plain, (v4_membered(k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_15185, plain, (v4_membered(k3_taylor_1)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_15695, plain, (r1_xxreal_0(X1,k1_seq_1(k19_sin_cos,k1_sin_cos4(X2)))|r1_xxreal_0(k1_seq_1(k19_sin_cos,k1_sin_cos4(X2)),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_16083, plain, (r1_xxreal_0(k1_seq_1(k19_sin_cos,k1_sin_cos4(X1)),X2)|r1_xxreal_0(X2,k1_seq_1(k19_sin_cos,k1_sin_cos4(X1)))|~v1_xxreal_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2283, plain, (v1_rat_1(k7_xcmplx_0(X1,k21_sin_cos(X2)))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2284, plain, (v1_int_1(k7_xcmplx_0(X1,k21_sin_cos(X2)))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 268
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 348
# Removed in clause preprocessing      : 28
# Initial clauses in saturation        : 320
# Processed clauses                    : 2739
# ...of these trivial                  : 65
# ...subsumed                          : 606
# ...remaining for further processing  : 2068
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 42
# Backward-rewritten                   : 77
# Generated clauses                    : 15781
# ...of the previous two non-trivial   : 14017
# Contextual simplify-reflections      : 755
# Paramodulations                      : 15773
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
# Propositional unsat check successes  : 0
# Current number of processed clauses  : 1949
#    Positive orientable unit clauses  : 274
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 29
#    Non-unit-clauses                  : 1644
# Current number of unprocessed clauses: 11321
# ...number of literals in the above   : 64101
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2187
# Proof object given clauses           : 123
# Proof search given clauses           : 2068
# Clause-clause subsumption calls (NU) : 810114
# Rec. Clause-clause subsumption calls : 128977
# Non-unit clause-clause subsumptions  : 1353
# Unit Clause-clause subsumption calls : 8148
# Rewrite failures with RHS unbound    : 2
# BW rewrite match attempts            : 83
# BW rewrite match successes           : 39
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 320358

# -------------------------------------------------
# User time                : 3.456 s
# System time              : 0.104 s
# Total time               : 3.560 s
# ...preprocessing         : 0.488 s
# ...main loop             : 3.072 s
# Maximum resident set size: 64232 pages
