2020-06-01 20:12:49.330111: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_339
2020-06-01 20:12:49.411851: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-01 20:12:49.533189: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-01 20:12:49.878246: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-01 20:12:51.307634: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1977590 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_339' loaded (query=256; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.811 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_l5_algstr_0, axiom, ![X1]:(l5_algstr_0(X1)=>(l4_algstr_0(X1)&l4_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_l5_algstr_0)).
fof(dt_l6_algstr_0, axiom, ![X1]:(l6_algstr_0(X1)=>(l2_algstr_0(X1)&l5_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_l6_algstr_0)).
fof(d3_matrix14, axiom, ![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(v7_ordinal1(X3)=>k4_matrix14(X1,X2,X3)=k2_finseq_7(k3_finseq_2(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1,X3)))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', d3_matrix14)).
fof(redefinition_k3_finseq_2, axiom, ![X1]:k3_finseq_2(X1)=k13_finseq_1(X1), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', redefinition_k3_finseq_2)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', cc8_ordinal1)).
fof(cc3_struct_0, axiom, ![X1]:(l4_struct_0(X1)=>(~(v6_struct_0(X1))=>~(v7_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', cc3_struct_0)).
fof(dt_l2_algstr_0, axiom, ![X1]:(l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_l2_algstr_0)).
fof(dt_k2_finseq_7, axiom, ![X1, X2, X3, X4]:((((~(v1_xboole_0(X1))&m1_finseq_1(X2,X1))&v7_ordinal1(X3))&v7_ordinal1(X4))=>m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k2_finseq_7)).
fof(fc12_finseq_1, axiom, ![X1]:~(v1_xboole_0(k13_finseq_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', fc12_finseq_1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', redefinition_k5_numbers)).
fof(cc2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>(~(v7_struct_0(X1))=>~(v2_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', cc2_struct_0)).
fof(dt_l2_struct_0, axiom, ![X1]:(l2_struct_0(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_l2_struct_0)).
fof(t30_finseq_7, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m2_finseq_1(X2,X1)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>![X5]:(v7_ordinal1(X5)=>((r1_xxreal_0(np__1,X4)&r1_xxreal_0(X4,k3_finseq_1(X2)))=>(X3=X4|X5=X4|k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X5),X4)=k7_partfun1(X1,X2,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t30_finseq_7)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', spc1_numerals)).
fof(t37_matrix_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>![X5]:(~(v1_xboole_0(X5))=>![X6]:(m1_matrix_1(X6,X5,X1,X2)=>((((r1_xxreal_0(np__1,X3)&r1_xxreal_0(X3,X1))&r1_xxreal_0(np__1,X4))&r1_xxreal_0(X4,X2))=>r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))))))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t37_matrix_1)).
fof(dt_k4_matrix14, axiom, ![X1, X2, X3]:((((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))&m1_subset_1(X2,k5_numbers))&v7_ordinal1(X3))=>m1_matrix_1(k4_matrix14(X1,X2,X3),u1_struct_0(X1),X2,X2)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k4_matrix14)).
fof(dt_m1_matrix_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&v7_ordinal1(X2))&v7_ordinal1(X3))=>![X4]:(m1_matrix_1(X4,X1,X2,X3)=>(v1_matrix_1(X4)&m2_finseq_1(X4,k3_finseq_2(X1))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_m1_matrix_1)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', redefinition_m2_finseq_1)).
fof(d5_matrix_1, axiom, ![X1, X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>(r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))=>![X5]:(m1_subset_1(X5,X1)=>(X5=k3_matrix_1(X1,X2,X3,X4)<=>?[X6]:((m2_finseq_1(X6,X1)&X6=k1_funct_1(X2,X3))&X5=k1_funct_1(X6,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', d5_matrix_1)).
fof(t15_finseq_4, axiom, ![X1]:(v7_ordinal1(X1)=>![X2, X3]:(m2_finseq_1(X3,X2)=>((r1_xxreal_0(np__1,X1)&r1_xxreal_0(X1,k3_finseq_1(X3)))=>k7_partfun1(X2,X3,X1)=k1_funct_1(X3,X1)))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t15_finseq_4)).
fof(t24_matrix_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_matrix_1(X3,X2,X1,X1)=>((k3_finseq_1(X3)=X1&k1_matrix_1(X3)=X1)&k2_matrix_1(X3)=k2_zfmisc_1(k2_finseq_1(X1),k2_finseq_1(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t24_matrix_1)).
fof(d2_matrix_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>![X4]:((v1_matrix_1(X4)&m2_finseq_1(X4,k3_finseq_2(X1)))=>(m1_matrix_1(X4,X1,X2,X3)<=>(k3_finseq_1(X4)=X2&![X5]:(m2_finseq_1(X5,X1)=>(r2_hidden(X5,k10_xtuple_0(X4))=>k3_finseq_1(X5)=X3)))))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', d2_matrix_1)).
fof(dt_k12_matrix_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l6_algstr_0(X1))&v7_ordinal1(X2))=>m1_matrix_1(k12_matrix_1(X1,X2),u1_struct_0(X1),X2,X2)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k12_matrix_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', fc2_struct_0)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', connectedness_r1_xxreal_0)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', cc2_xxreal_0)).
fof(t31_finseq_7, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m2_finseq_1(X2,X1)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>((((r1_xxreal_0(np__1,X3)&r1_xxreal_0(X3,k3_finseq_1(X2)))&r1_xxreal_0(np__1,X4))&r1_xxreal_0(X4,k3_finseq_1(X2)))=>(k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)&k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3))))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t31_finseq_7)).
fof(t43_matrix14, conjecture, ![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(v7_ordinal1(X3)=>![X4]:(m1_matrix_1(X4,u1_struct_0(X1),X2,X2)=>(((r1_xxreal_0(np__1,X3)&r1_xxreal_0(X3,X2))&X4=k4_matrix14(X1,X2,X3))=>![X5]:(v7_ordinal1(X5)=>![X6]:(v7_ordinal1(X6)=>((((r1_xxreal_0(np__1,X5)&r1_xxreal_0(X5,X2))&r1_xxreal_0(np__1,X6))&r1_xxreal_0(X6,X2))=>(X3=np__1|((((((X5=np__1&X6=X3)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1))&((X5=X3&X6=np__1)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1)))&((X5=np__1&X6=np__1)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k4_struct_0(X1)))&((X5=X3&X6=X3)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k4_struct_0(X1)))&(~(((X5=np__1|X5=X3)&(X6=np__1|X6=X3)))=>((X5=X6=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1))&(X5!=X6=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k4_struct_0(X1)))))))))))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t43_matrix14)).
fof(t26_matrix14, axiom, ![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>((r1_xxreal_0(np__1,X2)&r1_xxreal_0(X2,X3))=>k1_funct_1(k12_matrix_1(X1,X3),X2)=k3_matrix14(X1,X3,X2))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t26_matrix14)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r1, axiom, r1_xxreal_0(np__1,np__1), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', rqLessOrEqual__r1_xxreal_0__r1_r1)).
fof(dt_k3_matrix14, axiom, ![X1, X2, X3]:((((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))&v7_ordinal1(X2))&v7_ordinal1(X3))=>m2_finseq_1(k3_matrix14(X1,X2,X3),u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k3_matrix14)).
fof(t2_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X3))=>r1_xxreal_0(X1,X3))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t2_xxreal_0)).
fof(t1_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X1))=>X1=X2))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t1_xxreal_0)).
fof(t25_matrix14, axiom, ![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>((r1_xxreal_0(np__1,X3)&r1_xxreal_0(X3,X4))=>(X2=X3|k1_funct_1(k3_matrix14(X1,X4,X2),X3)=k4_struct_0(X1))))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t25_matrix14)).
fof(redefinition_k2_finseq_7, axiom, ![X1, X2, X3, X4]:((((~(v1_xboole_0(X1))&m1_finseq_1(X2,X1))&v7_ordinal1(X3))&v7_ordinal1(X4))=>k2_finseq_7(X1,X2,X3,X4)=k10_funct_7(X2,X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', redefinition_k2_finseq_7)).
fof(dt_k4_struct_0, axiom, ![X1]:(l2_struct_0(X1)=>m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k4_struct_0)).
fof(dt_m1_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_m1_finseq_1)).
fof(d11_finseq_1, axiom, ![X1, X2]:(X2=k13_finseq_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>m2_finseq_1(X3,X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', d11_finseq_1)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', fc1_struct_0)).
fof(dt_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)=>((v1_funct_1(X2)&v1_finseq_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_m2_finseq_1)).
fof(dt_k1_matrix_1, axiom, ![X1]:((((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))&v1_matrix_1(X1))=>m1_subset_1(k1_matrix_1(X1),k5_numbers)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k1_matrix_1)).
fof(t24_matrix14, axiom, ![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>((r1_xxreal_0(np__1,X2)&r1_xxreal_0(X2,X3))=>k1_funct_1(k3_matrix14(X1,X3,X2),X2)=k5_struct_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', t24_matrix14)).
fof(dt_k5_struct_0, axiom, ![X1]:(l3_struct_0(X1)=>m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k5_struct_0)).
fof(dt_l4_algstr_0, axiom, ![X1]:(l4_algstr_0(X1)=>(l3_struct_0(X1)&l3_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_l4_algstr_0)).
fof(dt_k3_matrix_1, axiom, ![X1, X2, X3, X4]:((((v1_matrix_1(X2)&m1_finseq_1(X2,k3_finseq_2(X1)))&v7_ordinal1(X3))&v7_ordinal1(X4))=>m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)), file('/scratch/mptp/allpred/preds__0.005/t43_matrix14', dt_k3_matrix_1)).
fof(c_0_45, plain, ![X94]:((l4_algstr_0(X94)|~l5_algstr_0(X94))&(l4_struct_0(X94)|~l5_algstr_0(X94))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l5_algstr_0])])])).
fof(c_0_46, plain, ![X23]:((l2_algstr_0(X23)|~l6_algstr_0(X23))&(l5_algstr_0(X23)|~l6_algstr_0(X23))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l6_algstr_0])])])).
fof(c_0_47, plain, ![X61, X62, X63]:(v2_struct_0(X61)|v6_struct_0(X61)|~v13_algstr_0(X61)|~v33_algstr_0(X61)|~v2_rlvect_1(X61)|~v3_rlvect_1(X61)|~v4_rlvect_1(X61)|~v3_group_1(X61)|~v5_group_1(X61)|~v4_vectsp_1(X61)|~v5_vectsp_1(X61)|~l6_algstr_0(X61)|(~m1_subset_1(X62,k5_numbers)|(~v7_ordinal1(X63)|k4_matrix14(X61,X62,X63)=k2_finseq_7(k3_finseq_2(u1_struct_0(X61)),k12_matrix_1(X61,X62),np__1,X63)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_matrix14])])])])).
fof(c_0_48, plain, ![X271]:k3_finseq_2(X271)=k13_finseq_1(X271), inference(variable_rename,[status(thm)],[redefinition_k3_finseq_2])).
fof(c_0_49, plain, ![X81]:(~m1_subset_1(X81,k4_ordinal1)|v7_ordinal1(X81)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_50, plain, ![X197]:(~l4_struct_0(X197)|(v6_struct_0(X197)|~v7_struct_0(X197))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_struct_0])])])).
cnf(c_0_51, plain, (l4_struct_0(X1)|~l5_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_53, plain, ![X80]:((l2_struct_0(X80)|~l2_algstr_0(X80))&(l1_algstr_0(X80)|~l2_algstr_0(X80))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
fof(c_0_54, plain, ![X167, X168, X169, X170]:(v1_xboole_0(X167)|~m1_finseq_1(X168,X167)|~v7_ordinal1(X169)|~v7_ordinal1(X170)|m2_finseq_1(k2_finseq_7(X167,X168,X169,X170),X167)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_finseq_7])])])).
cnf(c_0_55, plain, (v2_struct_0(X1)|v6_struct_0(X1)|k4_matrix14(X1,X2,X3)=k2_finseq_7(k3_finseq_2(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1,X3)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (k3_finseq_2(X1)=k13_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_57, plain, ![X263]:~v1_xboole_0(k13_finseq_1(X263)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc12_finseq_1])])).
cnf(c_0_58, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_60, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v7_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (l4_struct_0(X1)|~l5_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_52])).
fof(c_0_63, plain, ![X290]:(~l1_struct_0(X290)|(v7_struct_0(X290)|~v2_struct_0(X290))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_struct_0])])])).
fof(c_0_64, plain, ![X159]:(~l2_struct_0(X159)|l1_struct_0(X159)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_struct_0])])).
cnf(c_0_65, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_66, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_67, plain, ![X54, X55, X56, X57, X58]:(v1_xboole_0(X54)|(~m2_finseq_1(X55,X54)|(~v7_ordinal1(X56)|(~v7_ordinal1(X57)|(~v7_ordinal1(X58)|(~r1_xxreal_0(np__1,X57)|~r1_xxreal_0(X57,k3_finseq_1(X55))|(X56=X57|X58=X57|k7_partfun1(X54,k2_finseq_7(X54,X55,X56,X58),X57)=k7_partfun1(X54,X55,X57)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_finseq_7])])])])).
cnf(c_0_68, plain, (v1_xboole_0(X1)|m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_69, plain, (k4_matrix14(X1,X2,X3)=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1,X3)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)), inference(rw,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_70, plain, (~v1_xboole_0(k13_finseq_1(X1))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_71, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_72, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_73, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_74, plain, ![X37, X38, X39, X40, X41, X42]:(~v7_ordinal1(X37)|(~v7_ordinal1(X38)|(~v7_ordinal1(X39)|(~v7_ordinal1(X40)|(v1_xboole_0(X41)|(~m1_matrix_1(X42,X41,X37,X38)|(~r1_xxreal_0(np__1,X39)|~r1_xxreal_0(X39,X37)|~r1_xxreal_0(np__1,X40)|~r1_xxreal_0(X40,X38)|r2_hidden(k4_tarski(X39,X40),k2_matrix_1(X42))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t37_matrix_1])])])])).
fof(c_0_75, plain, ![X70, X71, X72]:(v2_struct_0(X70)|v6_struct_0(X70)|~v13_algstr_0(X70)|~v33_algstr_0(X70)|~v2_rlvect_1(X70)|~v3_rlvect_1(X70)|~v4_rlvect_1(X70)|~v3_group_1(X70)|~v5_group_1(X70)|~v4_vectsp_1(X70)|~v5_vectsp_1(X70)|~l6_algstr_0(X70)|~m1_subset_1(X71,k5_numbers)|~v7_ordinal1(X72)|m1_matrix_1(k4_matrix14(X70,X71,X72),u1_struct_0(X70),X71,X71)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_matrix14])])])).
cnf(c_0_76, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v7_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_77, plain, (l4_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_78, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_79, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_80, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_81, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_66])).
fof(c_0_82, plain, ![X119, X120, X121, X122]:((v1_matrix_1(X122)|~m1_matrix_1(X122,X119,X120,X121)|(v1_xboole_0(X119)|~v7_ordinal1(X120)|~v7_ordinal1(X121)))&(m2_finseq_1(X122,k3_finseq_2(X119))|~m1_matrix_1(X122,X119,X120,X121)|(v1_xboole_0(X119)|~v7_ordinal1(X120)|~v7_ordinal1(X121)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_matrix_1])])])])])).
cnf(c_0_83, plain, (v1_xboole_0(X1)|X3=X4|X5=X4|k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X5),X4)=k7_partfun1(X1,X2,X4)|~m2_finseq_1(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k3_finseq_1(X2))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_84, plain, (v1_xboole_0(X1)|m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_85, plain, (k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_86, plain, (~v1_xboole_0(k13_finseq_1(X1))), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_87, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_88, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_73])).
fof(c_0_89, plain, ![X136, X137]:((~m2_finseq_1(X137,X136)|m1_finseq_1(X137,X136))&(~m1_finseq_1(X137,X136)|m2_finseq_1(X137,X136))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
fof(c_0_90, plain, ![X73, X74, X75, X76, X77, X79]:((((m2_finseq_1(esk9_5(X73,X74,X75,X76,X77),X73)|X77!=k3_matrix_1(X73,X74,X75,X76)|~m1_subset_1(X77,X73)|~r2_hidden(k4_tarski(X75,X76),k2_matrix_1(X74))|~v7_ordinal1(X76)|~v7_ordinal1(X75)|(~v1_matrix_1(X74)|~m2_finseq_1(X74,k3_finseq_2(X73))))&(esk9_5(X73,X74,X75,X76,X77)=k1_funct_1(X74,X75)|X77!=k3_matrix_1(X73,X74,X75,X76)|~m1_subset_1(X77,X73)|~r2_hidden(k4_tarski(X75,X76),k2_matrix_1(X74))|~v7_ordinal1(X76)|~v7_ordinal1(X75)|(~v1_matrix_1(X74)|~m2_finseq_1(X74,k3_finseq_2(X73)))))&(X77=k1_funct_1(esk9_5(X73,X74,X75,X76,X77),X76)|X77!=k3_matrix_1(X73,X74,X75,X76)|~m1_subset_1(X77,X73)|~r2_hidden(k4_tarski(X75,X76),k2_matrix_1(X74))|~v7_ordinal1(X76)|~v7_ordinal1(X75)|(~v1_matrix_1(X74)|~m2_finseq_1(X74,k3_finseq_2(X73)))))&(~m2_finseq_1(X79,X73)|X79!=k1_funct_1(X74,X75)|X77!=k1_funct_1(X79,X76)|X77=k3_matrix_1(X73,X74,X75,X76)|~m1_subset_1(X77,X73)|~r2_hidden(k4_tarski(X75,X76),k2_matrix_1(X74))|~v7_ordinal1(X76)|~v7_ordinal1(X75)|(~v1_matrix_1(X74)|~m2_finseq_1(X74,k3_finseq_2(X73))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_matrix_1])])])])])).
cnf(c_0_91, plain, (v1_xboole_0(X5)|r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_matrix_1(X6,X5,X1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_92, plain, (v2_struct_0(X1)|v6_struct_0(X1)|m1_matrix_1(k4_matrix14(X1,X2,X3),u1_struct_0(X1),X2,X2)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_93, plain, (v6_struct_0(X1)|~v7_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77])])).
cnf(c_0_94, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_95, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_96, plain, (l2_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81])])).
cnf(c_0_97, plain, (m2_finseq_1(X1,k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_82])).
fof(c_0_98, plain, ![X82, X83, X84]:(~v7_ordinal1(X82)|(~m2_finseq_1(X84,X83)|(~r1_xxreal_0(np__1,X82)|~r1_xxreal_0(X82,k3_finseq_1(X84))|k7_partfun1(X83,X84,X82)=k1_funct_1(X84,X82)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_finseq_4])])])).
cnf(c_0_99, plain, (X4=X5|X3=X4|k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X5),X4)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k3_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_100, plain, (m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])])).
cnf(c_0_101, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88])])).
cnf(c_0_102, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_103, plain, ![X43, X44, X45]:(((k3_finseq_1(X45)=X43|~m1_matrix_1(X45,X44,X43,X43)|v1_xboole_0(X44)|~v7_ordinal1(X43))&(k1_matrix_1(X45)=X43|~m1_matrix_1(X45,X44,X43,X43)|v1_xboole_0(X44)|~v7_ordinal1(X43)))&(k2_matrix_1(X45)=k2_zfmisc_1(k2_finseq_1(X43),k2_finseq_1(X43))|~m1_matrix_1(X45,X44,X43,X43)|v1_xboole_0(X44)|~v7_ordinal1(X43))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_matrix_1])])])])])).
cnf(c_0_104, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|~m2_finseq_1(X1,X2)|X1!=k1_funct_1(X3,X4)|X5!=k1_funct_1(X1,X6)|~m1_subset_1(X5,X2)|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3))|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_105, plain, (v1_xboole_0(X5)|r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~m1_matrix_1(X6,X5,X1,X2)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_106, plain, (v2_struct_0(X1)|v6_struct_0(X1)|m1_matrix_1(k4_matrix14(X1,X2,X3),u1_struct_0(X1),X2,X2)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_107, plain, (v6_struct_0(X1)|~l6_algstr_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94])])).
cnf(c_0_108, plain, (l1_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96])])).
cnf(c_0_109, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k13_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(rw,[status(thm)],[c_0_97, c_0_56])).
cnf(c_0_110, plain, (k7_partfun1(X3,X2,X1)=k1_funct_1(X2,X1)|~v7_ordinal1(X1)|~m2_finseq_1(X2,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_111, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X4)|np__1=X4|X4=X3|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_85]), c_0_86])])).
cnf(c_0_112, plain, (m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_113, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_114, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~m1_matrix_1(X1,X3,X2,X2)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_115, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|X5!=k1_funct_1(X1,X6)|X1!=k1_funct_1(X3,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X5,X2)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k13_finseq_1(X2))|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3))), inference(rw,[status(thm)],[c_0_104, c_0_56])).
cnf(c_0_116, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_tarski(X2,X3),k2_matrix_1(k4_matrix14(X1,X4,X5)))|v6_struct_0(X1)|v2_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106])])).
cnf(c_0_117, plain, (v6_struct_0(X1)|~l6_algstr_0(X1)|~v2_struct_0(X1)), inference(csr,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_118, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k13_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_119, plain, (k7_partfun1(X3,X2,X1)=k1_funct_1(X2,X1)|~v7_ordinal1(X1)|~m2_finseq_1(X2,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_120, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X4)|np__1=X4|X4=X3|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_101])])).
cnf(c_0_121, plain, (m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113])])).
cnf(c_0_122, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_123, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|X5!=k1_funct_1(X1,X6)|X1!=k1_funct_1(X3,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X5,X2)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k13_finseq_1(X2))|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3))), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_124, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_tarski(X2,X3),k2_matrix_1(k4_matrix14(X1,X4,X5)))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_116, c_0_87]), c_0_117])).
cnf(c_0_125, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_106])])).
cnf(c_0_126, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X4)=k1_funct_1(k4_matrix14(X1,X2,X3),X4)|np__1=X4|X4=X3|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_127, plain, (m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_121, c_0_117])).
cnf(c_0_128, plain, (k3_finseq_1(k4_matrix14(X1,X2,X3))=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_106])])).
fof(c_0_129, plain, ![X31, X32, X33, X34, X35]:(((k3_finseq_1(X34)=X32|~m1_matrix_1(X34,X31,X32,X33)|(~v1_matrix_1(X34)|~m2_finseq_1(X34,k3_finseq_2(X31)))|~v7_ordinal1(X33)|~v7_ordinal1(X32)|v1_xboole_0(X31))&(~m2_finseq_1(X35,X31)|(~r2_hidden(X35,k10_xtuple_0(X34))|k3_finseq_1(X35)=X33)|~m1_matrix_1(X34,X31,X32,X33)|(~v1_matrix_1(X34)|~m2_finseq_1(X34,k3_finseq_2(X31)))|~v7_ordinal1(X33)|~v7_ordinal1(X32)|v1_xboole_0(X31)))&((m2_finseq_1(esk8_4(X31,X32,X33,X34),X31)|k3_finseq_1(X34)!=X32|m1_matrix_1(X34,X31,X32,X33)|(~v1_matrix_1(X34)|~m2_finseq_1(X34,k3_finseq_2(X31)))|~v7_ordinal1(X33)|~v7_ordinal1(X32)|v1_xboole_0(X31))&((r2_hidden(esk8_4(X31,X32,X33,X34),k10_xtuple_0(X34))|k3_finseq_1(X34)!=X32|m1_matrix_1(X34,X31,X32,X33)|(~v1_matrix_1(X34)|~m2_finseq_1(X34,k3_finseq_2(X31)))|~v7_ordinal1(X33)|~v7_ordinal1(X32)|v1_xboole_0(X31))&(k3_finseq_1(esk8_4(X31,X32,X33,X34))!=X33|k3_finseq_1(X34)!=X32|m1_matrix_1(X34,X31,X32,X33)|(~v1_matrix_1(X34)|~m2_finseq_1(X34,k3_finseq_2(X31)))|~v7_ordinal1(X33)|~v7_ordinal1(X32)|v1_xboole_0(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_matrix_1])])])])])])).
cnf(c_0_130, plain, (X1=k3_matrix_1(X2,k4_matrix14(X3,X4,X5),X6,X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X1!=k1_funct_1(X8,X7)|X8!=k1_funct_1(k4_matrix14(X3,X4,X5),X6)|~v1_matrix_1(k4_matrix14(X3,X4,X5))|~m2_finseq_1(k4_matrix14(X3,X4,X5),k13_finseq_1(X2))|~m2_finseq_1(X8,X2)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~m1_subset_1(X1,X2)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,X4)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_131, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_125, c_0_87])).
cnf(c_0_132, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k4_matrix14(X1,X2,X4)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_126, c_0_127]), c_0_117])).
cnf(c_0_133, plain, (k3_finseq_1(k4_matrix14(X1,X2,X3))=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_128, c_0_87])).
cnf(c_0_134, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~m1_matrix_1(X1,X3,X2,X4)|~v1_matrix_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X3))|~v7_ordinal1(X4)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
fof(c_0_135, plain, ![X99, X100]:(v2_struct_0(X99)|~l6_algstr_0(X99)|~v7_ordinal1(X100)|m1_matrix_1(k12_matrix_1(X99,X100),u1_struct_0(X99),X100,X100)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k12_matrix_1])])])).
cnf(c_0_136, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_137, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v6_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_138, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v1_matrix_1(X1)|~m2_finseq_1(X1,k13_finseq_1(X3))|~m1_matrix_1(X1,X3,X2,X4)), inference(rw,[status(thm)],[c_0_134, c_0_56])).
cnf(c_0_139, plain, (v1_matrix_1(X1)|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_140, plain, (v2_struct_0(X1)|m1_matrix_1(k12_matrix_1(X1,X2),u1_struct_0(X1),X2,X2)|~l6_algstr_0(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_135])).
fof(c_0_141, plain, ![X150]:(v2_struct_0(X150)|~l1_struct_0(X150)|~v1_xboole_0(u1_struct_0(X150))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_142, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_136, c_0_117])).
cnf(c_0_143, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_137, c_0_117])).
cnf(c_0_144, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v1_matrix_1(X1)|~m2_finseq_1(X1,k13_finseq_1(X3))|~m1_matrix_1(X1,X3,X2,X4)), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_145, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_146, plain, (v2_struct_0(X1)|m1_matrix_1(k12_matrix_1(X1,X2),u1_struct_0(X1),X2,X2)|~l6_algstr_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_147, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_148, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|np__1=X5|X5=X4|X7!=k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143])])).
cnf(c_0_149, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~m2_finseq_1(X1,k13_finseq_1(X3))|~m1_matrix_1(X1,X3,X2,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_150, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_146])])).
cnf(c_0_151, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_152, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_119])])).
cnf(c_0_153, plain, (v1_matrix_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_106])])).
cnf(c_0_154, plain, (k3_finseq_1(k12_matrix_1(X1,X2))=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_matrix_1(k12_matrix_1(X1,X2),u1_struct_0(X1),X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_155, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_108])])).
cnf(c_0_156, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_150])])).
cnf(c_0_157, plain, (v1_matrix_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_153, c_0_87]), c_0_117])).
cnf(c_0_158, plain, (k3_finseq_1(k12_matrix_1(X1,X2))=X3|v2_struct_0(X1)|~m1_matrix_1(k12_matrix_1(X1,X2),u1_struct_0(X1),X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(csr,[status(thm)],[c_0_154, c_0_155])).
fof(c_0_159, plain, ![X231, X232]:(~v1_xxreal_0(X231)|~v1_xxreal_0(X232)|(r1_xxreal_0(X231,X232)|r1_xxreal_0(X232,X231))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
fof(c_0_160, plain, ![X255]:(~v7_ordinal1(X255)|v1_xxreal_0(X255)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
fof(c_0_161, plain, ![X4, X5, X6, X3, X1]:(epred1_5(X1,X6,X5,X3,X4)<=>((((X5=np__1&X6=X3)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1))&((X5=X3&X6=np__1)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1)))&((X5=np__1&X6=np__1)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k4_struct_0(X1)))), introduced(definition)).
cnf(c_0_162, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_156, c_0_87]), c_0_157]), c_0_117])).
cnf(c_0_163, plain, (k3_finseq_1(k12_matrix_1(X1,X2))=X2|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_146])])).
fof(c_0_164, plain, ![X46, X47, X48, X49]:((k7_partfun1(X46,k2_finseq_7(X46,X47,X48,X49),X48)=k7_partfun1(X46,X47,X49)|(~r1_xxreal_0(np__1,X48)|~r1_xxreal_0(X48,k3_finseq_1(X47))|~r1_xxreal_0(np__1,X49)|~r1_xxreal_0(X49,k3_finseq_1(X47)))|~v7_ordinal1(X49)|~v7_ordinal1(X48)|~m2_finseq_1(X47,X46)|v1_xboole_0(X46))&(k7_partfun1(X46,k2_finseq_7(X46,X47,X48,X49),X49)=k7_partfun1(X46,X47,X48)|(~r1_xxreal_0(np__1,X48)|~r1_xxreal_0(X48,k3_finseq_1(X47))|~r1_xxreal_0(np__1,X49)|~r1_xxreal_0(X49,k3_finseq_1(X47)))|~v7_ordinal1(X49)|~v7_ordinal1(X48)|~m2_finseq_1(X47,X46)|v1_xboole_0(X46))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_finseq_7])])])])])).
cnf(c_0_165, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_166, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_160])).
fof(c_0_167, negated_conjecture, ~(![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_group_1(X1))&v5_group_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(v7_ordinal1(X3)=>![X4]:(m1_matrix_1(X4,u1_struct_0(X1),X2,X2)=>(((r1_xxreal_0(np__1,X3)&r1_xxreal_0(X3,X2))&X4=k4_matrix14(X1,X2,X3))=>![X5]:(v7_ordinal1(X5)=>![X6]:(v7_ordinal1(X6)=>((((r1_xxreal_0(np__1,X5)&r1_xxreal_0(X5,X2))&r1_xxreal_0(np__1,X6))&r1_xxreal_0(X6,X2))=>(X3=np__1|((epred1_5(X1,X6,X5,X3,X4)&((X5=X3&X6=X3)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k4_struct_0(X1)))&(~(((X5=np__1|X5=X3)&(X6=np__1|X6=X3)))=>((X5=X6=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1))&(X5!=X6=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k4_struct_0(X1))))))))))))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[t43_matrix14]), c_0_161])).
cnf(c_0_168, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
fof(c_0_169, plain, ![X17, X18, X19]:(v2_struct_0(X17)|v6_struct_0(X17)|~v13_algstr_0(X17)|~v33_algstr_0(X17)|~v2_rlvect_1(X17)|~v3_rlvect_1(X17)|~v4_rlvect_1(X17)|~v3_group_1(X17)|~v5_group_1(X17)|~v4_vectsp_1(X17)|~v5_vectsp_1(X17)|~l6_algstr_0(X17)|(~v7_ordinal1(X18)|(~v7_ordinal1(X19)|(~r1_xxreal_0(np__1,X18)|~r1_xxreal_0(X18,X19)|k1_funct_1(k12_matrix_1(X17,X19),X18)=k3_matrix14(X17,X19,X18))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_matrix14])])])])).
cnf(c_0_170, plain, (k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_164])).
cnf(c_0_171, plain, (r1_xxreal_0(np__1,np__1)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r1])).
cnf(c_0_172, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_173, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_166])).
fof(c_0_174, negated_conjecture, ((((((((((((~v2_struct_0(esk1_0)&~v6_struct_0(esk1_0))&v13_algstr_0(esk1_0))&v33_algstr_0(esk1_0))&v2_rlvect_1(esk1_0))&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&v3_group_1(esk1_0))&v5_group_1(esk1_0))&v4_vectsp_1(esk1_0))&v5_vectsp_1(esk1_0))&l6_algstr_0(esk1_0))&(m1_subset_1(esk2_0,k5_numbers)&(v7_ordinal1(esk3_0)&(m1_matrix_1(esk4_0,u1_struct_0(esk1_0),esk2_0,esk2_0)&(((r1_xxreal_0(np__1,esk3_0)&r1_xxreal_0(esk3_0,esk2_0))&esk4_0=k4_matrix14(esk1_0,esk2_0,esk3_0))&(v7_ordinal1(esk5_0)&(v7_ordinal1(esk6_0)&((((r1_xxreal_0(np__1,esk5_0)&r1_xxreal_0(esk5_0,esk2_0))&r1_xxreal_0(np__1,esk6_0))&r1_xxreal_0(esk6_0,esk2_0))&(esk3_0!=np__1&((((((esk6_0!=np__1|esk5_0!=np__1|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(esk6_0!=esk3_0|esk5_0!=np__1|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))))&((esk6_0!=np__1|esk5_0!=esk3_0|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(esk6_0!=esk3_0|esk5_0!=esk3_0|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))))&(((esk5_0!=esk6_0|esk5_0=esk6_0|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|esk5_0=esk6_0|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))))&((esk5_0!=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|(esk5_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))))))&((((esk6_0!=np__1|esk5_0!=np__1|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(esk6_0!=esk3_0|esk5_0!=np__1|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))))&((esk6_0!=np__1|esk5_0!=esk3_0|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(esk6_0!=esk3_0|esk5_0!=esk3_0|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))))&(((esk5_0!=esk6_0|esk5_0=esk6_0|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|esk5_0=esk6_0|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))))&((esk5_0!=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|(esk6_0=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))))))&((((esk6_0!=np__1|esk5_0!=np__1|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(esk6_0!=esk3_0|esk5_0!=np__1|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))))&((esk6_0!=np__1|esk5_0!=esk3_0|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(esk6_0!=esk3_0|esk5_0!=esk3_0|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))))&(((esk5_0!=esk6_0|esk5_0=esk6_0|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|esk5_0=esk6_0|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))))&((esk5_0!=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))&(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)))))))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_167])])])])])).
cnf(c_0_175, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_168, c_0_87]), c_0_117])).
cnf(c_0_176, plain, (v2_struct_0(X1)|v6_struct_0(X1)|k1_funct_1(k12_matrix_1(X1,X3),X2)=k3_matrix14(X1,X3,X2)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_177, plain, (k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_178, plain, (r1_xxreal_0(np__1,np__1)), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_179, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173])])).
cnf(c_0_180, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_181, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k1_funct_1(k12_matrix_1(X2,X3),X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_175])])).
cnf(c_0_182, plain, (k1_funct_1(k12_matrix_1(X1,X3),X2)=k3_matrix14(X1,X3,X2)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(np__1,X2)), inference(evalgc,[status(thm)],[c_0_176])).
fof(c_0_183, plain, ![X64, X65, X66]:(v2_struct_0(X64)|v6_struct_0(X64)|~v13_algstr_0(X64)|~v33_algstr_0(X64)|~v2_rlvect_1(X64)|~v3_rlvect_1(X64)|~v4_rlvect_1(X64)|~v3_group_1(X64)|~v5_group_1(X64)|~v4_vectsp_1(X64)|~v5_vectsp_1(X64)|~l6_algstr_0(X64)|~v7_ordinal1(X65)|~v7_ordinal1(X66)|m2_finseq_1(k3_matrix14(X64,X65,X66),u1_struct_0(X64))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_matrix14])])])).
cnf(c_0_184, negated_conjecture, (m1_matrix_1(esk4_0,u1_struct_0(esk1_0),esk2_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_185, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_85]), c_0_178])]), c_0_86])])).
cnf(c_0_186, negated_conjecture, (esk4_0=k4_matrix14(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_187, negated_conjecture, (r1_xxreal_0(np__1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_188, negated_conjecture, (v7_ordinal1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_189, negated_conjecture, (l6_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_190, negated_conjecture, (v5_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_191, negated_conjecture, (v4_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_192, negated_conjecture, (v5_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_193, negated_conjecture, (v3_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_194, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_195, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_196, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_197, negated_conjecture, (v33_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_198, negated_conjecture, (v13_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_199, negated_conjecture, (~v6_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
fof(c_0_200, plain, ![X91, X92, X93]:(~v1_xxreal_0(X91)|(~v1_xxreal_0(X92)|(~v1_xxreal_0(X93)|(~r1_xxreal_0(X91,X92)|~r1_xxreal_0(X92,X93)|r1_xxreal_0(X91,X93))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_xxreal_0])])])).
cnf(c_0_201, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_173])])).
cnf(c_0_202, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_203, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~m2_finseq_1(k3_matrix14(X2,X3,X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182])])).
cnf(c_0_204, plain, (v2_struct_0(X1)|v6_struct_0(X1)|m2_finseq_1(k3_matrix14(X1,X2,X3),u1_struct_0(X1))|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_205, negated_conjecture, (m1_matrix_1(esk4_0,u1_struct_0(esk1_0),esk2_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_206, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_207, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_101])]), c_0_117])).
cnf(c_0_208, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,esk3_0)=esk4_0), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_209, negated_conjecture, (r1_xxreal_0(np__1,esk3_0)), inference(evalgc,[status(thm)],[c_0_187])).
cnf(c_0_210, negated_conjecture, (v7_ordinal1(esk3_0)), inference(evalgc,[status(thm)],[c_0_188])).
cnf(c_0_211, negated_conjecture, (l6_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_212, negated_conjecture, (v5_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_213, negated_conjecture, (v4_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_214, negated_conjecture, (v5_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_215, negated_conjecture, (v3_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_216, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_217, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_218, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_219, negated_conjecture, (v33_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_220, negated_conjecture, (v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_221, negated_conjecture, (~v6_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_199])).
cnf(c_0_222, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_200])).
cnf(c_0_223, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_101])])).
cnf(c_0_224, negated_conjecture, (v7_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_202])])).
fof(c_0_225, plain, ![X52, X53]:(~v1_xxreal_0(X52)|(~v1_xxreal_0(X53)|(~r1_xxreal_0(X52,X53)|~r1_xxreal_0(X53,X52)|X52=X53))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xxreal_0])])])).
cnf(c_0_226, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~m2_finseq_1(k3_matrix14(X2,X3,X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_203, c_0_87]), c_0_117])).
cnf(c_0_227, plain, (v2_struct_0(X1)|v6_struct_0(X1)|m2_finseq_1(k3_matrix14(X1,X2,X3),u1_struct_0(X1))|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_204])).
fof(c_0_228, plain, ![X13, X14, X15, X16]:(v2_struct_0(X13)|v6_struct_0(X13)|~v13_algstr_0(X13)|~v33_algstr_0(X13)|~v2_rlvect_1(X13)|~v3_rlvect_1(X13)|~v4_rlvect_1(X13)|~v3_group_1(X13)|~v5_group_1(X13)|~v4_vectsp_1(X13)|~v5_vectsp_1(X13)|~l6_algstr_0(X13)|(~v7_ordinal1(X14)|(~v7_ordinal1(X15)|(~v7_ordinal1(X16)|(~r1_xxreal_0(np__1,X15)|~r1_xxreal_0(X15,X16)|(X14=X15|k1_funct_1(k3_matrix14(X13,X16,X14),X15)=k4_struct_0(X13))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_matrix14])])])])).
fof(c_0_229, plain, ![X129, X130, X131, X132]:(v1_xboole_0(X129)|~m1_finseq_1(X130,X129)|~v7_ordinal1(X131)|~v7_ordinal1(X132)|k2_finseq_7(X129,X130,X131,X132)=k10_funct_7(X130,X131,X132)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k2_finseq_7])])])).
cnf(c_0_230, negated_conjecture, (v1_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_205])])).
cnf(c_0_231, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_232, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_209]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221])])).
cnf(c_0_233, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_234, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk2_0,np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224])])).
cnf(c_0_235, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_236, negated_conjecture, (esk3_0!=np__1), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_237, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227])])).
cnf(c_0_238, plain, (v2_struct_0(X1)|v6_struct_0(X1)|X2=X3|k1_funct_1(k3_matrix14(X1,X4,X2),X3)=k4_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_239, plain, (v1_xboole_0(X1)|k2_finseq_7(X1,X2,X3,X4)=k10_funct_7(X2,X3,X4)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_229])).
cnf(c_0_240, negated_conjecture, (v1_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_230, c_0_224])])).
cnf(c_0_241, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_211]), c_0_231])])).
cnf(c_0_242, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(esk1_0)))|~v7_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_205])])).
cnf(c_0_243, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_232]), c_0_178]), c_0_101])])])).
cnf(c_0_244, negated_conjecture, (r1_xxreal_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_245, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234])])).
cnf(c_0_246, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_247, negated_conjecture, (esk3_0!=np__1), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_248, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_237, c_0_87]), c_0_117])).
cnf(c_0_249, plain, (X2=X3|k1_funct_1(k3_matrix14(X1,X4,X2),X3)=k4_struct_0(X1)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(np__1,X3)), inference(evalgc,[status(thm)],[c_0_238])).
fof(c_0_250, plain, ![X182]:(~l2_struct_0(X182)|m1_subset_1(k4_struct_0(X182),u1_struct_0(X182))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_struct_0])])).
cnf(c_0_251, plain, (k2_finseq_7(X1,X2,X3,X4)=k10_funct_7(X2,X3,X4)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_252, negated_conjecture, (v1_matrix_1(esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_240, c_0_241])])).
cnf(c_0_253, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_242, c_0_224])])).
cnf(c_0_254, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_150]), c_0_224]), c_0_211])]), c_0_241]), c_0_231])])).
cnf(c_0_255, negated_conjecture, (r1_xxreal_0(esk3_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_256, negated_conjecture, (k3_finseq_1(esk4_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_205])])).
cnf(c_0_257, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_173]), c_0_224])])])).
cnf(c_0_258, negated_conjecture, (~v1_xxreal_0(np__1)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_209]), c_0_247])])).
cnf(c_0_259, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X5=X6|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_249])])).
cnf(c_0_260, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_250])).
cnf(c_0_261, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_85]), c_0_86])])).
cnf(c_0_262, plain, (k1_funct_1(k2_finseq_7(X1,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_177])])).
cnf(c_0_263, negated_conjecture, (X1=k3_matrix_1(X2,esk4_0,X3,X4)|X5!=k1_funct_1(esk4_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_finseq_1(esk4_0,k13_finseq_1(X2))|~m2_finseq_1(X5,X2)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_208]), c_0_252]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_241]), c_0_221])])).
cnf(c_0_264, negated_conjecture, (m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_253, c_0_241])])).
cnf(c_0_265, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_163]), c_0_255]), c_0_224]), c_0_211])]), c_0_231])])).
cnf(c_0_266, negated_conjecture, (k3_finseq_1(esk4_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_256, c_0_224])])).
cnf(c_0_267, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_173]), c_0_101])])])).
cnf(c_0_268, negated_conjecture, (~v1_xxreal_0(np__1)|~r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_173]), c_0_210])])])).
cnf(c_0_269, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X6|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_259, c_0_87]), c_0_117])).
cnf(c_0_270, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_260])).
fof(c_0_271, plain, ![X4, X5, X6, X3, X1]:(((((X5=np__1&X6=X3)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1))&((X5=X3&X6=np__1)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k5_struct_0(X1)))&((X5=np__1&X6=np__1)=>k3_matrix_1(u1_struct_0(X1),X4,X5,X6)=k4_struct_0(X1)))=>epred1_5(X1,X6,X5,X3,X4)), inference(split_equiv,[status(thm)],[c_0_161])).
cnf(c_0_272, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_261, c_0_101])]), c_0_117])).
cnf(c_0_273, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_85]), c_0_178]), c_0_101])]), c_0_86])])).
cnf(c_0_274, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X2,X3)|X4!=k1_funct_1(esk4_0,X2)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264])])).
cnf(c_0_275, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k3_matrix14(esk1_0,esk2_0,np__1)|~r1_xxreal_0(np__1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_265]), c_0_178]), c_0_224]), c_0_101]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_276, negated_conjecture, (k3_finseq_1(esk4_0)=esk2_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_266, c_0_241])])).
cnf(c_0_277, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_173])])).
cnf(c_0_278, negated_conjecture, (~r1_xxreal_0(esk3_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_173]), c_0_101])])])).
cnf(c_0_279, negated_conjecture, (esk5_0=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_280, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k4_struct_0(X1)|np__1=X4|X4=X5|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270])])).
fof(c_0_281, plain, ![X349, X350, X351, X352, X353]:(((((((X350=np__1|(X350=X352|X350=np__1)|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(X350=X352|X350=np__1)|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(X350=X352|X350=np__1)|epred1_5(X353,X351,X350,X352,X349)))&(((X350=np__1|(X351=np__1|X350=np__1)|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(X351=np__1|X350=np__1)|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(X351=np__1|X350=np__1)|epred1_5(X353,X351,X350,X352,X349))))&(((X350=np__1|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|X350=np__1)|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|X350=np__1)|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|X350=np__1)|epred1_5(X353,X351,X350,X352,X349))))&(((((X350=np__1|(X350=X352|X351=X352)|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(X350=X352|X351=X352)|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(X350=X352|X351=X352)|epred1_5(X353,X351,X350,X352,X349)))&(((X350=np__1|(X351=np__1|X351=X352)|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(X351=np__1|X351=X352)|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(X351=np__1|X351=X352)|epred1_5(X353,X351,X350,X352,X349))))&(((X350=np__1|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|X351=X352)|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|X351=X352)|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|X351=X352)|epred1_5(X353,X351,X350,X352,X349)))))&(((((X350=np__1|(X350=X352|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(X350=X352|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(X350=X352|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349)))&(((X350=np__1|(X351=np__1|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(X351=np__1|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(X351=np__1|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349))))&(((X350=np__1|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349))&(X351=np__1|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349)))&(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k4_struct_0(X353)|(k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353)|k3_matrix_1(u1_struct_0(X353),X349,X350,X351)!=k5_struct_0(X353))|epred1_5(X353,X351,X350,X352,X349))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_271])])])).
cnf(c_0_282, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|np__1=X1|X1=esk3_0|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_208]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_283, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_113])])).
fof(c_0_284, plain, ![X162, X163]:(((v1_relat_1(X163)|~m1_finseq_1(X163,X162))&(v1_funct_1(X163)|~m1_finseq_1(X163,X162)))&(v1_finseq_1(X163)|~m1_finseq_1(X163,X162))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_finseq_1])])])).
fof(c_0_285, plain, ![X24, X25, X26, X27, X28, X29]:(((~r2_hidden(X26,X25)|m2_finseq_1(X26,X24)|X25!=k13_finseq_1(X24))&(~m2_finseq_1(X27,X24)|r2_hidden(X27,X25)|X25!=k13_finseq_1(X24)))&((~r2_hidden(esk7_2(X28,X29),X29)|~m2_finseq_1(esk7_2(X28,X29),X28)|X29=k13_finseq_1(X28))&(r2_hidden(esk7_2(X28,X29),X29)|m2_finseq_1(esk7_2(X28,X29),X28)|X29=k13_finseq_1(X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_finseq_1])])])])])])).
fof(c_0_286, plain, ![X336]:(~v2_struct_0(X336)|~l1_struct_0(X336)|v1_xboole_0(u1_struct_0(X336))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_287, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_273, c_0_127]), c_0_117])).
cnf(c_0_288, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|X3!=k1_funct_1(esk4_0,X1)|k1_funct_1(X3,X2)!=k4_struct_0(esk1_0)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_270])])).
cnf(c_0_289, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,np__1)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_275]), c_0_264]), c_0_276]), c_0_255]), c_0_209]), c_0_210])])])).
cnf(c_0_290, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_255]), c_0_210])]), c_0_278])])).
cnf(c_0_291, negated_conjecture, (esk6_0=esk5_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[c_0_279])).
cnf(c_0_292, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k4_struct_0(X1)|np__1=X4|X4=X3|X4=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_280, c_0_96])).
cnf(c_0_293, negated_conjecture, (r1_xxreal_0(np__1,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_294, negated_conjecture, (r1_xxreal_0(np__1,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_295, negated_conjecture, (r1_xxreal_0(esk5_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_296, negated_conjecture, (r1_xxreal_0(esk6_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_297, negated_conjecture, (v7_ordinal1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_298, negated_conjecture, (v7_ordinal1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_299, plain, (X1=np__1|X1=X2|X1=np__1|epred1_5(X3,X4,X1,X2,X5)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_300, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|X1=esk3_0|np__1=X1|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_282])])).
cnf(c_0_301, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_150])])).
cnf(c_0_302, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_284])).
fof(c_0_303, plain, ![X199, X200]:(((v1_funct_1(X200)|~m2_finseq_1(X200,X199))&(v1_finseq_1(X200)|~m2_finseq_1(X200,X199)))&(m1_subset_1(X200,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X199)))|~m2_finseq_1(X200,X199))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_1])])])).
cnf(c_0_304, plain, (r2_hidden(X1,X3)|~m2_finseq_1(X1,X2)|X3!=k13_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_285])).
cnf(c_0_305, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_286])).
cnf(c_0_306, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_133])])).
cnf(c_0_307, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|k1_funct_1(X3,X2)!=k4_struct_0(esk1_0)|X3!=k1_funct_1(esk4_0,X1)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_96]), c_0_211])])])).
cnf(c_0_308, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1)=k4_struct_0(esk1_0)|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_289]), c_0_224]), c_0_101]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_309, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,np__1)=k1_funct_1(esk4_0,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_289, c_0_290])])])).
cnf(c_0_310, negated_conjecture, (esk6_0=esk5_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_311, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|np__1=X1|X1=esk3_0|X1=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_208]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_241]), c_0_221])])).
cnf(c_0_312, negated_conjecture, (r1_xxreal_0(np__1,esk6_0)), inference(evalgc,[status(thm)],[c_0_293])).
cnf(c_0_313, negated_conjecture, (r1_xxreal_0(np__1,esk5_0)), inference(evalgc,[status(thm)],[c_0_294])).
cnf(c_0_314, negated_conjecture, (r1_xxreal_0(esk5_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_295])).
cnf(c_0_315, negated_conjecture, (r1_xxreal_0(esk6_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_316, negated_conjecture, (v7_ordinal1(esk6_0)), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_317, negated_conjecture, (v7_ordinal1(esk5_0)), inference(evalgc,[status(thm)],[c_0_298])).
cnf(c_0_318, plain, (X1=X2|X1=np__1|epred1_5(X3,X4,X1,X2,X5)), inference(cn,[status(thm)],[c_0_299])).
cnf(c_0_319, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|np__1=X1|X1=esk3_0|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_150]), c_0_224]), c_0_211])]), c_0_241]), c_0_231])])).
cnf(c_0_320, plain, (v1_xboole_0(X1)|m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_251])])).
cnf(c_0_321, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_301, c_0_87]), c_0_117])).
fof(c_0_322, plain, ![X246]:(~v1_relat_1(X246)|~v1_funct_1(X246)|~v1_finseq_1(X246)|~v1_matrix_1(X246)|m1_subset_1(k1_matrix_1(X246),k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_matrix_1])])).
cnf(c_0_323, plain, (k1_matrix_1(X1)=X2|v1_xboole_0(X3)|~m1_matrix_1(X1,X3,X2,X2)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_324, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_302])).
cnf(c_0_325, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_303])).
cnf(c_0_326, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_303])).
cnf(c_0_327, plain, (r2_hidden(X1,X3)|X3!=k13_finseq_1(X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_304])).
cnf(c_0_328, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_305])).
cnf(c_0_329, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_306, c_0_117])).
cnf(c_0_330, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,X1),X2)!=k4_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_307])])).
cnf(c_0_331, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1)=k4_struct_0(esk1_0)|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_308, c_0_290])])])).
cnf(c_0_332, negated_conjecture, (m2_finseq_1(k1_funct_1(esk4_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_309]), c_0_101]), c_0_224]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_333, negated_conjecture, (esk6_0=esk5_0|esk3_0=esk5_0|esk5_0=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311]), c_0_312]), c_0_313]), c_0_314]), c_0_315]), c_0_316]), c_0_317])])])).
cnf(c_0_334, plain, (X1=X2|X1=np__1|epred1_5(X3,X4,X1,X2,X5)), inference(evalgc,[status(thm)],[c_0_318])).
cnf(c_0_335, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|X1=esk3_0|np__1=X1|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_163]), c_0_224]), c_0_211])]), c_0_231])])).
fof(c_0_336, plain, ![X20, X21, X22]:(v2_struct_0(X20)|v6_struct_0(X20)|~v13_algstr_0(X20)|~v33_algstr_0(X20)|~v2_rlvect_1(X20)|~v3_rlvect_1(X20)|~v4_rlvect_1(X20)|~v3_group_1(X20)|~v5_group_1(X20)|~v4_vectsp_1(X20)|~v5_vectsp_1(X20)|~l6_algstr_0(X20)|(~v7_ordinal1(X21)|(~v7_ordinal1(X22)|(~r1_xxreal_0(np__1,X21)|~r1_xxreal_0(X21,X22)|k1_funct_1(k3_matrix14(X20,X22,X21),X21)=k5_struct_0(X20))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_matrix14])])])])).
fof(c_0_337, plain, ![X174]:(~l3_struct_0(X174)|m1_subset_1(k5_struct_0(X174),u1_struct_0(X174))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_struct_0])])).
fof(c_0_338, plain, ![X101]:((l3_struct_0(X101)|~l4_algstr_0(X101))&(l3_algstr_0(X101)|~l4_algstr_0(X101))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l4_algstr_0])])])).
cnf(c_0_339, plain, (v1_xboole_0(X1)|m2_finseq_1(k4_matrix14(X2,X3,X4),X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_321]), c_0_101])])])).
cnf(c_0_340, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_matrix_1(X1)), inference(split_conjunct,[status(thm)],[c_0_322])).
cnf(c_0_341, plain, (k1_matrix_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2)), inference(evalgc,[status(thm)],[c_0_323])).
cnf(c_0_342, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_113])])).
cnf(c_0_343, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_325])).
cnf(c_0_344, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_326])).
cnf(c_0_345, plain, (m2_finseq_1(X1,X3)|~r2_hidden(X1,X2)|X2!=k13_finseq_1(X3)), inference(split_conjunct,[status(thm)],[c_0_285])).
cnf(c_0_346, plain, (r2_hidden(k12_matrix_1(X1,X2),X3)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|X3!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_150])])).
cnf(c_0_347, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_108])])).
cnf(c_0_348, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_142])])).
cnf(c_0_349, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_232]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_241]), c_0_221])])).
cnf(c_0_350, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k4_struct_0(esk1_0)|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_331]), c_0_332]), c_0_209]), c_0_255]), c_0_210])])])).
cnf(c_0_351, negated_conjecture, (esk5_0=np__1|esk3_0=esk5_0|esk6_0=esk5_0), inference(csr,[status(thm)],[c_0_333, c_0_334])).
cnf(c_0_352, plain, (X1=np__1|X2=np__1|X1=np__1|epred1_5(X3,X2,X1,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_353, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|np__1=X3|X3=X2|X3=esk3_0|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X3),X4)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_335]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_241]), c_0_221])])).
cnf(c_0_354, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)=k3_matrix14(esk1_0,esk2_0,X1)|np__1=X1|X1=esk3_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_335]), c_0_224]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_355, plain, (v2_struct_0(X1)|v6_struct_0(X1)|k1_funct_1(k3_matrix14(X1,X3,X2),X2)=k5_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_336])).
cnf(c_0_356, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l3_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_337])).
cnf(c_0_357, plain, (l3_struct_0(X1)|~l4_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_338])).
cnf(c_0_358, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_359, plain, (v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|m2_finseq_1(k4_matrix14(X1,X3,X4),X2)|v6_struct_0(X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X3),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_113])])).
cnf(c_0_360, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_340])).
cnf(c_0_361, plain, (k1_matrix_1(k12_matrix_1(X1,X2))=X2|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_146])])).
cnf(c_0_362, plain, (v1_matrix_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_363, plain, (v1_relat_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_150])])).
cnf(c_0_364, plain, (v1_funct_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_150])])).
cnf(c_0_365, plain, (v1_finseq_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_150])])).
cnf(c_0_366, plain, (m2_finseq_1(X1,X3)|X2!=k13_finseq_1(X3)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_345])).
cnf(c_0_367, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k12_matrix_1(X1,X2),X3)|X3!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(csr,[status(thm)],[c_0_346, c_0_347])).
cnf(c_0_368, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),X6)|~m2_finseq_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_348, c_0_157])).
cnf(c_0_369, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_150]), c_0_224]), c_0_211])]), c_0_241]), c_0_231])])).
cnf(c_0_370, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k3_matrix14(esk1_0,esk2_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_275, c_0_290])])])).
cnf(c_0_371, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,X1)=k4_struct_0(esk1_0)|np__1=X1|esk5_0=np__1|esk6_0=esk5_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_351])])).
cnf(c_0_372, plain, (X2=np__1|X1=np__1|epred1_5(X3,X2,X1,X4,X5)), inference(cn,[status(thm)],[c_0_352])).
cnf(c_0_373, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|X3=esk3_0|np__1=X3|X3=X2|X1!=k1_funct_1(k3_matrix14(esk1_0,esk2_0,X3),X4)|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_354])])).
cnf(c_0_374, plain, (k1_funct_1(k3_matrix14(X1,X3,X2),X2)=k5_struct_0(X1)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(np__1,X2)), inference(evalgc,[status(thm)],[c_0_355])).
cnf(c_0_375, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l3_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_356])).
cnf(c_0_376, plain, (l3_struct_0(X1)|~l4_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_357])).
cnf(c_0_377, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_358])).
cnf(c_0_378, plain, (k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_164])).
cnf(c_0_379, plain, (X1=k3_matrix_1(X2,k4_matrix14(X3,X4,X5),X6,X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X8!=k1_funct_1(k4_matrix14(X3,X4,X5),X6)|X1!=k1_funct_1(X8,X7)|~v1_matrix_1(k4_matrix14(X3,X4,X5))|~m2_finseq_1(X8,X2)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,X4)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~m2_finseq_1(k12_matrix_1(X3,X4),k13_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_359]), c_0_86])])).
cnf(c_0_380, plain, (m1_subset_1(X2,k5_numbers)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v1_matrix_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_361])])).
cnf(c_0_381, plain, (v1_matrix_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(csr,[status(thm)],[c_0_362, c_0_347])).
cnf(c_0_382, plain, (v1_relat_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(csr,[status(thm)],[c_0_363, c_0_347])).
cnf(c_0_383, plain, (v1_funct_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(csr,[status(thm)],[c_0_364, c_0_347])).
cnf(c_0_384, plain, (v1_finseq_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(csr,[status(thm)],[c_0_365, c_0_347])).
cnf(c_0_385, plain, (m2_finseq_1(k12_matrix_1(X1,X2),X3)|v1_xboole_0(u1_struct_0(X1))|X4!=k13_finseq_1(X3)|X4!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_367])])).
cnf(c_0_386, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),X3)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_369]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_241]), c_0_221])])).
cnf(c_0_387, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k1_funct_1(esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_370, c_0_309])).
cnf(c_0_388, negated_conjecture, (esk5_0=esk3_0|esk5_0!=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_389, negated_conjecture, (esk6_0=esk5_0|esk5_0=np__1|esk6_0=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_371]), c_0_312]), c_0_315]), c_0_316])])])).
cnf(c_0_390, plain, (X2=np__1|X1=np__1|epred1_5(X3,X2,X1,X4,X5)), inference(evalgc,[status(thm)],[c_0_372])).
cnf(c_0_391, plain, (X1=np__1|epred1_5(X2,X1,X4,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X4,X1)!=k5_struct_0(X2)|k3_matrix_1(u1_struct_0(X2),X3,X4,X1)!=k5_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_392, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X3)|np__1=X3|X3=esk3_0|X3=X2|X1!=k5_struct_0(esk1_0)|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_374]), c_0_224]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_393, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l4_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_376])])).
cnf(c_0_394, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_62])])).
cnf(c_0_395, plain, (k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_378])).
cnf(c_0_396, plain, (X1=k3_matrix_1(X2,k4_matrix14(X3,X4,X5),X6,X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X8!=k1_funct_1(k4_matrix14(X3,X4,X5),X6)|X1!=k1_funct_1(X8,X7)|~m2_finseq_1(k12_matrix_1(X3,X4),k13_finseq_1(X2))|~m2_finseq_1(X8,X2)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,X4)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)), inference(csr,[status(thm)],[c_0_379, c_0_157])).
cnf(c_0_397, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_380, c_0_381]), c_0_382]), c_0_383]), c_0_384]), c_0_347])).
cnf(c_0_398, plain, (k1_funct_1(k4_matrix14(X1,X2,X4),X3)=k1_funct_1(k12_matrix_1(X1,X2),X3)|np__1=X3|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_143])])).
cnf(c_0_399, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k12_matrix_1(X1,X2),X3)|k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))!=k13_finseq_1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_385])])).
cnf(c_0_400, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X3)|~m2_finseq_1(k1_funct_1(esk4_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_386, c_0_387]), c_0_387])])).
cnf(c_0_401, negated_conjecture, (esk3_0=esk5_0|esk6_0!=esk5_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_388])).
cnf(c_0_402, negated_conjecture, (esk6_0=np__1|esk5_0=np__1|esk6_0=esk5_0), inference(csr,[status(thm)],[c_0_389, c_0_390])).
cnf(c_0_403, plain, (X1=np__1|epred1_5(X2,X1,X4,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X4,X1)!=k5_struct_0(X2)), inference(cn,[status(thm)],[c_0_391])).
cnf(c_0_404, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X3)|X3=esk3_0|np__1=X3|X3=X2|X1!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_227]), c_0_224]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_405, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_394])])).
cnf(c_0_406, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_85]), c_0_178])]), c_0_86])])).
cnf(c_0_407, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_150])])).
cnf(c_0_408, negated_conjecture, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_211]), c_0_241])])).
cnf(c_0_409, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X4)=k1_funct_1(k12_matrix_1(X1,X2),X4)|np__1=X4|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_399])])).
cnf(c_0_410, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X3)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_400, c_0_332])])).
cnf(c_0_411, negated_conjecture, (esk3_0=esk5_0|esk6_0=np__1|esk5_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk5_0)!=k5_struct_0(esk1_0)|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_402])])).
cnf(c_0_412, plain, (X1=np__1|epred1_5(X2,X1,X4,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X4,X1)!=k5_struct_0(X2)), inference(evalgc,[status(thm)],[c_0_403])).
cnf(c_0_413, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X2,X2)=k5_struct_0(esk1_0)|np__1=X2|X2=esk3_0|X2=X1|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_405]), c_0_211])])])).
cnf(c_0_414, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_406, c_0_101])])).
cnf(c_0_415, plain, (X1=k1_funct_1(esk9_5(X2,X3,X4,X5,X1),X5)|X1!=k3_matrix_1(X2,X3,X4,X5)|~m1_subset_1(X1,X2)|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_416, plain, (esk9_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~m1_subset_1(X5,X1)|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_90])).
fof(c_0_417, plain, ![X187, X188, X189, X190]:(~v1_matrix_1(X188)|~m1_finseq_1(X188,k3_finseq_2(X187))|~v7_ordinal1(X189)|~v7_ordinal1(X190)|m1_subset_1(k3_matrix_1(X187,X188,X189,X190),X187)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_matrix_1])])).
cnf(c_0_418, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_407, c_0_408]), c_0_117])).
cnf(c_0_419, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X4)=k1_funct_1(k12_matrix_1(X1,X2),X4)|np__1=X4|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_409, c_0_408])).
cnf(c_0_420, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X3)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_163]), c_0_290]), c_0_255]), c_0_224]), c_0_211])]), c_0_231])])).
cnf(c_0_421, negated_conjecture, (esk5_0=np__1|esk6_0=np__1|esk3_0=esk5_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk5_0)!=k5_struct_0(esk1_0)), inference(csr,[status(thm)],[c_0_411, c_0_412])).
cnf(c_0_422, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_208]), c_0_210])])])).
cnf(c_0_423, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_208]), c_0_209]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_424, plain, (X1=k1_funct_1(esk9_5(X2,X3,X4,X5,X1),X5)|X1!=k3_matrix_1(X2,X3,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k13_finseq_1(X2))|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3))), inference(rw,[status(thm)],[c_0_415, c_0_56])).
cnf(c_0_425, plain, (esk9_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k13_finseq_1(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))), inference(rw,[status(thm)],[c_0_416, c_0_56])).
cnf(c_0_426, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_417])).
cnf(c_0_427, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|np__1=X5|X5=X4|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_419])])).
cnf(c_0_428, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|np__1=X3|X1!=k4_struct_0(esk1_0)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_331])])).
cnf(c_0_429, negated_conjecture, (esk5_0=np__1|esk6_0=np__1|esk3_0=esk5_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_422]), c_0_313]), c_0_314]), c_0_317])])])).
cnf(c_0_430, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_423]), c_0_209]), c_0_210])])])).
cnf(c_0_431, plain, (k1_funct_1(esk9_5(X2,X3,X4,X5,X1),X5)=X1|X1!=k3_matrix_1(X2,X3,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k13_finseq_1(X2))|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3))), inference(evalgc,[status(thm)],[c_0_424])).
cnf(c_0_432, plain, (esk9_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k13_finseq_1(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))), inference(evalgc,[status(thm)],[c_0_425])).
cnf(c_0_433, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k13_finseq_1(X2))), inference(rw,[status(thm)],[c_0_426, c_0_56])).
cnf(c_0_434, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_163])])).
cnf(c_0_435, negated_conjecture, (esk6_0!=esk3_0|esk5_0!=esk3_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_436, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k4_struct_0(esk1_0)|np__1=X2|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_270])])).
cnf(c_0_437, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,esk5_0)=esk4_0|esk6_0=np__1|esk5_0=np__1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_429])])).
cnf(c_0_438, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_150]), c_0_224]), c_0_211])]), c_0_241]), c_0_231])])).
cnf(c_0_439, plain, (k1_funct_1(esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7),X6)=X7|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_124])])).
cnf(c_0_440, plain, (esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7)=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_124])])).
cnf(c_0_441, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k13_finseq_1(X2))), inference(evalgc,[status(thm)],[c_0_433])).
cnf(c_0_442, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_434, c_0_117])).
cnf(c_0_443, negated_conjecture, (esk3_0!=esk5_0|esk3_0!=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_435])).
cnf(c_0_444, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,X1)=k4_struct_0(esk1_0)|np__1=X1|esk5_0=np__1|esk6_0=np__1|~l2_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_436, c_0_437]), c_0_314]), c_0_313]), c_0_317])])])).
cnf(c_0_445, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~r1_xxreal_0(np__1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_163]), c_0_255]), c_0_224]), c_0_211])]), c_0_231])])).
cnf(c_0_446, plain, (k1_funct_1(esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7),X6)=X7|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X7,X1)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_439, c_0_157])).
cnf(c_0_447, plain, (esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7)=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X7,X1)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_440, c_0_157])).
cnf(c_0_448, plain, (m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k13_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_113])])).
cnf(c_0_449, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k1_funct_1(k12_matrix_1(X2,X3),X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_442])])).
cnf(c_0_450, negated_conjecture, (esk6_0=np__1|esk5_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk5_0)!=k4_struct_0(esk1_0)|esk3_0!=esk5_0|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_402])])).
cnf(c_0_451, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,X1)=k4_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_96]), c_0_211])])])).
cnf(c_0_452, plain, (X1=np__1|epred1_5(X2,X4,X1,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X1,X4)!=k5_struct_0(X2)|k3_matrix_1(u1_struct_0(X2),X3,X1,X4)!=k5_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_453, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k5_struct_0(esk1_0)|X3!=k1_funct_1(esk4_0,X1)|k1_funct_1(X3,X2)!=k5_struct_0(esk1_0)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_405]), c_0_211])])])).
cnf(c_0_454, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_445, c_0_290])])])).
cnf(c_0_455, plain, (k1_funct_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),X6)=X7|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X7,X1)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_446, c_0_447])])).
cnf(c_0_456, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_matrix_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_448, c_0_131])])).
cnf(c_0_457, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~m2_finseq_1(k3_matrix14(X2,X3,X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_182])])).
cnf(c_0_458, negated_conjecture, (esk6_0=esk3_0|esk6_0!=np__1|esk5_0!=esk3_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_459, negated_conjecture, (esk5_0=np__1|esk6_0=np__1|esk3_0!=esk5_0|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_451]), c_0_313]), c_0_314]), c_0_317])])])).
cnf(c_0_460, plain, (X1=np__1|epred1_5(X2,X4,X1,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X1,X4)!=k5_struct_0(X2)), inference(cn,[status(thm)],[c_0_452])).
cnf(c_0_461, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k5_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,X1),X2)!=k5_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_453])])).
cnf(c_0_462, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk3_0),np__1)=k5_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_309]), c_0_178]), c_0_290]), c_0_224]), c_0_101]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_463, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)=k3_matrix14(esk1_0,esk2_0,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_454]), c_0_209]), c_0_255]), c_0_224]), c_0_210]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_464, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=X3|X3!=k3_matrix_1(X4,esk4_0,X1,X2)|~m2_finseq_1(esk4_0,k13_finseq_1(X4))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,X4)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_455, c_0_208]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_241]), c_0_221])])).
cnf(c_0_465, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5),u1_struct_0(X1))|v6_struct_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_456, c_0_157]), c_0_117])).
cnf(c_0_466, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_457, c_0_227]), c_0_117])).
cnf(c_0_467, negated_conjecture, (esk3_0=esk6_0|esk3_0!=esk5_0|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_458])).
cnf(c_0_468, negated_conjecture, (esk6_0=np__1|esk5_0=np__1|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[c_0_459, c_0_429])).
cnf(c_0_469, plain, (X1=np__1|epred1_5(X2,X4,X1,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X1,X4)!=k5_struct_0(X2)), inference(evalgc,[status(thm)],[c_0_460])).
cnf(c_0_470, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,np__1)=k5_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462]), c_0_332]), c_0_178]), c_0_209]), c_0_290]), c_0_255]), c_0_101]), c_0_210])])])).
cnf(c_0_471, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,esk3_0)=k1_funct_1(esk4_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_463]), c_0_264]), c_0_276]), c_0_290]), c_0_178]), c_0_101])])])).
cnf(c_0_472, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=k3_matrix_1(X3,esk4_0,X1,X2)|~m2_finseq_1(esk4_0,k13_finseq_1(X3))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(X3,esk4_0,X1,X2),X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_464])])).
cnf(c_0_473, negated_conjecture, (m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2),u1_struct_0(esk1_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_208]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_241]), c_0_221])])).
cnf(c_0_474, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X5=X6|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_249])])).
cnf(c_0_475, negated_conjecture, (esk6_0=esk5_0|esk5_0=np__1|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk5_0,esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_467, c_0_351])])).
cnf(c_0_476, plain, (esk6_0=np__1|esk5_0=np__1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_390])])).
cnf(c_0_477, plain, (epred1_5(esk1_0,np__1,esk3_0,X1,esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_469, c_0_470]), c_0_247])])).
cnf(c_0_478, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X1!=k1_funct_1(k1_funct_1(esk4_0,np__1),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_471]), c_0_209]), c_0_255]), c_0_210]), c_0_202]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_247]), c_0_241]), c_0_221])])).
cnf(c_0_479, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_472, c_0_473]), c_0_264])])])).
cnf(c_0_480, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X6|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_474, c_0_117])).
cnf(c_0_481, negated_conjecture, (esk5_0=np__1|~epred1_5(esk1_0,np__1,esk5_0,esk5_0,esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_475, c_0_476])])).
cnf(c_0_482, negated_conjecture, (epred1_5(esk1_0,np__1,esk5_0,X1,esk4_0)|esk5_0=np__1|esk6_0=esk5_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_351])])).
cnf(c_0_483, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_478, c_0_479]), c_0_178]), c_0_290]), c_0_101])])])).
cnf(c_0_484, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k4_struct_0(X1)|np__1=X4|X4=X5|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_270])])).
cnf(c_0_485, negated_conjecture, (esk5_0=np__1|esk6_0=esk5_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_482])])).
cnf(c_0_486, negated_conjecture, (esk5_0=esk3_0|esk6_0!=np__1|esk5_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_487, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X1)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)|esk3_0=X2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_483])])).
cnf(c_0_488, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k4_struct_0(X1)|np__1=X4|X4=X3|X4=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_484, c_0_96])).
cnf(c_0_489, plain, (esk5_0=np__1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_485])])).
cnf(c_0_490, negated_conjecture, (esk3_0=esk5_0|esk5_0!=np__1|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_486])).
cnf(c_0_491, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2)=k4_struct_0(esk1_0)|esk3_0=X1|esk3_0=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_488]), c_0_209]), c_0_255]), c_0_210]), c_0_224]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_247]), c_0_241]), c_0_221])])).
cnf(c_0_492, negated_conjecture, (esk6_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_310, c_0_489]), c_0_489]), c_0_489])])).
cnf(c_0_493, negated_conjecture, (esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_490, c_0_489]), c_0_489]), c_0_489])]), c_0_247])])).
cnf(c_0_494, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k4_struct_0(esk1_0)|esk3_0=X1|esk3_0=X2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_473]), c_0_101])])])).
cnf(c_0_495, negated_conjecture, (esk5_0=esk3_0|esk6_0!=esk3_0|esk5_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_496, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0)), inference(csr,[status(thm)],[c_0_492, c_0_493])).
cnf(c_0_497, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k4_struct_0(esk1_0)|esk3_0=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_101]), c_0_247])])).
cnf(c_0_498, negated_conjecture, (esk3_0=esk5_0|esk5_0!=np__1|esk3_0!=esk6_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_495])).
cnf(c_0_499, plain, (X3=X5|X4=X5|epred1_5(X1,X4,X3,X5,X2)|k3_matrix_1(u1_struct_0(X1),X2,X3,X4)!=k4_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_500, negated_conjecture, (esk3_0=esk6_0|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_497]), c_0_312]), c_0_315]), c_0_316])])])).
cnf(c_0_501, negated_conjecture, (esk3_0!=esk6_0|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_498, c_0_489]), c_0_489]), c_0_489])]), c_0_247])])).
cnf(c_0_502, plain, (X4=X5|X3=X5|epred1_5(X1,X4,X3,X5,X2)|k3_matrix_1(u1_struct_0(X1),X2,X3,X4)!=k4_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_499])).
cnf(c_0_503, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,np__1),esk3_0)=k5_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_471]), c_0_209]), c_0_255]), c_0_224]), c_0_210]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_504, negated_conjecture, (m2_finseq_1(k1_funct_1(esk4_0,np__1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_471]), c_0_210]), c_0_224]), c_0_211]), c_0_212]), c_0_213]), c_0_214]), c_0_215]), c_0_216]), c_0_217]), c_0_218]), c_0_219]), c_0_220])]), c_0_221]), c_0_231])])).
cnf(c_0_505, negated_conjecture, (~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0)), inference(csr,[status(thm)],[c_0_500, c_0_501])).
cnf(c_0_506, plain, (np__1=X1|X2=X1|epred1_5(esk1_0,X2,np__1,X1,esk4_0)|esk3_0=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_497])])).
cnf(c_0_507, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk3_0)=k5_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_503]), c_0_504]), c_0_209]), c_0_178]), c_0_255]), c_0_290]), c_0_210]), c_0_101])])])).
cnf(c_0_508, negated_conjecture, (esk3_0=esk6_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_505, c_0_506]), c_0_312]), c_0_315]), c_0_316])]), c_0_247])])).
cnf(c_0_509, plain, (epred1_5(esk1_0,esk3_0,np__1,X1,esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_507]), c_0_247])])).
cnf(c_0_510, negated_conjecture, (~epred1_5(esk1_0,esk6_0,np__1,esk6_0,esk4_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_505, c_0_508])])).
cnf(c_0_511, plain, (epred1_5(esk1_0,esk6_0,np__1,X1,esk4_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_509, c_0_508])])).
cnf(c_0_512, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_510, c_0_511])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 513
# Proof object clause steps            : 422
# Proof object formula steps           : 91
# Proof object conjectures             : 176
# Proof object clause conjectures      : 173
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 89
# Proof object initial formulas used   : 45
# Proof object generating inferences   : 165
# Proof object simplifying inferences  : 626
# Training examples: 266 positive, 2464 negative
# Training: Positive examples begin
cnf(c_0_72, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_247, negated_conjecture, (esk3_0!=np__1)).# trainpos
cnf(c_0_220, negated_conjecture, (v13_algstr_0(esk1_0))).# trainpos
cnf(c_0_219, negated_conjecture, (v33_algstr_0(esk1_0))).# trainpos
cnf(c_0_218, negated_conjecture, (v2_rlvect_1(esk1_0))).# trainpos
cnf(c_0_217, negated_conjecture, (v3_rlvect_1(esk1_0))).# trainpos
cnf(c_0_216, negated_conjecture, (v4_rlvect_1(esk1_0))).# trainpos
cnf(c_0_215, negated_conjecture, (v3_group_1(esk1_0))).# trainpos
cnf(c_0_214, negated_conjecture, (v5_group_1(esk1_0))).# trainpos
cnf(c_0_213, negated_conjecture, (v4_vectsp_1(esk1_0))).# trainpos
cnf(c_0_231, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_221, negated_conjecture, (~v6_struct_0(esk1_0))).# trainpos
cnf(c_0_212, negated_conjecture, (v5_vectsp_1(esk1_0))).# trainpos
cnf(c_0_211, negated_conjecture, (l6_algstr_0(esk1_0))).# trainpos
cnf(c_0_490, negated_conjecture, (esk3_0=esk5_0|esk5_0!=np__1|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_210, negated_conjecture, (v7_ordinal1(esk3_0))).# trainpos
cnf(c_0_317, negated_conjecture, (v7_ordinal1(esk5_0))).# trainpos
cnf(c_0_316, negated_conjecture, (v7_ordinal1(esk6_0))).# trainpos
cnf(c_0_202, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_209, negated_conjecture, (r1_xxreal_0(np__1,esk3_0))).# trainpos
cnf(c_0_88, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_498, negated_conjecture, (esk3_0=esk5_0|esk5_0!=np__1|esk3_0!=esk6_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_178, plain, (r1_xxreal_0(np__1,np__1))).# trainpos
cnf(c_0_313, negated_conjecture, (r1_xxreal_0(np__1,esk5_0))).# trainpos
cnf(c_0_312, negated_conjecture, (r1_xxreal_0(np__1,esk6_0))).# trainpos
cnf(c_0_255, negated_conjecture, (r1_xxreal_0(esk3_0,esk2_0))).# trainpos
cnf(c_0_314, negated_conjecture, (r1_xxreal_0(esk5_0,esk2_0))).# trainpos
cnf(c_0_315, negated_conjecture, (r1_xxreal_0(esk6_0,esk2_0))).# trainpos
cnf(c_0_86, plain, (~v1_xboole_0(k13_finseq_1(X1)))).# trainpos
cnf(c_0_208, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,esk3_0)=esk4_0)).# trainpos
cnf(c_0_205, negated_conjecture, (m1_matrix_1(esk4_0,u1_struct_0(esk1_0),esk2_0,esk2_0))).# trainpos
cnf(c_0_443, negated_conjecture, (esk3_0!=esk5_0|esk3_0!=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_467, negated_conjecture, (esk3_0=esk6_0|esk3_0!=esk5_0|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_81, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_62, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_173, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_80, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_377, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1))).# trainpos
cnf(c_0_61, plain, (l4_struct_0(X1)|~l5_algstr_0(X1))).# trainpos
cnf(c_0_376, plain, (l3_struct_0(X1)|~l4_algstr_0(X1))).# trainpos
cnf(c_0_95, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_310, negated_conjecture, (esk6_0=esk5_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_401, negated_conjecture, (esk3_0=esk5_0|esk6_0!=esk5_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_71, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_94, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_76, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v7_struct_0(X1))).# trainpos
cnf(c_0_324, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_343, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_344, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_113, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_375, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l3_struct_0(X1))).# trainpos
cnf(c_0_270, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).# trainpos
cnf(c_0_151, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_328, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_327, plain, (r2_hidden(X1,X3)|X3!=k13_finseq_1(X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_233, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_122, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2))).# trainpos
cnf(c_0_390, plain, (X2=np__1|X1=np__1|epred1_5(X3,X2,X1,X4,X5))).# trainpos
cnf(c_0_366, plain, (m2_finseq_1(X1,X3)|X2!=k13_finseq_1(X3)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_334, plain, (X1=X2|X1=np__1|epred1_5(X3,X4,X1,X2,X5))).# trainpos
cnf(c_0_172, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_360, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_246, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_145, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_106, plain, (v2_struct_0(X1)|v6_struct_0(X1)|m1_matrix_1(k4_matrix14(X1,X2,X3),u1_struct_0(X1),X2,X2)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_469, plain, (X1=np__1|epred1_5(X2,X4,X1,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X1,X4)!=k5_struct_0(X2))).# trainpos
cnf(c_0_341, plain, (k1_matrix_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2))).# trainpos
cnf(c_0_146, plain, (v2_struct_0(X1)|m1_matrix_1(k12_matrix_1(X1,X2),u1_struct_0(X1),X2,X2)|~l6_algstr_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_144, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v1_matrix_1(X1)|~m2_finseq_1(X1,k13_finseq_1(X3))|~m1_matrix_1(X1,X3,X2,X4))).# trainpos
cnf(c_0_412, plain, (X1=np__1|epred1_5(X2,X1,X4,X5,X3)|k3_matrix_1(u1_struct_0(X2),X3,X4,X1)!=k5_struct_0(X2))).# trainpos
cnf(c_0_227, plain, (v2_struct_0(X1)|v6_struct_0(X1)|m2_finseq_1(k3_matrix14(X1,X2,X3),u1_struct_0(X1))|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_119, plain, (k7_partfun1(X3,X2,X1)=k1_funct_1(X2,X1)|~v7_ordinal1(X1)|~m2_finseq_1(X2,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2)))).# trainpos
cnf(c_0_118, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k13_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_441, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k13_finseq_1(X2)))).# trainpos
cnf(c_0_502, plain, (X4=X5|X3=X5|epred1_5(X1,X4,X3,X5,X2)|k3_matrix_1(u1_struct_0(X1),X2,X3,X4)!=k4_struct_0(X1))).# trainpos
cnf(c_0_84, plain, (v1_xboole_0(X1)|m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_251, plain, (k2_finseq_7(X1,X2,X3,X4)=k10_funct_7(X2,X3,X4)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_432, plain, (esk9_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k13_finseq_1(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2)))).# trainpos
cnf(c_0_431, plain, (k1_funct_1(esk9_5(X2,X3,X4,X5,X1),X5)=X1|X1!=k3_matrix_1(X2,X3,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k13_finseq_1(X2))|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3)))).# trainpos
cnf(c_0_374, plain, (k1_funct_1(k3_matrix14(X1,X3,X2),X2)=k5_struct_0(X1)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(np__1,X2))).# trainpos
cnf(c_0_85, plain, (k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_182, plain, (k1_funct_1(k12_matrix_1(X1,X3),X2)=k3_matrix14(X1,X3,X2)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(np__1,X2))).# trainpos
cnf(c_0_249, plain, (X2=X3|k1_funct_1(k3_matrix14(X1,X4,X2),X3)=k4_struct_0(X1)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(np__1,X3))).# trainpos
cnf(c_0_123, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|X5!=k1_funct_1(X1,X6)|X1!=k1_funct_1(X3,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X5,X2)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k13_finseq_1(X2))|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3)))).# trainpos
cnf(c_0_177, plain, (k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2)))).# trainpos
cnf(c_0_105, plain, (v1_xboole_0(X5)|r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~m1_matrix_1(X6,X5,X1,X2))).# trainpos
cnf(c_0_395, plain, (k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2)))).# trainpos
cnf(c_0_99, plain, (X4=X5|X3=X4|k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X5),X4)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,k3_finseq_1(X2)))).# trainpos
cnf(c_0_128, plain, (k3_finseq_1(k4_matrix14(X1,X2,X3))=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_224, negated_conjecture, (v7_ordinal1(esk2_0))).# trainpos
cnf(c_0_101, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_96, plain, (l2_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_394, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_77, plain, (l4_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_256, negated_conjecture, (k3_finseq_1(esk4_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(esk2_0))).# trainpos
cnf(c_0_258, negated_conjecture, (~v1_xxreal_0(np__1)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_230, negated_conjecture, (v1_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(esk2_0))).# trainpos
cnf(c_0_361, plain, (k1_matrix_1(k12_matrix_1(X1,X2))=X2|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_342, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_393, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l4_algstr_0(X1))).# trainpos
cnf(c_0_179, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_108, plain, (l1_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_150, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_242, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(esk1_0)))|~v7_ordinal1(esk2_0))).# trainpos
cnf(c_0_125, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_268, negated_conjecture, (~v1_xxreal_0(np__1)|~r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_93, plain, (v6_struct_0(X1)|~v7_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_100, plain, (m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_201, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_111, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X4)|np__1=X4|X4=X3|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_448, plain, (m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k13_finseq_1(X1)))).# trainpos
cnf(c_0_405, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l6_algstr_0(X1))).# trainpos
cnf(c_0_262, plain, (k1_funct_1(k2_finseq_7(X1,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_406, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_107, plain, (v6_struct_0(X1)|~l6_algstr_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).# trainpos
cnf(c_0_278, negated_conjecture, (~r1_xxreal_0(esk3_0,np__1))).# trainpos
cnf(c_0_155, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1))).# trainpos
cnf(c_0_154, plain, (k3_finseq_1(k12_matrix_1(X1,X2))=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_matrix_1(k12_matrix_1(X1,X2),u1_struct_0(X1),X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_121, plain, (m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_223, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_347, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_364, plain, (v1_funct_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_363, plain, (v1_relat_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_365, plain, (v1_finseq_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_423, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_126, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X4)=k1_funct_1(k4_matrix14(X1,X2,X3),X4)|np__1=X4|X4=X3|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_116, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_tarski(X2,X3),k2_matrix_1(k4_matrix14(X1,X4,X5)))|v6_struct_0(X1)|v2_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_163, plain, (k3_finseq_1(k12_matrix_1(X1,X2))=X2|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_241, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_234, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk2_0,np__1))).# trainpos
cnf(c_0_137, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v6_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_185, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_261, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_273, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_283, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_276, negated_conjecture, (k3_finseq_1(esk4_0)=esk2_0)).# trainpos
cnf(c_0_264, negated_conjecture, (m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_252, negated_conjecture, (v1_matrix_1(esk4_0))).# trainpos
cnf(c_0_306, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)|v6_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_398, plain, (k1_funct_1(k4_matrix14(X1,X2,X4),X3)=k1_funct_1(k12_matrix_1(X1,X2),X3)|np__1=X3|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_282, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|np__1=X1|X1=esk3_0|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_130, plain, (X1=k3_matrix_1(X2,k4_matrix14(X3,X4,X5),X6,X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X1!=k1_funct_1(X8,X7)|X8!=k1_funct_1(k4_matrix14(X3,X4,X5),X6)|~v1_matrix_1(k4_matrix14(X3,X4,X5))|~m2_finseq_1(k4_matrix14(X3,X4,X5),k13_finseq_1(X2))|~m2_finseq_1(X8,X2)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~m1_subset_1(X1,X2)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,X4)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3))).# trainpos
cnf(c_0_300, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|X1=esk3_0|np__1=X1|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_301, plain, (k10_funct_7(k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_263, negated_conjecture, (X1=k3_matrix_1(X2,esk4_0,X3,X4)|X5!=k1_funct_1(esk4_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_finseq_1(esk4_0,k13_finseq_1(X2))|~m2_finseq_1(X5,X2)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_136, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_148, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|np__1=X5|X5=X4|X7!=k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3))))).# trainpos
cnf(c_0_274, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X2,X3)|X4!=k1_funct_1(esk4_0,X2)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_319, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|np__1=X1|X1=esk3_0|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_453, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k5_struct_0(esk1_0)|X3!=k1_funct_1(esk4_0,X1)|k1_funct_1(X3,X2)!=k5_struct_0(esk1_0)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_152, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_335, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|X1=esk3_0|np__1=X1|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_153, plain, (v1_matrix_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_288, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|X3!=k1_funct_1(esk4_0,X1)|k1_funct_1(X3,X2)!=k4_struct_0(esk1_0)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).# trainpos
cnf(c_0_307, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|k1_funct_1(X3,X2)!=k4_struct_0(esk1_0)|X3!=k1_funct_1(esk4_0,X1)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_354, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)=k3_matrix14(esk1_0,esk2_0,X1)|np__1=X1|X1=esk3_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_461, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k5_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,X1),X2)!=k5_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_232, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_348, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_330, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,X1),X2)!=k4_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_243, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_440, plain, (esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7)=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_349, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_254, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_369, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_265, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0))).# trainpos
cnf(c_0_386, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),X3)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_362, plain, (v1_matrix_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_430, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_156, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_438, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).# trainpos
cnf(c_0_168, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_275, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k3_matrix14(esk1_0,esk2_0,np__1)|~r1_xxreal_0(np__1,esk2_0))).# trainpos
cnf(c_0_181, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k1_funct_1(k12_matrix_1(X2,X3),X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_289, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,np__1)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0))).# trainpos
cnf(c_0_445, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|~r1_xxreal_0(np__1,esk2_0))).# trainpos
cnf(c_0_353, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|np__1=X3|X3=X2|X3=esk3_0|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X3),X4)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_373, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|X3=esk3_0|np__1=X3|X3=X2|X1!=k1_funct_1(k3_matrix14(esk1_0,esk2_0,X3),X4)|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_245, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk2_0))).# trainpos
cnf(c_0_392, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X3)|np__1=X3|X3=esk3_0|X3=X2|X1!=k5_struct_0(esk1_0)|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_320, plain, (v1_xboole_0(X1)|m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_203, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~m2_finseq_1(k3_matrix14(X2,X3,X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_257, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk2_0))).# trainpos
cnf(c_0_308, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1)=k4_struct_0(esk1_0)|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,esk2_0))).# trainpos
cnf(c_0_339, plain, (v1_xboole_0(X1)|m2_finseq_1(k4_matrix14(X2,X3,X4),X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_237, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_404, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X3)|X3=esk3_0|np__1=X3|X3=X2|X1!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_259, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X5=X6|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_359, plain, (v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|m2_finseq_1(k4_matrix14(X1,X3,X4),X2)|v6_struct_0(X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X3),X2))).# trainpos
cnf(c_0_267, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk2_0))).# trainpos
cnf(c_0_413, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X2,X2)=k5_struct_0(esk1_0)|np__1=X2|X2=esk3_0|X2=X1|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_422, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_280, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k4_struct_0(X1)|np__1=X4|X4=X5|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_311, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|np__1=X1|X1=esk3_0|X1=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_333, negated_conjecture, (esk6_0=esk5_0|esk3_0=esk5_0|esk5_0=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_475, negated_conjecture, (esk6_0=esk5_0|esk5_0=np__1|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk5_0,esk4_0))).# trainpos
cnf(c_0_277, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_290, negated_conjecture, (r1_xxreal_0(np__1,esk2_0))).# trainpos
cnf(c_0_309, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,np__1)=k1_funct_1(esk4_0,esk3_0))).# trainpos
cnf(c_0_370, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k3_matrix14(esk1_0,esk2_0,np__1))).# trainpos
cnf(c_0_454, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1))).# trainpos
cnf(c_0_332, negated_conjecture, (m2_finseq_1(k1_funct_1(esk4_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_463, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)=k3_matrix14(esk1_0,esk2_0,esk3_0))).# trainpos
cnf(c_0_471, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,esk3_0)=k1_funct_1(esk4_0,np__1))).# trainpos
cnf(c_0_380, plain, (m1_subset_1(X2,k5_numbers)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v1_matrix_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_462, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk3_0),np__1)=k5_struct_0(esk1_0))).# trainpos
cnf(c_0_400, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X3)|~m2_finseq_1(k1_funct_1(esk4_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_470, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,np__1)=k5_struct_0(esk1_0))).# trainpos
cnf(c_0_504, negated_conjecture, (m2_finseq_1(k1_funct_1(esk4_0,np__1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_408, negated_conjecture, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_503, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,np__1),esk3_0)=k5_struct_0(esk1_0))).# trainpos
cnf(c_0_507, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk3_0)=k5_struct_0(esk1_0))).# trainpos
cnf(c_0_477, plain, (epred1_5(esk1_0,np__1,esk3_0,X1,esk4_0))).# trainpos
cnf(c_0_478, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X1!=k1_funct_1(k1_funct_1(esk4_0,np__1),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_420, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X3)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_509, plain, (epred1_5(esk1_0,esk3_0,np__1,X1,esk4_0))).# trainpos
cnf(c_0_482, negated_conjecture, (epred1_5(esk1_0,np__1,esk5_0,X1,esk4_0)|esk5_0=np__1|esk6_0=esk5_0)).# trainpos
cnf(c_0_346, plain, (r2_hidden(k12_matrix_1(X1,X2),X3)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|X3!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_385, plain, (m2_finseq_1(k12_matrix_1(X1,X2),X3)|v1_xboole_0(u1_struct_0(X1))|X4!=k13_finseq_1(X3)|X4!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_331, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1)=k4_struct_0(esk1_0)|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_350, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k4_struct_0(esk1_0)|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_428, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|np__1=X3|X1!=k4_struct_0(esk1_0)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_399, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k12_matrix_1(X1,X2),X3)|k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))!=k13_finseq_1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_409, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X4)=k1_funct_1(k12_matrix_1(X1,X2),X4)|np__1=X4|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_371, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,X1)=k4_struct_0(esk1_0)|np__1=X1|esk5_0=np__1|esk6_0=esk5_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_436, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k4_struct_0(esk1_0)|np__1=X2|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).# trainpos
cnf(c_0_389, negated_conjecture, (esk6_0=esk5_0|esk5_0=np__1|esk6_0=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_411, negated_conjecture, (esk3_0=esk5_0|esk6_0=np__1|esk5_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk5_0)!=k5_struct_0(esk1_0)|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_450, negated_conjecture, (esk6_0=np__1|esk5_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk5_0)!=k4_struct_0(esk1_0)|esk3_0!=esk5_0|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_429, negated_conjecture, (esk5_0=np__1|esk6_0=np__1|esk3_0=esk5_0)).# trainpos
cnf(c_0_437, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,esk5_0)=esk4_0|esk6_0=np__1|esk5_0=np__1)).# trainpos
cnf(c_0_456, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_matrix_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_444, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,X1)=k4_struct_0(esk1_0)|np__1=X1|esk5_0=np__1|esk6_0=np__1|~l2_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_451, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,X1)=k4_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_473, negated_conjecture, (m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2),u1_struct_0(esk1_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_459, negated_conjecture, (esk5_0=np__1|esk6_0=np__1|esk3_0!=esk5_0|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_476, plain, (esk6_0=np__1|esk5_0=np__1)).# trainpos
cnf(c_0_481, negated_conjecture, (esk5_0=np__1|~epred1_5(esk1_0,np__1,esk5_0,esk5_0,esk4_0))).# trainpos
cnf(c_0_485, negated_conjecture, (esk5_0=np__1|esk6_0=esk5_0)).# trainpos
cnf(c_0_489, plain, (esk5_0=np__1)).# trainpos
cnf(c_0_501, negated_conjecture, (esk3_0!=esk6_0|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0))).# trainpos
cnf(c_0_493, negated_conjecture, (esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0))).# trainpos
cnf(c_0_492, negated_conjecture, (esk6_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0))).# trainpos
cnf(c_0_439, plain, (k1_funct_1(esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7),X6)=X7|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_455, plain, (k1_funct_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),X6)=X7|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X7,X1)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_464, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=X3|X3!=k3_matrix_1(X4,esk4_0,X1,X2)|~m2_finseq_1(esk4_0,k13_finseq_1(X4))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,X4))).# trainpos
cnf(c_0_472, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=k3_matrix_1(X3,esk4_0,X1,X2)|~m2_finseq_1(esk4_0,k13_finseq_1(X3))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(X3,esk4_0,X1,X2),X3))).# trainpos
cnf(c_0_479, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_483, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_487, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X1)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)|esk3_0=X2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_379, plain, (X1=k3_matrix_1(X2,k4_matrix14(X3,X4,X5),X6,X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X8!=k1_funct_1(k4_matrix14(X3,X4,X5),X6)|X1!=k1_funct_1(X8,X7)|~v1_matrix_1(k4_matrix14(X3,X4,X5))|~m2_finseq_1(X8,X2)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,X4)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~m2_finseq_1(k12_matrix_1(X3,X4),k13_finseq_1(X2)))).# trainpos
cnf(c_0_407, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_427, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|np__1=X5|X5=X4|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3))))).# trainpos
cnf(c_0_434, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k1_funct_1(k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_449, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k1_funct_1(k12_matrix_1(X2,X3),X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_457, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X5),X6)|~m2_finseq_1(k3_matrix14(X2,X3,X5),u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_474, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X5=X6|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_484, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k4_struct_0(X1)|np__1=X4|X4=X5|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_491, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2)=k4_struct_0(esk1_0)|esk3_0=X1|esk3_0=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_494, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k4_struct_0(esk1_0)|esk3_0=X1|esk3_0=X2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_497, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k4_struct_0(esk1_0)|esk3_0=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_500, negated_conjecture, (esk3_0=esk6_0|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0))).# trainpos
cnf(c_0_506, plain, (np__1=X1|X2=X1|epred1_5(esk1_0,X2,np__1,X1,esk4_0)|esk3_0=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_508, negated_conjecture, (esk3_0=esk6_0)).# trainpos
cnf(c_0_510, negated_conjecture, (~epred1_5(esk1_0,esk6_0,np__1,esk6_0,esk4_0))).# trainpos
cnf(c_0_511, plain, (epred1_5(esk1_0,esk6_0,np__1,X1,esk4_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_133, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_280, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_8, negated_conjecture, (esk5_0!=np__1|esk6_0!=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_257, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_153, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_98, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_166, plain, (v1_xboole_0(o_0_0_xboole_0))).#trainneg
cnf(i_0_347, plain, (v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_345, plain, (v1_xxreal_0(esk13_0))).#trainneg
cnf(i_0_129, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_188, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_346, plain, (v1_xcmplx_0(esk13_0))).#trainneg
cnf(i_0_152, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_315, plain, (v3_ordinal1(esk11_0))).#trainneg
cnf(i_0_344, plain, (v1_xreal_0(esk13_0))).#trainneg
cnf(i_0_350, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_316, plain, (~v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_224, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_7, negated_conjecture, (esk5_0!=np__1|esk3_0!=esk6_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_267, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_176, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_301, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_157, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_247, plain, (v1_relat_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_246, plain, (v1_funct_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_227, plain, (v1_finseq_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_234, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_332, plain, (v4_funct_1(k13_finseq_1(X1)))).#trainneg
cnf(i_0_225, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_265, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_215, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_235, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_338, plain, (k6_xcmplx_0(np__2,np__2)=np__0)).#trainneg
cnf(i_0_126, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_186, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_185, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_309, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_6, negated_conjecture, (esk3_0!=esk5_0|esk6_0!=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_308, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_16, negated_conjecture, (esk3_0=esk6_0|esk5_0!=np__1|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_312, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_291, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_261, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_259, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_266, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_320, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_304, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_335, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_289, plain, (~v1_xboole_0(k9_finseq_1(X1)))).#trainneg
cnf(i_0_282, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_214, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_317, plain, (m1_subset_1(esk11_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_325, plain, (v3_card_1(k1_tarski(X1),np__1))).#trainneg
cnf(i_0_281, plain, (m1_subset_1(esk10_1(X1),X1))).#trainneg
cnf(i_0_303, plain, (m1_finseq_2(k13_finseq_1(X1),X1))).#trainneg
cnf(i_0_128, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_322, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_187, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_310, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_288, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_263, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_262, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107, plain, (l2_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_86, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_106, plain, (l3_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_100, plain, (l3_algstr_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_105, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_1, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_252, plain, (l1_struct_0(X1)|~l3_struct_0(X1))).#trainneg
cnf(i_0_10, negated_conjecture, (esk3_0=esk6_0|esk6_0!=esk5_0|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk6_0)!=k5_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_244, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_295, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_223, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_275, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_292, plain, (v1_relat_1(k1_tarski(k4_tarski(X1,X2))))).#trainneg
cnf(i_0_184, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_287, plain, (v1_xboole_0(k10_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_256, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_151, plain, (v1_xboole_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_248, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_271, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_217, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_270, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_70, plain, (X1=k9_finseq_1(X2)|k3_finseq_1(X1)!=np__1|k1_tarski(X2)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_228, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_314, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_269, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_285, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_258, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_294, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_221, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_337, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_138, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_137, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_311, plain, (v3_card_1(k1_finseq_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_279, plain, (v9_struct_0(k4_struct_0(X1),X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_323, plain, (v9_struct_0(esk12_1(X1),X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_226, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_272, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_250, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_189, plain, (m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_283, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_328, plain, (v1_finseq_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_286, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_299, plain, (v2_struct_0(X1)|v1_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_298, plain, (v2_struct_0(X1)|v2_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_242, plain, (v2_struct_0(X1)|v3_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_117, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_274, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_136, plain, (v1_relat_1(k2_finseq_2(X1,X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_135, plain, (v1_funct_1(k2_finseq_2(X1,X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_134, plain, (v1_finseq_1(k2_finseq_2(X1,X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_324, plain, (m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_245, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_302, plain, (v1_xboole_0(k1_finseq_1(X1))|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_229, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v9_struct_0(k5_struct_0(X1),X1))).#trainneg
cnf(i_0_241, plain, (v2_struct_0(X1)|v6_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_333, plain, (v2_struct_0(X1)|v1_group_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_321, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k10_xtuple_0(X1)))).#trainneg
cnf(i_0_290, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k9_xtuple_0(X1)))).#trainneg
cnf(i_0_145, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_268, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_326, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_318, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_319, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_167, plain, (m1_finseq_2(k4_finseq_2(X1,X2),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_181, plain, (v3_card_1(k2_finseq_2(X1,X2),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_172, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_168, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_348, plain, (v7_struct_0(X1)|~l2_struct_0(X1)|~v9_struct_0(esk14_1(X1),X1))).#trainneg
cnf(i_0_173, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_254, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_264, plain, (v1_xreal_0(k1_xreal_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_273, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_156, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3))).#trainneg
cnf(i_0_219, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_77, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__1)|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_341, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_237, plain, (v1_xboole_0(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k4_finseq_2(X1,X2)))).#trainneg
cnf(i_0_307, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_155, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3))).#trainneg
cnf(i_0_260, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_165, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_161, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_365, plain, (X1=X2|X1=np__1|epred1_5(X3,X1,X4,X2,X5))).#trainneg
cnf(i_0_349, plain, (v7_struct_0(X1)|m1_subset_1(esk14_1(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_146, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_102, plain, (k1_xreal_0(X1,X2)=k7_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109, plain, (k12_finseq_1(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_119, plain, (m1_subset_1(k7_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_97, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_175, plain, (v1_relat_1(k2_funct_7(X1,X2,X3))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_352, plain, (epred1_5(X1,X4,X3,X5,X2)|k3_matrix_1(u1_struct_0(X1),X2,X3,X4)!=k5_struct_0(X1)|k3_matrix_1(u1_struct_0(X1),X2,X3,X4)!=k4_struct_0(X1))).#trainneg
cnf(i_0_59, plain, (v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|k3_finseq_1(X4)!=X2|k3_finseq_1(esk8_4(X1,X2,X3,X4))!=X3|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1)))).#trainneg
cnf(i_0_72, plain, (k3_finseq_1(X1)=np__1|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_121, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1))).#trainneg
cnf(i_0_296, plain, (v1_xboole_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_297, plain, (v1_xboole_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_342, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_76, plain, (k1_funct_1(X1,np__1)=X2|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_71, plain, (k1_tarski(X2)=k10_xtuple_0(X1)|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_174, plain, (v1_funct_1(k2_funct_7(X1,X2,X3))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_253, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_104, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_96, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_108, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_125, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_120, plain, (v1_xboole_0(X1)|m2_finseq_1(k12_finseq_1(X1,X2),X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_313, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_158, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_131, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_56, plain, (X2=k13_finseq_1(X1)|~r2_hidden(esk7_2(X1,X2),X2)|~m2_finseq_1(esk7_2(X1,X2),X1))).#trainneg
cnf(i_0_75, plain, (X1=k9_finseq_1(X2)|k1_funct_1(X1,np__1)!=X2|k9_xtuple_0(X1)!=k2_finseq_1(np__1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_232, plain, (v1_relat_1(k1_rfinseq(X1,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2))).#trainneg
cnf(i_0_171, plain, (v1_relat_1(k16_finseq_1(X1,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2))).#trainneg
cnf(i_0_231, plain, (v1_funct_1(k1_rfinseq(X1,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2))).#trainneg
cnf(i_0_170, plain, (v1_funct_1(k16_finseq_1(X1,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2))).#trainneg
cnf(i_0_113, plain, (k2_rfinseq(X3,X1,X2)=k1_rfinseq(X1,X2)|~v7_ordinal1(X1)|~m1_finseq_1(X2,X3))).#trainneg
cnf(i_0_114, plain, (k17_finseq_1(X3,X1,X2)=k16_finseq_1(X1,X2)|~v7_ordinal1(X1)|~m1_finseq_1(X2,X3))).#trainneg
cnf(i_0_142, plain, (m2_finseq_1(k2_rfinseq(X3,X1,X2),X3)|~v7_ordinal1(X1)|~m1_finseq_1(X2,X3))).#trainneg
cnf(i_0_140, plain, (m2_finseq_1(k17_finseq_1(X3,X1,X2),X3)|~v7_ordinal1(X1)|~m1_finseq_1(X2,X3))).#trainneg
cnf(i_0_351, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55, plain, (X2=k13_finseq_1(X1)|r2_hidden(esk7_2(X1,X2),X2)|m2_finseq_1(esk7_2(X1,X2),X1))).#trainneg
cnf(i_0_116, plain, (k8_finseq_1(X2,X1,X3)=k7_finseq_1(X1,X3)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_220, plain, (m2_finseq_1(k8_finseq_1(X2,X1,X3),X2)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_238, plain, (v1_xboole_0(X1)|v3_card_1(X3,X2)|~v7_ordinal1(X2)|~m1_subset_1(X3,k4_finseq_2(X2,X1)))).#trainneg
cnf(i_0_230, plain, (v1_finseq_1(k1_rfinseq(X1,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2))).#trainneg
cnf(i_0_169, plain, (v1_finseq_1(k16_finseq_1(X1,X2))|~v7_ordinal1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2))).#trainneg
cnf(i_0_160, plain, (k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_179, plain, (v1_relat_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_122, plain, (k5_finseq_2(X1,X2,X3)=k2_finseq_2(X2,X3)|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_178, plain, (v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_177, plain, (v1_finseq_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_123, plain, (k2_nat_1(k3_finseq_1(X1),k3_finseq_1(X2))=k3_finseq_1(k7_finseq_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_233, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_218, plain, (m1_subset_1(k7_partfun1(X2,X1,X3),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_65, plain, (k2_matrix_1(X1)=k2_zfmisc_1(k2_finseq_1(X2),k2_finseq_1(X2))|v1_xboole_0(X3)|~v7_ordinal1(X2)|~m1_matrix_1(X1,X3,X2,X2))).#trainneg
cnf(i_0_149, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_148, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_239, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_103, plain, (k7_finseq_1(k7_finseq_1(X1,X2),X3)=k7_finseq_1(X1,k7_finseq_1(X2,X3))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_92, plain, (k1_funct_1(k7_finseq_1(X1,X2),X3)=k1_funct_1(X1,X3)|~v7_ordinal1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,k3_finseq_1(X1)))).#trainneg
cnf(i_0_144, plain, (v1_xboole_0(X1)|m2_finseq_2(k5_finseq_2(X1,X2,X3),X1,k4_finseq_2(X2,X1))|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_110, plain, (v1_xboole_0(X1)|m2_finseq_1(k1_finseq_7(X1,X2,X3,X4),X1)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_130, plain, (k1_finseq_7(X1,X2,X3,X4)=k2_funct_7(X2,X3,X4)|v1_xboole_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_61, plain, (v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|m2_finseq_1(esk8_4(X1,X2,X3,X4),X1)|k3_finseq_1(X4)!=X2|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1)))).#trainneg
cnf(i_0_62, plain, (k3_finseq_1(X1)=X4|v1_xboole_0(X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X1,X2)|~r2_hidden(X1,k10_xtuple_0(X3))|~m2_finseq_1(X3,k13_finseq_1(X2))|~m1_matrix_1(X3,X2,X5,X4))).#trainneg
cnf(i_0_60, plain, (v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|r2_hidden(esk8_4(X1,X2,X3,X4),k10_xtuple_0(X4))|k3_finseq_1(X4)!=X2|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1)))).#trainneg
cnf(i_0_85, plain, (m2_finseq_1(esk9_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k13_finseq_1(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_80, plain, (k1_finseq_7(u1_struct_0(X1),k5_finseq_2(u1_struct_0(X1),X2,k4_struct_0(X1)),X3,k5_struct_0(X1))=k3_matrix14(X1,X2,X3)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_90, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2)))).#trainneg
cnf(i_0_379, plain, (esk13_0=k1_xboole_0)).#trainneg
cnf(i_0_639, plain, (k1_matrix_1(k4_matrix14(X1,X2,X3))=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_546, plain, (k3_finseq_1(k9_finseq_1(X1))=np__1)).#trainneg
cnf(i_0_587, negated_conjecture, (esk2_0=esk5_0|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,esk5_0))).#trainneg
cnf(i_0_586, negated_conjecture, (esk2_0=esk6_0|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,esk6_0))).#trainneg
cnf(i_0_561, plain, (k1_funct_1(k9_finseq_1(X1),np__1)=X1)).#trainneg
cnf(i_0_588, negated_conjecture, (esk2_0=esk3_0|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,esk3_0))).#trainneg
cnf(i_0_671, plain, (r2_hidden(k3_matrix14(X1,X2,X3),X4)|v6_struct_0(X1)|v2_struct_0(X1)|X4!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_592, negated_conjecture, (esk5_0=np__1|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_490, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_384, plain, (v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_380, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_428, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_385, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_467, negated_conjecture, (v3_ordinal1(esk2_0))).#trainneg
cnf(i_0_476, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_491, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_591, negated_conjecture, (esk6_0=np__1|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_497, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_498, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_578, plain, (v1_xboole_0(X1)|m2_finseq_1(k9_finseq_1(X2),X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_492, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_390, plain, (l1_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_396, plain, (v7_ordinal1(np__2))).#trainneg
cnf(i_0_470, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_595, plain, (k2_xcmplx_0(X1,X2)=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_389, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_391, plain, (l1_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_392, plain, (v7_ordinal1(esk10_1(k5_numbers)))).#trainneg
cnf(i_0_399, plain, (v3_ordinal1(k1_card_1(X1)))).#trainneg
cnf(i_0_547, plain, (m2_finseq_1(X1,X2)|~m1_finseq_2(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_677, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_574, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_650, plain, (k3_finseq_1(k12_finseq_1(k13_finseq_1(X1),X2))=X3|v1_xboole_0(X1)|~m1_matrix_1(k12_finseq_1(k13_finseq_1(X1),X2),X1,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k13_finseq_1(X1)))).#trainneg
cnf(i_0_641, plain, (k3_finseq_1(k12_matrix_1(X1,X2))=X2|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_537, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_638, negated_conjecture, (k1_matrix_1(esk4_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(esk2_0))).#trainneg
cnf(i_0_439, plain, (~v1_xboole_0(k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_447, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk10_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_472, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_401, plain, (v7_ordinal1(k10_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_434, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_403, plain, (v7_ordinal1(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_405, plain, (v7_ordinal1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_499, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_682, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_598, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_670, plain, (k2_matrix_1(k4_matrix14(X1,X2,X3))=k2_zfmisc_1(k2_finseq_1(X2),k2_finseq_1(X2))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_515, plain, (v1_xboole_0(X1)|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_417, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_516, plain, (m2_finseq_1(X1,X2)|v1_xboole_0(X3)|X3!=k13_finseq_1(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_445, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l4_struct_0(X1))).#trainneg
cnf(i_0_563, plain, (v7_ordinal1(k6_xcmplx_0(X1,X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_448, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_461, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_478, plain, (v1_relat_1(esk11_0)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_436, plain, (~v1_xboole_0(k1_numbers)|~r2_hidden(X1,esk11_0))).#trainneg
cnf(i_0_407, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_481, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_487, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_435, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_409, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_412, plain, (v1_xreal_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_575, plain, (r2_hidden(k12_finseq_1(X1,X2),X3)|v1_xboole_0(X1)|X3!=k13_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_706, plain, (k2_finseq_7(X1,X2,X3,X4)=k2_finseq_7(X5,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X5))).#trainneg
cnf(i_0_379, plain, (esk13_0=np__0)).#trainneg
cnf(i_0_415, plain, (X1=k1_card_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_133, plain, (o_0_0_xboole_0=np__0)).#trainneg
cnf(i_0_345, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_418, plain, (X1=k9_xtuple_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_419, plain, (X1=k10_xtuple_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_851, plain, (~v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_780, negated_conjecture, (esk2_0=esk3_0|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk2_0,esk3_0))).#trainneg
cnf(i_0_772, plain, (k3_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2))=k2_nat_1(np__1,k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_804, plain, (np__1=X2|v1_xboole_0(X3)|~m1_matrix_1(k9_finseq_1(X1),X3,X2,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~m1_subset_1(X1,k13_finseq_1(X3)))).#trainneg
cnf(i_0_872, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_786, negated_conjecture, (esk5_0=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_700, plain, (m2_finseq_1(k5_finseq_2(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_finseq_2(k4_finseq_2(X2,X1),X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_867, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_793, negated_conjecture, (esk6_0=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_346, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_314, plain, (v1_xboole_0(X1)|~v3_card_1(X1,np__0))).#trainneg
cnf(i_0_424, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_774, negated_conjecture, (esk2_0=esk5_0|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk2_0,esk5_0))).#trainneg
cnf(i_0_344, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_868, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_430, plain, (v3_card_1(k2_finseq_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_431, plain, (v7_ordinal1(k1_card_1(X1))|~v3_ordinal1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_848, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk4_0),k2_finseq_1(esk2_0))=k2_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0)|~v1_matrix_1(esk4_0))).#trainneg
cnf(i_0_802, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_798, negated_conjecture, (k2_nat_1(esk2_0,k3_finseq_1(X1))=k3_finseq_1(k7_finseq_1(esk4_0,X1))|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(X1)|~v1_finseq_1(esk4_0)|~v1_funct_1(X1)|~v1_funct_1(esk4_0)|~v1_relat_1(X1)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_800, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_738, plain, (k1_funct_1(k2_finseq_7(X1,X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X3,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_746, plain, (k7_partfun1(X1,k10_funct_7(X2,X3,X4),X5)=k7_partfun1(X1,X2,X5)|X3=X5|X5=X4|v1_xboole_0(X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X5,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_791, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_776, negated_conjecture, (esk2_0=esk6_0|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk2_0,esk6_0))).#trainneg
cnf(i_0_788, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_808, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_697, plain, (k7_finseq_1(X1,X2)=k9_finseq_1(X3)|k9_xtuple_0(k7_finseq_1(X1,X2))!=k2_finseq_1(np__1)|k1_funct_1(X1,np__1)!=X3|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_446, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk11_0))).#trainneg
cnf(i_0_933, plain, (v1_relat_1(k5_numbers)|~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_449, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_451, plain, (v7_ordinal1(k1_finseq_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_702, plain, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),X2,X3,X4))=X5|v1_xboole_0(X1)|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),X2,X3,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_finseq_1(X2,k13_finseq_1(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_853, plain, (k2_zfmisc_1(k4_finseq_1(k12_matrix_1(X1,X2)),k2_finseq_1(X2))=k2_matrix_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v1_matrix_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_765, plain, (k2_zfmisc_1(k4_finseq_1(k4_matrix14(X1,X2,X3)),k2_finseq_1(X2))=k2_matrix_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v1_matrix_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_838, plain, (k3_finseq_1(k7_finseq_1(X1,k12_matrix_1(X2,X3)))=k2_nat_1(k3_finseq_1(X1),X3)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|~v1_finseq_1(k12_matrix_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(X2,X3))|~v1_relat_1(X1)|~v7_ordinal1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_943, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_845, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_882, plain, (k2_matrix_1(k4_matrix14(X1,X2,X3))=k2_matrix_1(k4_matrix14(X4,X2,X5))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|v2_struct_0(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_929, plain, (~v1_xboole_0(k1_numbers)|~m1_subset_1(X1,esk11_0))).#trainneg
cnf(i_0_452, plain, (X1=k1_finseq_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_460, plain, (v3_xxreal_0(np__0)|~v1_xreal_0(np__0)|~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_944, plain, (v1_xcmplx_0(np__1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_947, plain, (v1_xreal_0(np__1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_474, plain, (v3_ordinal1(k4_struct_0(X1))|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_559, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_453, plain, (v1_xboole_0(k2_finseq_1(X1))|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_458, plain, (v3_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1030, negated_conjecture, (esk5_0=np__1|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_1028, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_1052, negated_conjecture, (esk6_0=np__1|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_1136, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_970, negated_conjecture, (k3_finseq_1(esk4_0)=X1|v1_xboole_0(u1_struct_0(esk1_0))|~m1_matrix_1(esk4_0,u1_struct_0(esk1_0),X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_975, plain, (k3_finseq_1(k4_matrix14(X1,X2,X3))=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_matrix_1(k4_matrix14(X1,X2,X3),u1_struct_0(X1),X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1093, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_1079, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1080, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|r1_xxreal_0(esk5_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_968, negated_conjecture, (v1_funct_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1022, negated_conjecture, (esk2_0=esk3_0|~r1_xxreal_0(esk2_0,esk3_0))).#trainneg
cnf(i_0_459, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1063, negated_conjecture, (esk2_0=esk5_0|~r1_xxreal_0(esk2_0,esk5_0))).#trainneg
cnf(i_0_1015, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_969, negated_conjecture, (v1_finseq_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1039, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_464, plain, (v3_ordinal1(esk12_1(X1))|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_466, plain, (v3_ordinal1(esk10_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1114, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_1134, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_1104, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_1116, negated_conjecture, (esk2_0=esk6_0|~r1_xxreal_0(esk2_0,esk6_0))).#trainneg
cnf(i_0_1081, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(esk3_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1077, plain, (v7_ordinal1(k1_card_1(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1221, plain, (~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_1140, plain, (k7_finseq_1(k9_finseq_1(X1),X2)=k9_finseq_1(X1)|k9_xtuple_0(k7_finseq_1(k9_finseq_1(X1),X2))!=k2_finseq_1(np__1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_473, plain, (v3_ordinal1(k2_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_797, negated_conjecture, (m2_subset_1(esk2_0,k1_numbers,k5_numbers)|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_945, plain, (v1_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_948, plain, (v1_xreal_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_967, negated_conjecture, (v1_relat_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1082, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk2_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_479, plain, (v1_relat_1(esk10_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_480, plain, (v1_relat_1(k2_finseq_1(X1))|~v1_relat_1(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_977, plain, (k2_finseq_7(X4,k12_matrix_1(X1,X2),np__1,X3)=k4_matrix14(X1,X2,X3)|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),X4)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1002, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_981, plain, (k7_partfun1(X1,k2_finseq_7(X5,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|v1_xboole_0(X5)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X5)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_1074, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1055, plain, (v1_xboole_0(k1_finseq_1(np__0)))).#trainneg
cnf(i_0_1075, plain, (v1_xboole_0(k2_finseq_1(np__0)))).#trainneg
cnf(i_0_1138, plain, (k7_finseq_1(X1,X2)=k9_finseq_1(k1_funct_1(X1,np__1))|k9_xtuple_0(k7_finseq_1(X1,X2))!=k2_finseq_1(np__1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(X1)))).#trainneg
cnf(i_0_494, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(np__2)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_495, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_887, negated_conjecture, (v1_xboole_0(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1228, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1150, plain, (v3_ordinal1(k5_struct_0(X1))|~v3_ordinal1(u1_struct_0(X1))|~l6_algstr_0(X1))).#trainneg
cnf(i_0_500, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(np__2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_503, plain, (k9_xtuple_0(k9_finseq_1(X1))=k2_finseq_1(np__1))).#trainneg
cnf(i_0_982, plain, (k7_partfun1(X1,k2_finseq_7(X5,X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X5)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_1163, plain, (k2_zfmisc_1(k9_xtuple_0(k12_matrix_1(X1,X2)),k2_finseq_1(X2))=k2_matrix_1(k12_matrix_1(X1,X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v1_matrix_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_1209, negated_conjecture, (k2_matrix_1(k4_matrix14(X1,esk2_0,X2))=k2_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1092, negated_conjecture, (k2_zfmisc_1(k9_xtuple_0(esk4_0),k2_finseq_1(esk2_0))=k2_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_1106, plain, (k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3)=k1_funct_1(k4_matrix14(X1,X2,X3),np__1)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1190, plain, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_1056, plain, (v1_xboole_0(k2_finseq_2(np__0,X1)))).#trainneg
cnf(i_0_1091, negated_conjecture, (v1_relat_1(k2_matrix_1(esk4_0))|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_1139, plain, (k7_finseq_1(k7_finseq_1(X1,X2),X3)=k9_finseq_1(X4)|k9_xtuple_0(k7_finseq_1(k7_finseq_1(X1,X2),X3))!=k2_finseq_1(np__1)|k1_funct_1(X1,np__1)!=X4|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,X2)))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(X1)))).#trainneg
cnf(i_0_504, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_1183, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_668, negated_conjecture, (k2_zfmisc_1(k2_finseq_1(esk2_0),k2_finseq_1(esk2_0))=k2_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(esk2_0))).#trainneg
cnf(i_0_1099, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,esk4_0))=k2_nat_1(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_543, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_925, plain, (m2_subset_1(X1,X2,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_510, plain, (k3_finseq_1(k1_card_1(X1))=k1_card_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_511, plain, (v1_card_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1268, negated_conjecture, (r1_xxreal_0(esk5_0,esk5_0))).#trainneg
cnf(i_0_1253, negated_conjecture, (r1_xxreal_0(esk6_0,esk6_0))).#trainneg
cnf(i_0_1286, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk5_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1297, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0))).#trainneg
cnf(i_0_1254, negated_conjecture, (r1_xxreal_0(esk6_0,esk5_0)|r1_xxreal_0(esk5_0,esk6_0))).#trainneg
cnf(i_0_1244, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1289, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1245, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk3_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1392, plain, (k1_finseq_1(np__0)=np__0)).#trainneg
cnf(i_0_1482, plain, (k7_finseq_1(X1,k7_finseq_1(X2,X3))=k9_finseq_1(X4)|k9_xtuple_0(k7_finseq_1(X1,k7_finseq_1(X2,X3)))!=k2_finseq_1(np__1)|k1_funct_1(X1,np__1)!=X4|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(X1)))).#trainneg
cnf(i_0_512, plain, (k1_card_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1240, negated_conjecture, (r1_xxreal_0(np__0,esk3_0))).#trainneg
cnf(i_0_1434, plain, (~v1_xboole_0(k2_finseq_1(np__1)))).#trainneg
cnf(i_0_1349, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0))).#trainneg
cnf(i_0_513, plain, (v1_xboole_0(k3_finseq_1(X1))|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_514, plain, (v1_finset_1(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1269, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(esk3_0,esk5_0))).#trainneg
cnf(i_0_1287, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk5_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1247, negated_conjecture, (k3_finseq_1(esk4_0)=X1|v1_xboole_0(u1_struct_0(esk1_0))|~m1_matrix_1(esk4_0,u1_struct_0(esk1_0),X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1250, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1361, plain, (k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)=k4_matrix14(X2,X3,X4)|v1_xboole_0(X1)|v6_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_1241, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1292, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1315, negated_conjecture, (m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(esk1_0)))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1255, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(esk3_0,esk6_0))).#trainneg
cnf(i_0_1251, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1290, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk6_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1242, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1400, plain, (k2_finseq_1(np__0)=np__0)).#trainneg
cnf(i_0_1404, plain, (k7_finseq_1(X1,k7_finseq_1(X2,X3))=k9_finseq_1(k1_funct_1(k7_finseq_1(X1,X2),np__1))|k9_xtuple_0(k7_finseq_1(X1,k7_finseq_1(X2,X3)))!=k2_finseq_1(np__1)|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,X2)))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_517, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1249, negated_conjecture, (r1_xxreal_0(np__0,esk5_0))).#trainneg
cnf(i_0_1237, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_519, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_529, plain, (v3_ordinal1(esk14_1(X1))|v7_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1293, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1291, negated_conjecture, (r1_xxreal_0(np__0,esk6_0))).#trainneg
cnf(i_0_740, plain, (k7_partfun1(X1,k10_funct_7(X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_734, plain, (k7_partfun1(X1,k10_funct_7(X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_1444, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),X1)=k1_funct_1(esk4_0,X1)|np__1=X1|X1=esk3_0|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(esk4_0))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1374, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1474, plain, (v7_ordinal1(k2_finseq_2(np__0,X1)))).#trainneg
cnf(i_0_1412, negated_conjecture, (u1_struct_0(esk1_0)=np__0|v1_xboole_0(esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1481, plain, (k7_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)),X4)=k9_finseq_1(X5)|k9_xtuple_0(k7_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)),X4))!=k2_finseq_1(np__1)|k1_funct_1(k7_finseq_1(X1,X2),np__1)!=X5|~v1_finseq_1(X4)|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3))))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,X2)))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_531, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1748, negated_conjecture, (m1_matrix_1(esk4_0,np__0,esk2_0,esk2_0)|v1_xboole_0(esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1514, plain, (v1_card_1(np__1))).#trainneg
cnf(i_0_1409, negated_conjecture, (v7_ordinal1(u1_struct_0(esk1_0))|v1_xboole_0(esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1364, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_532, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_533, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_969, negated_conjecture, (v1_finseq_1(esk4_0))).#trainneg
cnf(i_0_968, negated_conjecture, (v1_funct_1(esk4_0))).#trainneg
cnf(i_0_967, negated_conjecture, (v1_relat_1(esk4_0))).#trainneg
cnf(i_0_1571, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1607, negated_conjecture, (esk2_0=esk5_0|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_638, negated_conjecture, (k1_matrix_1(esk4_0)=esk2_0)).#trainneg
cnf(i_0_534, plain, (v1_xreal_0(k7_nat_d(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1612, negated_conjecture, (esk2_0=esk3_0|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_1650, plain, (v1_relat_1(np__0)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_1547, plain, (v3_card_1(np__0,np__0))).#trainneg
cnf(i_0_1417, negated_conjecture, (v3_xxreal_0(esk6_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1418, negated_conjecture, (v3_xxreal_0(esk5_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk5_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_535, plain, (v3_ordinal1(k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_536, plain, (v7_ordinal1(k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1516, plain, (v1_card_1(X2)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1611, plain, (k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)=k4_matrix14(X2,X3,X4)|v1_xboole_0(X1)|v6_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1522, plain, (k3_finseq_1(k7_finseq_1(X1,k12_matrix_1(X2,X3)))=k2_nat_1(k3_finseq_1(X1),X3)|v2_struct_0(X2)|~v1_finseq_1(k12_matrix_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(X2,X3))|~v1_relat_1(X1)|~v7_ordinal1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_1593, plain, (v1_finset_1(np__1)|~v1_finset_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_1618, negated_conjecture, (esk6_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1566, plain, (k1_card_1(np__1)=np__1)).#trainneg
cnf(i_0_1544, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk3_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1616, negated_conjecture, (esk2_0=esk6_0|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_1091, negated_conjecture, (v1_relat_1(k2_matrix_1(esk4_0)))).#trainneg
cnf(i_0_1099, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,esk4_0))=k2_nat_1(esk2_0,esk2_0))).#trainneg
cnf(i_0_544, plain, (v3_ordinal1(k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_1243, negated_conjecture, (r1_xxreal_0(np__0,esk2_0)|~r1_xxreal_0(np__0,esk3_0))).#trainneg
cnf(i_0_1419, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1420, negated_conjecture, (v3_xxreal_0(esk2_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_545, plain, (v7_ordinal1(k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_555, plain, (v3_xxreal_0(k6_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1594, plain, (v1_finset_1(X2)|v2_struct_0(X1)|~v1_finset_1(k12_matrix_1(X1,X2))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_1687, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|esk3_0=X1|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1169, negated_conjecture, (k2_zfmisc_1(k4_finseq_1(esk4_0),k2_finseq_1(esk2_0))=k2_matrix_1(esk4_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0)|~v1_matrix_1(esk4_0))).#trainneg
cnf(i_0_744, plain, (k1_funct_1(k2_finseq_7(X1,X2,X3,X4),X5)=k7_partfun1(X1,X2,X5)|X3=X5|X5=X4|v1_xboole_0(X1)|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X5,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X5)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X5,k3_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1614, negated_conjecture, (esk3_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1688, plain, (k1_funct_1(k4_matrix14(X1,X2,X5),X3)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|X3=X5|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1609, negated_conjecture, (esk5_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1113, plain, (k7_partfun1(X1,k2_finseq_7(X6,X2,X3,X4),X5)=k7_partfun1(X1,X2,X5)|X5=X4|X3=X5|v1_xboole_0(X1)|v1_xboole_0(X6)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X5,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X6))).#trainneg
cnf(i_0_1605, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk5_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1652, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_1092, negated_conjecture, (k2_zfmisc_1(k9_xtuple_0(esk4_0),k2_finseq_1(esk2_0))=k2_matrix_1(esk4_0))).#trainneg
cnf(i_0_556, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1569, negated_conjecture, (np__0=esk3_0|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,np__0))).#trainneg
cnf(i_0_1430, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1866, negated_conjecture, (esk2_0=esk5_0|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_557, plain, (v3_xxreal_0(k1_xreal_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k1_xreal_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_558, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2055, negated_conjecture, (esk5_0=X1|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~v1_xreal_0(esk5_0))).#trainneg
cnf(i_0_2054, negated_conjecture, (esk5_0=X1|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1954, negated_conjecture, (esk6_0=X1|~r1_xxreal_0(esk6_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_1953, negated_conjecture, (esk6_0=X1|~r1_xxreal_0(esk6_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1960, plain, (k3_finseq_1(np__1)=np__1|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1))).#trainneg
cnf(i_0_2051, negated_conjecture, (esk3_0=X1|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~v1_xreal_0(esk3_0))).#trainneg
cnf(i_0_2050, negated_conjecture, (esk3_0=X1|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1247, negated_conjecture, (esk2_0=X1|~m1_matrix_1(esk4_0,u1_struct_0(esk1_0),X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1638, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,esk6_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1880, negated_conjecture, (esk2_0=esk3_0|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_1923, negated_conjecture, (v1_card_1(esk2_0))).#trainneg
cnf(i_0_668, negated_conjecture, (k2_zfmisc_1(k2_finseq_1(esk2_0),k2_finseq_1(esk2_0))=k2_matrix_1(esk4_0))).#trainneg
cnf(i_0_560, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2088, negated_conjecture, (np__0=esk3_0|~r1_xxreal_0(esk3_0,np__0))).#trainneg
cnf(i_0_1853, negated_conjecture, (v1_xboole_0(esk2_0)|~v1_funct_1(esk4_0)|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_1968, negated_conjecture, (esk2_0=esk6_0|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_1884, negated_conjecture, (v3_xxreal_0(esk6_0)|v2_xxreal_0(esk5_0)|r1_xxreal_0(esk5_0,esk6_0))).#trainneg
cnf(i_0_562, plain, (k10_xtuple_0(k9_finseq_1(X1))=k1_tarski(X1))).#trainneg
cnf(i_0_564, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1846, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2040, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),esk3_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|esk3_0=X1|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1444, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),X1)=k1_funct_1(esk4_0,X1)|X1=esk3_0|np__1=X1|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2066, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X2),X1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|np__1=X1|X1=X2|X1=esk3_0|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2067, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1)|X1=esk3_0|np__1=X1|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_798, negated_conjecture, (k2_nat_1(esk2_0,k3_finseq_1(X1))=k3_finseq_1(k7_finseq_1(esk4_0,X1))|~v1_finseq_1(X1)|~v1_funct_1(esk4_0)|~v1_funct_1(X1)|~v1_relat_1(esk4_0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1630, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2093, negated_conjecture, (v3_xxreal_0(np__1)|v2_xxreal_0(esk3_0))).#trainneg
cnf(i_0_565, plain, (r1_xxreal_0(X1,k6_xcmplx_0(X1,X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1659, negated_conjecture, (np__0=esk5_0|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk5_0,np__0))).#trainneg
cnf(i_0_1885, negated_conjecture, (v3_xxreal_0(esk6_0)|v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,esk6_0))).#trainneg
cnf(i_0_1886, negated_conjecture, (v3_xxreal_0(esk6_0)|v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,esk6_0))).#trainneg
cnf(i_0_566, plain, (v1_xboole_0(X1)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_567, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v1_xboole_0(k6_xcmplx_0(X2,X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_823, plain, (m2_finseq_1(X1,X2)|~m1_subset_1(X1,k13_finseq_1(X2)))).#trainneg
cnf(i_0_1859, negated_conjecture, (r2_hidden(esk4_0,X1)|X1!=k13_finseq_1(k13_finseq_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_972, plain, (v1_relat_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_773, plain, (k3_finseq_1(k7_finseq_1(X1,k9_finseq_1(X2)))=k2_nat_1(k3_finseq_1(X1),np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1309, plain, (r2_hidden(k4_matrix14(X1,X2,X3),X4)|v6_struct_0(X1)|X4!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1455, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),np__1)=k1_funct_1(k12_matrix_1(X1,X2),X3)|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1646, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1684, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1899, negated_conjecture, (v3_xxreal_0(esk5_0)|v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,esk5_0))).#trainneg
cnf(i_0_797, negated_conjecture, (m2_subset_1(esk2_0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_568, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__1))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1690, negated_conjecture, (np__0=esk6_0|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,np__0))).#trainneg
cnf(i_0_1647, plain, (m1_subset_1(np__0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1901, negated_conjecture, (v3_xxreal_0(esk5_0)|v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,esk5_0))).#trainneg
cnf(i_0_1902, negated_conjecture, (v3_xxreal_0(esk5_0)|v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,esk5_0))).#trainneg
cnf(i_0_569, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_570, plain, (k6_xcmplx_0(X1,np__0)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1705, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),np__1)=k1_funct_1(esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(esk4_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_979, plain, (v1_xboole_0(X1)|m2_finseq_1(k2_finseq_7(X5,X2,X3,X4),X1)|v1_xboole_0(X5)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X5))).#trainneg
cnf(i_0_974, plain, (v1_finseq_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1108, plain, (k1_funct_1(k2_finseq_7(X5,X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~m2_finseq_1(k2_finseq_7(X5,X2,X3,X4),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(k2_finseq_7(X5,X2,X3,X4)))|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_finseq_1(X2,X1)|~m1_finseq_1(X2,X5))).#trainneg
cnf(i_0_2130, negated_conjecture, (esk6_0=esk5_0|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2151, negated_conjecture, (esk6_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2309, negated_conjecture, (r2_hidden(esk4_0,k13_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2174, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk3_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2289, negated_conjecture, (v2_xxreal_0(esk3_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2204, plain, (k9_finseq_1(X1)=k9_finseq_1(X2)|k1_tarski(X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_571, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2296, plain, (~v1_xcmplx_0(np__2)|~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2297, plain, (~v1_xcmplx_0(np__1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1983, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,esk3_0))).#trainneg
cnf(i_0_1984, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(esk5_0)|r1_xxreal_0(esk5_0,esk3_0))).#trainneg
cnf(i_0_576, plain, (v1_funct_1(k12_finseq_1(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2206, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|~m2_finseq_1(X1,X3)|~r2_hidden(X1,k1_tarski(X4))|~m1_matrix_1(k9_finseq_1(X4),X3,X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2173, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(esk6_0,esk3_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_2237, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X3,X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_577, plain, (v1_finseq_1(k12_finseq_1(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2298, negated_conjecture, (np__0=esk5_0|~r1_xxreal_0(esk5_0,np__0))).#trainneg
cnf(i_0_1986, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,esk3_0))).#trainneg
cnf(i_0_2091, negated_conjecture, (v3_xxreal_0(np__1)|v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_579, plain, (v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_580, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2281, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X6!=k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),np__1)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3))))).#trainneg
cnf(i_0_2373, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k1_funct_1(esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_2390, plain, (v1_card_1(X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~v1_funct_1(k4_matrix14(X2,X1,X3))|~v1_relat_1(k4_matrix14(X2,X1,X3))|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2233, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2201, negated_conjecture, (v2_xxreal_0(esk5_0)|r1_xxreal_0(esk5_0,esk6_0)|~v2_xxreal_0(esk6_0)|~v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_2355, negated_conjecture, (v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,esk5_0)|~v2_xxreal_0(esk5_0)|~v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_2260, plain, (k7_partfun1(X1,k4_matrix14(X2,X3,X4),X5)=k7_partfun1(X1,k12_matrix_1(X2,X3),X5)|X5=X4|np__1=X5|v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2267, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,X1))=k2_xcmplx_0(k3_finseq_1(X1),esk2_0)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2194, negated_conjecture, (esk2_0=np__0|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_581, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2345, negated_conjecture, (np__0=esk6_0|~r1_xxreal_0(esk6_0,np__0))).#trainneg
cnf(i_0_1854, negated_conjecture, (v1_finset_1(esk2_0)|~v1_finset_1(esk4_0)|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_2092, negated_conjecture, (v3_xxreal_0(np__1)|v2_xxreal_0(esk5_0)|r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2094, negated_conjecture, (v3_xxreal_0(np__1)|v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_582, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_583, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2265, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,X1))=k2_xcmplx_0(esk2_0,k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2274, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,k9_finseq_1(X1)))=k2_nat_1(esk2_0,np__1))).#trainneg
cnf(i_0_704, plain, (v1_funct_1(k2_finseq_7(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_735, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_tarski(X1,X2),k2_matrix_1(esk4_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(esk2_0))).#trainneg
cnf(i_0_736, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_tarski(X2,X3),k2_matrix_1(k12_matrix_1(X1,X4)))|v2_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_769, plain, (k3_finseq_1(k7_finseq_1(X1,k4_matrix14(X2,X3,X4)))=k2_nat_1(k3_finseq_1(X1),X3)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|~v1_finseq_1(k4_matrix14(X2,X3,X4))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(X2,X3,X4))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(X2,X3,X4))|~v1_relat_1(X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1948, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k1_funct_1(k12_matrix_1(X1,X2),np__1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1519, plain, (v1_card_1(X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_2045, plain, (k1_funct_1(k10_funct_7(X2,X3,X4),X5)=k7_partfun1(X1,X2,X5)|X5=X4|X3=X5|v1_xboole_0(X1)|~m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X5,k3_finseq_1(k10_funct_7(X2,X3,X4)))|~r1_xxreal_0(X5,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_1262, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_584, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2406, negated_conjecture, (~v1_xxreal_0(esk3_0)|~v1_xboole_0(esk3_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2190, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_585, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_597, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2397, negated_conjecture, (esk6_0=esk5_0|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2398, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|esk6_0=esk5_0|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2415, negated_conjecture, (esk3_0=esk6_0|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_2403, negated_conjecture, (esk3_0=esk5_0|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2404, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|esk3_0=esk5_0|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2416, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|esk3_0=esk6_0|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_2450, negated_conjecture, (v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,np__1)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_599, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2561, negated_conjecture, (v7_ordinal1(k1_card_1(esk2_0))|~v1_finset_1(esk4_0))).#trainneg
cnf(i_0_1542, negated_conjecture, (r1_xxreal_0(np__2,esk2_0)|~r1_xxreal_0(np__2,esk3_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_600, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_602, plain, (k9_finseq_1(k1_funct_1(X1,np__1))=X1|k9_xtuple_0(X1)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2657, plain, (v1_card_1(X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_2630, negated_conjecture, (esk9_5(X1,esk4_0,X2,X3,X4)=k1_funct_1(esk4_0,X2)|X4!=k3_matrix_1(X1,esk4_0,X2,X3)|~m2_finseq_1(esk4_0,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X4,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0))).#trainneg
cnf(i_0_2517, plain, (k7_partfun1(X1,k12_matrix_1(X2,X3),X5)=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v1_xboole_0(X1)|v6_struct_0(X2)|~m2_finseq_1(k4_matrix14(X2,X3,X4),X1)|~r1_xxreal_0(X5,k3_finseq_1(k4_matrix14(X2,X3,X4)))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X5)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2409, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2490, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X6!=k1_funct_1(k12_matrix_1(X2,X3),np__1)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2653, negated_conjecture, (k2_nat_1(k3_finseq_1(X1),esk2_0)=k3_finseq_1(k7_finseq_1(X1,esk4_0))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2562, negated_conjecture, (v2_xxreal_0(esk5_0)|r1_xxreal_0(esk5_0,np__1)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2444, negated_conjecture, (esk12_1(esk1_0)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)|X3!=k1_funct_1(esk4_0,X1)|esk12_1(esk1_0)!=k1_funct_1(X3,X2)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_1603, negated_conjecture, (r1_xxreal_0(np__2,esk2_0)|~r1_xxreal_0(np__2,esk5_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1483, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_607, plain, (r2_hidden(k2_rfinseq(X1,X2,X3),X4)|X4!=k13_finseq_1(X1)|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_608, plain, (v1_funct_1(k2_rfinseq(X1,X2,X3))|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2418, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5),X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5),u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2423, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X5=X8|X7!=k1_funct_1(k4_matrix14(X2,X3,X8),X5)|X1!=k1_funct_1(X7,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X8))).#trainneg
cnf(i_0_2519, negated_conjecture, (k7_partfun1(X1,k12_matrix_1(esk1_0,esk2_0),X2)=k7_partfun1(X1,esk4_0,X2)|np__1=X2|X2=esk3_0|v1_xboole_0(X1)|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2430, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2400, negated_conjecture, (m1_subset_1(esk4_0,k13_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2489, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),np__1),X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),np__1),u1_struct_0(X2))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2426, negated_conjecture, (k2_xcmplx_0(X1,X2)=esk5_0|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk5_0,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2632, negated_conjecture, (k1_funct_1(esk9_5(X1,esk4_0,X2,X3,X4),X3)=X4|X4!=k3_matrix_1(X1,esk4_0,X2,X3)|~m2_finseq_1(esk4_0,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X4,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0))).#trainneg
cnf(i_0_2425, negated_conjecture, (k2_xcmplx_0(X1,X2)=esk6_0|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2607, negated_conjecture, (k2_xcmplx_0(esk2_0,k3_finseq_1(k12_matrix_1(X1,X2)))=k2_nat_1(esk2_0,X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2))))).#trainneg
cnf(i_0_2494, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_2563, negated_conjecture, (v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,np__1)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2445, negated_conjecture, (esk10_1(u1_struct_0(esk1_0))=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)|X3!=k1_funct_1(esk4_0,X1)|esk10_1(u1_struct_0(esk1_0))!=k1_funct_1(X3,X2)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_609, plain, (v1_finseq_1(k2_rfinseq(X1,X2,X3))|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1636, negated_conjecture, (r1_xxreal_0(np__2,esk2_0)|~r1_xxreal_0(np__2,esk6_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2694, negated_conjecture, (~v1_xxreal_0(np__1)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_610, plain, (m2_finseq_1(k1_rfinseq(X2,X3),X1)|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_611, plain, (r2_hidden(k17_finseq_1(X1,X2,X3),X4)|X4!=k13_finseq_1(X1)|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2735, negated_conjecture, (esk6_0=esk5_0|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk6_0,esk5_0)|~r1_xxreal_0(esk5_0,np__1)|~v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_2734, negated_conjecture, (esk6_0=esk5_0|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk6_0,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2740, negated_conjecture, (esk3_0=esk5_0|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk3_0,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2736, negated_conjecture, (esk6_0=esk5_0|r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_2742, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_2936, negated_conjecture, (esk5_0=np__1|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk5_0,np__0))).#trainneg
cnf(i_0_2773, plain, (~v1_xboole_0(X1)|k9_xtuple_0(X1)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_612, plain, (v1_funct_1(k17_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2925, negated_conjecture, (v3_ordinal1(esk4_0)|~v3_ordinal1(k13_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2738, negated_conjecture, (esk3_0=esk6_0|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk3_0,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_2109, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|esk2_0=esk5_0|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_613, plain, (v1_finseq_1(k17_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_614, plain, (m2_finseq_1(k16_finseq_1(X2,X3),X1)|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2921, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_2747, negated_conjecture, (v2_xxreal_0(esk5_0)|r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(esk6_0,np__1)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2776, plain, (k3_finseq_1(X1)=np__1|k9_xtuple_0(X1)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2790, plain, (k9_finseq_1(k3_matrix14(X1,X2,np__1))=k12_matrix_1(X1,X2)|v6_struct_0(X1)|v2_struct_0(X1)|k9_xtuple_0(k12_matrix_1(X1,X2))!=k2_finseq_1(np__1)|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1952, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k1_funct_1(k4_matrix14(X1,X2,X4),X4)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1450, negated_conjecture, (X1=k3_matrix_1(X2,k4_matrix14(X3,esk2_0,X4),X5,X6)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X1!=k1_funct_1(X7,X6)|X7!=k1_funct_1(k4_matrix14(X3,esk2_0,X4),X5)|~v1_matrix_1(k4_matrix14(X3,esk2_0,X4))|~m2_finseq_1(k4_matrix14(X3,esk2_0,X4),k13_finseq_1(X2))|~m2_finseq_1(X7,X2)|~r2_hidden(k4_tarski(X5,X6),k2_matrix_1(esk4_0))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X1,X2)|~v7_ordinal1(X4)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_2022, negated_conjecture, (r1_xxreal_0(esk6_0,k1_matrix_1(X1))|r1_xxreal_0(k1_matrix_1(X1),esk6_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_2956, negated_conjecture, (esk6_0=np__1|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,np__0))).#trainneg
cnf(i_0_1276, negated_conjecture, (r1_xxreal_0(esk5_0,np__2)|r1_xxreal_0(np__2,esk5_0))).#trainneg
cnf(i_0_615, plain, (v7_ordinal1(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2754, negated_conjecture, (v7_ordinal1(k1_card_1(esk2_0))|~m1_subset_1(esk4_0,k5_numbers))).#trainneg
cnf(i_0_2755, negated_conjecture, (r1_xxreal_0(np__2,esk2_0)|~r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_2183, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|esk2_0=esk3_0|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_2200, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|esk2_0=esk6_0|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_616, plain, (v1_xreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_617, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2606, negated_conjecture, (k2_nat_1(esk2_0,np__1)=k2_xcmplx_0(esk2_0,np__1))).#trainneg
cnf(i_0_2023, negated_conjecture, (r1_xxreal_0(esk5_0,k1_matrix_1(X1))|r1_xxreal_0(k1_matrix_1(X1),esk5_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_1185, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1615, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2622, plain, (v1_funct_1(k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_708, plain, (k3_finseq_1(k1_finseq_7(k13_finseq_1(X1),X2,X3,X4))=X5|v1_xboole_0(X1)|~m1_matrix_1(k1_finseq_7(k13_finseq_1(X1),X2,X3,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_finseq_1(X2,k13_finseq_1(X1))|~v7_ordinal1(X3)|~m1_subset_1(X4,k13_finseq_1(X1)))).#trainneg
cnf(i_0_2392, plain, (k7_partfun1(X4,k12_matrix_1(X1,X2),X3)=k1_funct_1(k4_matrix14(X1,X2,X3),np__1)|v1_xboole_0(X4)|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),X4)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2429, plain, (k2_xcmplx_0(X1,X2)=X3|~v1_xxreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3011, negated_conjecture, (esk6_0=esk5_0|~r1_xxreal_0(esk6_0,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_3022, negated_conjecture, (esk5_0=np__1|~r1_xxreal_0(esk5_0,np__0))).#trainneg
cnf(i_0_3109, negated_conjecture, (esk6_0=np__1|~r1_xxreal_0(esk6_0,np__0))).#trainneg
cnf(i_0_2987, negated_conjecture, (~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_2274, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,k9_finseq_1(X1)))=k2_xcmplx_0(esk2_0,np__1))).#trainneg
cnf(i_0_618, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v1_xboole_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3013, negated_conjecture, (esk3_0=esk5_0|~r1_xxreal_0(esk3_0,esk5_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_2862, negated_conjecture, (r1_xxreal_0(np__2,esk2_0)|~r1_xxreal_0(np__2,esk5_0))).#trainneg
cnf(i_0_3026, negated_conjecture, (esk3_0=esk6_0|~r1_xxreal_0(esk3_0,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_619, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v1_xboole_0(k4_xcmplx_0(k2_xcmplx_0(X2,X1)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_620, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3015, negated_conjecture, (esk6_0=esk5_0|r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_2990, plain, (k3_finseq_1(k1_rfinseq(X1,X2))=X3|v1_xboole_0(X4)|~m1_matrix_1(k1_rfinseq(X1,X2),X4,X3,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~m1_finseq_1(X2,k13_finseq_1(X4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3032, plain, (k3_finseq_1(k16_finseq_1(X1,X2))=X3|v1_xboole_0(X4)|~m1_matrix_1(k16_finseq_1(X1,X2),X4,X3,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~m1_finseq_1(X2,k13_finseq_1(X4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3017, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_3233, plain, (k7_partfun1(X1,k12_matrix_1(X2,X3),X4)=k1_funct_1(k4_matrix14(X2,X3,X4),np__1)|v1_xboole_0(X1)|v6_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X2,X3,X4)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_3235, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1),X3)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_3020, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,X1)=k1_funct_1(esk4_0,X1)|X1=esk3_0|np__1=X1|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3108, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k1_funct_1(esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_2992, plain, (r2_hidden(k1_rfinseq(X1,X2),X3)|X3!=k13_finseq_1(X4)|~m1_finseq_1(X2,X4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1472, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2860, negated_conjecture, (esk12_1(esk1_0)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)|k1_funct_1(k1_funct_1(esk4_0,X1),X2)!=esk12_1(esk1_0)|~l2_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_621, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(np__1)))=k2_xcmplx_0(k4_xcmplx_0(X1),np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1980, negated_conjecture, (esk2_0=np__0|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_2985, negated_conjecture, (r1_xxreal_0(np__2,esk2_0)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_3028, negated_conjecture, (esk2_0=esk5_0|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_3157, negated_conjecture, (esk2_0=esk3_0|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_622, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_623, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,np__0))=k2_xcmplx_0(k4_xcmplx_0(X1),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3092, negated_conjecture, (X1=k3_matrix_1(X2,esk4_0,X3,X4)|X5!=k1_funct_1(esk4_0,X3)|X1!=k1_funct_1(X5,X4)|~m2_finseq_1(esk4_0,k13_finseq_1(X2))|~m2_finseq_1(X5,X2)|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(esk4_0))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_3008, plain, (r2_hidden(k16_finseq_1(X2,X3),X4)|X4!=k13_finseq_1(X1)|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3105, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),X3)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),u1_struct_0(esk1_0))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_2336, negated_conjecture, (v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,esk5_0)|~v2_xxreal_0(esk5_0)|~v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_2640, plain, (X1=k3_matrix_1(X2,k12_matrix_1(X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X3))|X1!=k1_funct_1(X7,X6)|X7!=k1_funct_1(k12_matrix_1(X3,X4),X5)|~v1_matrix_1(k12_matrix_1(X3,X4))|~m2_finseq_1(k12_matrix_1(X3,X4),k13_finseq_1(X2))|~m2_finseq_1(X7,X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X1,X2)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X4)|~r1_xxreal_0(X5,X4)|~v7_ordinal1(X4)|~l6_algstr_0(X3))).#trainneg
cnf(i_0_768, plain, (k3_finseq_1(k7_finseq_1(k4_matrix14(X1,X2,X3),X4))=k2_nat_1(X2,k3_finseq_1(X4))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(X4)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(X4)|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(X4)|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1570, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_1660, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_1304, negated_conjecture, (r1_xxreal_0(esk3_0,np__2)|r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_2983, negated_conjecture, (esk10_1(u1_struct_0(esk1_0))=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)|k1_funct_1(k1_funct_1(esk4_0,X1),X2)!=esk10_1(u1_struct_0(esk1_0))|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_624, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(np__1),X1))=k2_xcmplx_0(np__1,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3406, negated_conjecture, (esk2_0=np__0|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_625, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_626, plain, (k4_xcmplx_0(k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__0,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_3352, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3349, negated_conjecture, (esk6_0=esk5_0|r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3353, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3350, negated_conjecture, (esk6_0=esk5_0|r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3523, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_3525, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_3361, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_627, plain, (m1_subset_1(esk7_2(X1,X2),X2)|X2=k13_finseq_1(X1)|m2_finseq_1(esk7_2(X1,X2),X1))).#trainneg
cnf(i_0_1356, negated_conjecture, (r1_xxreal_0(esk2_0,np__2)|r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_3411, negated_conjecture, (esk2_0=esk5_0|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_3413, negated_conjecture, (esk2_0=esk3_0|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_628, plain, (m2_finseq_1(esk7_2(X1,X2),X3)|X2=k13_finseq_1(X1)|m2_finseq_1(esk7_2(X1,X2),X1)|X2!=k13_finseq_1(X3))).#trainneg
cnf(i_0_630, plain, (r2_hidden(esk7_2(X1,X2),X3)|X2=k13_finseq_1(X1)|r2_hidden(esk7_2(X1,X2),X2)|X3!=k13_finseq_1(X1))).#trainneg
cnf(i_0_3457, plain, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3507, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|X6!=k1_funct_1(k12_matrix_1(X2,X3),X4)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k12_matrix_1(X2,X3))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2))).#trainneg
cnf(i_0_3358, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X1)=k5_struct_0(esk1_0)|np__1=X1|X1=esk3_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3359, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=k4_struct_0(esk1_0)|X1=X2|np__1=X1|X1=esk3_0|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_3346, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3773, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X1)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3778, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|X1=X2|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3771, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3780, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|X1=X2|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3689, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3692, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2529, negated_conjecture, (k2_nat_1(esk2_0,np__1)=k2_xcmplx_0(np__1,esk2_0))).#trainneg
cnf(i_0_3776, negated_conjecture, (k9_finseq_1(k4_struct_0(esk1_0))=k1_funct_1(esk4_0,X1)|X1=esk3_0|np__1=X1|k9_xtuple_0(k1_funct_1(esk4_0,X1))!=k2_finseq_1(np__1)|~v1_funct_1(k1_funct_1(esk4_0,X1))|~v1_relat_1(k1_funct_1(esk4_0,X1))|~r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_631, plain, (v1_funct_1(esk7_2(X1,X2))|X2=k13_finseq_1(X1)|r2_hidden(esk7_2(X1,X2),X2))).#trainneg
cnf(i_0_3185, plain, (v1_xreal_0(k4_xcmplx_0(np__2))|~v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_2746, negated_conjecture, (r1_xxreal_0(esk6_0,np__1)|~v1_xxreal_0(esk6_0)|~v1_xboole_0(esk6_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_3718, negated_conjecture, (esk2_0=esk5_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3721, negated_conjecture, (esk2_0=esk3_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_632, plain, (v1_finseq_1(esk7_2(X1,X2))|X2=k13_finseq_1(X1)|r2_hidden(esk7_2(X1,X2),X2))).#trainneg
cnf(i_0_634, plain, (r2_hidden(k8_finseq_1(X1,X2,X3),X4)|X4!=k13_finseq_1(X1)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_3770, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k12_matrix_1(X2,X3))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3761, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_3360, negated_conjecture, (r2_hidden(k1_funct_1(esk4_0,X1),X2)|np__1=X1|X1=esk3_0|X2!=k13_finseq_1(u1_struct_0(esk1_0))|~v7_ordinal1(X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_3502, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X3)|~m2_finseq_1(k1_funct_1(esk4_0,esk3_0),u1_struct_0(esk1_0))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3694, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_3354, plain, (k7_partfun1(X1,k12_matrix_1(X2,X3),X4)=k1_funct_1(k4_matrix14(X2,X3,X4),np__1)|v1_xboole_0(X1)|v6_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X2,X3,X4)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3343, plain, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1526, negated_conjecture, (esk2_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,esk5_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2531, negated_conjecture, (k2_nat_1(esk2_0,esk2_0)=k2_xcmplx_0(esk2_0,esk2_0))).#trainneg
cnf(i_0_2778, plain, (k7_finseq_1(X1,X2)=X1|k9_xtuple_0(k7_finseq_1(X1,X2))!=k2_finseq_1(np__1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k9_xtuple_0(X1)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_635, plain, (v1_funct_1(k8_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_2859, negated_conjecture, (r1_xxreal_0(esk5_0,np__1)|~v1_xxreal_0(esk5_0)|~v1_xboole_0(esk5_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2982, negated_conjecture, (r1_xxreal_0(esk2_0,np__1)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(esk2_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_3719, negated_conjecture, (esk2_0=esk5_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3722, negated_conjecture, (esk2_0=esk3_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_636, plain, (v1_finseq_1(k8_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_637, plain, (m2_finseq_1(k7_finseq_1(X2,X3),X1)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_3362, plain, (r2_hidden(k1_rfinseq(X1,X2),X3)|X3!=k13_finseq_1(X4)|~v7_ordinal1(X1)|~m2_finseq_1(X2,X4))).#trainneg
cnf(i_0_1622, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,esk3_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_1597, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|r1_xxreal_0(esk5_0,esk3_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_4077, plain, (k3_finseq_1(k7_finseq_1(X1,X2))=X3|v1_xboole_0(X4)|~m1_matrix_1(k7_finseq_1(X1,X2),X4,X3,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~m1_finseq_1(X2,k13_finseq_1(X4))|~m1_finseq_1(X1,k13_finseq_1(X4)))).#trainneg
cnf(i_0_2606, negated_conjecture, (k2_nat_1(esk2_0,np__1)=k2_xcmplx_0(np__1,esk2_0))).#trainneg
cnf(i_0_3845, negated_conjecture, (v7_ordinal1(k2_xcmplx_0(np__1,esk2_0)))).#trainneg
cnf(i_0_1099, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,esk4_0))=k2_xcmplx_0(esk2_0,esk2_0))).#trainneg
cnf(i_0_640, plain, (v3_card_1(esk10_1(k4_finseq_2(X1,X2)),X1)|v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3862, negated_conjecture, (~v1_xboole_0(k2_xcmplx_0(np__1,esk2_0)))).#trainneg
cnf(i_0_644, plain, (v1_relat_1(k2_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_4043, plain, (v1_funct_1(k7_finseq_1(X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_3864, negated_conjecture, (r1_xxreal_0(np__1,k2_xcmplx_0(np__1,esk2_0)))).#trainneg
cnf(i_0_4035, negated_conjecture, (esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,esk5_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4026, plain, (k4_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4105, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_3847, negated_conjecture, (r1_xxreal_0(esk2_0,k2_xcmplx_0(np__1,esk2_0)))).#trainneg
cnf(i_0_3999, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X4!=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|X1!=k1_funct_1(X4,X3)|~v1_matrix_1(k4_matrix14(esk1_0,esk2_0,X2))|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_3849, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(np__1,esk2_0),X1)=k2_xcmplx_0(esk2_0,k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(esk2_0))).#trainneg
cnf(i_0_3985, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1),np__1)=k1_funct_1(esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))))).#trainneg
cnf(i_0_4060, negated_conjecture, (esk2_0=esk5_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3675, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2300, negated_conjecture, (v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,esk6_0)|~v2_xxreal_0(esk6_0)|~v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_2274, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,k9_finseq_1(X1)))=k2_xcmplx_0(np__1,esk2_0))).#trainneg
cnf(i_0_645, plain, (k2_zfmisc_1(k9_xtuple_0(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_3894, plain, (v1_xreal_0(np__2)|~v1_xcmplx_0(np__2)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3895, negated_conjecture, (r1_xxreal_0(esk6_0,np__1)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_3159, negated_conjecture, (esk2_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_646, plain, (k3_finseq_1(esk7_2(k13_finseq_1(X1),X2))=X3|v1_xboole_0(X1)|X2=k13_finseq_1(k13_finseq_1(X1))|r2_hidden(esk7_2(k13_finseq_1(X1),X2),X2)|~m1_matrix_1(esk7_2(k13_finseq_1(X1),X2),X1,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_647, plain, (k3_finseq_1(k8_finseq_1(k13_finseq_1(X1),X2,X3))=X4|v1_xboole_0(X1)|~m1_matrix_1(k8_finseq_1(k13_finseq_1(X1),X2,X3),X1,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_finseq_1(X3,k13_finseq_1(X1))|~m1_finseq_1(X2,k13_finseq_1(X1)))).#trainneg
cnf(i_0_4075, plain, (v1_finseq_1(k7_finseq_1(X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_3986, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),np__1)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2))))).#trainneg
cnf(i_0_3984, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))))).#trainneg
cnf(i_0_3992, plain, (k1_funct_1(k4_matrix14(X2,X3,X4),np__1)=k1_funct_1(k12_matrix_1(X2,X3),X4)|v1_xboole_0(u1_struct_0(X2))|v1_xboole_0(X1)|v6_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X2,X3,X4)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3998, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)!=X2|~v1_finseq_1(X1)|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_3996, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,np__1)=k9_finseq_1(X1)|k9_xtuple_0(k4_matrix14(esk1_0,esk2_0,np__1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)!=X1|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_718, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(X3)|v1_xboole_0(k10_xtuple_0(X4))|~m2_finseq_1(X1,X3)|~m1_matrix_1(X4,X3,X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k10_xtuple_0(X4)))).#trainneg
cnf(i_0_3988, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X6!=k1_funct_1(k4_matrix14(X2,X3,np__1),np__1)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X2,X3,np__1))))).#trainneg
cnf(i_0_4073, negated_conjecture, (esk2_0=esk3_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3687, negated_conjecture, (esk6_0=esk5_0|r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk5_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2301, negated_conjecture, (v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,esk6_0)|~v2_xxreal_0(esk6_0)|~v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_648, plain, (k3_finseq_1(k17_finseq_1(k13_finseq_1(X1),X2,X3))=X4|v1_xboole_0(X1)|~m1_matrix_1(k17_finseq_1(k13_finseq_1(X1),X2,X3),X1,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_finseq_1(X3,k13_finseq_1(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4044, negated_conjecture, (r1_xxreal_0(esk5_0,np__1)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk5_0))).#trainneg
cnf(i_0_4046, negated_conjecture, (r1_xxreal_0(esk2_0,np__1)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_4107, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_649, plain, (k3_finseq_1(k2_rfinseq(k13_finseq_1(X1),X2,X3))=X4|v1_xboole_0(X1)|~m1_matrix_1(k2_rfinseq(k13_finseq_1(X1),X2,X3),X1,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_finseq_1(X3,k13_finseq_1(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_651, plain, (k2_nat_1(k3_finseq_1(X1),k3_finseq_1(X2))=k3_finseq_1(k7_finseq_1(X2,X1))|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4307, plain, (v1_finseq_1(k7_finseq_1(X1,X2))|~m1_finseq_1(X1,X3)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_4153, plain, (v1_funct_1(k7_finseq_1(X1,X2))|~m1_finseq_1(X1,X3)|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_4337, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_4148, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X6!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_4232, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_4308, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),np__1)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4115, negated_conjecture, (k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),X1)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X1)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4169, negated_conjecture, (esk2_0=np__1|r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_4301, negated_conjecture, (r1_xxreal_0(esk6_0,np__1)|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_4110, negated_conjecture, (k12_matrix_1(esk1_0,esk2_0)=k9_finseq_1(X1)|k9_xtuple_0(k12_matrix_1(esk1_0,esk2_0))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)!=X1|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_652, plain, (k2_xcmplx_0(k3_finseq_1(X1),k3_finseq_1(X2))=k3_finseq_1(k7_finseq_1(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4303, negated_conjecture, (esk2_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_653, plain, (k2_nat_1(k3_finseq_1(X2),k3_finseq_1(X1))=k3_finseq_1(k7_finseq_1(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_654, plain, (m2_subset_1(k3_finseq_1(k7_finseq_1(X1,X2)),k1_numbers,k5_numbers)|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4269, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1),np__1)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4341, plain, (k2_xcmplx_0(k3_finseq_1(X1),k3_finseq_1(X2))=k3_finseq_1(k7_finseq_1(X2,X1))|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4309, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4116, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X3)|~m2_finseq_1(k1_funct_1(esk4_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4317, plain, (k3_finseq_1(X1)=X2|v1_xboole_0(k10_xtuple_0(k12_matrix_1(X3,X2)))|v1_xboole_0(u1_struct_0(X3))|v2_struct_0(X3)|~m2_finseq_1(X1,u1_struct_0(X3))|~v7_ordinal1(X2)|~m1_subset_1(X1,k10_xtuple_0(k12_matrix_1(X3,X2)))|~l6_algstr_0(X3))).#trainneg
cnf(i_0_4316, negated_conjecture, (k3_finseq_1(X1)=esk2_0|v1_xboole_0(k10_xtuple_0(esk4_0))|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k10_xtuple_0(esk4_0)))).#trainneg
cnf(i_0_4347, negated_conjecture, (k2_nat_1(k3_finseq_1(X1),esk2_0)=k3_finseq_1(k7_finseq_1(esk4_0,X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_4118, negated_conjecture, (k2_nat_1(np__1,esk2_0)=k2_xcmplx_0(np__1,esk2_0))).#trainneg
cnf(i_0_4333, negated_conjecture, (r1_xxreal_0(esk5_0,np__1)|~v1_xboole_0(esk5_0))).#trainneg
cnf(i_0_655, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4299, plain, (v1_xreal_0(np__2)|~v1_xcmplx_0(np__2)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_4335, negated_conjecture, (r1_xxreal_0(esk2_0,np__1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_4472, negated_conjecture, (esk2_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4473, negated_conjecture, (esk2_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_656, plain, (v7_ordinal1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_657, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4377, plain, (v1_finseq_1(k7_finseq_1(X1,X2))|~m2_finseq_1(X2,X3)|~m2_finseq_1(X1,X3))).#trainneg
cnf(i_0_4378, plain, (v1_funct_1(k7_finseq_1(X1,X2))|~m2_finseq_1(X2,X3)|~m2_finseq_1(X1,X3))).#trainneg
cnf(i_0_4397, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|k3_matrix14(X2,X6,X7)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k3_matrix14(X2,X6,X7),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X7)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_4404, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k3_matrix14(X2,X3,X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4381, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk3_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4400, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,esk3_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4380, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3690, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3693, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4414, negated_conjecture, (esk6_0=np__1|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_4312, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0))|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_658, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4618, negated_conjecture, (esk5_0=np__1|~v1_xboole_0(esk5_0))).#trainneg
cnf(i_0_659, plain, (v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xboole_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_660, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4481, negated_conjecture, (k2_nat_1(k3_finseq_1(X1),esk2_0)=k3_finseq_1(k7_finseq_1(esk4_0,X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_4399, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4843, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|k3_matrix14(X2,X6,X5)!=k3_matrix14(X2,X3,X4)|X1!=k5_struct_0(X2)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~r1_xxreal_0(X5,X6))).#trainneg
cnf(i_0_4453, plain, (k3_finseq_1(k7_finseq_1(X1,k12_matrix_1(X2,X3)))=k2_xcmplx_0(k3_finseq_1(X1),X3)|v2_struct_0(X2)|~v1_finseq_1(k12_matrix_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(X2,X3))|~v1_relat_1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_4535, plain, (k2_xcmplx_0(k3_finseq_1(X2),k3_finseq_1(X1))=k3_finseq_1(k7_finseq_1(X2,X1))|~v1_xreal_0(k3_finseq_1(X2))|~v1_xreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_3977, negated_conjecture, (X2=esk3_0|np__1=X2|~v1_xboole_0(X1)|X1!=k13_finseq_1(u1_struct_0(esk1_0))|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_661, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)))|~v7_ordinal1(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4683, negated_conjecture, (esk2_0=np__1|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_4684, negated_conjecture, (esk2_0=esk5_0|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_662, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_663, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_5131, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=k5_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix14(X1,X5,X4)!=k3_matrix14(X1,X2,X3)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X4,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5130, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=k4_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix14(X1,X5,X4)!=k3_matrix14(X1,X2,X3)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X4,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_5206, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,X1))=k2_xcmplx_0(esk2_0,k3_finseq_1(X1))|~v1_xreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_5221, plain, (k3_finseq_1(k7_finseq_1(np__1,X1))=k2_xcmplx_0(np__1,k3_finseq_1(X1))|~v1_xreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_finseq_1(np__1)|~v1_funct_1(X1)|~v1_funct_1(np__1)|~v1_relat_1(X1)|~v1_relat_1(np__1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_5100, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,X1))=k3_finseq_1(k7_finseq_1(X1,esk4_0))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_5187, plain, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X2,X1)))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v1_xreal_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_4919, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4685, negated_conjecture, (esk2_0=esk3_0|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_5383, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,X1)=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X2,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_664, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_4686, negated_conjecture, (esk2_0=esk6_0|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_5256, negated_conjecture, (esk2_0=np__1|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_4891, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(esk6_0,esk5_0)|esk3_0=esk6_0|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_5117, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(esk5_0,esk6_0)|esk3_0=esk5_0|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_665, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_666, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5381, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X3)=k5_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5386, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X3)=k4_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5389, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k4_struct_0(esk1_0)|np__1=X3|X3=esk3_0|k1_funct_1(esk4_0,X3)!=k3_matrix14(esk1_0,X1,X2)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5387, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)=k4_struct_0(esk1_0)|np__1=X1|X1=esk3_0|k3_matrix14(esk1_0,X3,X2)!=k1_funct_1(esk4_0,X1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5384, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k5_struct_0(esk1_0)|np__1=X3|X3=esk3_0|k1_funct_1(esk4_0,X3)!=k3_matrix14(esk1_0,X1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5390, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,np__1)=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,esk3_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(np__1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5388, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,X1)=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X2,X1)!=k1_funct_1(esk4_0,esk3_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5382, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)=k5_struct_0(esk1_0)|np__1=X1|X1=esk3_0|k3_matrix14(esk1_0,X3,X2)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5385, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,np__1)=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(np__1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4933, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5257, negated_conjecture, (esk2_0=esk5_0|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_4112, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)!=X2|~v1_finseq_1(X1)|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_667, plain, (m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_5546, negated_conjecture, (esk2_0=esk3_0|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_5590, negated_conjecture, (esk2_0=esk6_0|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_4408, negated_conjecture, (esk2_0=np__1|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_669, plain, (k2_matrix_1(k12_matrix_1(X1,X2))=k2_zfmisc_1(k2_finseq_1(X2),k2_finseq_1(X2))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_672, plain, (v1_funct_1(k3_matrix14(X1,X2,X3))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5591, negated_conjecture, (esk6_0=esk5_0|esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk5_0,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_5743, plain, (X1=np__1|epred1_5(esk1_0,X1,X2,X3,k12_matrix_1(esk1_0,X4))|X1=esk3_0|k1_funct_1(esk4_0,X1)!=k3_matrix14(esk1_0,X4,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_5726, plain, (X1=np__1|epred1_5(X2,X1,X1,X3,k12_matrix_1(X2,X4))|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_5725, plain, (epred1_5(X1,X2,X2,X3,k12_matrix_1(X1,X4))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4707, negated_conjecture, (esk2_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5380, plain, (k2_xcmplx_0(np__0,np__2)=np__2|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_5767, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0))|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_673, plain, (v1_finseq_1(k3_matrix14(X1,X2,X3))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4415, negated_conjecture, (esk6_0=esk5_0|~r1_xxreal_0(esk5_0,np__1)|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_675, plain, (v1_xboole_0(X1)|v1_xboole_0(esk10_1(k1_zfmisc_1(X1)))|m1_subset_1(X2,esk10_1(k1_zfmisc_1(X1)))|~m2_subset_1(X2,X1,esk10_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_678, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,X1))|v1_xboole_0(X2)|m1_subset_1(X3,X2)|~m2_subset_1(X3,k2_zfmisc_1(k5_numbers,X1),X2)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_5753, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|np__1=X2|X2=esk3_0|k1_funct_1(esk4_0,X2)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5739, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|np__1=X2|X2=esk3_0|k1_funct_1(esk4_0,X2)!=k1_funct_1(esk4_0,X1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5785, plain, (k7_partfun1(u1_struct_0(X1),X2,X3)=k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X4),X5,X6)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix14(X1,X7,X6)!=k3_matrix14(X1,X4,X5)|k7_partfun1(u1_struct_0(X1),X2,X3)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X4)|~r1_xxreal_0(X5,X4)|~r1_xxreal_0(X6,X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X7)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5592, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_5548, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,np__1)=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_5190, negated_conjecture, (k3_finseq_1(k7_finseq_1(X1,esk4_0))=k2_xcmplx_0(k3_finseq_1(X1),esk2_0)|~v1_xreal_0(esk2_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_4470, plain, (k3_finseq_1(k7_finseq_1(np__1,X1))=k2_xcmplx_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_finseq_1(np__1)|~v1_funct_1(X1)|~v1_funct_1(np__1)|~v1_relat_1(X1)|~v1_relat_1(np__1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_5593, negated_conjecture, (r1_xxreal_0(esk6_0,k2_xcmplx_0(X1,X2))|esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk5_0,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_4495, plain, (k2_nat_1(k3_finseq_1(X1),np__1)=k3_finseq_1(k7_finseq_1(np__1,X1))|~v1_finseq_1(X1)|~v1_finseq_1(np__1)|~v1_funct_1(X1)|~v1_funct_1(np__1)|~v1_relat_1(X1)|~v1_relat_1(np__1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_5644, plain, (k2_xcmplx_0(np__2,np__0)=np__2|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_4313, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|k1_funct_1(esk4_0,esk3_0)!=X2|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_680, plain, (m2_subset_1(X1,X2,esk10_1(k1_zfmisc_1(X2)))|v1_xboole_0(esk10_1(k1_zfmisc_1(X2)))|v1_xboole_0(X2)|~m1_subset_1(X1,esk10_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_683, plain, (m2_subset_1(X1,k2_zfmisc_1(k5_numbers,X2),X3)|v1_xboole_0(X3)|v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~m1_subset_1(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_685, plain, (v1_xboole_0(X1)|v1_xboole_0(esk10_1(k1_zfmisc_1(X1)))|m1_subset_1(X2,X1)|~m2_subset_1(X2,X1,esk10_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_5164, plain, (k3_finseq_1(k7_finseq_1(X1,X2))=k3_finseq_1(X1)|~v1_xxreal_0(k3_finseq_1(k7_finseq_1(X1,X2)))|~r1_xxreal_0(k3_finseq_1(k7_finseq_1(X1,X2)),k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v1_xreal_0(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_4842, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X2),X3,X4)|np__1=X5|X5=esk3_0|k1_funct_1(esk4_0,X5)!=k3_matrix14(esk1_0,X2,X3)|X1!=k1_funct_1(k1_funct_1(esk4_0,X5),X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(X5,esk2_0)|~r1_xxreal_0(np__1,X5))).#trainneg
cnf(i_0_5203, plain, (k3_finseq_1(k7_finseq_1(X1,k12_matrix_1(X2,X3)))=k2_xcmplx_0(k3_finseq_1(X1),X3)|v2_struct_0(X2)|~v1_xreal_0(X3)|~v1_finseq_1(k12_matrix_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(X2,X3))|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1))|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_5784, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X3,X4)|X5!=k1_funct_1(esk4_0,X3)|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k1_funct_1(X5,X4)|~m2_finseq_1(X5,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5876, plain, (epred1_5(esk1_0,np__1,np__1,X1,k12_matrix_1(esk1_0,esk2_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_5817, negated_conjecture, (esk2_0=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_5646, plain, (k2_xcmplx_0(np__2,np__1)=k2_xcmplx_0(np__1,np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_688, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,X1))|v1_xboole_0(X2)|m1_subset_1(X3,k2_zfmisc_1(k5_numbers,X1))|~m2_subset_1(X3,k2_zfmisc_1(k5_numbers,X1),X2)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_5833, negated_conjecture, (esk3_0=esk6_0|esk6_0=esk5_0|r1_xxreal_0(esk5_0,esk3_0)|~v1_xxreal_0(esk5_0)|~r1_xxreal_0(esk5_0,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_4409, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|esk2_0=np__1|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_693, plain, (v1_relat_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)))|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_694, plain, (v1_funct_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)))|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5752, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,np__1)=k4_struct_0(esk1_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_4538, plain, (k2_xcmplx_0(k3_finseq_1(X2),k3_finseq_1(X1))=k3_finseq_1(k7_finseq_1(X2,X1))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X2),k5_numbers)|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5788, plain, (v1_xboole_0(X1)|m2_finseq_1(k7_partfun1(X1,X2,X3),X4)|X1!=k13_finseq_1(X4)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_5756, plain, (X1=np__1|epred1_5(esk1_0,np__1,X1,X2,k12_matrix_1(esk1_0,X3))|k3_matrix14(esk1_0,X3,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5596, negated_conjecture, (esk3_0=esk6_0|esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,esk6_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_4439, plain, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X2,X1)))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_5161, plain, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X1,X2)))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v1_xreal_0(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_5502, negated_conjecture, (k3_finseq_1(k7_finseq_1(X1,esk4_0))=k2_xcmplx_0(esk2_0,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_6127, negated_conjecture, (esk6_0=esk5_0|esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|~r1_xxreal_0(esk5_0,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_6121, negated_conjecture, (esk2_0=np__1|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_1691, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_2356, negated_conjecture, (v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,esk5_0)|~v2_xxreal_0(esk5_0)|~v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_4587, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|k3_matrix14(esk1_0,esk2_0,np__1)!=X2|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_695, plain, (v1_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)))|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5832, negated_conjecture, (v1_funct_1(k1_funct_1(esk4_0,esk3_0))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_5845, negated_conjecture, (v1_finseq_1(k1_funct_1(esk4_0,esk3_0))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_6129, negated_conjecture, (esk2_0=np__1|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_6261, negated_conjecture, (esk3_0=esk5_0|esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk6_0)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk3_0,esk6_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_696, plain, (k7_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)),X4)=k7_finseq_1(k7_finseq_1(X1,X2),k7_finseq_1(X3,X4))|~v1_finseq_1(X4)|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_698, plain, (k1_funct_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)),X4)=k1_funct_1(k7_finseq_1(X1,X2),X4)|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~r1_xxreal_0(X4,k3_finseq_1(k7_finseq_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6196, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,X1))=k2_xcmplx_0(esk2_0,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_6118, negated_conjecture, (k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),X1,X2)|X1=esk3_0|np__1=X1|X1=X3|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1),X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6119, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|X3=esk3_0|np__1=X3|X3=X2|X1!=k1_funct_1(k1_funct_1(esk4_0,X3),X4)|~m2_finseq_1(k1_funct_1(esk4_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6277, negated_conjecture, (r1_xxreal_0(esk2_0,k3_finseq_1(k7_finseq_1(esk4_0,X1)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_6233, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_4609, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,k12_matrix_1(X1,X2)))=k2_nat_1(X2,esk2_0)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_5503, negated_conjecture, (k3_finseq_1(k7_finseq_1(X1,esk4_0))=k2_xcmplx_0(k3_finseq_1(X1),esk2_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_1981, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_1230, plain, (v7_ordinal1(k2_finseq_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2449, negated_conjecture, (v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,esk3_0)|~v2_xxreal_0(esk3_0)|~v1_xxreal_0(esk3_0))).#trainneg
cnf(i_0_5768, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|k1_funct_1(esk4_0,esk3_0)!=X2|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_699, plain, (m1_subset_1(k5_finseq_2(X1,X2,X3),k4_finseq_2(X2,X1))|v1_xboole_0(X1)|~m1_finseq_2(k4_finseq_2(X2,X1),X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_1530, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|r1_xxreal_0(esk5_0,esk6_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_701, plain, (m2_finseq_2(k2_finseq_2(X2,X3),X1,k4_finseq_2(X2,X1))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_703, plain, (r2_hidden(k2_finseq_7(X1,X2,X3,X4),X5)|v1_xboole_0(X1)|X5!=k13_finseq_1(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6395, negated_conjecture, (esk3_0=esk6_0|esk6_0=esk5_0|r1_xxreal_0(esk5_0,esk3_0)|~r1_xxreal_0(esk5_0,esk6_0)|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_6537, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X5=X7|X1!=k1_funct_1(k1_funct_1(k4_matrix14(X2,X3,X7),X5),X6)|~m1_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~v7_ordinal1(X7))).#trainneg
cnf(i_0_6531, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X5),X6)|~m1_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_6506, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|np__1=X3|X3=esk3_0|X3=X2|X3=X4|X1!=k4_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6511, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|np__1=X3|X3=esk3_0|X3=X2|X3=X4|X1!=k4_struct_0(esk1_0)|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6507, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X3)|np__1=X3|X3=esk3_0|X3=X2|X1!=k5_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6405, negated_conjecture, (esk3_0=esk6_0|esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|~r1_xxreal_0(esk3_0,esk6_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_6571, negated_conjecture, (r1_xxreal_0(k3_finseq_1(k12_matrix_1(X1,X2)),k2_nat_1(X2,esk2_0))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_6575, negated_conjecture, (k2_nat_1(X2,esk2_0)=k2_nat_1(esk2_0,X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_6574, negated_conjecture, (k2_nat_1(X2,esk2_0)=k2_xcmplx_0(esk2_0,X2)|v2_struct_0(X1)|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_6398, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_705, plain, (v1_finseq_1(k2_finseq_7(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3504, negated_conjecture, (r1_xxreal_0(esk6_0,esk5_0)|~v1_xxreal_0(esk6_0)|~v1_xboole_0(esk6_0)|~v2_xxreal_0(esk5_0)|~v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_6403, negated_conjecture, (esk2_0=np__1|r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_6647, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_709, plain, (r2_hidden(k1_finseq_7(X1,X2,X3,X4),X5)|v1_xboole_0(X1)|X5!=k13_finseq_1(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_6573, negated_conjecture, (k2_nat_1(X2,esk2_0)=k2_xcmplx_0(esk2_0,X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_6536, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1),X3)|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_6652, plain, (m2_finseq_1(k2_finseq_2(X1,X2),X3)|v1_xboole_0(X3)|~m1_finseq_2(k4_finseq_2(X1,X3),X3)|~v7_ordinal1(X1)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_4394, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|k9_finseq_1(X6)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k9_finseq_1(X6),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_subset_1(X6,u1_struct_0(X2)))).#trainneg
cnf(i_0_6269, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X1,esk4_0)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_6504, negated_conjecture, (k1_funct_1(k3_matrix14(esk1_0,esk2_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),X1,X2)|np__1=X1|X1=esk3_0|X1=X3|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k1_funct_1(k3_matrix14(esk1_0,esk2_0,X1),X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4578, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_4391, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|k12_finseq_1(u1_struct_0(X2),X6)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k12_finseq_1(u1_struct_0(X2),X6),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_subset_1(X6,u1_struct_0(X2)))).#trainneg
cnf(i_0_4525, plain, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X2,X1)))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_5547, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,X1)=k5_struct_0(esk1_0)|np__1=X1|X1=esk3_0|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6627, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_3229, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5843, plain, (k2_xcmplx_0(np__0,np__2)=np__2|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_5770, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|k3_matrix14(esk1_0,esk2_0,np__1)!=X2|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_710, plain, (v1_funct_1(k1_finseq_7(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_2744, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_711, plain, (v1_finseq_1(k1_finseq_7(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_712, plain, (v1_xboole_0(X1)|m2_finseq_1(k2_funct_7(X2,X3,X4),X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_6827, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|X1=X2|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k5_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6796, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(esk4_0,X1)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_6681, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X5),X6)|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_6828, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),X1,X2)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|X1=X3|X1=X2|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6717, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6888, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6004, plain, (k2_xcmplx_0(np__2,np__0)=np__2|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_714, plain, (r2_hidden(esk8_4(X1,X2,X3,X4),X5)|v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|X5!=k13_finseq_1(X1)|k3_finseq_1(X4)!=X2|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6727, negated_conjecture, (esk2_0=np__1|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_715, plain, (v1_funct_1(esk8_4(X1,X2,X3,X4))|v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|k3_finseq_1(X4)!=X2|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_716, plain, (v1_finseq_1(esk8_4(X1,X2,X3,X4))|v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|k3_finseq_1(X4)!=X2|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6171, plain, (k2_xcmplx_0(k3_finseq_1(X1),k3_finseq_1(X2))=k3_finseq_1(k7_finseq_1(X2,X1))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X2),k5_numbers)|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5747, plain, (epred1_5(esk1_0,np__1,X1,X2,k12_matrix_1(esk1_0,X3))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X3,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4545, plain, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X1,X2)))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_6718, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6889, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_717, plain, (k3_finseq_1(esk7_2(X1,k10_xtuple_0(X2)))=X3|v1_xboole_0(X4)|k10_xtuple_0(X2)=k13_finseq_1(X1)|m2_finseq_1(esk7_2(X1,k10_xtuple_0(X2)),X1)|~m2_finseq_1(esk7_2(X1,k10_xtuple_0(X2)),X4)|~m1_matrix_1(X2,X4,X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6728, negated_conjecture, (esk2_0=np__1|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6729, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_719, plain, (m1_subset_1(esk8_4(X1,X2,X3,X4),k10_xtuple_0(X4))|v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|k3_finseq_1(X4)!=X2|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_720, plain, (m2_finseq_1(esk8_4(X1,X2,X3,X4),X5)|v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|k10_xtuple_0(X4)!=k13_finseq_1(X5)|k3_finseq_1(X4)!=X2|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7004, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X5)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k5_struct_0(X2)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7110, plain, (k3_finseq_1(k7_finseq_1(X2,X1))=k3_finseq_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_6954, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X2,X2)=k5_struct_0(esk1_0)|np__1=X2|X2=esk3_0|X2=X1|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k5_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_722, plain, (k3_finseq_1(esk8_4(X1,X2,X3,X4))=X5|v1_xboole_0(X6)|v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|~m2_finseq_1(esk8_4(X1,X2,X3,X4),X6)|~m1_matrix_1(X4,X6,X7,X5)|~v7_ordinal1(X7)|~v7_ordinal1(X5)|k3_finseq_1(X4)!=X2|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4331, negated_conjecture, (r1_xxreal_0(esk2_0,esk6_0)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(esk2_0)|~v2_xxreal_0(esk6_0)|~v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_6893, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(np__1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk2_0))).#trainneg
cnf(i_0_723, plain, (esk9_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k13_finseq_1(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_724, plain, (m2_finseq_1(esk9_5(X1,X2,X3,X4,X5),X1)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k13_finseq_1(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_7082, plain, (k3_finseq_1(k7_finseq_1(X2,X1))=k3_finseq_1(X1)|~v1_xxreal_0(k3_finseq_1(k7_finseq_1(X2,X1)))|~r1_xxreal_0(k3_finseq_1(k7_finseq_1(X2,X1)),k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_7286, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X4)=k5_struct_0(X1)|np__1=X4|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7293, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k5_struct_0(X1)|np__1=X4|X4=X5|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7260, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7418, plain, (k3_finseq_1(k4_matrix14(X1,X2,X3))=X4|v1_xboole_0(X5)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~m1_matrix_1(k4_matrix14(X1,X2,X3),X5,X4,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(X5))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7227, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1882, negated_conjecture, (v3_xxreal_0(esk6_0)|v2_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),esk6_0))).#trainneg
cnf(i_0_725, plain, (k1_funct_1(esk9_5(X1,X2,X3,X4,X5),X4)=X5|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k13_finseq_1(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_6267, negated_conjecture, (r1_xxreal_0(esk2_0,k2_xcmplx_0(esk2_0,esk2_0)))).#trainneg
cnf(i_0_6400, negated_conjecture, (r1_xxreal_0(esk2_0,esk5_0)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(esk2_0)|~v2_xxreal_0(esk5_0)|~v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_7261, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk6_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7439, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_726, plain, (k3_matrix14(X1,X2,np__1)=k9_finseq_1(X3)|v6_struct_0(X1)|v2_struct_0(X1)|k9_xtuple_0(k3_matrix14(X1,X2,np__1))!=k2_finseq_1(np__1)|k5_struct_0(X1)!=X3|~v1_funct_1(k3_matrix14(X1,X2,np__1))|~v1_relat_1(k3_matrix14(X1,X2,np__1))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(np__1)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_729, plain, (k12_matrix_1(X1,X2)=k9_finseq_1(X3)|v6_struct_0(X1)|v2_struct_0(X1)|k9_xtuple_0(k12_matrix_1(X1,X2))!=k2_finseq_1(np__1)|k3_matrix14(X1,X2,np__1)!=X3|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(np__1)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7007, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X2,X3)=k4_struct_0(esk1_0)|np__1=X2|X2=esk3_0|X2=X3|X2=X1|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7454, plain, (X1=np__1|epred1_5(esk1_0,X1,X1,X2,k4_matrix14(esk1_0,esk2_0,X3))|X1=esk3_0|X1=X3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7432, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(esk4_0,X1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))).#trainneg
cnf(i_0_7012, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X2,X2)=k4_struct_0(esk1_0)|np__1=X2|X2=esk3_0|X2=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_7241, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1898, negated_conjecture, (v3_xxreal_0(esk5_0)|v2_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),esk5_0))).#trainneg
cnf(i_0_4237, negated_conjecture, (k2_xcmplx_0(np__1,esk2_0)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(np__1,esk2_0),esk5_0))).#trainneg
cnf(i_0_731, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|v1_xboole_0(k2_matrix_1(X3))|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k13_finseq_1(X2))|~m2_finseq_1(X6,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,X2)|~m1_subset_1(k4_tarski(X4,X5),k2_matrix_1(X3)))).#trainneg
cnf(i_0_7442, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7438, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_730, plain, (k3_matrix14(X1,X2,X3)=k9_finseq_1(X4)|X3=np__1|v6_struct_0(X1)|v2_struct_0(X1)|k9_xtuple_0(k3_matrix14(X1,X2,X3))!=k2_finseq_1(np__1)|k4_struct_0(X1)!=X4|~v1_funct_1(k3_matrix14(X1,X2,X3))|~v1_relat_1(k3_matrix14(X1,X2,X3))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(np__1)|~v7_ordinal1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_741, plain, (k2_funct_7(k5_finseq_2(u1_struct_0(X1),X2,k4_struct_0(X1)),X3,k5_struct_0(X1))=k3_matrix14(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k5_finseq_2(u1_struct_0(X1),X2,k4_struct_0(X1)),u1_struct_0(X1))|~v7_ordinal1(X3)|~m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7285, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X4)=k4_struct_0(X1)|np__1=X4|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_7420, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v1_xboole_0(u1_struct_0(X6))|v6_struct_0(X6)|k4_matrix14(X6,X7,X8)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k4_matrix14(X6,X7,X8),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X6,X7),u1_struct_0(X2))|~v7_ordinal1(X8)|~m1_subset_1(X7,k5_numbers)|~l6_algstr_0(X6)|~v5_vectsp_1(X6)|~v4_vectsp_1(X6)|~v5_group_1(X6)|~v3_group_1(X6)|~v4_rlvect_1(X6)|~v3_rlvect_1(X6)|~v2_rlvect_1(X6)|~v33_algstr_0(X6)|~v13_algstr_0(X6))).#trainneg
cnf(i_0_7329, negated_conjecture, (r1_xxreal_0(esk2_0,k3_finseq_1(k7_finseq_1(X1,esk4_0)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_7350, plain, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X1,X2)))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_6956, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,X1))=k3_finseq_1(X1)|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(k3_finseq_1(k7_finseq_1(esk4_0,X1)))|~r1_xxreal_0(k3_finseq_1(k7_finseq_1(esk4_0,X1)),k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_7491, negated_conjecture, (k3_finseq_1(k7_finseq_1(X1,esk4_0))=k3_finseq_1(X1)|~v1_xxreal_0(k3_finseq_1(k7_finseq_1(X1,esk4_0)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(k7_finseq_1(X1,esk4_0)),k3_finseq_1(X1))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_7258, negated_conjecture, (esk2_0=np__1|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6931, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_7584, negated_conjecture, (k9_finseq_1(X1)=k1_funct_1(esk4_0,esk3_0)|k9_xtuple_0(k1_funct_1(esk4_0,esk3_0))!=k2_finseq_1(np__1)|k5_struct_0(esk1_0)!=X1|~v1_funct_1(k1_funct_1(esk4_0,esk3_0))|~v1_relat_1(k1_funct_1(esk4_0,esk3_0))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_748, plain, (m2_finseq_1(k2_finseq_7(X1,X2,np__1,X3),X1)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7635, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__2,esk2_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_7638, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk6_0,esk3_0))).#trainneg
cnf(i_0_743, plain, (k1_finseq_7(u1_struct_0(X1),k2_finseq_2(X2,k4_struct_0(X1)),X3,k5_struct_0(X1))=k3_matrix14(X1,X2,X3)|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_7678, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X1)=k4_struct_0(esk1_0)|np__1=X1|X1=esk3_0|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7535, negated_conjecture, (esk2_0=X1|v1_xboole_0(X2)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(X2))|~m1_matrix_1(esk4_0,X2,X1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7511, plain, (X1=np__1|epred1_5(esk1_0,X1,X1,X2,esk4_0)|X1=esk3_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7639, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk5_0,esk3_0))).#trainneg
cnf(i_0_7586, negated_conjecture, (k12_matrix_1(esk1_0,esk2_0)=k9_finseq_1(X1)|k9_xtuple_0(k12_matrix_1(esk1_0,esk2_0))!=k2_finseq_1(np__1)|k1_funct_1(esk4_0,esk3_0)!=X1|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_749, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_6642, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_0)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(esk2_0)|~v2_xxreal_0(esk3_0)|~v1_xxreal_0(esk3_0))).#trainneg
cnf(i_0_7641, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk2_0,esk3_0))).#trainneg
cnf(i_0_747, plain, (k7_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_750, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X1,X2,np__1,X3),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,np__1,X3)))),k12_finseq_1(X1,k7_partfun1(X1,X2,X3))),k2_rfinseq(X1,X4,k2_finseq_7(X1,X2,np__1,X3)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,np__1,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k2_finseq_7(X1,X2,np__1,X3),X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3)))|~v7_ordinal1(X4)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_7929, plain, (X1=np__1|epred1_5(esk1_0,X2,X1,X3,esk4_0)|X1=esk3_0|X1=X2|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8128, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X1,X2,np__1,X3),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,np__1,X3)))),k12_finseq_1(X1,k1_funct_1(X2,X3))),k2_rfinseq(X1,X4,k2_finseq_7(X1,X2,np__1,X3)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,np__1,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k2_finseq_7(X1,X2,np__1,X3),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7523, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k5_struct_0(esk1_0)|np__1=X1|X1=esk3_0|X1=X2|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7519, plain, (epred1_5(X1,X2,X3,X4,k4_matrix14(X1,X5,X6))|np__1=X3|X3=X6|X3=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X5)|~r1_xxreal_0(X2,X5)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~m1_subset_1(X5,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7856, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1982, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),esk3_0))).#trainneg
cnf(i_0_753, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k1_funct_1(X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_7909, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_7526, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(esk3_0,esk5_0)|r1_xxreal_0(esk6_0,esk3_0))).#trainneg
cnf(i_0_751, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X1,X2,X3,np__1),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,X3,np__1)))),k12_finseq_1(X1,k7_partfun1(X1,X2,X3))),k2_rfinseq(X1,X4,k2_finseq_7(X1,X2,X3,np__1)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,X3,np__1),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k2_finseq_7(X1,X2,X3,np__1),X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,X3,np__1)))|~v7_ordinal1(X4)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_752, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),X5)),k17_finseq_1(X1,k7_nat_d(X5,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X5,k2_finseq_7(X1,X2,X3,X4)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,X3,X4),np__1,X5)|v1_xboole_0(X1)|r1_xxreal_0(X5,np__1)|X3=np__1|np__1=X4|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X5,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~v7_ordinal1(X5)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8156, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,esk5_0)=esk4_0|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_8192, negated_conjecture, (esk6_0=esk5_0|esk5_0=np__1|~r1_xxreal_0(esk5_0,esk6_0)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_8182, negated_conjecture, (v2_xxreal_0(esk5_0)|esk5_0=np__1|esk6_0=esk5_0|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_8179, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)=k1_funct_1(esk4_0,esk5_0)|esk5_0=np__1|esk6_0=esk5_0|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_7857, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8169, negated_conjecture, (k9_finseq_1(X1)=k1_funct_1(esk4_0,esk5_0)|esk5_0=np__1|esk6_0=esk5_0|k9_xtuple_0(k1_funct_1(esk4_0,esk5_0))!=k2_finseq_1(np__1)|k5_struct_0(esk1_0)!=X1|~v1_relat_1(k1_funct_1(esk4_0,esk5_0))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_754, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_755, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_756, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k16_finseq_1(k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_finseq_1(k2_rfinseq(X1,np__1,X2),X1)|~v7_ordinal1(k7_nat_d(X3,np__2)))).#trainneg
cnf(i_0_8160, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk5_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|esk5_0=np__1|esk6_0=esk5_0|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_8200, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)=k3_matrix14(esk1_0,esk2_0,np__1)|esk5_0=np__1|esk6_0=esk5_0|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_7673, plain, (X1=np__1|epred1_5(X2,X1,X1,X3,k4_matrix14(X2,X4,X5))|X1=X5|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|k4_struct_0(X2)!=k5_struct_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2308, plain, (m2_finseq_1(esk10_1(k13_finseq_1(X1)),X1))).#trainneg
cnf(i_0_8196, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(X2)|esk5_0=np__1|esk6_0=esk5_0|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)!=X2|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk5_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_760, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_8328, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_8329, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__2,esk5_0))).#trainneg
cnf(i_0_757, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X4)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),np__1))),k2_rfinseq(X1,X3,k2_finseq_7(X1,X2,X3,X4)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,X3,X4),np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X3,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_758, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),np__1))),k2_rfinseq(X1,X4,k2_finseq_7(X1,X2,X3,X4)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,X3,X4),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~v7_ordinal1(X4)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7515, plain, (X1=np__1|epred1_5(X2,X1,X1,X3,k4_matrix14(X2,X4,X5))|X1=X5|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_8641, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_8189, negated_conjecture, (esk6_0=esk5_0|r1_xxreal_0(esk6_0,esk5_0)|esk5_0=np__1|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_8530, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8199, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(X2)|esk5_0=np__1|esk6_0=esk5_0|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)!=X2|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_761, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1))).#trainneg
cnf(i_0_8330, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_759, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X5)),k17_finseq_1(X1,k7_nat_d(X5,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X1,X2,X3,X4),np__1))),k2_rfinseq(X1,X5,k2_finseq_7(X1,X2,X3,X4)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,X3,X4),np__1,X5)|v1_xboole_0(X1)|r1_xxreal_0(X5,np__1)|X3=X5|X5=X4|~m2_finseq_1(k2_finseq_7(X1,X2,X3,X4),X1)|~r1_xxreal_0(X5,k3_finseq_1(k2_finseq_7(X1,X2,X3,X4)))|~v7_ordinal1(X5)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X5,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_762, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k7_finseq_1(k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_finseq_1(k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2)),X1)|~m1_finseq_1(k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),X1))).#trainneg
cnf(i_0_8930, negated_conjecture, (X1=np__1|X1=esk5_0|epred1_5(esk1_0,X1,X1,X2,esk4_0)|esk6_0=esk5_0|esk5_0=np__1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8770, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_9014, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_7904, plain, (v1_xboole_0(X1)|m2_finseq_1(k2_finseq_7(X1,X2,np__1,X3),X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_8142, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(k2_finseq_7(X1,X2,np__1,X3),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,np__1,X3)))),k12_finseq_1(X1,k7_partfun1(X1,X2,X3))),k2_rfinseq(X1,X4,k2_finseq_7(X1,X2,np__1,X3)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,np__1,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k2_finseq_7(X1,X2,np__1,X3),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,X4))).#trainneg
cnf(i_0_9273, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_8144, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,np__1)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,np__1,X3)))),k12_finseq_1(X1,k7_partfun1(X1,X2,X3))),k2_rfinseq(X1,X3,k2_finseq_7(X1,X2,np__1,X3)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,np__1,X3),np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(k2_finseq_7(X1,X2,np__1,X3),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7657, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,esk5_0))).#trainneg
cnf(i_0_8734, plain, (v1_relat_1(esk10_1(k13_finseq_1(X1))))).#trainneg
cnf(i_0_8797, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk3_0,np__2))).#trainneg
cnf(i_0_8741, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0))|esk6_0=esk5_0|esk5_0=np__1|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk5_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,esk2_0))).#trainneg
cnf(i_0_783, plain, (m1_subset_1(k3_matrix14(X1,X2,X3),X4)|v6_struct_0(X1)|v2_struct_0(X1)|X4!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_763, plain, (k8_finseq_1(X1,k7_finseq_1(k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_finseq_1(k12_finseq_1(X1,k7_partfun1(X1,X2,np__1)),X1)|~m1_finseq_1(k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),X1))).#trainneg
cnf(i_0_767, plain, (m2_subset_1(X2,k1_numbers,k5_numbers)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8769, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k1_funct_1(X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~r1_xxreal_0(np__1,k3_finseq_1(X2)))).#trainneg
cnf(i_0_9636, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(k12_matrix_1(X2,X3),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|v2_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_9653, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k1_funct_1(k12_matrix_1(X2,X3),np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|v2_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_9013, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_8031, plain, (k3_finseq_1(k2_rfinseq(X1,X3,X2))=k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3))|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~v1_xxreal_0(k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3)))|~v1_finseq_1(k2_rfinseq(X1,X3,X2))|~v1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))))|~v1_funct_1(k2_rfinseq(X1,X3,X2))|~v1_funct_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))))|~v1_relat_1(k2_rfinseq(X1,X3,X2))|~v1_relat_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))))|~r1_xxreal_0(k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3)),k3_finseq_1(k2_rfinseq(X1,X3,X2)))|~m1_subset_1(k3_finseq_1(k2_rfinseq(X1,X3,X2)),k5_numbers)|~m1_subset_1(k3_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1)))),k5_numbers)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8747, plain, (k7_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9218, plain, (k8_finseq_1(X1,k7_finseq_1(k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1)|~m1_finseq_1(k9_finseq_1(k7_partfun1(X1,X2,np__1)),X1)|~m1_finseq_1(k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),X1))).#trainneg
cnf(i_0_8200, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)=k3_matrix14(esk1_0,esk2_0,np__1)|esk6_0=esk5_0|esk5_0=np__1)).#trainneg
cnf(i_0_4874, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk3_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8735, plain, (v1_funct_1(esk10_1(k13_finseq_1(X1))))).#trainneg
cnf(i_0_5845, negated_conjecture, (v1_finseq_1(k1_funct_1(esk4_0,esk3_0)))).#trainneg
cnf(i_0_7584, negated_conjecture, (k9_finseq_1(X1)=k1_funct_1(esk4_0,esk3_0)|k9_xtuple_0(k1_funct_1(esk4_0,esk3_0))!=k2_finseq_1(np__1)|k5_struct_0(esk1_0)!=X1|~v1_relat_1(k1_funct_1(esk4_0,esk3_0)))).#trainneg
cnf(i_0_784, plain, (m2_finseq_1(k3_matrix14(X1,X2,X3),X4)|v6_struct_0(X1)|v2_struct_0(X1)|X5!=k13_finseq_1(X4)|X5!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5832, negated_conjecture, (v1_funct_1(k1_funct_1(esk4_0,esk3_0)))).#trainneg
cnf(i_0_1532, negated_conjecture, (esk2_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_782, plain, (k3_finseq_1(k3_matrix14(X1,X2,X3))=X4|v1_xboole_0(X5)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k3_matrix14(X1,X2,X3),X5)|~m1_matrix_1(X6,X5,X7,X4)|~v7_ordinal1(X7)|~v7_ordinal1(X4)|k10_xtuple_0(X6)!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_805, plain, (r2_hidden(k9_finseq_1(X1),X2)|v1_xboole_0(X3)|X2!=k13_finseq_1(X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_3937, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0))).#trainneg
cnf(i_0_5752, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,np__1)=k4_struct_0(esk1_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_5876, plain, (epred1_5(esk1_0,np__1,np__1,X1,k12_matrix_1(esk1_0,esk2_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_5548, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,np__1)=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_8179, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)=k1_funct_1(esk4_0,esk5_0)|esk6_0=esk5_0|esk5_0=np__1)).#trainneg
cnf(i_0_8160, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk5_0)=k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)|esk6_0=esk5_0|esk5_0=np__1)).#trainneg
cnf(i_0_5756, plain, (X1=np__1|epred1_5(esk1_0,np__1,X1,X2,k12_matrix_1(esk1_0,X3))|k3_matrix14(esk1_0,X3,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5747, plain, (epred1_5(esk1_0,np__1,X1,X2,k12_matrix_1(esk1_0,X3))|k3_matrix14(esk1_0,X3,X1)!=k1_funct_1(esk4_0,esk3_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_5385, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,np__1)=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4888, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(X1,esk3_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8736, plain, (v1_finseq_1(esk10_1(k13_finseq_1(X1))))).#trainneg
cnf(i_0_1999, negated_conjecture, (v3_xxreal_0(esk2_0)|v2_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),esk2_0))).#trainneg
cnf(i_0_3996, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,np__1)=k9_finseq_1(X1)|k9_xtuple_0(k4_matrix14(esk1_0,esk2_0,np__1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)!=X1|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_813, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,X2))=k4_xcmplx_0(X1)|~v1_xxreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~r1_xxreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)),k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1535, negated_conjecture, (esk2_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9969, negated_conjecture, (esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_815, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~r1_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k2_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_817, plain, (k6_xcmplx_0(X1,X2)=X1|~v1_xxreal_0(k6_xcmplx_0(X1,X2))|~r1_xxreal_0(k6_xcmplx_0(X1,X2),X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3937, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k1_funct_1(esk4_0,esk3_0))).#trainneg
cnf(i_0_9950, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,np__1)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,X1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9944, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,np__1)=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,esk3_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9967, plain, (m2_finseq_1(k3_matrix14(X1,X2,X3),X4)|v6_struct_0(X1)|k13_finseq_1(u1_struct_0(X1))!=k13_finseq_1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_10177, negated_conjecture, (esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2369, plain, (v7_ordinal1(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8169, negated_conjecture, (k9_finseq_1(X1)=k1_funct_1(esk4_0,esk5_0)|esk6_0=esk5_0|esk5_0=np__1|k9_xtuple_0(k1_funct_1(esk4_0,esk5_0))!=k2_finseq_1(np__1)|k5_struct_0(esk1_0)!=X1|~v1_relat_1(k1_funct_1(esk4_0,esk5_0)))).#trainneg
cnf(i_0_821, plain, (v3_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_822, plain, (m2_finseq_1(X1,X2)|~m1_subset_1(X1,k4_finseq_2(X3,X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8160, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk5_0)=k3_matrix14(esk1_0,esk2_0,esk3_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_4401, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k3_matrix14(esk1_0,esk2_0,esk3_0))).#trainneg
cnf(i_0_10328, negated_conjecture, (m2_finseq_1(k1_funct_1(esk4_0,esk5_0),u1_struct_0(esk1_0))|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_9951, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,np__1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10395, plain, (X1=np__1|epred1_5(esk1_0,np__1,X1,X2,k12_matrix_1(esk1_0,esk2_0))|X1=esk3_0|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10417, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|X3=esk3_0|np__1=X3|X3=X4|X3=X2|X1!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7586, negated_conjecture, (k12_matrix_1(esk1_0,esk2_0)=k9_finseq_1(X1)|k9_xtuple_0(k12_matrix_1(esk1_0,esk2_0))!=k2_finseq_1(np__1)|k1_funct_1(esk4_0,esk3_0)!=X1|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10319, negated_conjecture, (v1_relat_1(k1_funct_1(esk4_0,esk3_0)))).#trainneg
cnf(i_0_922, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m2_subset_1(X2,X1,X1))).#trainneg
cnf(i_0_854, plain, (v1_relat_1(esk8_4(X1,X2,X3,X4))|v1_xboole_0(X1)|m1_matrix_1(X4,X1,X2,X3)|k3_finseq_1(X4)!=X2|~v1_matrix_1(X4)|~m2_finseq_1(X4,k13_finseq_1(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5388, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,X1)=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X2,X1)!=k1_funct_1(esk4_0,esk3_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9691, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|v2_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_5383, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,X1)=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X2,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5547, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),np__1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10318, negated_conjecture, (r2_hidden(k1_funct_1(esk4_0,esk3_0),X1)|X1!=k13_finseq_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8200, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)=k1_funct_1(esk4_0,esk3_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_4312, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0))|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_855, plain, (v1_relat_1(esk7_2(X1,X2))|X2=k13_finseq_1(X1)|r2_hidden(esk7_2(X1,X2),X2))).#trainneg
cnf(i_0_8343, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(esk5_0,esk6_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_856, plain, (v1_relat_1(k8_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_857, plain, (v1_relat_1(k17_finseq_1(X1,X2,X3))|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1214, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),esk3_0)=k3_matrix14(esk1_0,esk2_0,esk3_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_3235, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X1!=k1_funct_1(k3_matrix14(esk1_0,esk2_0,esk3_0),X3)|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10186, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),esk3_0)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2040, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),esk3_0)=k3_matrix14(esk1_0,esk2_0,esk3_0)|esk3_0=X1|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10551, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2371, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10539, negated_conjecture, (v1_finseq_1(k1_funct_1(esk4_0,np__1)))).#trainneg
cnf(i_0_7584, negated_conjecture, (k9_finseq_1(X1)=k1_funct_1(esk4_0,esk3_0)|k9_xtuple_0(k1_funct_1(esk4_0,esk3_0))!=k2_finseq_1(np__1)|k5_struct_0(esk1_0)!=X1)).#trainneg
cnf(i_0_858, plain, (v1_relat_1(k2_rfinseq(X1,X2,X3))|~m1_finseq_1(X3,X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10540, negated_conjecture, (v1_funct_1(k1_funct_1(esk4_0,np__1)))).#trainneg
cnf(i_0_10671, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_3902, negated_conjecture, (esk2_0=X1|esk2_0=esk5_0|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_859, plain, (v1_relat_1(k12_finseq_1(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_860, plain, (v1_relat_1(k1_finseq_7(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_4401, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk3_0)=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_10550, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10577, plain, (m2_finseq_1(k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4),X1)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m1_finseq_1(k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),X1)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10553, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,esk3_0)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|k1_funct_1(esk4_0,np__1)!=k1_funct_1(esk4_0,X1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10631, negated_conjecture, (k1_funct_1(esk4_0,esk3_0)=k1_funct_1(esk4_0,esk5_0)|esk6_0=esk5_0|esk5_0=np__1)).#trainneg
cnf(i_0_10554, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,esk3_0)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|k1_funct_1(esk4_0,np__1)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10576, plain, (k7_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m1_finseq_1(k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),X1)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10584, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k1_rfinseq(np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1))).#trainneg
cnf(i_0_10557, plain, (X1=np__1|epred1_5(esk1_0,np__1,X1,X2,k12_matrix_1(esk1_0,esk2_0))|X1=esk3_0|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7569, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk6_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10458, plain, (v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2090, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),np__1))).#trainneg
cnf(i_0_5767, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0))|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_862, plain, (v1_relat_1(k2_finseq_7(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3910, negated_conjecture, (esk2_0=X1|esk2_0=esk3_0|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4692, negated_conjecture, (esk2_0=X1|esk2_0=esk6_0|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_863, plain, (v1_relat_1(k3_matrix14(X1,X2,X3))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_864, plain, (m1_subset_1(esk7_2(X1,esk10_1(k1_zfmisc_1(X2))),X2)|esk10_1(k1_zfmisc_1(X2))=k13_finseq_1(X1)|m2_finseq_1(esk7_2(X1,esk10_1(k1_zfmisc_1(X2))),X1))).#trainneg
cnf(i_0_10770, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,np__1)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_10776, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk5_0)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10773, plain, (epred1_5(esk1_0,esk3_0,X1,X2,k12_matrix_1(esk1_0,X3))|k3_matrix14(esk1_0,X3,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(esk3_0,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10835, negated_conjecture, (k3_finseq_1(k2_rfinseq(X1,X2,X3))=k3_finseq_1(k2_finseq_7(X1,X3,np__1,X2))|v1_xboole_0(X1)|r1_xxreal_0(X2,np__1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X3,X2)),k17_finseq_1(X1,k7_nat_d(X2,np__2),k2_rfinseq(X1,np__1,X3))),k12_finseq_1(X1,k7_partfun1(X1,X3,np__1))),X1)|~m1_finseq_1(k2_rfinseq(X1,X2,X3),X1)|~v1_xxreal_0(k3_finseq_1(k2_finseq_7(X1,X3,np__1,X2)))|~m2_finseq_1(X3,X1)|~r1_xxreal_0(k3_finseq_1(k2_finseq_7(X1,X3,np__1,X2)),k3_finseq_1(k2_rfinseq(X1,X2,X3)))|~r1_xxreal_0(X2,k3_finseq_1(X3))|~v7_ordinal1(X2)|~m1_subset_1(k3_finseq_1(k2_rfinseq(X1,X2,X3)),k5_numbers)|~v7_ordinal1(k3_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X3,X2)),k17_finseq_1(X1,k7_nat_d(X2,np__2),k2_rfinseq(X1,np__1,X3))),k12_finseq_1(X1,k7_partfun1(X1,X3,np__1))))))).#trainneg
cnf(i_0_10834, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk5_0)=k1_funct_1(esk4_0,np__1)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_10891, plain, (X1=np__1|epred1_5(esk1_0,esk3_0,X1,X2,k12_matrix_1(esk1_0,esk2_0))|X1=esk3_0|k1_funct_1(esk4_0,np__1)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1193, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_10836, negated_conjecture, (~v7_ordinal1(k5_numbers))).#trainneg
cnf(i_0_8741, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0))|esk5_0=np__1|esk6_0=esk5_0|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk5_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_865, plain, (m1_subset_1(k3_matrix14(X1,X2,X3),X4)|v6_struct_0(X1)|v2_struct_0(X1)|esk10_1(k1_zfmisc_1(X4))!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_10779, negated_conjecture, (v1_relat_1(k1_funct_1(esk4_0,np__1)))).#trainneg
cnf(i_0_10811, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk3_0))).#trainneg
cnf(i_0_10813, negated_conjecture, (esk2_0=esk5_0|esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_866, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk10_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk10_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_10578, plain, (r2_hidden(k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4),X5)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|X5!=k13_finseq_1(X1)|~m1_finseq_1(k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),X1)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10887, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk5_0),np__1)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_10911, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,np__1),esk5_0)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_1405, plain, (r1_xxreal_0(X1,np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_10806, plain, (r1_xxreal_0(X1,k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3998, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)!=X2|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_871, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11006, negated_conjecture, (esk2_0=esk3_0|esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11008, negated_conjecture, (esk2_0=esk6_0|esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_876, plain, (m2_subset_1(k3_finseq_1(k7_finseq_1(X1,X2)),k1_numbers,k5_numbers)|~v7_ordinal1(k3_finseq_1(X1))|~m1_subset_1(k3_finseq_1(X2),k5_numbers)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11047, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk5_0)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_11055, plain, (esk5_0=np__1|epred1_5(esk1_0,esk5_0,X1,X2,k12_matrix_1(esk1_0,X3))|esk6_0=esk5_0|k3_matrix14(esk1_0,X3,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(esk5_0,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11208, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk5_0)=k4_struct_0(esk1_0)|esk6_0=esk5_0|esk5_0=np__1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_4112, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(X2)|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk3_0)!=X2|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_11133, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk3_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_881, plain, (v1_relat_1(k2_matrix_1(k4_matrix14(X2,X1,X3)))|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_11132, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11121, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,esk5_0)=k1_funct_1(esk4_0,np__1)|esk6_0=esk5_0|esk5_0=np__1)).#trainneg
cnf(i_0_11158, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,esk3_0)|esk3_0=X2|X1!=k5_struct_0(esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11129, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11195, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,np__1)|X1!=k5_struct_0(esk1_0)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10548, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X2),X3,X4)|k3_matrix14(esk1_0,X2,X3)!=k1_funct_1(esk4_0,np__1)|X1!=k1_funct_1(k1_funct_1(esk4_0,np__1),X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10848, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk5_0)=k4_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10102, negated_conjecture, (k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,np__1)=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_891, plain, (k1_finseq_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5874, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_883, plain, (esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7)=k1_funct_1(k4_matrix14(X2,X3,X4),X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(k2_finseq_1(X3),k2_finseq_1(X3)))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_884, plain, (m2_finseq_1(esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7),X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(k2_finseq_1(X3),k2_finseq_1(X3)))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10616, plain, (X1=np__1|epred1_5(esk1_0,X1,np__1,X2,k12_matrix_1(esk1_0,esk2_0))|X1=esk3_0|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9196, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_11130, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10541, negated_conjecture, (r2_hidden(k1_funct_1(esk4_0,np__1),X1)|X1!=k13_finseq_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7594, plain, (X1=np__1|epred1_5(esk1_0,X2,X1,X3,k4_matrix14(esk1_0,esk2_0,X4))|X1=esk3_0|X1=X4|X1=X2|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11212, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_892, plain, (k1_card_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7073, negated_conjecture, (esk2_0=X1|esk2_0=np__1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1527, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,esk5_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_885, plain, (k1_funct_1(esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7),X6)=X7|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(k2_finseq_1(X3),k2_finseq_1(X3)))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_886, plain, (X1=k3_matrix_1(X2,k4_matrix14(X3,X4,X5),X6,X7)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|v2_struct_0(X3)|X1!=k1_funct_1(X8,X7)|X8!=k1_funct_1(k4_matrix14(X3,X4,X5),X6)|~v1_matrix_1(k4_matrix14(X3,X4,X5))|~m2_finseq_1(k4_matrix14(X3,X4,X5),k13_finseq_1(X2))|~m2_finseq_1(X8,X2)|~r2_hidden(k4_tarski(X6,X7),k2_zfmisc_1(k2_finseq_1(X4),k2_finseq_1(X4)))|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~m1_subset_1(X1,X2)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_10428, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),esk6_0)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11341, negated_conjecture, (epred1_5(esk1_0,esk5_0,np__1,X1,esk4_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_11418, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk3_0,esk3_0)=k5_struct_0(esk1_0)|esk3_0=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11443, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X2),X3,esk3_0)|k3_matrix14(esk1_0,X2,X3)!=k1_funct_1(esk4_0,np__1)|X1!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(esk3_0,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11440, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,np__1)=k5_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8196, negated_conjecture, (k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1)=k9_finseq_1(X2)|esk6_0=esk5_0|esk5_0=np__1|k9_xtuple_0(k7_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)!=X2|~v1_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk5_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_893, plain, (k9_xtuple_0(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_890, plain, (v1_xboole_0(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_894, plain, (k10_xtuple_0(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11393, negated_conjecture, (k9_finseq_1(X1)=k1_funct_1(esk4_0,np__1)|esk5_0=np__1|esk6_0=esk5_0|k9_xtuple_0(k1_funct_1(esk4_0,np__1))!=k2_finseq_1(np__1)|k4_struct_0(esk1_0)!=X1)).#trainneg
cnf(i_0_11437, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),X3,np__1)|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11442, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X2),X3,esk5_0)|esk6_0=esk5_0|esk5_0=np__1|k3_matrix14(esk1_0,X2,X3)!=k1_funct_1(esk4_0,np__1)|X1!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(esk5_0,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11415, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),esk3_0,esk3_0)|esk3_0=X3|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k5_struct_0(esk1_0)|~v7_ordinal1(X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11545, negated_conjecture, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(esk4_0,esk5_0)),k17_finseq_1(X1,k7_nat_d(esk3_0,np__2),k1_rfinseq(np__1,esk4_0))),k9_finseq_1(k7_partfun1(X1,esk4_0,np__1))),k2_rfinseq(X1,esk3_0,esk4_0))=k2_finseq_7(X1,esk4_0,np__1,esk3_0)|v1_xboole_0(X1)|esk5_0=np__1|esk6_0=esk5_0|~m1_finseq_1(esk4_0,X1))).#trainneg
cnf(i_0_4895, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(np__2,esk5_0)|r1_xxreal_0(esk3_0,np__2))).#trainneg
cnf(i_0_8199, negated_conjecture, (k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)=k9_finseq_1(X2)|esk6_0=esk5_0|esk5_0=np__1|k9_xtuple_0(k7_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1))!=k2_finseq_1(np__1)|k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,esk5_0)!=X2|~v1_finseq_1(k12_matrix_1(esk1_0,esk2_0))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_895, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(esk14_1(X1),X2)|v7_struct_0(X1)|u1_struct_0(X1)!=k13_finseq_1(X2)|~l2_struct_0(X1))).#trainneg
cnf(i_0_896, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(esk12_1(X1),X2)|u1_struct_0(X1)!=k13_finseq_1(X2)|~l2_struct_0(X1))).#trainneg
cnf(i_0_897, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|m2_finseq_1(esk11_0,X1)|k1_zfmisc_1(k1_numbers)!=k13_finseq_1(X1))).#trainneg
cnf(i_0_3999, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X4!=k1_funct_1(esk4_0,esk3_0)|X1!=k1_funct_1(X4,X3)|~v1_matrix_1(k4_matrix14(esk1_0,esk2_0,X2))|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10543, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X2,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11439, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,np__1)=k4_struct_0(esk1_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_11441, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X2),X3,X1)=k1_funct_1(k1_funct_1(esk4_0,np__1),X1)|k3_matrix14(esk1_0,X2,X3)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k1_funct_1(esk4_0,np__1),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11210, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,esk5_0)|esk3_0=X2|esk6_0=esk5_0|esk5_0=np__1|X1!=k5_struct_0(esk1_0)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5120, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(np__2,esk6_0)|r1_xxreal_0(esk3_0,np__2))).#trainneg
cnf(i_0_898, plain, (v1_xboole_0(X1)|m2_finseq_1(esk10_1(X1),X2)|X1!=k13_finseq_1(X2))).#trainneg
cnf(i_0_11422, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_11424, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_900, plain, (v1_xboole_0(k1_numbers)|m2_finseq_1(np__0,X1)|k13_finseq_1(X1)!=k1_numbers)).#trainneg
cnf(i_0_903, plain, (v1_xboole_0(k1_numbers)|m2_finseq_1(np__2,X1)|k13_finseq_1(X1)!=k1_numbers)).#trainneg
cnf(i_0_8752, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_11417, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk3_0,esk3_0)=k4_struct_0(esk1_0)|esk3_0=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_7725, plain, (r1_xxreal_0(k3_finseq_1(X1),k3_finseq_1(k7_finseq_1(X2,X1)))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_11806, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,esk3_0)=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_11652, plain, (epred1_5(esk1_0,esk3_0,esk3_0,X1,k4_matrix14(esk1_0,esk2_0,X2))|esk3_0=X2|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11803, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,esk5_0)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_11821, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk3_0,esk5_0)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0|esk3_0=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_905, plain, (v1_xboole_0(k1_zfmisc_1(k5_numbers))|m2_finseq_1(k2_finseq_1(X1),X2)|k1_zfmisc_1(k5_numbers)!=k13_finseq_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11840, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk6_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_11841, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk3_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_11624, negated_conjecture, (esk2_0=np__1|esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11504, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_907, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k4_struct_0(X1),X2)|u1_struct_0(X1)!=k13_finseq_1(X2)|~l2_struct_0(X1))).#trainneg
cnf(i_0_908, plain, (v1_xboole_0(k1_numbers)|m2_finseq_1(np__1,X1)|k13_finseq_1(X1)!=k1_numbers)).#trainneg
cnf(i_0_11656, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk5_0,esk5_0)=k5_struct_0(esk1_0)|esk5_0=X1|esk5_0=np__1|esk6_0=esk5_0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11667, plain, (X1=np__1|epred1_5(esk1_0,np__1,X1,X2,k4_matrix14(esk1_0,esk2_0,X1))|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11865, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k1_funct_1(X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2)))).#trainneg
cnf(i_0_11661, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=esk10_1(u1_struct_0(esk1_0))|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|esk10_1(u1_struct_0(esk1_0))!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11956, plain, (r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X2,k9_finseq_1(X1))))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_11626, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,esk5_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11581, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3776, negated_conjecture, (k9_finseq_1(k4_struct_0(esk1_0))=k1_funct_1(esk4_0,X1)|np__1=X1|X1=esk3_0|k9_xtuple_0(k1_funct_1(esk4_0,X1))!=k2_finseq_1(np__1)|~v1_funct_1(k1_funct_1(esk4_0,X1))|~v1_relat_1(k1_funct_1(esk4_0,X1))|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_910, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|m2_finseq_1(k5_numbers,X1)|k1_zfmisc_1(k1_numbers)!=k13_finseq_1(X1))).#trainneg
cnf(i_0_11425, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk5_0)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_1533, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,esk3_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1536, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,esk6_0)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_911, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|m2_finseq_1(X2,X3)|k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))!=k13_finseq_1(X3)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_913, plain, (v7_ordinal1(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11804, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k5_struct_0(esk1_0)|np__1=X1|X1=esk3_0|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11867, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11662, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,esk3_0)|k3_matrix14(esk1_0,X3,X4)!=k1_funct_1(esk4_0,np__1)|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(esk3_0,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11961, plain, (epred1_5(esk1_0,esk3_0,esk3_0,X1,k12_matrix_1(esk1_0,esk2_0)))).#trainneg
cnf(i_0_11973, plain, (esk5_0=np__1|epred1_5(esk1_0,esk5_0,esk3_0,X1,k12_matrix_1(esk1_0,esk2_0))|esk6_0=esk5_0)).#trainneg
cnf(i_0_11966, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk5_0,esk5_0)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_4578, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1),np__1)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(esk1_0,esk2_0,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_916, plain, (m1_subset_1(esk7_2(X1,k2_finseq_1(X2)),k5_numbers)|k2_finseq_1(X2)=k13_finseq_1(X1)|m2_finseq_1(esk7_2(X1,k2_finseq_1(X2)),X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12137, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_917, plain, (m1_subset_1(k3_matrix14(X1,X2,X3),k5_numbers)|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X4)|k2_finseq_1(X4)!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_918, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_11969, negated_conjecture, (esk5_0=X1|epred1_5(esk1_0,esk5_0,esk5_0,X2,k4_matrix14(esk1_0,esk2_0,X1))|esk5_0=np__1|esk6_0=esk5_0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12002, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk6_0,np__2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_919, plain, (v1_xboole_0(k1_zfmisc_1(X1))|m2_finseq_1(X1,X2)|k1_zfmisc_1(X1)!=k13_finseq_1(X2))).#trainneg
cnf(i_0_920, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_934, plain, (v3_xxreal_0(np__2)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_11979, plain, (epred1_5(esk1_0,esk5_0,esk3_0,X1,k4_matrix14(esk1_0,esk2_0,X2))|esk6_0=esk5_0|esk5_0=np__1|esk3_0=X2|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12180, plain, (X1=np__1|epred1_5(esk1_0,X1,esk3_0,X2,k12_matrix_1(esk1_0,esk2_0))|X1=esk3_0|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12162, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,esk6_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12124, plain, (r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(k9_finseq_1(X2),X1)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_12120, plain, (r1_xxreal_0(np__1,k2_nat_1(np__1,np__1)))).#trainneg
cnf(i_0_12006, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk3_0,np__2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_951, plain, (m1_subset_1(k12_finseq_1(X1,X2),X3)|v1_xboole_0(X1)|esk10_1(k1_zfmisc_1(X3))!=k13_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12159, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk5_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_12030, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_952, plain, (m1_subset_1(k12_finseq_1(X1,X2),k5_numbers)|v1_xboole_0(X1)|~v7_ordinal1(X3)|k2_finseq_1(X3)!=k13_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_953, plain, (k3_finseq_1(k12_finseq_1(X1,X2))=X3|v1_xboole_0(X4)|v1_xboole_0(X1)|~m2_finseq_1(k12_finseq_1(X1,X2),X4)|~m1_matrix_1(X5,X4,X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|k10_xtuple_0(X5)!=k13_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12247, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,esk3_0)|k3_matrix14(esk1_0,X3,X4)!=k1_funct_1(esk4_0,np__1)|k1_funct_1(X1,X2)!=k5_struct_0(esk1_0)|~m1_finseq_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(esk3_0,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(X2,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12160, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,esk3_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12140, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_11805, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,np__1)=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_11659, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=esk14_1(esk1_0)|v7_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|esk14_1(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_7284, plain, (k7_partfun1(u1_struct_0(X1),X2,X3)=k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X4,X5),X6,X6)|np__1=X6|X6=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k7_partfun1(u1_struct_0(X1),X2,X3)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_9138, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k7_finseq_1(k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2)),X1)|~m1_finseq_1(k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_11542, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k9_finseq_1(k1_funct_1(X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2)))).#trainneg
cnf(i_0_7011, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),X4,X4)|np__1=X4|X4=esk3_0|X4=X3|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11668, plain, (X1=X2|np__1=X2|epred1_5(esk1_0,np__1,X1,X2,k4_matrix14(esk1_0,esk2_0,X1))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7528, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(np__2,esk5_0)|r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_4115, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X1)=k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1)|~m2_finseq_1(k1_funct_1(esk4_0,esk3_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_955, plain, (m1_subset_1(k12_finseq_1(X1,X2),X3)|v1_xboole_0(X1)|X3!=k13_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12276, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12277, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_956, plain, (m2_finseq_1(k12_finseq_1(X1,X2),X3)|v1_xboole_0(X1)|X4!=k13_finseq_1(X3)|X4!=k13_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12248, negated_conjecture, (epred1_5(esk1_0,esk5_0,esk5_0,X1,k12_matrix_1(esk1_0,esk2_0))|esk5_0=np__1|esk6_0=esk5_0)).#trainneg
cnf(i_0_12338, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk5_0,np__2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_11598, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_971, plain, (r2_hidden(k4_matrix14(X1,X2,X3),X4)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|X4!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_973, plain, (v1_funct_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_12407, plain, (k1_funct_1(X2,X3)=k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X4,X5),X6,X6)|np__1=X6|X6=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_funct_1(X2,X3)!=k5_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_12356, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,esk3_0)|k3_matrix14(esk1_0,X3,X4)!=k1_funct_1(esk4_0,np__1)|k1_funct_1(X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X2,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(esk3_0,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10906, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X3,esk3_0)|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,np__1)!=k1_funct_1(esk4_0,X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7328, negated_conjecture, (k3_finseq_1(k7_finseq_1(X1,esk4_0))=k2_xcmplx_0(esk2_0,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_11600, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_11601, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(np__2,esk5_0)|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_11801, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X4!=k1_funct_1(esk4_0,esk3_0)|X1!=k1_funct_1(X4,X3)|~v1_matrix_1(k4_matrix14(esk1_0,esk2_0,X2))|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9946, negated_conjecture, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X2)),k17_finseq_1(X1,k7_nat_d(X2,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(esk1_0,esk2_0)))),k12_finseq_1(X1,k1_funct_1(esk4_0,esk3_0))),k2_rfinseq(X1,X2,k12_matrix_1(esk1_0,esk2_0)))=k2_finseq_7(X1,k12_matrix_1(esk1_0,esk2_0),np__1,X2)|v1_xboole_0(X1)|r1_xxreal_0(X2,np__1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10831, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|X4!=k1_funct_1(esk4_0,np__1)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8781, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k7_partfun1(X1,k12_matrix_1(X2,X3),np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|v2_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_11660, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=esk12_1(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|esk12_1(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_11866, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_12730, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),k2_xcmplx_0(esk2_0,k3_finseq_1(X1)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1))|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_12735, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),k2_xcmplx_0(esk2_0,k3_finseq_1(X1)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_10878, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=esk5_0|k1_funct_1(k1_funct_1(esk4_0,esk5_0),X1)!=k5_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,esk5_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12355, plain, (k3_finseq_1(k12_finseq_1(X1,X2))=X3|v1_xboole_0(X1)|k10_xtuple_0(X4)!=k13_finseq_1(X1)|~m1_matrix_1(X4,X1,X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9170, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k7_finseq_1(k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2)),X1)|~m1_finseq_1(k9_finseq_1(k7_partfun1(X1,X2,X3)),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1))).#trainneg
cnf(i_0_12588, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3522, negated_conjecture, (k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1)=k1_funct_1(esk4_0,esk3_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_996, plain, (k1_finseq_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_997, plain, (k1_card_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13015, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k5_struct_0(esk1_0)|np__1=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13021, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk5_0),X1)=k4_struct_0(esk1_0)|np__1=X1|esk5_0=np__1|esk6_0=esk5_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13019, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k4_struct_0(esk1_0)|np__1=X2|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13014, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X3,X4)|np__1=X4|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k4_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_12900, negated_conjecture, (r1_xxreal_0(k3_finseq_1(k12_matrix_1(X1,X2)),k2_nat_1(esk2_0,X2))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_11702, plain, (k2_finseq_1(np__1)!=np__0|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_999, plain, (k9_xtuple_0(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8003, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(esk3_0)|~v1_xboole_0(esk2_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_12572, negated_conjecture, (esk2_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1000, plain, (k10_xtuple_0(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1003, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_12996, negated_conjecture, (k3_finseq_1(k12_finseq_1(u1_struct_0(esk1_0),X1))=esk2_0|k10_xtuple_0(esk4_0)!=k13_finseq_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12997, plain, (k3_finseq_1(k12_finseq_1(u1_struct_0(X1),X2))=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|k10_xtuple_0(k12_matrix_1(X1,X3))!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1))).#trainneg
cnf(i_0_12998, plain, (k3_finseq_1(k12_finseq_1(u1_struct_0(X1),X2))=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|k10_xtuple_0(k4_matrix14(X1,X3,X4))!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9217, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k7_finseq_1(k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1)|~m1_finseq_1(k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2)),X1)|~m1_finseq_1(k12_finseq_1(X1,k1_funct_1(X2,X3)),X1))).#trainneg
cnf(i_0_1852, negated_conjecture, (k1_card_1(esk2_0)=esk2_0|~v1_finseq_1(esk4_0)|~v1_funct_1(esk4_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_13404, negated_conjecture, (esk2_0=np__1|k10_xtuple_0(esk4_0)!=k13_finseq_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1004, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13323, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_13325, negated_conjecture, (esk2_0=X1|~r1_xxreal_0(esk2_0,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12029, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(X1,esk5_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1005, plain, (k1_finseq_1(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13175, plain, (X1=np__1|epred1_5(esk1_0,X1,esk3_0,X2,esk4_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13248, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k4_struct_0(esk1_0)|np__1=X2|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13233, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),np__1)=k1_funct_1(k12_matrix_1(X1,X2),X3)|v6_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13235, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X4)=k1_funct_1(k4_matrix14(X1,X2,X5),X4)|np__1=X4|X4=X3|X4=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13238, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k1_funct_1(k4_matrix14(X1,X2,X4),X4)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13215, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k5_struct_0(esk1_0)|np__1=X2|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13237, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),X3)=k1_funct_1(k12_matrix_1(X1,X2),np__1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13469, negated_conjecture, (k3_finseq_1(esk2_0)=esk2_0|~v1_finseq_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_13505, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_0)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1610, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(esk5_0)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1008, plain, (k9_xtuple_0(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1009, plain, (k10_xtuple_0(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13723, plain, (X1=np__1|epred1_5(esk1_0,X1,X2,X3,k4_matrix14(esk1_0,esk2_0,X2))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1010, plain, (k1_finseq_1(X1)=k10_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1619, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12614, plain, (m2_finseq_1(k4_matrix14(X1,X2,X3),X4)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|X5!=k13_finseq_1(X4)|X5!=k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9012, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k1_funct_1(X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2)))).#trainneg
cnf(i_0_12522, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),X4,X4)|X4=esk3_0|np__1=X4|X4=X3|k1_funct_1(X1,X2)!=k5_struct_0(esk1_0)|~m1_finseq_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(X2,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10575, plain, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4))=X5|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),X4,k12_matrix_1(X2,X3)),k13_finseq_1(X1))|~m1_finseq_1(k8_finseq_1(k13_finseq_1(X1),k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,X4)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,np__1))),k13_finseq_1(X1))|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13803, negated_conjecture, (r1_xxreal_0(esk2_0,esk3_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_13805, negated_conjecture, (r1_xxreal_0(X1,esk5_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13969, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1014, plain, (k10_xtuple_0(X1)=k10_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1024, plain, (m2_subset_1(k2_nat_1(np__1,k3_finseq_1(X2)),k1_numbers,k5_numbers)|~v1_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_funct_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_relat_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_14046, plain, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4))=X5|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m1_finseq_1(k7_finseq_1(k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,X4)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,np__1))),k13_finseq_1(X1))|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),X4,k12_matrix_1(X2,X3)),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,np__1)),k13_finseq_1(X1))|~m1_finseq_1(k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,X4)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(X2,X3)))),k13_finseq_1(X1)))).#trainneg
cnf(i_0_14047, plain, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4))=X5|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),X4,k12_matrix_1(X2,X3)),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k8_finseq_1(k13_finseq_1(X1),k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,X4)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,np__1))),k13_finseq_1(X1)))).#trainneg
cnf(i_0_1473, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11502, plain, (v3_card_1(np__0,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6123, plain, (k2_xcmplx_0(np__2,np__1)=k2_xcmplx_0(np__1,np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1025, plain, (k2_nat_1(k2_nat_1(np__1,k3_finseq_1(X2)),k3_finseq_1(X3))=k3_finseq_1(k7_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2),X3))|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_14216, plain, (~v1_xboole_0(k2_xcmplx_0(np__1,np__2))|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1026, plain, (v1_xboole_0(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_funct_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_relat_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_xboole_0(k2_nat_1(np__1,k3_finseq_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1027, plain, (k2_nat_1(k3_finseq_1(X1),k2_nat_1(np__1,k3_finseq_1(X3)))=k3_finseq_1(k7_finseq_1(X1,k7_finseq_1(k9_finseq_1(X2),X3)))|~v1_finseq_1(k7_finseq_1(k9_finseq_1(X2),X3))|~v1_finseq_1(X1)|~v1_funct_1(k7_finseq_1(k9_finseq_1(X2),X3))|~v1_funct_1(X1)|~v1_relat_1(k7_finseq_1(k9_finseq_1(X2),X3))|~v1_relat_1(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_14190, negated_conjecture, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,esk5_0))=X2|v1_xboole_0(X1)|r1_xxreal_0(esk5_0,np__1)|esk5_0=np__1|esk6_0=esk5_0|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),esk5_0,k12_matrix_1(esk1_0,esk2_0)),k13_finseq_1(X1))|~m2_finseq_1(k8_finseq_1(k13_finseq_1(X1),k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,np__1)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(esk5_0,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(esk1_0,esk2_0)))),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,esk3_0))),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk5_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,esk5_0),X1,X2,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14062, negated_conjecture, (esk6_0=np__1|esk5_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk5_0)!=k5_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,esk5_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk5_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_1034, plain, (r2_hidden(k5_finseq_2(X1,X2,X3),X4)|v1_xboole_0(X1)|X4!=k13_finseq_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_5597, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk6_0)|~r1_xxreal_0(esk5_0,np__1))).#trainneg
cnf(i_0_12619, plain, (v1_card_1(X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1035, plain, (v1_relat_1(k5_finseq_2(X1,X2,X3))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_1036, plain, (v1_funct_1(k5_finseq_2(X1,X2,X3))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_14179, negated_conjecture, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,esk5_0))=X2|v1_xboole_0(X1)|r1_xxreal_0(esk5_0,np__1)|esk5_0=np__1|esk6_0=esk5_0|~m1_finseq_1(k7_finseq_1(k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,np__1)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(esk5_0,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(esk1_0,esk2_0)))),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,esk3_0))),k13_finseq_1(X1))|~m1_finseq_1(k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,np__1)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(esk5_0,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(esk1_0,esk2_0)))),k13_finseq_1(X1))|~m1_finseq_1(k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,esk3_0)),k13_finseq_1(X1))|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),esk5_0,k12_matrix_1(esk1_0,esk2_0)),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk5_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,esk5_0),X1,X2,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14197, plain, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4))=X5|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),X4,k12_matrix_1(X2,X3)),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,np__1)),k13_finseq_1(X1))|~m1_finseq_1(k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,X4)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(X2,X3)))),k13_finseq_1(X1)))).#trainneg
cnf(i_0_14035, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),X4,X4)|np__1=X4|X4=esk3_0|X4=X3|k1_funct_1(X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X2,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1037, plain, (v1_finseq_1(k5_finseq_2(X1,X2,X3))|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_14568, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_14570, plain, (v1_card_1(X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1054, plain, (k2_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13978, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k4_matrix14(X1,X2,X3),X4)|v6_struct_0(X1)|k13_finseq_1(k13_finseq_1(u1_struct_0(X1)))!=k13_finseq_1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_14188, plain, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4))=X5|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),X4,k12_matrix_1(X2,X3)),k13_finseq_1(X1))|~m2_finseq_1(k8_finseq_1(k13_finseq_1(X1),k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,X4)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(X4,np__2),k1_rfinseq(np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(X2,X3,np__1))),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(X2,X3),np__1,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(X1)))).#trainneg
cnf(i_0_13991, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,k12_matrix_1(X2,X3),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k3_matrix14(X2,X3,np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|v2_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X4)|~m1_subset_1(k7_partfun1(X1,k12_matrix_1(X2,X3),X4),X1)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_14019, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X3,np__2),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),np__1,k4_matrix14(X1,X2,X3)))),k12_finseq_1(k13_finseq_1(u1_struct_0(X1)),k1_funct_1(k4_matrix14(X1,X2,X3),np__1))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X3,k4_matrix14(X1,X2,X3)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1,X3)|r1_xxreal_0(X3,np__1)|v6_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(X3,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1),k13_finseq_1(u1_struct_0(X1)))|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_14193, negated_conjecture, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,esk3_0))=X2|v1_xboole_0(X1)|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),esk3_0,k12_matrix_1(esk1_0,esk2_0)),k13_finseq_1(X1))|~m2_finseq_1(k8_finseq_1(k13_finseq_1(X1),k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,np__1)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(esk3_0,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(esk1_0,esk2_0)))),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,esk3_0))),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,esk3_0),X1,X2,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11733, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk5_0)=esk10_1(u1_struct_0(esk1_0))|esk5_0=np__1|esk6_0=esk5_0|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|esk10_1(u1_struct_0(esk1_0))!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14211, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14213, plain, (v7_ordinal1(k2_xcmplx_0(np__1,np__2))|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_12618, plain, (v1_xboole_0(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1060, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1061, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_12623, plain, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(X3),k4_matrix14(X3,X4,X5),X6,X6)|np__1=X6|X6=X5|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|k1_funct_1(X1,X2)!=k5_struct_0(X3)|~r1_xxreal_0(X2,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X6,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~m2_finseq_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_12960, negated_conjecture, (r1_xxreal_0(X2,k2_xcmplx_0(esk2_0,X2))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_14736, plain, (epred1_5(esk1_0,esk5_0,esk5_0,X1,k12_matrix_1(esk1_0,esk2_0))|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_14724, plain, (epred1_5(esk1_0,esk5_0,np__1,X1,esk4_0)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_14718, plain, (epred1_5(esk1_0,np__1,esk5_0,X1,esk4_0)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_14674, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk5_0,esk5_0)=k5_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_14243, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(np__1,np__2))|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1066, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14888, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_14748, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk6_0))).#trainneg
cnf(i_0_14878, plain, (v7_ordinal1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1078, plain, (r1_xxreal_0(X1,esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14716, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,np__1)=k5_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_14723, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk5_0)=k5_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_14719, plain, (epred1_5(esk1_0,esk5_0,X1,X2,k12_matrix_1(esk1_0,X3))|esk6_0=np__1|esk5_0=np__1|k3_matrix14(esk1_0,X3,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(esk5_0,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14713, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk5_0)=k5_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14727, plain, (esk5_0=X1|epred1_5(esk1_0,esk5_0,esk5_0,X2,k4_matrix14(esk1_0,esk2_0,X1))|esk6_0=np__1|esk5_0=np__1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14729, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk5_0,esk5_0)=k5_struct_0(esk1_0)|esk5_0=X1|esk6_0=np__1|esk5_0=np__1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14709, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,esk5_0)=k1_funct_1(esk4_0,np__1)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_14684, negated_conjecture, (v2_xxreal_0(esk5_0)|esk6_0=np__1|esk5_0=np__1|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_1083, plain, (r1_xxreal_0(X1,k6_xcmplx_0(X2,X3))|r1_xxreal_0(k6_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14884, plain, (X1=k4_matrix14(X2,X3,X4)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1587, plain, (v1_xboole_0(X2)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_xboole_0(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1084, plain, (r1_xxreal_0(X1,k1_card_1(X2))|r1_xxreal_0(k1_card_1(X2),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_14710, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk5_0)=k1_funct_1(esk4_0,np__1)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_11666, plain, (epred1_5(esk1_0,np__1,X1,X2,k4_matrix14(esk1_0,esk2_0,X1))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9710, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),np__1,k4_matrix14(X1,X2,X3)))),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X4,k4_matrix14(X1,X2,X3)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1,X4)|r1_xxreal_0(X4,np__1)|v6_struct_0(X1)|v2_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~v7_ordinal1(X4)|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4),k13_finseq_1(u1_struct_0(X1)))|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_11213, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15096, plain, (r1_xxreal_0(esk10_1(k5_numbers),esk10_1(k5_numbers)))).#trainneg
cnf(i_0_11207, negated_conjecture, (esk12_1(esk1_0)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk5_0)|esk6_0=esk5_0|esk5_0=np__1|esk12_1(esk1_0)!=k5_struct_0(esk1_0)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_1086, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15093, negated_conjecture, (v7_ordinal1(esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_15074, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_15078, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk6_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15077, negated_conjecture, (esk3_0=esk5_0|r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(X1,esk3_0)|~r1_xxreal_0(esk5_0,np__1)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1087, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1088, plain, (r1_xxreal_0(X1,k9_xtuple_0(X2))|r1_xxreal_0(k9_xtuple_0(X2),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1089, plain, (r1_xxreal_0(X1,k10_xtuple_0(X2))|r1_xxreal_0(k10_xtuple_0(X2),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3154, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(np__2)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_3305, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(np__2)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk5_0))).#trainneg
cnf(i_0_15082, negated_conjecture, (v3_xxreal_0(esk6_0)|v2_xxreal_0(k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk6_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15083, negated_conjecture, (v3_xxreal_0(esk5_0)|v2_xxreal_0(k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk5_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1107, plain, (k1_funct_1(k10_funct_7(X2,X3,X4),X3)=k7_partfun1(X1,X2,X4)|v1_xboole_0(X1)|~m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(k10_funct_7(X2,X3,X4)))|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_15180, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,esk5_0)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=np__1)).#trainneg
cnf(i_0_15035, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,np__1)=k4_struct_0(esk1_0)|esk5_0=np__1|esk6_0=np__1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_15163, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk5_0)=k4_struct_0(esk1_0)|esk5_0=np__1|esk6_0=np__1|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk5_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1122, plain, (m1_subset_1(k3_matrix_1(X1,esk7_2(k13_finseq_1(X1),X2),X3,X4),X1)|X2=k13_finseq_1(k13_finseq_1(X1))|r2_hidden(esk7_2(k13_finseq_1(X1),X2),X2)|~v1_matrix_1(esk7_2(k13_finseq_1(X1),X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3408, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(np__2)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_15084, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk3_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15085, negated_conjecture, (v3_xxreal_0(esk2_0)|v2_xxreal_0(k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1111, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X5)),k17_finseq_1(X1,k7_nat_d(X5,np__2),k2_rfinseq(X1,np__1,k10_funct_7(X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,k10_funct_7(X2,X3,X4),np__1))),k2_rfinseq(X1,X5,k10_funct_7(X2,X3,X4)))=k2_finseq_7(X1,k10_funct_7(X2,X3,X4),np__1,X5)|v1_xboole_0(X1)|r1_xxreal_0(X5,np__1)|X5=X4|X3=X5|~m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~r1_xxreal_0(X5,k3_finseq_1(k10_funct_7(X2,X3,X4)))|~v7_ordinal1(X5)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X5,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1112, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k10_funct_7(X2,X3,X4),X5)),k17_finseq_1(X1,k7_nat_d(X5,np__2),k2_rfinseq(X1,np__1,k10_funct_7(X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X5,k10_funct_7(X2,X3,X4)))=k2_finseq_7(X1,k10_funct_7(X2,X3,X4),np__1,X5)|v1_xboole_0(X1)|r1_xxreal_0(X5,np__1)|np__1=X4|X3=np__1|~m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~r1_xxreal_0(X5,k3_finseq_1(k10_funct_7(X2,X3,X4)))|~v7_ordinal1(X5)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_15446, plain, (epred1_5(esk1_0,esk5_0,esk3_0,X1,k12_matrix_1(esk1_0,esk2_0))|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_15454, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,np__1)=k4_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_15300, plain, (r1_xxreal_0(X1,np__2)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1124, plain, (m1_subset_1(k3_matrix_1(X1,k8_finseq_1(k13_finseq_1(X1),X2,X3),X4,X5),X1)|~v1_matrix_1(k8_finseq_1(k13_finseq_1(X1),X2,X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_finseq_1(X3,k13_finseq_1(X1))|~m1_finseq_1(X2,k13_finseq_1(X1)))).#trainneg
cnf(i_0_15364, plain, (v1_xreal_0(np__1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_15086, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),np__1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1125, plain, (m1_subset_1(k3_matrix_1(X1,k17_finseq_1(k13_finseq_1(X1),X2,X3),X4,X5),X1)|~v1_matrix_1(k17_finseq_1(k13_finseq_1(X1),X2,X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_finseq_1(X3,k13_finseq_1(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1126, plain, (m1_subset_1(k3_matrix_1(X1,k2_rfinseq(k13_finseq_1(X1),X2,X3),X4,X5),X1)|~v1_matrix_1(k2_rfinseq(k13_finseq_1(X1),X2,X3))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_finseq_1(X3,k13_finseq_1(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15177, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,esk5_0)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|esk5_0=np__1|esk6_0=np__1|k1_funct_1(esk4_0,np__1)!=k1_funct_1(esk4_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15010, plain, (X1=np__1|X1=esk5_0|epred1_5(esk1_0,X1,X1,X2,esk4_0)|esk5_0=np__1|esk6_0=np__1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15004, negated_conjecture, (k1_funct_1(esk4_0,esk3_0)=k1_funct_1(esk4_0,esk5_0)|esk5_0=np__1|esk6_0=np__1|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk5_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_10489, plain, (m2_finseq_1(esk10_1(k4_finseq_2(X1,X2)),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11209, negated_conjecture, (esk10_1(u1_struct_0(esk1_0))=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk5_0)|esk6_0=esk5_0|esk5_0=np__1|esk10_1(u1_struct_0(esk1_0))!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_1132, plain, (m1_subset_1(k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_matrix_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_15367, plain, (v1_xreal_0(np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_15431, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(np__2)|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_15080, plain, (r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|r1_xxreal_0(k4_matrix14(X2,X3,X4),X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~v7_ordinal1(X1)|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15088, negated_conjecture, (r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk6_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1141, plain, (k7_finseq_1(k12_matrix_1(X1,X2),X3)=k9_finseq_1(X4)|v6_struct_0(X1)|v2_struct_0(X1)|k9_xtuple_0(k7_finseq_1(k12_matrix_1(X1,X2),X3))!=k2_finseq_1(np__1)|k3_matrix14(X1,X2,np__1)!=X4|~v1_finseq_1(X3)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k12_matrix_1(X1,X2))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15262, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),np__1,k4_matrix14(X1,X2,X3)))),k9_finseq_1(k1_funct_1(k12_matrix_1(X1,X2),X3))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X4,k4_matrix14(X1,X2,X3)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1,X4)|r1_xxreal_0(X4,np__1)|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(k1_funct_1(k12_matrix_1(X1,X2),X3),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15155, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk5_0,X1)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=np__1|k3_matrix14(esk1_0,X2,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15002, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)=k1_funct_1(esk4_0,X1)|np__1=X1|X1=esk5_0|esk5_0=np__1|esk6_0=np__1|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15282, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X4)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X4,np__2),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),np__1,k4_matrix14(X1,X2,X3)))),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X4,k4_matrix14(X1,X2,X3)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1,X4)|r1_xxreal_0(X4,np__1)|v6_struct_0(X1)|np__1=X4|X4=X3|v2_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X4),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X4))).#trainneg
cnf(i_0_15666, plain, (v1_relat_1(esk10_1(k4_finseq_2(X1,X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1145, plain, (m1_subset_1(esk7_2(X1,esk11_0),k1_numbers)|k13_finseq_1(X1)=esk11_0|m2_finseq_1(esk7_2(X1,esk11_0),X1))).#trainneg
cnf(i_0_15723, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_15433, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(np__2)|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk5_0))).#trainneg
cnf(i_0_15471, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(np__2)|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_15726, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(k4_matrix14(X1,esk2_0,X2),X3)|r1_xxreal_0(X3,k4_matrix14(X1,esk2_0,X2))|v6_struct_0(X1)|~v1_xboole_0(esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1142, plain, (k7_finseq_1(k3_matrix14(X1,X2,np__1),X3)=k9_finseq_1(X4)|v6_struct_0(X1)|v2_struct_0(X1)|k9_xtuple_0(k7_finseq_1(k3_matrix14(X1,X2,np__1),X3))!=k2_finseq_1(np__1)|k5_struct_0(X1)!=X4|~v1_finseq_1(X3)|~v1_finseq_1(k3_matrix14(X1,X2,np__1))|~v1_funct_1(X3)|~v1_funct_1(k3_matrix14(X1,X2,np__1))|~v1_relat_1(X3)|~v1_relat_1(k3_matrix14(X1,X2,np__1))|~r1_xxreal_0(np__1,k3_finseq_1(k3_matrix14(X1,X2,np__1)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1143, plain, (k7_finseq_1(k3_matrix14(X1,X2,X3),X4)=k9_finseq_1(X5)|X3=np__1|v6_struct_0(X1)|v2_struct_0(X1)|k9_xtuple_0(k7_finseq_1(k3_matrix14(X1,X2,X3),X4))!=k2_finseq_1(np__1)|k4_struct_0(X1)!=X5|~v1_finseq_1(X4)|~v1_finseq_1(k3_matrix14(X1,X2,X3))|~v1_funct_1(X4)|~v1_funct_1(k3_matrix14(X1,X2,X3))|~v1_relat_1(X4)|~v1_relat_1(k3_matrix14(X1,X2,X3))|~r1_xxreal_0(np__1,k3_finseq_1(k3_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15826, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk5_0,esk3_0)=k5_struct_0(esk1_0)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_15658, negated_conjecture, (k1_funct_1(esk4_0,esk3_0)=k1_funct_1(esk4_0,esk5_0)|esk6_0=np__1|esk5_0=np__1)).#trainneg
cnf(i_0_15680, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X4),esk3_0,esk3_0)|esk3_0=X4|k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)!=k5_struct_0(esk1_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15681, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X4),X4,np__1)|k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15685, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X4),X4,X5)|np__1=X5|k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)!=k4_struct_0(esk1_0)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,esk2_0)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15684, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X4),X5,esk3_0)|k3_matrix14(esk1_0,X4,X5)!=k1_funct_1(esk4_0,np__1)|k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(esk3_0,X4)|~r1_xxreal_0(X5,X4)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16009, negated_conjecture, (esk5_0=np__1|esk6_0=np__1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_15667, plain, (v1_funct_1(esk10_1(k4_finseq_2(X1,X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1146, plain, (m1_subset_1(k12_finseq_1(X1,X2),k1_numbers)|v1_xboole_0(X1)|k13_finseq_1(X1)!=esk11_0|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_15888, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk5_0))).#trainneg
cnf(i_0_15089, negated_conjecture, (r1_xxreal_0(esk5_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk5_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1147, plain, (m1_subset_1(k3_matrix14(X1,X2,X3),k1_numbers)|v6_struct_0(X1)|v2_struct_0(X1)|k13_finseq_1(u1_struct_0(X1))!=esk11_0|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1149, plain, (v1_xboole_0(u1_struct_0(X1))|m2_finseq_1(k5_struct_0(X1),X2)|u1_struct_0(X1)!=k13_finseq_1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_15912, plain, (epred1_5(esk1_0,esk3_0,esk5_0,X1,k12_matrix_1(esk1_0,esk2_0))|esk5_0=np__1|esk6_0=np__1)).#trainneg
cnf(i_0_16004, negated_conjecture, (esk5_0=np__1|esk6_0=np__1|esk3_0!=esk5_0|esk3_0!=esk6_0|~epred1_5(esk1_0,esk6_0,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_15668, plain, (v1_finseq_1(esk10_1(k4_finseq_2(X1,X2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15436, negated_conjecture, (v3_xxreal_0(esk6_0)|v2_xxreal_0(esk4_0)|r1_xxreal_0(esk4_0,esk6_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_10605, negated_conjecture, (k9_finseq_1(k5_struct_0(esk1_0))=k1_funct_1(esk4_0,esk3_0)|k9_xtuple_0(k1_funct_1(esk4_0,esk3_0))!=k2_finseq_1(np__1))).#trainneg
cnf(i_0_1151, plain, (m1_subset_1(esk7_2(X1,k5_numbers),k1_numbers)|k13_finseq_1(X1)=k5_numbers|m2_finseq_1(esk7_2(X1,k5_numbers),X1))).#trainneg
cnf(i_0_15890, negated_conjecture, (esk2_0=np__2|~r1_xxreal_0(esk2_0,np__2)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_15693, negated_conjecture, (m1_subset_1(k5_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15090, negated_conjecture, (r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk3_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15091, negated_conjecture, (r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1152, plain, (m1_subset_1(k12_finseq_1(X1,X2),k1_numbers)|v1_xboole_0(X1)|k13_finseq_1(X1)!=k5_numbers|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1153, plain, (m1_subset_1(k3_matrix14(X1,X2,X3),k1_numbers)|v6_struct_0(X1)|v2_struct_0(X1)|k13_finseq_1(u1_struct_0(X1))!=k5_numbers|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16122, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),esk3_0,esk3_0)|esk3_0=X3|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k5_struct_0(esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15951, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,esk5_0),np__1)=k5_struct_0(esk1_0)|esk5_0=np__1|esk6_0=np__1)).#trainneg
cnf(i_0_8737, plain, (v1_xboole_0(esk10_1(k13_finseq_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15440, negated_conjecture, (v3_xxreal_0(esk5_0)|v2_xxreal_0(esk4_0)|r1_xxreal_0(esk4_0,esk5_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_13094, negated_conjecture, (k9_finseq_1(k1_funct_1(esk4_0,esk3_0))=k4_matrix14(esk1_0,esk2_0,np__1)|k9_xtuple_0(k4_matrix14(esk1_0,esk2_0,np__1))!=k2_finseq_1(np__1)|~v1_funct_1(k4_matrix14(esk1_0,esk2_0,np__1))|~v1_relat_1(k4_matrix14(esk1_0,esk2_0,np__1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_1155, plain, (r1_xxreal_0(X1,k1_finseq_1(X2))|r1_xxreal_0(k1_finseq_1(X2),X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_15092, plain, (r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),np__1)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15720, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1165, plain, (k1_funct_1(k10_funct_7(X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|~m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(k10_funct_7(X2,X3,X4)))|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_15965, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk3_0,esk3_0)=k4_struct_0(esk1_0)|esk3_0=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15950, negated_conjecture, (m2_finseq_1(k1_funct_1(esk4_0,esk5_0),u1_struct_0(esk1_0))|esk5_0=np__1|esk6_0=np__1)).#trainneg
cnf(i_0_16328, plain, (v7_ordinal1(esk10_1(k13_finseq_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15474, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(esk4_0)|r1_xxreal_0(esk4_0,esk3_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_16198, negated_conjecture, (k9_finseq_1(k5_struct_0(esk1_0))=k1_funct_1(esk4_0,esk5_0)|esk5_0=np__1|esk6_0=esk5_0|k9_xtuple_0(k1_funct_1(esk4_0,esk5_0))!=k2_finseq_1(np__1))).#trainneg
cnf(i_0_1178, plain, (r1_xxreal_0(k6_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k6_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16230, negated_conjecture, (v3_ordinal1(k5_struct_0(esk1_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16367, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(esk6_0,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1166, plain, (k1_funct_1(k2_finseq_7(X5,X2,X3,X4),X4)=k7_partfun1(X1,X2,X3)|v1_xboole_0(X1)|v1_xboole_0(X5)|~m2_finseq_1(k2_finseq_7(X5,X2,X3,X4),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X5,X2,X3,X4)))|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1)|~m1_finseq_1(X2,X5))).#trainneg
cnf(i_0_16231, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X4,X3)!=k3_matrix14(esk1_0,X1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_16232, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X1,X2),X3,X3)=k5_struct_0(esk1_0)|np__1=X3|X3=X2|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_16124, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X3),esk3_0,esk5_0)|esk5_0=np__1|esk6_0=esk5_0|esk3_0=X3|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k5_struct_0(esk1_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1809, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15478, negated_conjecture, (v3_xxreal_0(esk2_0)|v2_xxreal_0(esk4_0)|r1_xxreal_0(esk4_0,esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1179, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16369, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(esk5_0,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1170, plain, (k2_zfmisc_1(k9_xtuple_0(k4_matrix14(X1,X2,X3)),k2_finseq_1(X2))=k2_matrix_1(k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v1_matrix_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16115, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,X5)|k3_matrix14(esk1_0,X6,X5)!=k3_matrix14(esk1_0,X3,X4)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X5,X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16126, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,esk3_0)|k3_matrix14(esk1_0,X3,X4)!=k1_funct_1(esk4_0,np__1)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(esk3_0,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16175, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v1_xboole_0(u1_struct_0(X6))|k5_struct_0(X6)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k5_struct_0(X6),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|u1_struct_0(X6)!=k13_finseq_1(u1_struct_0(X2))|~l6_algstr_0(X6))).#trainneg
cnf(i_0_16125, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,esk5_0)|esk5_0=np__1|esk6_0=esk5_0|k3_matrix14(esk1_0,X3,X4)!=k1_funct_1(esk4_0,np__1)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(esk5_0,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16440, negated_conjecture, (esk5_0=np__1|esk3_0!=esk5_0|~epred1_5(esk1_0,np__1,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_14215, plain, (r1_xxreal_0(np__2,k2_xcmplx_0(np__1,np__2))|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_16441, negated_conjecture, (esk5_0=np__1|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk5_0,np__1)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,np__1,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_1180, plain, (r1_xxreal_0(k1_xreal_0(X1,X2),X3)|r1_xxreal_0(X3,k1_xreal_0(X1,X2))|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16544, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15081, plain, (v3_xxreal_0(k4_matrix14(X1,X2,X3))|v2_xxreal_0(X4)|r1_xxreal_0(X4,k4_matrix14(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v7_ordinal1(X4)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1182, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16578, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X2)=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16588, plain, (X1=np__1|epred1_5(esk1_0,X1,X1,X2,k4_matrix14(esk1_0,X3,X4))|X1=X4|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_16607, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk3_0,esk5_0)=k4_struct_0(esk1_0)|esk6_0=esk5_0|esk5_0=np__1|esk3_0=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15347, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1186, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|~v7_ordinal1(X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15739, negated_conjecture, (r1_xxreal_0(esk5_0,esk6_0)|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk5_0)|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16247, negated_conjecture, (r1_xxreal_0(esk5_0,esk3_0)|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk5_0)|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1187, plain, (k2_xcmplx_0(X1,X2)=X3|~v1_xxreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1188, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15284, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X3,np__2),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),np__1,k4_matrix14(X1,X2,X3)))),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X3,k4_matrix14(X1,X2,X3)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1,X3)|r1_xxreal_0(X3,np__1)|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16741, negated_conjecture, (esk5_0=np__1|esk3_0=esk5_0|~epred1_5(esk1_0,np__1,esk5_0,esk3_0,esk4_0))).#trainneg
cnf(i_0_16787, plain, (X1=np__1|epred1_5(esk1_0,X1,X1,X2,k12_matrix_1(esk1_0,X3))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1189, plain, (r1_xxreal_0(X1,k4_xcmplx_0(k2_xcmplx_0(X2,X3)))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1191, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v7_ordinal1(X4)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1194, plain, (r1_xxreal_0(X1,k6_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_16786, plain, (epred1_5(esk1_0,X1,X1,X2,k12_matrix_1(esk1_0,X3))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_16992, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_16839, plain, (k2_xcmplx_0(np__1,X1)=np__1|~v1_xxreal_0(k2_xcmplx_0(np__1,X1))|~r1_xxreal_0(k2_xcmplx_0(np__1,X1),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15275, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X4,np__2),k1_rfinseq(np__1,k4_matrix14(X1,X2,X3)))),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X4,k4_matrix14(X1,X2,X3)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1,X4)|r1_xxreal_0(X4,np__1)|v6_struct_0(X1)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),X4),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3),k13_finseq_1(u1_struct_0(X1)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_15643, negated_conjecture, (v3_xxreal_0(np__1)|v2_xxreal_0(esk4_0)|r1_xxreal_0(esk4_0,np__1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_10388, negated_conjecture, (k9_finseq_1(k1_funct_1(esk4_0,esk3_0))=k12_matrix_1(esk1_0,esk2_0)|k9_xtuple_0(k12_matrix_1(esk1_0,esk2_0))!=k2_finseq_1(np__1)|~v1_funct_1(k12_matrix_1(esk1_0,esk2_0))|~v1_relat_1(k12_matrix_1(esk1_0,esk2_0)))).#trainneg
cnf(i_0_1216, plain, (k1_finseq_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_16761, negated_conjecture, (v3_xxreal_0(k4_matrix14(X1,esk2_0,X2))|v2_xxreal_0(X3)|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(X3,k4_matrix14(X1,esk2_0,X2))|v6_struct_0(X1)|~v1_xboole_0(esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16015, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,np__1)=k4_struct_0(esk1_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7003, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X3,X4)|np__1=X3|X3=X2|X3=esk3_0|X1!=k1_funct_1(k1_funct_1(esk4_0,X3),X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X4,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14186, negated_conjecture, (k3_finseq_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,X2))=X3|v1_xboole_0(X1)|r1_xxreal_0(X2,np__1)|~m1_finseq_1(k2_rfinseq(k13_finseq_1(X1),X2,k12_matrix_1(esk1_0,esk2_0)),k13_finseq_1(X1))|~m2_finseq_1(k8_finseq_1(k13_finseq_1(X1),k8_finseq_1(k13_finseq_1(X1),k12_finseq_1(k13_finseq_1(X1),k3_matrix14(esk1_0,esk2_0,X2)),k17_finseq_1(k13_finseq_1(X1),k7_nat_d(X2,np__2),k2_rfinseq(k13_finseq_1(X1),np__1,k12_matrix_1(esk1_0,esk2_0)))),k12_finseq_1(k13_finseq_1(X1),k1_funct_1(esk4_0,esk3_0))),k13_finseq_1(X1))|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(X1))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,esk2_0)|~m1_matrix_1(k2_finseq_7(k13_finseq_1(X1),k12_matrix_1(esk1_0,esk2_0),np__1,X2),X1,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15686, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X5),X6,X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix14(X1,X8,X7)!=k3_matrix14(X1,X5,X6)|k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,X5)|~r1_xxreal_0(X6,X5)|~r1_xxreal_0(X7,X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X8)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11878, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k12_finseq_1(X1,k7_partfun1(X1,k12_matrix_1(X2,X3),np__1))),k1_rfinseq(X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|v2_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_12300, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk6_0,np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_13559, plain, (k1_finseq_1(X2)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_16248, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk6_0)|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15748, plain, (v3_ordinal1(k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_12034, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v1_xboole_0(u1_struct_0(X6))|k4_struct_0(X6)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k4_struct_0(X6),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|u1_struct_0(X6)!=k13_finseq_1(u1_struct_0(X2))|~l2_struct_0(X6))).#trainneg
cnf(i_0_13239, plain, (k1_funct_1(k4_matrix14(X1,X2,X3),np__1)=k1_funct_1(k4_matrix14(X1,X2,X4),X3)|np__1=X3|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5782, plain, (k3_finseq_1(k7_partfun1(k10_xtuple_0(k12_matrix_1(X1,X2)),X3,X4))=X2|v1_xboole_0(k10_xtuple_0(k12_matrix_1(X1,X2)))|v1_xboole_0(u1_struct_0(X1))|~m2_finseq_1(k7_partfun1(k10_xtuple_0(k12_matrix_1(X1,X2)),X3,X4),u1_struct_0(X1))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~m1_finseq_1(X3,k10_xtuple_0(k12_matrix_1(X1,X2))))).#trainneg
cnf(i_0_5823, plain, (esk9_5(X1,k12_matrix_1(X2,X3),X4,X5,X6)=k1_funct_1(k12_matrix_1(X2,X3),X4)|v1_xboole_0(u1_struct_0(X2))|X6!=k3_matrix_1(X1,k12_matrix_1(X2,X3),X4,X5)|~v1_matrix_1(k12_matrix_1(X2,X3))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(X1))|~r2_hidden(k4_tarski(X4,X5),k2_zfmisc_1(k2_finseq_1(X3),k2_finseq_1(X3)))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X6,X1)|~v7_ordinal1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_12364, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|r1_xxreal_0(esk3_0,np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_17573, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)!=k4_struct_0(esk1_0)|esk3_0!=esk6_0|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0))).#trainneg
cnf(i_0_10487, plain, (v3_ordinal1(X2)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15741, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(esk6_0,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v7_ordinal1(k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1222, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k6_xcmplx_0(X2,X3))|r1_xxreal_0(k6_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1223, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)!=k4_struct_0(esk1_0)|esk6_0!=np__1|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0))).#trainneg
cnf(i_0_1, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)!=k5_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)!=k4_struct_0(esk1_0)|~epred1_5(esk1_0,esk6_0,np__1,esk3_0,esk4_0))).#trainneg
cnf(i_0_17635, negated_conjecture, (r1_xxreal_0(esk3_0,np__2)|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1224, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k1_xreal_0(X2,X3))|r1_xxreal_0(k1_xreal_0(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_16252, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(esk3_0,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v7_ordinal1(k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_16360, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(np__1,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v7_ordinal1(k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1225, plain, (v3_xxreal_0(X1)|v2_xxreal_0(np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1226, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17630, negated_conjecture, (esk9_5(X1,k12_matrix_1(X2,esk2_0),X3,X4,X5)=k1_funct_1(k12_matrix_1(X2,esk2_0),X3)|v1_xboole_0(u1_struct_0(X2))|X5!=k3_matrix_1(X1,k12_matrix_1(X2,esk2_0),X3,X4)|~m2_finseq_1(k12_matrix_1(X2,esk2_0),k13_finseq_1(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(esk4_0))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_17574, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1229, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_16361, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(np__1,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v7_ordinal1(k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_16362, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(np__1,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v7_ordinal1(k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1231, plain, (k2_finseq_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1232, plain, (k2_finseq_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17529, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k4_struct_0(esk1_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X4,X3)!=k3_matrix14(esk1_0,X1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_7183, plain, (r1_xxreal_0(k3_finseq_1(X1),k2_nat_1(k3_finseq_1(X1),np__1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_17885, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X2)=k4_struct_0(esk1_0)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1233, plain, (k2_finseq_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17641, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_17651, plain, (v3_ordinal1(X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_17678, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1234, plain, (k2_finseq_1(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1235, plain, (k2_finseq_1(X1)=k10_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_17888, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X2,X1)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_17926, plain, (r1_xxreal_0(X2,k2_nat_1(X2,np__1))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_17927, plain, (r1_xxreal_0(X2,k2_nat_1(X2,np__1))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_17761, negated_conjecture, (esk9_5(X1,k12_matrix_1(X2,esk2_0),X3,X4,X5)=k1_funct_1(k12_matrix_1(X2,esk2_0),X3)|v1_xboole_0(u1_struct_0(X2))|X5!=k3_matrix_1(X1,k12_matrix_1(X2,esk2_0),X3,X4)|~m2_finseq_1(k12_matrix_1(X2,esk2_0),k13_finseq_1(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~l6_algstr_0(X2)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(X3,esk2_0))).#trainneg
cnf(i_0_16049, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k4_struct_0(esk1_0)|np__1=X2|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_2024, negated_conjecture, (r1_xxreal_0(esk3_0,k1_matrix_1(X1))|r1_xxreal_0(k1_matrix_1(X1),esk3_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_1236, plain, (X1=k2_finseq_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_17739, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_17742, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk6_0)|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1252, negated_conjecture, (r1_xxreal_0(esk6_0,esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),esk6_0))).#trainneg
cnf(i_0_1257, negated_conjecture, (r1_xxreal_0(esk6_0,k6_xcmplx_0(X1,X2))|r1_xxreal_0(k6_xcmplx_0(X1,X2),esk6_0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2838, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k1_funct_1(k12_matrix_1(X2,X3),np__1)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6263, plain, (k3_finseq_1(k7_finseq_1(X1,X2))=k3_finseq_1(X2)|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(k3_finseq_1(k7_finseq_1(X1,X2)))|~r1_xxreal_0(k3_finseq_1(k7_finseq_1(X1,X2)),k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X1),k5_numbers))).#trainneg
cnf(i_0_18151, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,np__1)=k4_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_4237, negated_conjecture, (k2_xcmplx_0(np__1,esk2_0)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(np__1,esk2_0),np__1))).#trainneg
cnf(i_0_1258, negated_conjecture, (r1_xxreal_0(esk6_0,k1_finseq_1(X1))|r1_xxreal_0(k1_finseq_1(X1),esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_18026, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_18188, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1259, negated_conjecture, (r1_xxreal_0(esk6_0,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1261, negated_conjecture, (r1_xxreal_0(esk6_0,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),esk6_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18150, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k4_struct_0(esk1_0)|np__1=X1|X1=esk3_0|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18222, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k1_funct_1(k12_matrix_1(X2,X3),np__1)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_18154, plain, (r1_xxreal_0(X1,k2_nat_1(X1,np__1))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_5113, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,k12_matrix_1(X1,X2)))=k2_nat_1(X2,esk2_0)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_10583, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(X2,X3,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(X2,X3)))),k9_finseq_1(k3_matrix14(X2,X3,np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_subset_1(k3_matrix14(X2,X3,np__1),X1))).#trainneg
cnf(i_0_17776, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17715, negated_conjecture, (r1_xxreal_0(esk3_0,np__2)|r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_1263, negated_conjecture, (r1_xxreal_0(esk6_0,k9_xtuple_0(X1))|r1_xxreal_0(k9_xtuple_0(X1),esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15337, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk6_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_18191, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk6_0)|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_17789, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk3_0)|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1264, negated_conjecture, (r1_xxreal_0(esk6_0,k10_xtuple_0(X1))|r1_xxreal_0(k10_xtuple_0(X1),esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12245, negated_conjecture, (k1_funct_1(k4_matrix14(X1,X2,X3),np__1)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X4),X5,esk3_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix14(esk1_0,X4,X5)!=k1_funct_1(esk4_0,np__1)|k1_funct_1(k4_matrix14(X1,X2,X3),np__1)!=k5_struct_0(esk1_0)|~m1_finseq_1(k12_matrix_1(X1,X2),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(esk3_0,X4)|~r1_xxreal_0(X5,X4)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7291, plain, (k7_partfun1(u1_struct_0(X1),X2,X3)=k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X4,X5),X6,X7)|np__1=X6|X6=X7|X6=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k7_partfun1(u1_struct_0(X1),X2,X3)!=k4_struct_0(X1)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,X4)|~r1_xxreal_0(X7,X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_9243, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_18288, plain, (esk3_0=X1|np__1=X1|epred1_5(esk1_0,np__1,esk3_0,X1,k12_matrix_1(esk1_0,esk2_0))|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_18286, plain, (epred1_5(esk1_0,np__1,esk3_0,X1,k12_matrix_1(esk1_0,esk2_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_2054, negated_conjecture, (np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1294, negated_conjecture, (r1_xxreal_0(esk3_0,esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),esk3_0))).#trainneg
cnf(i_0_15338, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_15341, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk3_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_18519, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk3_0)|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_17792, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1299, negated_conjecture, (r1_xxreal_0(esk3_0,k6_xcmplx_0(X1,X2))|r1_xxreal_0(k6_xcmplx_0(X1,X2),esk3_0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1300, negated_conjecture, (r1_xxreal_0(esk3_0,k1_finseq_1(X1))|r1_xxreal_0(k1_finseq_1(X1),esk3_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_18534, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=k1_funct_1(esk4_0,np__1)|k1_funct_1(esk4_0,np__1)!=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18396, negated_conjecture, (r1_xxreal_0(X1,k2_nat_1(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18394, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k3_matrix14(X2,X3,np__1)|X1!=k1_funct_1(X6,X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_18601, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_18480, negated_conjecture, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(esk4_0,np__1)),k17_finseq_1(X1,k7_nat_d(esk3_0,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(esk1_0,esk2_0)))),k9_finseq_1(k1_funct_1(esk4_0,esk3_0))),k2_rfinseq(X1,esk3_0,k12_matrix_1(esk1_0,esk2_0)))=k2_finseq_7(X1,k12_matrix_1(esk1_0,esk2_0),np__1,esk3_0)|v1_xboole_0(X1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~m1_subset_1(k1_funct_1(esk4_0,esk3_0),X1))).#trainneg
cnf(i_0_2426, negated_conjecture, (k2_xcmplx_0(X1,X2)=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1301, negated_conjecture, (r1_xxreal_0(esk3_0,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),esk3_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8528, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_15342, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),esk3_0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_15345, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),esk6_0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_18311, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1303, negated_conjecture, (r1_xxreal_0(esk3_0,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),esk3_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1305, negated_conjecture, (r1_xxreal_0(esk3_0,k9_xtuple_0(X1))|r1_xxreal_0(k9_xtuple_0(X1),esk3_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_18710, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=k1_funct_1(esk4_0,np__1)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|k1_funct_1(esk4_0,np__1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18741, negated_conjecture, (k8_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k8_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k8_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k9_finseq_1(k1_funct_1(k12_matrix_1(esk1_0,esk2_0),X1)),k17_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k7_nat_d(X1,np__2),k1_rfinseq(np__1,k12_matrix_1(esk1_0,esk2_0)))),k9_finseq_1(k1_funct_1(esk4_0,esk3_0))),k2_rfinseq(k13_finseq_1(u1_struct_0(esk1_0)),X1,k12_matrix_1(esk1_0,esk2_0)))=k2_finseq_7(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~v7_ordinal1(X1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_18719, negated_conjecture, (k2_nat_1(X1,np__1)=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(k2_nat_1(X1,np__1))|~r1_xxreal_0(k2_nat_1(X1,np__1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18850, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_13805, negated_conjecture, (r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1306, negated_conjecture, (r1_xxreal_0(esk3_0,k10_xtuple_0(X1))|r1_xxreal_0(k10_xtuple_0(X1),esk3_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_18020, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk6_0,np__2)|~v1_xxreal_0(np__2)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__2))).#trainneg
cnf(i_0_18666, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(k4_matrix14(X1,X2,X3),esk2_0)|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1316, plain, (r1_xxreal_0(np__1,esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),np__1))).#trainneg
cnf(i_0_1321, plain, (r1_xxreal_0(np__1,k6_xcmplx_0(X1,X2))|r1_xxreal_0(k6_xcmplx_0(X1,X2),np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18744, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k1_funct_1(k4_matrix14(X1,X2,X3),X4)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X4,np__2),k1_rfinseq(np__1,k4_matrix14(X1,X2,X3)))),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),X3))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X4,k4_matrix14(X1,X2,X3)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k4_matrix14(X1,X2,X3),np__1,X4)|r1_xxreal_0(X4,np__1)|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X4,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~v7_ordinal1(X4)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18728, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18720, negated_conjecture, (r1_xxreal_0(X1,k2_nat_1(X2,np__1))|~v1_xxreal_0(k2_nat_1(X2,np__1))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18729, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X2)|X3!=k1_funct_1(esk4_0,esk3_0)|X1!=k1_funct_1(X3,X2)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18711, negated_conjecture, (k2_nat_1(esk6_0,np__1)=esk6_0|~r1_xxreal_0(k2_nat_1(esk6_0,np__1),np__1)|~v7_ordinal1(k2_nat_1(esk6_0,np__1)))).#trainneg
cnf(i_0_18574, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k9_finseq_1(k1_funct_1(X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2)))).#trainneg
cnf(i_0_18889, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(np__1,np__2),esk3_0)|r1_xxreal_0(esk3_0,k2_xcmplx_0(np__1,np__2))|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_18457, negated_conjecture, (k2_nat_1(X2,esk2_0)=k2_xcmplx_0(esk2_0,X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_17636, negated_conjecture, (r1_xxreal_0(esk3_0,np__2)|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1322, plain, (r1_xxreal_0(np__1,k1_finseq_1(X1))|r1_xxreal_0(k1_finseq_1(X1),np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_18998, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk6_0,np__2)|~v1_xxreal_0(np__2)|~r1_xxreal_0(esk2_0,np__2))).#trainneg
cnf(i_0_18316, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1323, plain, (r1_xxreal_0(np__1,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1325, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19186, negated_conjecture, (k2_nat_1(X1,esk2_0)=k2_xcmplx_0(esk2_0,X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_19105, negated_conjecture, (r1_xxreal_0(X1,k2_nat_1(X2,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(k2_nat_1(X2,np__1)))).#trainneg
cnf(i_0_19106, negated_conjecture, (r1_xxreal_0(X1,k2_nat_1(X2,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v1_xreal_0(k2_nat_1(X2,np__1)))).#trainneg
cnf(i_0_19121, negated_conjecture, (k2_nat_1(np__1,esk6_0)=esk6_0|~r1_xxreal_0(k2_nat_1(np__1,esk6_0),np__1)|~v7_ordinal1(k2_nat_1(np__1,esk6_0))|~m1_subset_1(esk6_0,k5_numbers))).#trainneg
cnf(i_0_19119, negated_conjecture, (k2_nat_1(np__1,esk6_0)=esk6_0|~r1_xxreal_0(k2_nat_1(np__1,esk6_0),np__1)|~v7_ordinal1(k2_nat_1(np__1,esk6_0)))).#trainneg
cnf(i_0_19120, negated_conjecture, (k2_xcmplx_0(esk6_0,np__1)=esk6_0|~r1_xxreal_0(k2_xcmplx_0(esk6_0,np__1),np__1)|~v7_ordinal1(k2_xcmplx_0(esk6_0,np__1))|~m1_subset_1(esk6_0,k5_numbers))).#trainneg
cnf(i_0_19114, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k5_struct_0(esk1_0)|X2!=k1_funct_1(esk4_0,esk3_0)|k1_funct_1(X2,X1)!=k5_struct_0(esk1_0)|~m2_finseq_1(X2,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18548, plain, (k1_funct_1(X2,X3)=k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X4,X5),X6,X7)|np__1=X6|X6=X5|X6=X7|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_funct_1(X2,X3)!=k4_struct_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1))|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,X4)|~r1_xxreal_0(X7,X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m2_finseq_1(X2,u1_struct_0(X1))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_19201, negated_conjecture, (r1_xxreal_0(esk3_0,np__2)|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18974, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|esk3_0=esk6_0|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_17477, negated_conjecture, (v2_xxreal_0(esk4_0)|r1_xxreal_0(esk4_0,np__1)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1327, plain, (r1_xxreal_0(np__1,k9_xtuple_0(X1))|r1_xxreal_0(k9_xtuple_0(X1),np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_18968, negated_conjecture, (np__2=esk3_0|esk3_0=esk6_0|~v1_xxreal_0(np__2)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,np__2)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_18512, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk6_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18657, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk3_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1328, plain, (r1_xxreal_0(np__1,k10_xtuple_0(X1))|r1_xxreal_0(k10_xtuple_0(X1),np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1330, plain, (r1_xxreal_0(X1,k1_card_1(X2))|r1_xxreal_0(k1_card_1(X2),X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_18923, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=k1_funct_1(esk4_0,np__1)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|k1_funct_1(esk4_0,np__1)!=k5_struct_0(esk1_0)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18559, negated_conjecture, (k8_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k8_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k8_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k9_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),X1)),k17_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)),k7_nat_d(X1,np__2),k1_rfinseq(np__1,k12_matrix_1(esk1_0,esk2_0)))),k9_finseq_1(k1_funct_1(esk4_0,esk3_0))),k2_rfinseq(k13_finseq_1(u1_struct_0(esk1_0)),X1,k12_matrix_1(esk1_0,esk2_0)))=k2_finseq_7(k13_finseq_1(u1_struct_0(esk1_0)),k12_matrix_1(esk1_0,esk2_0),np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~v7_ordinal1(X1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_18739, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18801, negated_conjecture, (k2_xcmplx_0(np__1,esk2_0)=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(k2_xcmplx_0(np__1,esk2_0),np__1))).#trainneg
cnf(i_0_18725, negated_conjecture, (r1_xxreal_0(X1,k2_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19332, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k5_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19292, negated_conjecture, (k2_nat_1(X1,esk2_0)=k2_xcmplx_0(esk2_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19333, plain, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(X3),k4_matrix14(X3,X4,X5),X6,X7)|np__1=X6|X6=X7|X6=X5|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|k1_funct_1(X1,X2)!=k4_struct_0(X3)|~r1_xxreal_0(X2,k3_finseq_1(X1))|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X6,X4)|~r1_xxreal_0(X7,X4)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X2)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3)|~m2_finseq_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_19325, negated_conjecture, (k2_xcmplx_0(np__1,esk6_0)=esk6_0|~r1_xxreal_0(k2_xcmplx_0(np__1,esk6_0),np__1)|~v7_ordinal1(k2_xcmplx_0(np__1,esk6_0)))).#trainneg
cnf(i_0_19346, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1331, negated_conjecture, (r1_xxreal_0(esk6_0,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),esk6_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_19219, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk6_0,np__2)|~r1_xxreal_0(esk2_0,np__2))).#trainneg
cnf(i_0_18862, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_19000, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1333, negated_conjecture, (r1_xxreal_0(esk3_0,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),esk3_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1334, plain, (r1_xxreal_0(np__1,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),np__1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_19400, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,esk3_0)=k1_funct_1(esk4_0,np__1)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|k1_funct_1(esk4_0,np__1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k13_finseq_1(k13_finseq_1(u1_struct_0(esk1_0)))!=k13_finseq_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_19447, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k9_finseq_1(k1_funct_1(X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2)))).#trainneg
cnf(i_0_18475, negated_conjecture, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k3_matrix14(esk1_0,esk2_0,X2)),k17_finseq_1(X1,k7_nat_d(X2,np__2),k2_rfinseq(X1,np__1,k12_matrix_1(esk1_0,esk2_0)))),k9_finseq_1(k1_funct_1(esk4_0,esk3_0))),k2_rfinseq(X1,X2,k12_matrix_1(esk1_0,esk2_0)))=k2_finseq_7(X1,k12_matrix_1(esk1_0,esk2_0),np__1,X2)|v1_xboole_0(X1)|r1_xxreal_0(X2,np__1)|~m2_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(esk4_0,esk3_0),X1))).#trainneg
cnf(i_0_19453, negated_conjecture, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k1_funct_1(esk4_0,np__1)),k17_finseq_1(X1,k7_nat_d(esk3_0,np__2),k1_rfinseq(np__1,k12_matrix_1(esk1_0,esk2_0)))),k9_finseq_1(k7_partfun1(X1,k12_matrix_1(esk1_0,esk2_0),np__1))),k1_rfinseq(esk3_0,k12_matrix_1(esk1_0,esk2_0)))=k2_finseq_7(X1,k12_matrix_1(esk1_0,esk2_0),np__1,esk3_0)|v1_xboole_0(X1)|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),X1)|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_18813, negated_conjecture, (k2_xcmplx_0(np__1,np__2)=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(k2_xcmplx_0(np__1,np__2),np__1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_18393, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),np__1),X5)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k1_funct_1(k12_matrix_1(X2,X3),np__1),u1_struct_0(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3500, plain, (r2_hidden(k16_finseq_1(X1,X2),X3)|X3!=k13_finseq_1(X4)|~v7_ordinal1(X1)|~m2_finseq_1(X2,X4))).#trainneg
cnf(i_0_19347, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19570, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk6_0,np__2)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_19223, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_19369, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk6_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk6_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1336, plain, (v7_ordinal1(k3_finseq_1(X1))|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1345, negated_conjecture, (r1_xxreal_0(esk2_0,esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),esk2_0))).#trainneg
cnf(i_0_19505, negated_conjecture, (k2_nat_1(esk2_0,X1)=k2_xcmplx_0(esk2_0,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19752, plain, (r2_hidden(k16_finseq_1(X1,k17_finseq_1(X2,X3,X4)),X5)|X5!=k13_finseq_1(X2)|~v7_ordinal1(X1)|~m1_finseq_1(X4,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_19753, plain, (r2_hidden(k16_finseq_1(X1,k2_rfinseq(X2,X3,X4)),X5)|X5!=k13_finseq_1(X2)|~v7_ordinal1(X1)|~m1_finseq_1(X4,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_19766, plain, (r2_hidden(k16_finseq_1(X1,k3_matrix14(X2,X3,X4)),X5)|v6_struct_0(X2)|v2_struct_0(X2)|X5!=k13_finseq_1(u1_struct_0(X2))|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_19560, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),esk6_0)|r1_xxreal_0(esk6_0,k3_finseq_1(X1))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19772, plain, (r2_hidden(k16_finseq_1(X1,k4_matrix14(X2,X3,X4)),X5)|v1_xboole_0(u1_struct_0(X2))|v1_xboole_0(X6)|v6_struct_0(X2)|X5!=k13_finseq_1(X6)|~v7_ordinal1(X1)|~m2_finseq_1(k12_matrix_1(X2,X3),X6)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_19757, plain, (r2_hidden(k16_finseq_1(X1,k9_finseq_1(X2)),X3)|v1_xboole_0(X4)|X3!=k13_finseq_1(X4)|~v7_ordinal1(X1)|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_19562, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),esk6_0)|r1_xxreal_0(esk6_0,k3_finseq_1(X1))|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19590, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),esk3_0)|r1_xxreal_0(esk3_0,k3_finseq_1(X1))|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1350, negated_conjecture, (r1_xxreal_0(esk2_0,k6_xcmplx_0(X1,X2))|r1_xxreal_0(k6_xcmplx_0(X1,X2),esk2_0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19571, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk6_0,np__2)|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_19348, negated_conjecture, (r1_xxreal_0(esk4_0,np__1)|~v1_xxreal_0(esk4_0)|~v1_xboole_0(esk4_0)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_15346, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk6_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_19938, negated_conjecture, (r1_xxreal_0(esk6_0,X2)|r1_xxreal_0(X2,esk6_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_20057, negated_conjecture, (r1_xxreal_0(esk3_0,X2)|r1_xxreal_0(X2,esk3_0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1351, negated_conjecture, (r1_xxreal_0(esk2_0,k1_finseq_1(X1))|r1_xxreal_0(k1_finseq_1(X1),esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1352, negated_conjecture, (r1_xxreal_0(esk2_0,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_19504, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=k2_xcmplx_0(X1,esk2_0)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19508, negated_conjecture, (r1_xxreal_0(k3_finseq_1(k12_matrix_1(X1,X2)),k2_xcmplx_0(esk2_0,X2))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_19592, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),esk3_0)|r1_xxreal_0(esk3_0,k3_finseq_1(X1))|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19749, plain, (r2_hidden(k16_finseq_1(X1,k1_rfinseq(X2,X3)),X4)|X4!=k13_finseq_1(X5)|~v7_ordinal1(X1)|~m1_finseq_1(X3,X5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7009, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X2,X2)=esk12_1(esk1_0)|np__1=X2|X2=esk3_0|X2=X1|esk12_1(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_6793, negated_conjecture, (r1_xxreal_0(k3_finseq_1(k4_matrix14(X1,X2,X3)),k2_nat_1(X2,esk2_0))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(k3_finseq_1(k4_matrix14(X1,X2,X3)))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9589, plain, (k8_finseq_1(X1,k7_finseq_1(k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),X1)|~m1_finseq_1(k9_finseq_1(k7_partfun1(X1,X2,np__1)),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_19788, negated_conjecture, (esk3_0=esk6_0|r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5775, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k7_partfun1(k1_zfmisc_1(X2),X3,X4))|~m1_finseq_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1353, negated_conjecture, (r1_xxreal_0(esk2_0,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_20082, negated_conjecture, (r1_xxreal_0(esk4_0,np__1)|~v1_xxreal_0(np__1)|~v1_xboole_0(esk4_0)|~v7_ordinal1(esk4_0))).#trainneg
cnf(i_0_20223, negated_conjecture, (v7_ordinal1(k2_xcmplx_0(np__0,esk2_0))|~v1_xcmplx_0(esk2_0))).#trainneg
cnf(i_0_20090, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(X2,esk6_0)|r1_xxreal_0(esk6_0,X2)|v6_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_20118, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(X2,esk3_0)|r1_xxreal_0(esk3_0,X2)|v6_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1355, negated_conjecture, (r1_xxreal_0(esk2_0,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1357, negated_conjecture, (r1_xxreal_0(esk2_0,k9_xtuple_0(X1))|r1_xxreal_0(k9_xtuple_0(X1),esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_20214, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20207, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20255, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(X2,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20263, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),esk3_0)|~v7_ordinal1(k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20190, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=X1|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X1))|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20264, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),esk6_0)|~v7_ordinal1(k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20242, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk6_0|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),np__1)|~r1_xxreal_0(esk6_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20483, negated_conjecture, (r1_xxreal_0(esk3_0,X1)|r1_xxreal_0(X1,esk3_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_20482, negated_conjecture, (r1_xxreal_0(esk6_0,X1)|r1_xxreal_0(X1,esk6_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_20463, negated_conjecture, (r1_xxreal_0(esk4_0,np__1)|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_1358, negated_conjecture, (r1_xxreal_0(esk2_0,k10_xtuple_0(X1))|r1_xxreal_0(k10_xtuple_0(X1),esk2_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15348, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_15352, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk3_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_19372, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk3_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk3_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_19580, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v1_xreal_0(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1368, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X3,X4)),k17_finseq_1(X1,k7_nat_d(X5,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X2,X3,X4,X5)))),k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X2,X3,X4,X5),np__1))),k2_rfinseq(X1,X5,k2_finseq_7(X2,X3,X4,X5)))=k2_finseq_7(X1,k2_finseq_7(X2,X3,X4,X5),np__1,X5)|v1_xboole_0(X1)|r1_xxreal_0(X5,np__1)|v1_xboole_0(X2)|~m2_finseq_1(k2_finseq_7(X2,X3,X4,X5),X1)|~r1_xxreal_0(X5,k3_finseq_1(k2_finseq_7(X2,X3,X4,X5)))|~v7_ordinal1(X5)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X3,X1)|~r1_xxreal_0(X5,k3_finseq_1(X3))|~r1_xxreal_0(X4,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_20564, negated_conjecture, (r1_xxreal_0(X1,k2_nat_1(esk2_0,X3))|v2_struct_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X2,X3)))|~v1_finseq_1(k12_matrix_1(X2,X3))|~v1_funct_1(k12_matrix_1(X2,X3))|~v1_relat_1(k12_matrix_1(X2,X3))|~v7_ordinal1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_20577, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=np__1|~v1_xxreal_0(np__1)|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20593, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(X2,esk2_0))|~v1_xxreal_0(k2_xcmplx_0(X2,esk2_0))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20554, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X4))|~v1_xxreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X4),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_20620, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk2_0|r1_xxreal_0(esk3_0,k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20673, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk2_0|r1_xxreal_0(esk6_0,k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19365, negated_conjecture, (esk3_0=esk6_0|np__2=esk3_0|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,np__2)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_15356, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_15358, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_18861, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1369, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X2,X3,X4,np__1),X5)),k17_finseq_1(X1,k7_nat_d(X5,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X2,X3,X4,np__1)))),k12_finseq_1(X1,k7_partfun1(X1,X3,X4))),k2_rfinseq(X1,X5,k2_finseq_7(X2,X3,X4,np__1)))=k2_finseq_7(X1,k2_finseq_7(X2,X3,X4,np__1),np__1,X5)|v1_xboole_0(X1)|r1_xxreal_0(X5,np__1)|v1_xboole_0(X2)|~m2_finseq_1(k2_finseq_7(X2,X3,X4,np__1),X1)|~r1_xxreal_0(X5,k3_finseq_1(k2_finseq_7(X2,X3,X4,np__1)))|~v7_ordinal1(X5)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X3,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X3))|~r1_xxreal_0(X4,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_20553, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20575, negated_conjecture, (k2_nat_1(esk2_0,X2)=np__1|v2_struct_0(X1)|~v1_xxreal_0(np__1)|~r1_xxreal_0(k2_nat_1(esk2_0,X2),np__1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_20552, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X1))|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20629, negated_conjecture, (k3_finseq_1(k12_matrix_1(X1,X2))=k2_nat_1(esk2_0,X2)|v2_struct_0(X1)|~v1_xxreal_0(k2_nat_1(esk2_0,X2))|~r1_xxreal_0(k2_nat_1(esk2_0,X2),k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_20684, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=esk6_0|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),np__1)|~r1_xxreal_0(esk6_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3244, plain, (np__1=X1|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_20907, negated_conjecture, (np__2=esk3_0|esk3_0=esk6_0|~r1_xxreal_0(esk3_0,np__2)|~r1_xxreal_0(esk6_0,np__1)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_20708, negated_conjecture, (esk4_0=np__1|~v1_xxreal_0(esk4_0)|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,esk4_0)|~v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_15360, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),np__1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_19222, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_19368, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk6_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk6_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1416, plain, (v3_xxreal_0(esk10_1(k5_numbers))|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk10_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1421, plain, (v3_xxreal_0(k6_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k6_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20540, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=esk3_0|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),np__1)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X1))|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20650, negated_conjecture, (k2_xcmplx_0(esk2_0,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,X2)|~v1_xxreal_0(k2_xcmplx_0(esk2_0,k2_xcmplx_0(X1,X2)))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(k2_xcmplx_0(esk2_0,k2_xcmplx_0(X1,X2)),X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20555, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X4))|~v1_xxreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X4),X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_20996, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X2)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X2)))).#trainneg
cnf(i_0_21052, negated_conjecture, (k3_finseq_1(k12_matrix_1(X1,X2))=k2_xcmplx_0(esk2_0,X2)|v2_struct_0(X1)|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X2))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X2),k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_20997, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X2)|~v1_xreal_0(k2_xcmplx_0(esk2_0,X2)))).#trainneg
cnf(i_0_21247, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(k2_xcmplx_0(esk2_0,X2),X3))|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21053, negated_conjecture, (k3_finseq_1(k12_matrix_1(X1,X2))=k2_nat_1(X2,esk2_0)|v2_struct_0(X1)|~v1_xxreal_0(k2_nat_1(X2,esk2_0))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~r1_xxreal_0(k2_nat_1(X2,esk2_0),k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_21071, plain, (np__1=X1|~v1_xxreal_0(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0))).#trainneg
cnf(i_0_1422, plain, (v3_xxreal_0(k1_finseq_1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k1_finseq_1(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_21077, negated_conjecture, (esk4_0=np__1|~v1_xxreal_0(np__1)|~v1_xboole_0(esk4_0)|~r1_xxreal_0(np__1,esk4_0)|~v7_ordinal1(esk4_0))).#trainneg
cnf(i_0_19371, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk3_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk3_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_19579, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk2_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(np__1,k4_matrix14(X1,X2,X3))|v6_struct_0(X1)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk2_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1423, plain, (v3_xxreal_0(k1_card_1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k1_card_1(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1424, plain, (v3_xxreal_0(k1_card_1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k1_card_1(X1))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_21182, negated_conjecture, (k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1))=k2_xcmplx_0(esk2_0,X1)|~v1_xxreal_0(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1)))|~v7_ordinal1(X1)|~r1_xxreal_0(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1)),esk3_0)|~v7_ordinal1(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1))))).#trainneg
cnf(i_0_20549, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(esk2_0,X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X3))|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X3),X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_21493, plain, (np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21096, plain, (v3_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(np__0,esk10_1(k5_numbers)))).#trainneg
cnf(i_0_2023, negated_conjecture, (r1_xxreal_0(k1_matrix_1(X1),np__1)|r1_xxreal_0(np__1,k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_1425, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21506, negated_conjecture, (esk4_0=np__1|~v1_xboole_0(esk4_0)|~r1_xxreal_0(np__1,esk4_0))).#trainneg
cnf(i_0_20224, negated_conjecture, (r1_xxreal_0(esk2_0,k2_xcmplx_0(np__0,esk2_0))|~v1_xcmplx_0(esk2_0))).#trainneg
cnf(i_0_17583, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk6_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(esk6_0,esk3_0)|v6_struct_0(X1)|~v1_xxreal_0(k4_matrix14(X1,X2,X3))|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15737, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(esk6_0,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(k4_matrix14(X2,X3,X4))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1426, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1427, plain, (v3_xxreal_0(np__2)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21391, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X3),X2)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X2))|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_21600, negated_conjecture, (~r1_xxreal_0(esk3_0,np__0))).#trainneg
cnf(i_0_21616, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=np__1|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),np__0)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X1))|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20631, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=X1|~v1_xxreal_0(k2_xcmplx_0(X1,esk2_0))|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20578, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=np__1|r1_xxreal_0(np__1,k2_xcmplx_0(esk2_0,X1))|~v1_xxreal_0(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20494, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),esk2_0)|r1_xxreal_0(esk2_0,k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21655, negated_conjecture, (v3_xxreal_0(np__0)|v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,np__0))).#trainneg
cnf(i_0_1428, plain, (v3_xxreal_0(k9_xtuple_0(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k9_xtuple_0(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_21076, negated_conjecture, (esk3_0=esk6_0|np__2=esk3_0|r1_xxreal_0(esk6_0,np__2)|~r1_xxreal_0(esk3_0,np__2)|~r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_21604, negated_conjecture, (esk2_0=np__1|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_21691, negated_conjecture, (k4_matrix14(X1,X2,X3)=esk6_0|v1_xboole_0(u1_struct_0(X1))|r1_xxreal_0(esk3_0,k4_matrix14(X1,X2,X3))|r1_xxreal_0(esk6_0,esk3_0)|v6_struct_0(X1)|~v1_xxreal_0(esk6_0)|~v1_xboole_0(X2)|~r1_xxreal_0(esk6_0,k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(k4_matrix14(X1,X2,X3)))).#trainneg
cnf(i_0_1429, plain, (v3_xxreal_0(k10_xtuple_0(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k10_xtuple_0(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_20618, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk2_0|~v7_ordinal1(k2_xcmplx_0(X1,esk2_0))|~v7_ordinal1(X1)|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),np__1))).#trainneg
cnf(i_0_21389, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X2))|~v1_xxreal_0(k2_nat_1(X1,np__1))|~v1_xxreal_0(X1)|~r1_xxreal_0(k2_nat_1(X1,np__1),X2)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21000, negated_conjecture, (k2_nat_1(esk2_0,X1)=np__1|v2_struct_0(X2)|~v1_xxreal_0(np__1)|~v1_finseq_1(k12_matrix_1(X2,X1))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~r1_xxreal_0(k2_nat_1(esk2_0,X1),np__1)|~v7_ordinal1(X1)|~l6_algstr_0(X2)|~m1_subset_1(k12_matrix_1(X2,X1),k5_numbers))).#trainneg
cnf(i_0_21183, negated_conjecture, (k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1))=k2_xcmplx_0(esk2_0,X1)|~v1_xxreal_0(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1)))|~v7_ordinal1(X1)|~r1_xxreal_0(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1)),esk6_0)|~v7_ordinal1(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1))))).#trainneg
cnf(i_0_21468, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X2)))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20230, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk2_0|~v1_xxreal_0(k2_xcmplx_0(X1,esk2_0))|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20280, negated_conjecture, (r1_xxreal_0(X2,k2_xcmplx_0(esk2_0,X2))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_21475, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(esk2_0,X2)))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(esk2_0,X2),k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_21656, negated_conjecture, (v3_xxreal_0(np__0)|v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,np__0))).#trainneg
cnf(i_0_1437, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X3,X5)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X2,X3,X4,X5)))),k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X2,X3,X4,X5),np__1))),k2_rfinseq(X1,X4,k2_finseq_7(X2,X3,X4,X5)))=k2_finseq_7(X1,k2_finseq_7(X2,X3,X4,X5),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v1_xboole_0(X2)|~m2_finseq_1(k2_finseq_7(X2,X3,X4,X5),X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X2,X3,X4,X5)))|~v7_ordinal1(X4)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X3,X1)|~r1_xxreal_0(X5,k3_finseq_1(X3))|~r1_xxreal_0(X4,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_22107, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,esk2_0)))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(k2_xcmplx_0(X3,esk2_0))|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_21953, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),np__1)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22064, negated_conjecture, (k2_xcmplx_0(esk2_0,X1)=esk2_0|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X1))|~r1_xxreal_0(k2_xcmplx_0(esk2_0,X1),esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21657, negated_conjecture, (v3_xxreal_0(np__0)|v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_16245, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(esk3_0,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(k4_matrix14(X2,X3,X4))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1438, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k2_finseq_7(X2,X3,np__1,X4),X5)),k17_finseq_1(X1,k7_nat_d(X5,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X2,X3,np__1,X4)))),k12_finseq_1(X1,k7_partfun1(X1,X3,X4))),k2_rfinseq(X1,X5,k2_finseq_7(X2,X3,np__1,X4)))=k2_finseq_7(X1,k2_finseq_7(X2,X3,np__1,X4),np__1,X5)|v1_xboole_0(X1)|r1_xxreal_0(X5,np__1)|v1_xboole_0(X2)|~m2_finseq_1(k2_finseq_7(X2,X3,np__1,X4),X1)|~r1_xxreal_0(X5,k3_finseq_1(k2_finseq_7(X2,X3,np__1,X4)))|~v7_ordinal1(X5)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X3,X1)|~r1_xxreal_0(X4,k3_finseq_1(X3))|~r1_xxreal_0(np__1,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X4)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_22065, negated_conjecture, (k2_xcmplx_0(X1,esk2_0)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(X1,esk2_0),esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(k2_xcmplx_0(X1,esk2_0)))).#trainneg
cnf(i_0_22068, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X1))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_22015, negated_conjecture, (k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1))=esk3_0|~r1_xxreal_0(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1)),np__1)|~v7_ordinal1(k2_xcmplx_0(esk2_0,k2_xcmplx_0(esk2_0,X1)))|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X1)|~v7_ordinal1(k2_xcmplx_0(esk2_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21467, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,esk2_0),X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X2,esk2_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22349, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22255, negated_conjecture, (k2_nat_1(esk2_0,X2)=esk2_0|v2_struct_0(X1)|~v1_xxreal_0(k2_nat_1(esk2_0,X2))|~r1_xxreal_0(k2_nat_1(esk2_0,X2),esk2_0)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_22243, negated_conjecture, (k2_nat_1(esk2_0,X2)=esk2_0|v2_struct_0(X1)|~r1_xxreal_0(k2_nat_1(esk2_0,X2),np__1)|~v7_ordinal1(k2_nat_1(esk2_0,X2))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_21727, negated_conjecture, (v3_xxreal_0(np__2)|v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_16259, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(esk2_0,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(k4_matrix14(X2,X3,X4))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1443, plain, (v1_relat_1(k2_matrix_1(k12_matrix_1(X1,X2)))|v1_xboole_0(u1_struct_0(X1))|~v1_matrix_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_1459, plain, (m1_subset_1(k4_tarski(X1,X2),k2_matrix_1(k4_matrix14(X3,X4,X5)))|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_20594, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X4,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(X4,esk2_0),X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_20565, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(np__1,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_18876, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(np__1,esk2_0),esk3_0)|r1_xxreal_0(esk3_0,k2_xcmplx_0(np__1,esk2_0)))).#trainneg
cnf(i_0_18713, negated_conjecture, (k2_xcmplx_0(np__1,esk2_0)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(np__1,esk2_0),esk6_0))).#trainneg
cnf(i_0_2651, negated_conjecture, (k2_xcmplx_0(k3_finseq_1(k4_matrix14(X1,X2,X3)),esk2_0)=k2_nat_1(esk2_0,X2)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~m1_subset_1(k3_finseq_1(k4_matrix14(X1,X2,X3)),k5_numbers)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_11636, negated_conjecture, (X1=k3_matrix_1(X2,k4_matrix14(X3,esk2_0,X4),X5,X6)|v1_xboole_0(u1_struct_0(X3))|v6_struct_0(X3)|X7!=k1_funct_1(k4_matrix14(X3,esk2_0,X4),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(k4_matrix14(X3,esk2_0,X4),k13_finseq_1(X2))|~m2_finseq_1(X7,X2)|~r2_hidden(k4_tarski(X5,X6),k2_matrix_1(esk4_0))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,X2)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v33_algstr_0(X3)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_5579, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8145, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k2_finseq_7(X1,X2,np__1,X3)))),k12_finseq_1(X1,k7_partfun1(X1,X2,X3))),k2_rfinseq(X1,X4,k2_finseq_7(X1,X2,np__1,X3)))=k2_finseq_7(X1,k2_finseq_7(X1,X2,np__1,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|np__1=X4|X4=X3|~m2_finseq_1(k2_finseq_7(X1,X2,np__1,X3),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(k2_finseq_7(X1,X2,np__1,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4))).#trainneg
cnf(i_0_21494, plain, (np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21728, negated_conjecture, (v3_xxreal_0(np__2)|v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,np__2))).#trainneg
cnf(i_0_1475, plain, (k2_finseq_2(np__0,X1)=k1_card_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_21640, plain, (r1_xxreal_0(np__0,esk10_1(k5_numbers))|~v2_xxreal_0(esk10_1(k5_numbers))|~v1_xxreal_0(esk10_1(k5_numbers)))).#trainneg
cnf(i_0_16358, plain, (r1_xxreal_0(X1,np__1)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(np__1,k4_matrix14(X2,X3,X4))|v6_struct_0(X2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(k4_matrix14(X2,X3,X4))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1460, plain, (m2_finseq_1(k4_tarski(X1,X2),X3)|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|k2_matrix_1(k4_matrix14(X4,X5,X6))!=k13_finseq_1(X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,X5)|~r1_xxreal_0(X1,X5)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~m1_subset_1(X5,k5_numbers)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_1463, plain, (m2_finseq_1(esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7),X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~v1_matrix_1(k4_matrix14(X2,X3,X4))|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~m1_subset_1(X7,X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_22636, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,esk2_0,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k1_funct_1(k4_matrix14(X2,esk2_0,X3),X4)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(esk4_0))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_18459, negated_conjecture, (k2_nat_1(X2,esk2_0)=k2_nat_1(esk2_0,X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_4390, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|k2_rfinseq(u1_struct_0(X2),X6,X7)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k2_rfinseq(u1_struct_0(X2),X6,X7),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(X7,u1_struct_0(X2))|~v7_ordinal1(X6))).#trainneg
cnf(i_0_18742, plain, (k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k8_finseq_1(k13_finseq_1(u1_struct_0(X1)),k9_finseq_1(k1_funct_1(k12_matrix_1(X1,X2),X3)),k17_finseq_1(k13_finseq_1(u1_struct_0(X1)),k7_nat_d(X3,np__2),k1_rfinseq(np__1,k12_matrix_1(X1,X2)))),k9_finseq_1(k1_funct_1(k4_matrix14(X1,X2,np__1),np__1))),k2_rfinseq(k13_finseq_1(u1_struct_0(X1)),X3,k12_matrix_1(X1,X2)))=k2_finseq_7(k13_finseq_1(u1_struct_0(X1)),k12_matrix_1(X1,X2),np__1,X3)|r1_xxreal_0(X3,np__1)|v6_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(X3)|~m2_finseq_1(k12_matrix_1(X1,X2),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,np__1)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_19759, plain, (r2_hidden(k16_finseq_1(X1,k2_finseq_7(X2,X3,X4,X5)),X6)|v1_xboole_0(X2)|X6!=k13_finseq_1(X2)|~v7_ordinal1(X1)|~m1_finseq_1(X3,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_9528, plain, (k3_matrix14(X1,X2,X3)=k3_matrix_1(u1_struct_0(X4),k4_matrix14(X4,X5,X6),X7,X8)|np__1=X7|X7=X8|X7=X6|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|v6_struct_0(X1)|k3_matrix14(X1,X2,X3)!=k4_struct_0(X4)|~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,X5)|~r1_xxreal_0(X8,X5)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~m1_subset_1(X5,k5_numbers)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|u1_struct_0(X4)!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1878, plain, (r1_xxreal_0(k7_nat_d(X1,X2),k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21729, negated_conjecture, (v3_xxreal_0(np__2)|v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,np__2))).#trainneg
cnf(i_0_1476, plain, (k2_finseq_2(np__0,X1)=k1_finseq_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1468, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_tarski(X2,X3),k2_zfmisc_1(k2_finseq_1(X4),k2_finseq_1(X4)))|v6_struct_0(X1)|v2_struct_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~m1_subset_1(X4,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_6090, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X3,X4)|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=X7|esk9_5(X5,esk4_0,X6,X4,X7)!=k1_funct_1(esk4_0,X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0))|~m2_finseq_1(esk9_5(X5,esk4_0,X6,X4,X7),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|X7!=k3_matrix_1(X5,esk4_0,X6,X4)|~m2_finseq_1(esk4_0,k13_finseq_1(X5))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,esk2_0)|~v7_ordinal1(X6)|~m1_subset_1(X7,X5))).#trainneg
cnf(i_0_10833, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X2)|X3!=k1_funct_1(esk4_0,np__1)|X1!=k1_funct_1(X3,X2)|~v1_matrix_1(k12_matrix_1(esk1_0,esk2_0))|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12354, plain, (np__1=X3|v1_xboole_0(X1)|v1_xboole_0(X4)|k10_xtuple_0(X5)!=k13_finseq_1(X1)|~m2_finseq_1(k9_finseq_1(X2),X4)|~m1_matrix_1(X5,X4,X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_7094, plain, (k3_finseq_1(k7_finseq_1(X2,X1))=X3|~v1_xxreal_0(k3_finseq_1(k7_finseq_1(X2,X1)))|~v1_xxreal_0(X3)|~r1_xxreal_0(k3_finseq_1(k7_finseq_1(X2,X1)),X3)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X2))|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_22860, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k1_funct_1(esk4_0,esk3_0)|np__1=X4|X4=X3|X4=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_funct_1(esk4_0,esk3_0)!=k4_struct_0(X1)|u1_struct_0(X1)!=k13_finseq_1(u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_22817, negated_conjecture, (k2_nat_1(X1,esk2_0)=k2_nat_1(esk2_0,X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_22451, negated_conjecture, (v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,np__2)|~v1_xxreal_0(np__2))).#trainneg
cnf(i_0_1477, plain, (k2_finseq_2(np__0,X1)=np__0)).#trainneg
cnf(i_0_22744, negated_conjecture, (v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,np__2)|~v1_xxreal_0(np__2))).#trainneg
cnf(i_0_1493, plain, (r1_xxreal_0(k6_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k6_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1494, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22812, negated_conjecture, (m2_finseq_1(esk9_5(X1,esk4_0,X2,X3,X4),X1)|X4!=k3_matrix_1(X1,esk4_0,X2,X3)|~m2_finseq_1(esk4_0,k13_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_22861, negated_conjecture, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k1_funct_1(esk4_0,np__1)|np__1=X4|X4=X3|X4=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_funct_1(esk4_0,np__1)!=k4_struct_0(X1)|u1_struct_0(X1)!=k13_finseq_1(u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_22797, plain, (k3_finseq_1(esk9_5(k13_finseq_1(X1),k4_matrix14(X2,X3,X4),X5,X6,X7))=X8|v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m1_matrix_1(esk9_5(k13_finseq_1(X1),k4_matrix14(X2,X3,X4),X5,X6,X7),X1,X8,X9)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|X7!=k3_matrix_1(k13_finseq_1(X1),k4_matrix14(X2,X3,X4),X5,X6)|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(k13_finseq_1(X1)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X7,k13_finseq_1(X1))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_22803, plain, (r2_hidden(esk9_5(X1,k4_matrix14(X2,X3,X4),X5,X6,X7),X8)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X8!=k13_finseq_1(X1)|X7!=k3_matrix_1(X1,k4_matrix14(X2,X3,X4),X5,X6)|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(X1))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X7,X1)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_9139, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k7_finseq_1(k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2)),X1)|~m1_finseq_1(k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_8990, plain, (k7_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1))).#trainneg
cnf(i_0_22743, plain, (np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23057, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_1495, plain, (r1_xxreal_0(k1_xreal_0(X1,X2),X3)|r1_xxreal_0(X3,k1_xreal_0(X1,X2))|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23084, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1497, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23053, negated_conjecture, (k2_nat_1(X1,esk2_0)=k2_nat_1(esk2_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23044, plain, (X1=np__1|epred1_5(X2,X3,X1,X4,k4_matrix14(X2,X5,X6))|X1=X3|X1=X6|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|k1_funct_1(esk4_0,esk3_0)!=k5_struct_0(X2)|k1_funct_1(esk4_0,esk3_0)!=k4_struct_0(X2)|u1_struct_0(X2)!=k13_finseq_1(u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X5)|~r1_xxreal_0(X3,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~m1_subset_1(X5,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_23123, negated_conjecture, (k3_finseq_1(esk9_5(k13_finseq_1(X1),esk4_0,X2,X3,X4))=X5|v1_xboole_0(X1)|X4!=k3_matrix_1(k13_finseq_1(X1),esk4_0,X2,X3)|~m2_finseq_1(esk4_0,k13_finseq_1(k13_finseq_1(X1)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~m1_matrix_1(esk9_5(k13_finseq_1(X1),esk4_0,X2,X3,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X4,k13_finseq_1(X1)))).#trainneg
cnf(i_0_23140, plain, (k1_funct_1(k1_funct_1(k4_matrix14(X1,X2,X3),X4),X5)=X6|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|np__1=X4|X4=X3|X4=X5|X6!=k4_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X2)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X6,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_23138, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1),X2)=X3|np__1=X2|X3!=k4_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,esk2_0,X1),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_23097, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|esk9_5(u1_struct_0(X2),esk4_0,X6,X7,X8)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(esk9_5(u1_struct_0(X2),esk4_0,X6,X7,X8),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|X8!=k3_matrix_1(u1_struct_0(X2),esk4_0,X6,X7)|~m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,esk2_0)|~r1_xxreal_0(X6,esk2_0)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~m1_subset_1(X8,u1_struct_0(X2)))).#trainneg
cnf(i_0_23137, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1),X2)=X3|np__1=X2|X3!=k4_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,esk2_0,X1),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23135, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1),X2)=X3|np__1=X2|X3!=k4_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,esk2_0,X1),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_23124, plain, (k3_finseq_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5))=X8|v1_xboole_0(u1_struct_0(X2))|v1_xboole_0(X1)|v6_struct_0(X2)|X7!=k3_matrix_1(k13_finseq_1(X1),k4_matrix14(X2,X3,X4),X5,X6)|~m2_finseq_1(k4_matrix14(X2,X3,X4),k13_finseq_1(k13_finseq_1(X1)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~m1_matrix_1(k1_funct_1(k4_matrix14(X2,X3,X4),X5),X1,X8,X9)|~v7_ordinal1(X9)|~v7_ordinal1(X8)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X7,k13_finseq_1(X1))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_23058, plain, (v1_funct_1(np__0))).#trainneg
cnf(i_0_1498, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v1_xxreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_23059, plain, (v1_finseq_1(np__0))).#trainneg
cnf(i_0_1500, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1501, plain, (m2_subset_1(k1_card_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_23136, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1),X2)=X3|np__1=X2|X3!=k5_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,esk2_0,X1),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_23139, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,X1,X2),X3),X3)=X4|np__1=X3|X3=X2|X4!=k5_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,X1,X2),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23108, negated_conjecture, (r2_hidden(k1_rfinseq(X1,esk9_5(X2,esk4_0,X3,X4,X5)),X6)|X6!=k13_finseq_1(X2)|~v7_ordinal1(X1)|X5!=k3_matrix_1(X2,esk4_0,X3,X4)|~m2_finseq_1(esk4_0,k13_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X2))).#trainneg
cnf(i_0_23144, plain, (k1_funct_1(k1_funct_1(k4_matrix14(X1,X2,X3),X4),X4)=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|np__1=X4|X4=X3|X5!=k4_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X5,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1))).#trainneg
cnf(i_0_23096, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|esk9_5(u1_struct_0(X2),esk4_0,X7,X8,X9)!=k3_matrix14(X2,X3,X5)|X1!=k1_funct_1(esk9_5(u1_struct_0(X2),esk4_0,X7,X8,X9),X6)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|X9!=k3_matrix_1(u1_struct_0(X2),esk4_0,X7,X8)|~m2_finseq_1(esk4_0,k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X8,esk2_0)|~r1_xxreal_0(X7,esk2_0)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~m1_subset_1(X9,u1_struct_0(X2)))).#trainneg
cnf(i_0_23145, plain, (k1_funct_1(k1_funct_1(k4_matrix14(X1,X2,X3),X4),X4)=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|np__1=X4|X4=X3|X5!=k5_struct_0(X1)|~m2_finseq_1(k4_matrix14(X1,X2,X3),k13_finseq_1(u1_struct_0(X1)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X5,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_23113, plain, (X1=np__1|epred1_5(X2,X3,X1,X4,k4_matrix14(X2,X5,X6))|X1=X3|X1=X6|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|k1_funct_1(esk4_0,np__1)!=k5_struct_0(X2)|k1_funct_1(esk4_0,np__1)!=k4_struct_0(X2)|u1_struct_0(X2)!=k13_finseq_1(u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X5)|~r1_xxreal_0(X3,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X6)|~m1_subset_1(X5,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_23265, plain, (k7_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k9_finseq_1(k7_partfun1(X1,X2,np__1))),X1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~m1_subset_1(k7_partfun1(X1,X2,np__1),X1))).#trainneg
cnf(i_0_22866, negated_conjecture, (v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,np__2)|~v1_xxreal_0(np__2))).#trainneg
cnf(i_0_1502, plain, (k2_nat_1(k1_card_1(X1),k3_finseq_1(X2))=k3_finseq_1(k7_finseq_1(k1_card_1(X1),X2))|~v1_finseq_1(X2)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(X2)|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(X2)|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_23054, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|~v1_xxreal_0(esk6_0)|~v1_xboole_0(esk6_0)|~v1_xxreal_0(np__2))).#trainneg
cnf(i_0_23363, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xcmplx_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1505, plain, (k2_nat_1(k3_finseq_1(X1),k1_card_1(X2))=k3_finseq_1(k7_finseq_1(X1,k1_card_1(X2)))|~v1_finseq_1(k1_card_1(X2))|~v1_finseq_1(X1)|~v1_funct_1(k1_card_1(X2))|~v1_funct_1(X1)|~v1_relat_1(k1_card_1(X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_23518, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,X1,X2),X3),X3)=k5_struct_0(esk1_0)|np__1=X3|X3=X2|~m2_finseq_1(k4_matrix14(esk1_0,X1,X2),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_23520, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,X1,X2),X3),X3)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X4,X5)|np__1=X3|X3=X2|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X4,X5)!=k5_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,X1,X2),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_23485, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23405, negated_conjecture, (k2_nat_1(esk2_0,X1)=k2_xcmplx_0(X1,esk2_0)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23521, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,X1,X2),X3),X3)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X4),X5,X6)|np__1=X3|X3=X2|k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X4),X5,X6)!=k5_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,X1,X2),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_23266, plain, (k7_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k2_rfinseq(X1,X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k1_rfinseq(np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m1_finseq_1(k2_rfinseq(X1,X3,X2),X1)|~m2_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(k7_partfun1(X1,X2,X3),X1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_23683, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|~v1_xxreal_0(esk6_0)|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_1507, plain, (k3_finseq_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_finseq_1(k3_finseq_1(X1))|~v1_funct_1(k3_finseq_1(X1))|~v1_relat_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_23685, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1512, plain, (v1_card_1(k2_nat_1(np__1,k3_finseq_1(X2)))|~v1_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_funct_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_relat_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_23746, negated_conjecture, (k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)),esk4_0,X1,X2)=k1_funct_1(k1_funct_1(esk4_0,X1),X2)|~m2_finseq_1(esk4_0,k13_finseq_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)),esk4_0,X1,X2),X3))).#trainneg
cnf(i_0_23743, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X1)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_23811, negated_conjecture, (k2_nat_1(X1,esk2_0)=k2_xcmplx_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23747, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,X1),X2)=k3_matrix_1(k5_numbers,esk4_0,X1,X2)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,X1,X2)))).#trainneg
cnf(i_0_23134, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,esk2_0,X1),X1),np__1)=X2|X2!=k4_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,esk2_0,X1),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_1517, plain, (m2_subset_1(X2,k1_numbers,k5_numbers)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_23630, negated_conjecture, (r1_xxreal_0(esk2_0,np__2)|~v1_xxreal_0(esk2_0)|~v1_xboole_0(esk2_0)|~v1_xxreal_0(np__2))).#trainneg
cnf(i_0_17584, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|v1_xboole_0(u1_struct_0(X2))|r1_xxreal_0(esk3_0,k4_matrix14(X2,X3,X4))|r1_xxreal_0(esk6_0,esk3_0)|v6_struct_0(X2)|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(k4_matrix14(X2,X3,X4))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k4_matrix14(X2,X3,X4))|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1518, plain, (k3_finseq_1(k7_finseq_1(k12_matrix_1(X1,X2),X3))=k2_nat_1(X2,k3_finseq_1(X3))|v2_struct_0(X1)|~v1_finseq_1(X3)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(X3)|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_1520, plain, (v1_xboole_0(k12_matrix_1(X1,X2))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v1_xboole_0(X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24156, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k4_struct_0(esk1_0)|k3_matrix_1(k5_numbers,esk4_0,X1,X2)!=k4_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,X1,X2)))).#trainneg
cnf(i_0_24159, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(k5_numbers,esk4_0,X1,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,X1,X2)))).#trainneg
cnf(i_0_24154, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k5_struct_0(esk1_0)|k3_matrix_1(k5_numbers,esk4_0,X1,X2)!=k5_struct_0(esk1_0)|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,X1,X2)))).#trainneg
cnf(i_0_24146, negated_conjecture, (k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)=k4_struct_0(esk1_0)|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)))).#trainneg
cnf(i_0_24158, negated_conjecture, (k3_matrix_1(k5_numbers,esk4_0,X1,X2)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|X1=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,X1,X2)))).#trainneg
cnf(i_0_24061, plain, (k3_finseq_1(X2)=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(X2)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(X2)|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(X2)|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24147, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k5_struct_0(esk1_0)|k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)))).#trainneg
cnf(i_0_24063, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23973, negated_conjecture, (r1_xxreal_0(esk6_0,np__2)|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_24177, negated_conjecture, (r1_xxreal_0(esk2_0,np__2)|~v1_xxreal_0(np__2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1537, negated_conjecture, (r1_xxreal_0(k6_xcmplx_0(X1,X2),esk2_0)|~r1_xxreal_0(k6_xcmplx_0(X1,X2),esk3_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24161, negated_conjecture, (k3_matrix_1(k5_numbers,esk4_0,X1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,X1,X1)))).#trainneg
cnf(i_0_23961, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X3)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23960, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X2)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23963, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3),u1_struct_0(esk1_0))|~r1_xxreal_0(X3,esk2_0))).#trainneg
cnf(i_0_24100, negated_conjecture, (k3_finseq_1(k7_finseq_1(esk4_0,X1))=k2_xcmplx_0(k3_finseq_1(X1),esk2_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_24060, plain, (k3_finseq_1(X2)=X2|v2_struct_0(X1)|~v1_finseq_1(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24144, negated_conjecture, (k5_struct_0(esk1_0)=k3_matrix_1(k5_numbers,esk4_0,np__1,esk3_0)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,np__1,esk3_0)))).#trainneg
cnf(i_0_16613, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1538, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(X1),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(X1),esk3_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18495, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk3_0,np__2)|~v1_xxreal_0(np__2)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__2))).#trainneg
cnf(i_0_15209, plain, (k4_matrix14(X1,X2,X3)=k4_matrix14(X4,X5,X6)|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~v1_xboole_0(X5)|~v7_ordinal1(X6)|~m1_subset_1(X5,k5_numbers)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|~v1_xboole_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1539, negated_conjecture, (r1_xxreal_0(k1_xreal_0(X1,X2),esk2_0)|~r1_xxreal_0(k1_xreal_0(X1,X2),esk3_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1541, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),esk2_0)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),esk3_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24267, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)=k4_struct_0(esk1_0)|k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)!=k4_struct_0(esk1_0)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24268, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k4_struct_0(esk1_0)|k3_matrix_1(k5_numbers,esk4_0,np__1,X1)!=k4_struct_0(esk1_0)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24270, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k5_struct_0(esk1_0)|k3_matrix_1(k5_numbers,esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24348, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X1)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24485, plain, (k3_finseq_1(X1)=X1|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_finseq_1(X1)|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_funct_1(X1)|~v1_relat_1(k12_matrix_1(X2,X1))|~v1_relat_1(X1)|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_24459, negated_conjecture, (k2_xcmplx_0(k3_finseq_1(k12_matrix_1(X1,X2)),esk2_0)=k2_nat_1(esk2_0,X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2))))).#trainneg
cnf(i_0_24511, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xreal_0(k4_xcmplx_0(X1))|~r1_xxreal_0(X1,esk3_0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24313, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_24322, negated_conjecture, (r1_xxreal_0(esk2_0,np__2)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_13838, plain, (k9_xtuple_0(X2)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_24521, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk3_0,np__2)|~v1_xxreal_0(np__2)|~r1_xxreal_0(esk2_0,np__2))).#trainneg
cnf(i_0_24315, negated_conjecture, (np__2=esk6_0|~v1_xxreal_0(esk6_0)|~v1_xxreal_0(np__2)|~r1_xxreal_0(np__2,esk6_0)|~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_1549, plain, (k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4)))=k9_finseq_1(X5)|k9_xtuple_0(k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4))))!=k2_finseq_1(np__1)|k1_funct_1(X1,np__1)!=X5|~v1_finseq_1(X4)|~v1_finseq_1(k7_finseq_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(X4)|~v1_funct_1(k7_finseq_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(X4)|~v1_relat_1(k7_finseq_1(X2,X3))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3))))|~r1_xxreal_0(np__1,k3_finseq_1(X1))|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1550, plain, (k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4)))=k9_finseq_1(X5)|k9_xtuple_0(k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4))))!=k2_finseq_1(np__1)|k1_funct_1(k7_finseq_1(X1,X2),np__1)!=X5|~v1_finseq_1(X4)|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(k7_finseq_1(X1,X2),X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,X2)))|~v1_finseq_1(k7_finseq_1(X3,X4))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(k7_finseq_1(X3,X4))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(k7_finseq_1(X3,X4))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24353, plain, (epred1_5(esk1_0,X1,X2,X3,k12_matrix_1(esk1_0,X4))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X4,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24356, plain, (X1=X2|X3=X2|epred1_5(esk1_0,X3,X1,X2,k12_matrix_1(esk1_0,X4))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3)!=k4_struct_0(esk1_0)|k3_matrix14(esk1_0,X4,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24354, plain, (X1=np__1|epred1_5(esk1_0,X2,X1,X3,k12_matrix_1(esk1_0,X4))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2)!=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X4,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24355, plain, (X1=np__1|epred1_5(esk1_0,X1,X2,X3,k12_matrix_1(esk1_0,X4))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|k3_matrix14(esk1_0,X4,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24368, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24458, negated_conjecture, (k2_xcmplx_0(k3_finseq_1(k12_matrix_1(X1,X2)),esk2_0)=k2_xcmplx_0(esk2_0,X2)|v2_struct_0(X1)|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2))))).#trainneg
cnf(i_0_25075, negated_conjecture, (r1_xxreal_0(k3_finseq_1(k12_matrix_1(X1,X2)),k2_xcmplx_0(esk2_0,X2))|v2_struct_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24888, plain, (X1=X2|X3=X2|epred1_5(esk1_0,X3,X1,X2,k4_matrix14(esk1_0,esk2_0,X1))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X3)!=k4_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X3),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24885, plain, (epred1_5(esk1_0,X1,X2,X3,k4_matrix14(esk1_0,esk2_0,X2))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)!=k4_struct_0(esk1_0)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24992, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,esk3_0)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_24993, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,esk3_0)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_25030, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk3_0,np__2)|~r1_xxreal_0(esk2_0,np__2))).#trainneg
cnf(i_0_25094, negated_conjecture, (k2_xcmplx_0(esk2_0,k3_finseq_1(k12_matrix_1(X1,X2)))=k2_xcmplx_0(esk2_0,X2)|v2_struct_0(X1)|~m1_subset_1(k3_finseq_1(k12_matrix_1(X1,X2)),k5_numbers)|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24886, plain, (X1=np__1|epred1_5(esk1_0,X2,X1,X3,k4_matrix14(esk1_0,esk2_0,X1))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25114, negated_conjecture, (r1_xxreal_0(esk2_0,k2_xcmplx_0(esk2_0,X2))|v2_struct_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24887, plain, (X1=np__1|epred1_5(esk1_0,X1,X2,X3,k4_matrix14(esk1_0,esk2_0,X2))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25088, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk2_0,X2),esk6_0)|r1_xxreal_0(esk6_0,k2_xcmplx_0(esk2_0,X2))|v2_struct_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24498, plain, (r1_xxreal_0(X1,X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1564, plain, (k1_card_1(k2_nat_1(np__1,k3_finseq_1(X2)))=k2_nat_1(np__1,k3_finseq_1(X2))|~v1_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_funct_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_relat_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1567, plain, (k1_card_1(X2)=X2|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24904, plain, (esk3_0=X1|X2=X1|epred1_5(esk1_0,X2,esk3_0,X1,k4_matrix14(esk1_0,esk2_0,X3))|esk3_0=X3|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2)!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25107, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk2_0,X3))|v2_struct_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X2,X3)))|~v1_xreal_0(X3)|~v1_finseq_1(k12_matrix_1(X2,X3))|~v1_funct_1(k12_matrix_1(X2,X3))|~v1_relat_1(k12_matrix_1(X2,X3))|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_24901, plain, (epred1_5(esk1_0,X1,esk3_0,X2,k4_matrix14(esk1_0,esk2_0,X3))|esk3_0=X3|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24903, plain, (X1=np__1|epred1_5(esk1_0,X1,esk3_0,X2,k4_matrix14(esk1_0,esk2_0,X3))|esk3_0=X3|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24902, plain, (epred1_5(esk1_0,X1,esk3_0,X2,k4_matrix14(esk1_0,esk2_0,X3))|esk3_0=X3|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25089, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk2_0,X2),esk3_0)|r1_xxreal_0(esk3_0,k2_xcmplx_0(esk2_0,X2))|v2_struct_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24968, negated_conjecture, (k2_nat_1(esk2_0,X2)=esk6_0|v2_struct_0(X1)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(k2_nat_1(esk2_0,X2),np__1)|~r1_xxreal_0(esk6_0,k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_17831, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k1_card_1(X3))|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_13909, plain, (k10_xtuple_0(X2)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1575, plain, (k3_finseq_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_25184, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk3_0,np__2)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_1568, plain, (k1_card_1(X2)=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,X2,X3))|~v1_funct_1(k4_matrix14(X1,X2,X3))|~v1_relat_1(k4_matrix14(X1,X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1576, plain, (k3_finseq_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_25262, negated_conjecture, (epred1_5(esk1_0,X1,esk3_0,X2,esk4_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25236, negated_conjecture, (k2_xcmplx_0(esk2_0,X2)=np__1|r1_xxreal_0(np__1,k2_xcmplx_0(esk2_0,X2))|v2_struct_0(X1)|~v1_xxreal_0(np__1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v1_xreal_0(X2)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_25265, negated_conjecture, (r1_xxreal_0(esk2_0,k2_xcmplx_0(esk2_0,X1))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v1_finseq_1(k12_matrix_1(X2,X1))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~m1_subset_1(k12_matrix_1(X2,X1),k5_numbers))).#trainneg
cnf(i_0_25263, negated_conjecture, (r1_xxreal_0(esk2_0,k2_xcmplx_0(esk2_0,X1))|v2_struct_0(X2)|~v1_xreal_0(X1)|~v1_finseq_1(k12_matrix_1(X2,X1))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_1577, plain, (k3_finseq_1(X1)=np__0|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_25185, negated_conjecture, (esk2_0=np__2|r1_xxreal_0(esk3_0,np__2)|~r1_xxreal_0(esk2_0,np__0))).#trainneg
cnf(i_0_1578, plain, (k3_finseq_1(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1579, plain, (k3_finseq_1(X1)=k10_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_25736, negated_conjecture, (epred1_5(esk1_0,X1,esk3_0,X2,esk4_0)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_25537, plain, (k1_card_1(X1)=X1|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25305, negated_conjecture, (k2_nat_1(esk2_0,X1)=esk6_0|v2_struct_0(X2)|~v1_xxreal_0(esk6_0)|~v1_finseq_1(k12_matrix_1(X2,X1))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~r1_xxreal_0(esk6_0,X1)|~r1_xxreal_0(k2_nat_1(esk2_0,X1),np__1)|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_10191, negated_conjecture, (k2_xcmplx_0(np__1,esk2_0)=esk2_0|~r1_xxreal_0(k2_xcmplx_0(np__1,esk2_0),esk3_0))).#trainneg
cnf(i_0_23403, negated_conjecture, (k2_nat_1(esk2_0,k3_finseq_1(X1))=k3_finseq_1(k7_finseq_1(X1,esk4_0))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_1159, plain, (k3_finseq_1(k2_finseq_7(X7,X2,X3,X4))=X5|v1_xboole_0(X1)|v1_xboole_0(X7)|~m1_finseq_1(X2,k13_finseq_1(X1))|~m1_matrix_1(k2_finseq_7(X7,X2,X3,X4),X1,X5,X6)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X2,X7))).#trainneg
cnf(i_0_17871, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k1_finseq_1(X3))|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1580, plain, (X1=k3_finseq_1(X2)|~v1_xboole_0(X1)|~v1_funct_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_25032, negated_conjecture, (np__2=esk6_0|~v1_xxreal_0(esk6_0)|~v1_xboole_0(esk6_0)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_1583, plain, (v1_xboole_0(k2_nat_1(np__1,k3_finseq_1(X2)))|~v1_funct_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_xboole_0(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1586, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_xboole_0(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_24457, negated_conjecture, (k2_xcmplx_0(k3_finseq_1(k12_matrix_1(X1,X2)),esk2_0)=k2_xcmplx_0(esk2_0,X2)|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2))))).#trainneg
cnf(i_0_22853, plain, (v1_xboole_0(X3)|~v1_xboole_0(X1)|X1!=k13_finseq_1(X3)|~m1_finseq_1(X4,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_18454, negated_conjecture, (r1_xxreal_0(k3_finseq_1(k12_matrix_1(X1,X2)),k2_nat_1(X2,esk2_0))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_relat_1(k12_matrix_1(X1,X2))|~v7_ordinal1(k3_finseq_1(k12_matrix_1(X1,X2)))|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_19136, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k9_finseq_1(k7_partfun1(X1,k12_matrix_1(X2,X3),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k1_rfinseq(np__1,k12_matrix_1(X2,X3)))),k9_finseq_1(k3_matrix14(X2,X3,np__1))),k2_rfinseq(X1,X4,k12_matrix_1(X2,X3)))=k2_finseq_7(X1,k12_matrix_1(X2,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|v6_struct_0(X2)|v2_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_11843, plain, (k7_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),k1_rfinseq(X3,X2))=k2_finseq_7(X1,X2,np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m1_finseq_1(k1_rfinseq(X3,X2),X1)|~m1_finseq_1(k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k1_funct_1(X2,X3)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,X2))),k12_finseq_1(X1,k7_partfun1(X1,X2,np__1))),X1)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1588, plain, (v7_ordinal1(k3_finseq_1(X1))|~v3_ordinal1(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_26398, negated_conjecture, (np__2=esk6_0|~v1_xboole_0(esk6_0)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_1589, plain, (v7_ordinal1(k1_card_1(k3_finseq_1(X1)))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1591, plain, (v1_finset_1(k2_nat_1(np__1,k3_finseq_1(X2)))|~v1_finset_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_funct_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_relat_1(k7_finseq_1(k9_finseq_1(X1),X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4374, plain, (k3_finseq_1(k7_finseq_1(X3,k12_matrix_1(X1,X2)))=k2_nat_1(X2,k3_finseq_1(X3))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,X2))|~v1_finseq_1(X3)|~v1_funct_1(k12_matrix_1(X1,X2))|~v1_funct_1(X3)|~v1_relat_1(k12_matrix_1(X1,X2))|~v1_relat_1(X3)|~v7_ordinal1(k3_finseq_1(X3))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_10601, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X3,np__1)|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,X3)|~m1_finseq_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6051, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X2),X3,X4)|X5=esk3_0|np__1=X5|X5=X4|k1_funct_1(esk4_0,X5)!=k3_matrix14(esk1_0,X2,X3)|X1!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X5,esk2_0)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(X4,esk2_0))).#trainneg
cnf(i_0_23798, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_25072, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k5_struct_0(esk1_0)|np__1=X1|X1=esk3_0|k1_funct_1(esk4_0,X1)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4080, plain, (v1_relat_1(k7_finseq_1(X1,X2))|~m1_finseq_1(X2,X3)|~m1_finseq_1(X1,X3))).#trainneg
cnf(i_0_4395, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|k2_finseq_7(u1_struct_0(X2),X6,X7,X8)!=k3_matrix14(X2,X3,X4)|X1!=k1_funct_1(k2_finseq_7(u1_struct_0(X2),X6,X7,X8),X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(X6,u1_struct_0(X2))|~v7_ordinal1(X8)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_2385, plain, (v1_xboole_0(X2)|m2_finseq_1(k4_matrix14(X1,X3,X4),X2)|v6_struct_0(X1)|v2_struct_0(X1)|~m1_finseq_1(k12_matrix_1(X1,X3),X2)|~m1_finseq_1(k12_matrix_1(X1,X3),k13_finseq_1(u1_struct_0(X1)))|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24499, plain, (r1_xxreal_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_18084, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k9_xtuple_0(X3))|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1613, negated_conjecture, (k2_xcmplx_0(esk3_0,X1)=esk3_0|~v1_xxreal_0(k2_xcmplx_0(esk3_0,X1))|~v7_ordinal1(X1)|~r1_xxreal_0(k2_xcmplx_0(esk3_0,X1),np__1)|~v7_ordinal1(k2_xcmplx_0(esk3_0,X1)))).#trainneg
cnf(i_0_24555, plain, (m1_matrix_1(k4_matrix14(X4,X5,X6),u1_struct_0(X1),X2,X2)|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X4))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v1_xboole_0(X2)|~v1_xboole_0(X5)|~v7_ordinal1(X6)|~m1_subset_1(X5,k5_numbers)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_1617, negated_conjecture, (k2_xcmplx_0(esk6_0,X1)=esk6_0|~v1_xxreal_0(k2_xcmplx_0(esk6_0,X1))|~v7_ordinal1(X1)|~r1_xxreal_0(k2_xcmplx_0(esk6_0,X1),np__1)|~v7_ordinal1(k2_xcmplx_0(esk6_0,X1)))).#trainneg
cnf(i_0_1631, negated_conjecture, (r1_xxreal_0(k6_xcmplx_0(X1,X2),esk2_0)|~r1_xxreal_0(k6_xcmplx_0(X1,X2),esk6_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26776, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k5_struct_0(esk1_0)|np__1=X4|X4=esk3_0|X4=X3|k1_funct_1(esk4_0,X4)!=k3_matrix14(esk1_0,X1,X2)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_26778, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,X5)|np__1=X6|X6=esk3_0|X6=X5|k1_funct_1(esk4_0,X6)!=k3_matrix14(esk1_0,X3,X4)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X6,esk2_0)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,esk2_0)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26773, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,X5)|np__1=X6|X6=esk3_0|X6=X5|k1_funct_1(esk4_0,X6)!=k3_matrix14(esk1_0,X3,X4)|k7_partfun1(u1_struct_0(esk1_0),X1,X2)!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X6,esk2_0)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,esk2_0)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~m1_finseq_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_26780, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X1,X2),X3,X4)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X5),X6,X7)|np__1=X8|X8=esk3_0|X8=X7|k1_funct_1(esk4_0,X8)!=k3_matrix14(esk1_0,X5,X6)|k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X1,X2),X3,X4)!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X8,esk2_0)|~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(X7,esk2_0)|~r1_xxreal_0(X7,X5)|~r1_xxreal_0(X6,X5)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X8)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_26779, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X4),X5,X6)|np__1=X7|X7=esk3_0|X7=X6|k1_funct_1(esk4_0,X7)!=k3_matrix14(esk1_0,X4,X5)|k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X7,esk2_0)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X6,esk2_0)|~r1_xxreal_0(X6,X4)|~r1_xxreal_0(X5,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X7)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26763, negated_conjecture, (k7_partfun1(u1_struct_0(esk1_0),esk4_0,X1)=k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X2,np__1)|X1=esk3_0|np__1=X1|k7_partfun1(u1_struct_0(esk1_0),esk4_0,X1)!=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,X2)|~m1_finseq_1(k12_matrix_1(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26799, plain, (v1_xboole_0(X1)|m2_finseq_1(k4_matrix14(X2,X3,X4),X1)|v6_struct_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_1641, negated_conjecture, (r1_xxreal_0(k1_xreal_0(X1,X2),esk3_0)|~r1_xxreal_0(k1_xreal_0(X1,X2),np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1632, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(X1),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(X1),esk6_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4886, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(np__2,esk3_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1633, negated_conjecture, (r1_xxreal_0(k1_xreal_0(X1,X2),esk2_0)|~r1_xxreal_0(k1_xreal_0(X1,X2),esk6_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1635, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),esk2_0)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),esk6_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21249, negated_conjecture, (r1_xxreal_0(X1,k2_xcmplx_0(esk6_0,X2))|r1_xxreal_0(esk6_0,k2_xcmplx_0(esk2_0,X3))|~v1_xxreal_0(k2_xcmplx_0(esk2_0,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_26926, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|np__1=X5|X5=X4|X7!=k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_26923, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X6!=k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),np__1)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_26922, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|np__1=X5|X5=X8|X5=X4|X7!=k1_funct_1(k4_matrix14(X2,X3,X8),X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~v7_ordinal1(X8))).#trainneg
cnf(i_0_26925, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X6!=k1_funct_1(k4_matrix14(X2,X3,X7),X7)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X7,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,X3)|~v7_ordinal1(X7))).#trainneg
cnf(i_0_26924, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X6!=k1_funct_1(k12_matrix_1(X2,X3),np__1)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3))))).#trainneg
cnf(i_0_26920, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X4!=k1_funct_1(esk4_0,esk3_0)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(esk3_0,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0)))|~r1_xxreal_0(X2,k3_finseq_1(k12_matrix_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_27025, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xreal_0(k4_xcmplx_0(X1))|~r1_xxreal_0(X1,esk6_0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14106, negated_conjecture, (r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(esk2_0,esk6_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1639, negated_conjecture, (r1_xxreal_0(k6_xcmplx_0(X1,X2),esk3_0)|~r1_xxreal_0(k6_xcmplx_0(X1,X2),np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27035, negated_conjecture, (r1_xxreal_0(np__2,esk6_0)|r1_xxreal_0(esk6_0,esk3_0)|~r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_1640, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(X1),esk3_0)|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1643, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),esk3_0)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26992, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k4_struct_0(esk1_0)|X4=esk3_0|np__1=X4|X4=X3|k1_funct_1(esk4_0,X4)!=k3_matrix14(esk1_0,X1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_26982, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k4_struct_0(esk1_0)|X4=esk3_0|np__1=X4|X4=X3|k1_funct_1(esk4_0,X4)!=k3_matrix14(esk1_0,X1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_12520, negated_conjecture, (k1_funct_1(k4_matrix14(X1,X2,X3),np__1)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X4),X5,X5)|X5=esk3_0|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k1_funct_1(k4_matrix14(X1,X2,X3),np__1)!=k5_struct_0(esk1_0)|~m1_finseq_1(k12_matrix_1(X1,X2),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,esk2_0)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~r1_xxreal_0(np__1,k3_finseq_1(k4_matrix14(X1,X2,X3)))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(X3,k3_finseq_1(k12_matrix_1(X1,X2)))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24358, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2)=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(X2,esk2_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22960, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X2)|X3!=k1_funct_1(esk4_0,np__1)|X1!=k1_funct_1(X3,X2)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_20335, negated_conjecture, (k3_finseq_1(k4_matrix14(X2,X1,X3))=k2_nat_1(X1,esk2_0)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~v1_xxreal_0(k3_finseq_1(k4_matrix14(X2,X1,X3)))|~v1_xxreal_0(k2_nat_1(X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,esk2_0),k3_finseq_1(k4_matrix14(X2,X1,X3)))|~v7_ordinal1(k3_finseq_1(k4_matrix14(X2,X1,X3)))|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3768, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X4),X5)|~v1_matrix_1(k12_matrix_1(X2,X3))|~m2_finseq_1(k1_funct_1(k12_matrix_1(X2,X3),X4),u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2))).#trainneg
cnf(i_0_24365, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,np__1)=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,np__1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18136, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k10_xtuple_0(X3))|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_27370, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,np__1)=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_1653, plain, (k2_matrix_1(k4_matrix14(X1,np__0,X2))=k2_zfmisc_1(np__0,np__0)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_26854, plain, (v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X2))|m1_matrix_1(k4_matrix14(X2,X3,X4),u1_struct_0(X1),X5,X5)|v6_struct_0(X2)|v6_struct_0(X1)|~v1_xboole_0(X5)|~v1_xboole_0(X3)|~v7_ordinal1(X4)|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X1)|~l6_algstr_0(X2)|~v5_vectsp_1(X1)|~v5_vectsp_1(X2)|~v4_vectsp_1(X1)|~v4_vectsp_1(X2)|~v5_group_1(X1)|~v5_group_1(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v4_rlvect_1(X1)|~v4_rlvect_1(X2)|~v3_rlvect_1(X1)|~v3_rlvect_1(X2)|~v2_rlvect_1(X1)|~v2_rlvect_1(X2)|~v33_algstr_0(X1)|~v33_algstr_0(X2)|~v13_algstr_0(X1)|~v13_algstr_0(X2)|~v1_xboole_0(X6))).#trainneg
cnf(i_0_1654, plain, (k2_zfmisc_1(k4_finseq_1(k12_matrix_1(X1,np__0)),np__0)=k2_matrix_1(k12_matrix_1(X1,np__0))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v1_finseq_1(k12_matrix_1(X1,np__0))|~v1_funct_1(k12_matrix_1(X1,np__0))|~v1_relat_1(k12_matrix_1(X1,np__0))|~v1_matrix_1(k12_matrix_1(X1,np__0))|~l6_algstr_0(X1))).#trainneg
cnf(i_0_1655, plain, (k2_zfmisc_1(k4_finseq_1(k4_matrix14(X1,np__0,X2)),np__0)=k2_matrix_1(k4_matrix14(X1,np__0,X2))|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~v1_finseq_1(k4_matrix14(X1,np__0,X2))|~v1_funct_1(k4_matrix14(X1,np__0,X2))|~v1_relat_1(k4_matrix14(X1,np__0,X2))|~v1_matrix_1(k4_matrix14(X1,np__0,X2))|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27337, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k5_struct_0(esk1_0)|X2!=k1_funct_1(esk4_0,np__1)|k1_funct_1(X2,X1)!=k5_struct_0(esk1_0)|~m2_finseq_1(X2,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27116, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5),u1_struct_0(X2))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27336, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k4_struct_0(esk1_0)|X2!=k1_funct_1(esk4_0,np__1)|k1_funct_1(X2,X1)!=k4_struct_0(esk1_0)|~m2_finseq_1(X2,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_27369, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,X4),X5)|~m2_finseq_1(k3_matrix14(X2,X3,X4),u1_struct_0(X2))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27153, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k1_funct_1(k4_matrix14(X2,X3,X7),X7)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X7,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X7)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27368, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X2,X3)|np__1=X2|X2=esk3_0|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X2),X3)|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X2),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27180, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,esk6_0)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_27181, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,esk6_0)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_1656, plain, (k2_zfmisc_1(k9_xtuple_0(k12_matrix_1(X1,np__0)),np__0)=k2_matrix_1(k12_matrix_1(X1,np__0))|v1_xboole_0(u1_struct_0(X1))|~v1_matrix_1(k12_matrix_1(X1,np__0))|~l6_algstr_0(X1))).#trainneg
cnf(i_0_1657, plain, (k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4)))=k9_finseq_1(k1_funct_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)),np__1))|k9_xtuple_0(k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4))))!=k2_finseq_1(np__1)|~v1_finseq_1(X4)|~v1_finseq_1(k7_finseq_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(X4)|~v1_funct_1(k7_finseq_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(X4)|~v1_relat_1(k7_finseq_1(X2,X3))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3))))|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1658, plain, (k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4)))=k9_finseq_1(k1_funct_1(k7_finseq_1(k7_finseq_1(X1,X2),X3),np__1))|k9_xtuple_0(k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4))))!=k2_finseq_1(np__1)|~v1_finseq_1(X4)|~v1_finseq_1(X3)|~v1_finseq_1(k7_finseq_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k7_finseq_1(X1,X2))|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(k7_finseq_1(X1,X2))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(k7_finseq_1(X1,X2),X3)))|~v1_finseq_1(k7_finseq_1(X3,X4))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(k7_finseq_1(X3,X4))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(k7_finseq_1(X3,X4))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27161, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k1_funct_1(k12_matrix_1(X2,X3),np__1)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27317, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k4_struct_0(esk1_0)|np__1=X2|X2=esk3_0|X2=X1|k1_funct_1(esk4_0,X2)!=k1_funct_1(esk4_0,np__1)|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_27125, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),np__1),X5)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),np__1),u1_struct_0(X2))|~r1_xxreal_0(np__1,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(X4,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_828, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9526, plain, (k3_matrix14(X1,X2,X3)=k3_matrix_1(u1_struct_0(X4),k12_matrix_1(X4,X5),X6,X7)|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|v6_struct_0(X1)|k3_matrix14(X4,X8,X7)!=k3_matrix14(X4,X5,X6)|k3_matrix14(X1,X2,X3)!=k5_struct_0(X4)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X7,X5)|~r1_xxreal_0(X6,X5)|~r1_xxreal_0(X7,X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X8)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4)|u1_struct_0(X4)!=k13_finseq_1(u1_struct_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27237, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xreal_0(k4_xcmplx_0(X1))|~r1_xxreal_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15743, negated_conjecture, (r1_xxreal_0(esk4_0,esk6_0)|r1_xxreal_0(esk6_0,esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1663, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|m1_matrix_1(X5,X2,X3,X4)|~r2_hidden(X1,esk8_4(X2,X3,X4,X5))|k3_finseq_1(X5)!=X3|~v1_matrix_1(X5)|~m2_finseq_1(X5,k13_finseq_1(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_26576, negated_conjecture, (v7_ordinal1(k3_finseq_1(esk2_0))|~v1_finseq_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0)|~v1_finset_1(esk4_0))).#trainneg
cnf(i_0_1664, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|X3=k13_finseq_1(X2)|r2_hidden(esk7_2(X2,X3),X3)|~r2_hidden(X1,esk7_2(X2,X3)))).#trainneg
cnf(i_0_1665, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,k8_finseq_1(X2,X3,X4))|~m1_finseq_1(X4,X2)|~m1_finseq_1(X3,X2))).#trainneg
cnf(i_0_27565, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),esk3_0)|r1_xxreal_0(esk3_0,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27566, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X1),esk6_0)|r1_xxreal_0(esk6_0,k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27440, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k5_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,np__1),X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27459, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X4)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k5_struct_0(X2)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27561, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27454, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),esk3_0,X1)=k4_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,np__1),X1)!=k4_struct_0(esk1_0)|~l2_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27453, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X7!=k3_matrix14(X2,X3,X5)|X1!=k1_funct_1(X7,X6)|~m2_finseq_1(X7,u1_struct_0(X2))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27628, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_16255, negated_conjecture, (r1_xxreal_0(esk4_0,esk3_0)|r1_xxreal_0(esk3_0,esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1666, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,k17_finseq_1(X2,X3,X4))|~m1_finseq_1(X4,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11822, negated_conjecture, (np__2=esk6_0|r1_xxreal_0(esk3_0,np__2)|r1_xxreal_0(esk6_0,esk3_0)|~v1_xxreal_0(np__2)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_25010, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(np__2)|~r1_xxreal_0(np__2,esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1667, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,k2_rfinseq(X2,X3,X4))|~m1_finseq_1(X4,X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1668, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|~r2_hidden(X1,k12_finseq_1(X2,X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_27449, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(X2)),k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~m1_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_27927, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X3),X4,X4)|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27952, negated_conjecture, (k2_nat_1(esk2_0,X1)=np__1|v2_struct_0(X2)|~v1_xxreal_0(np__1)|~v1_finseq_1(k12_matrix_1(X2,X1))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~r1_xxreal_0(k2_nat_1(esk2_0,X1),np__1)|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_27904, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0))|k1_funct_1(k1_funct_1(esk4_0,np__1),X1)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_27441, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k1_funct_1(k12_matrix_1(X2,X3),X5),X6)|~m2_finseq_1(k1_funct_1(k12_matrix_1(X2,X3),X5),u1_struct_0(X2))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27460, plain, (X1=k3_matrix_1(u1_struct_0(X2),k12_matrix_1(X2,X3),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X4=X5|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27629, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_16261, negated_conjecture, (r1_xxreal_0(esk4_0,esk2_0)|r1_xxreal_0(esk2_0,esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1669, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|~r2_hidden(X1,k1_finseq_7(X2,X3,X4,X5))|~m1_finseq_1(X3,X2)|~v7_ordinal1(X4)|~m1_subset_1(X5,X2))).#trainneg
cnf(i_0_11614, plain, (v1_card_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_28263, plain, (v1_card_1(np__0))).#trainneg
cnf(i_0_28118, negated_conjecture, (esk2_0=np__2|~v1_xxreal_0(np__2)|~v1_xboole_0(esk2_0)|~r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_1670, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|~r2_hidden(X1,k5_finseq_2(X2,X3,X4))|~v7_ordinal1(X3)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_1671, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|~r2_hidden(X1,k9_finseq_1(X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_28205, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=k4_struct_0(X1)|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_27928, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X5),X6,X6)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X6,X5)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_27951, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|r1_xxreal_0(k3_finseq_1(X1),np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28207, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=k5_struct_0(X1)|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27892, plain, (epred1_5(esk1_0,X1,esk3_0,X2,k12_matrix_1(esk1_0,esk2_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,np__1),X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27893, plain, (epred1_5(esk1_0,X1,esk3_0,X2,k12_matrix_1(esk1_0,esk2_0))|k1_funct_1(k1_funct_1(esk4_0,np__1),X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16364, negated_conjecture, (r1_xxreal_0(esk4_0,np__1)|r1_xxreal_0(np__1,esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1672, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X2)|~r2_hidden(X1,k2_finseq_7(X2,X3,X4,X5))|~m1_finseq_1(X3,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_28116, negated_conjecture, (np__2=esk6_0|r1_xxreal_0(esk6_0,esk3_0)|r1_xxreal_0(esk3_0,np__2)|~v1_xxreal_0(esk6_0)|~r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_3880, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk3_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3881, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),esk3_0)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1673, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k13_finseq_1(u1_struct_0(X2))))|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k12_matrix_1(X2,X3))|~v7_ordinal1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_1674, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))|v6_struct_0(X2)|v2_struct_0(X2)|~r2_hidden(X1,k3_matrix14(X2,X3,X4))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_28472, negated_conjecture, (epred1_5(esk1_0,X1,esk3_0,X2,k12_matrix_1(esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28324, plain, (epred1_5(X1,X2,X3,X4,k12_matrix_1(X1,X5))|X3=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X2,X5)|~r1_xxreal_0(X3,X5)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15669, plain, (v1_xboole_0(esk10_1(k4_finseq_2(X1,X2)))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17990, negated_conjecture, (k2_zfmisc_1(np__0,np__0)=k2_matrix_1(esk4_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_28334, negated_conjecture, (k4_struct_0(esk1_0)=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_1677, negated_conjecture, (r1_xxreal_0(k6_xcmplx_0(X1,X2),esk6_0)|~r1_xxreal_0(k6_xcmplx_0(X1,X2),np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28512, negated_conjecture, (np__2=esk6_0|r1_xxreal_0(esk3_0,np__2)|r1_xxreal_0(esk6_0,esk3_0)|~r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_28270, negated_conjecture, (esk2_0=np__2|~v1_xboole_0(esk2_0)|~r1_xxreal_0(np__2,esk2_0))).#trainneg
cnf(i_0_28516, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk3_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_28349, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|X3=X4|v6_struct_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1675, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k13_finseq_1(u1_struct_0(X2))))|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|~r2_hidden(X1,k4_matrix14(X2,X3,X4))|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1678, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(X1),esk6_0)|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28327, plain, (X1=X2|X3=X2|epred1_5(X4,X3,X1,X2,k12_matrix_1(X4,X5))|X1=X3|v1_xboole_0(u1_struct_0(X4))|v6_struct_0(X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X3,X5)|~r1_xxreal_0(X1,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~l6_algstr_0(X4)|~v5_vectsp_1(X4)|~v4_vectsp_1(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v33_algstr_0(X4)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_27548, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X6!=k3_matrix14(X2,X3,np__1)|X1!=k1_funct_1(X6,X5)|~m2_finseq_1(X6,u1_struct_0(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3155, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(np__2,esk3_0))).#trainneg
cnf(i_0_3886, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_18315, negated_conjecture, (r1_xxreal_0(esk6_0,esk4_0)|r1_xxreal_0(esk4_0,esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1679, negated_conjecture, (r1_xxreal_0(k1_xreal_0(X1,X2),esk6_0)|~r1_xxreal_0(k1_xreal_0(X1,X2),np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28517, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk3_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_28520, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),esk3_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1681, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),esk6_0)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28749, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),X2,X3)|X4!=k1_funct_1(esk4_0,esk3_0)|X1!=k1_funct_1(X4,X3)|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28198, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=esk12_1(X1)|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|esk12_1(X1)!=k4_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_28738, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xreal_0(k4_xcmplx_0(X1))|~r1_xxreal_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27922, plain, (k7_partfun1(u1_struct_0(X1),X2,X3)=k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X4),X5,X5)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k7_partfun1(u1_struct_0(X1),X2,X3)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X4)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_finseq_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_28348, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,X1,X2),X3),X3)=k4_struct_0(esk1_0)|np__1=X3|X3=X2|X5=X6|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,X1,X2),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X4)|~r1_xxreal_0(X5,X4))).#trainneg
cnf(i_0_28750, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_28575, plain, (v7_ordinal1(esk10_1(k4_finseq_2(X1,X2)))|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28649, negated_conjecture, (X1=esk6_0|v1_xboole_0(u1_struct_0(X2))|m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|v6_struct_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(esk6_0,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1693, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X4)),k17_finseq_1(X1,k7_nat_d(X3,np__2),k2_rfinseq(X1,np__1,k10_funct_7(X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,k10_funct_7(X2,X3,X4),np__1))),k2_rfinseq(X1,X3,k10_funct_7(X2,X3,X4)))=k2_finseq_7(X1,k10_funct_7(X2,X3,X4),np__1,X3)|v1_xboole_0(X1)|r1_xxreal_0(X3,np__1)|~m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~r1_xxreal_0(X3,k3_finseq_1(k10_funct_7(X2,X3,X4)))|~v7_ordinal1(X3)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1694, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k10_funct_7(X2,np__1,X3),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k10_funct_7(X2,np__1,X3)))),k12_finseq_1(X1,k7_partfun1(X1,X2,X3))),k2_rfinseq(X1,X4,k10_funct_7(X2,np__1,X3)))=k2_finseq_7(X1,k10_funct_7(X2,np__1,X3),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k10_funct_7(X2,np__1,X3),X1)|~r1_xxreal_0(X4,k3_finseq_1(k10_funct_7(X2,np__1,X3)))|~v7_ordinal1(X4)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_28851, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k5_struct_0(esk1_0)|X3!=k1_funct_1(esk4_0,esk3_0)|k1_funct_1(X3,X2)!=k5_struct_0(esk1_0)|~m2_finseq_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28857, plain, (X1=np__1|epred1_5(X2,X3,X1,X4,k12_matrix_1(X2,X5))|X1=X3|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|esk12_1(X2)!=k5_struct_0(X2)|esk12_1(X2)!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X3,X5)|~r1_xxreal_0(X1,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_28197, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=esk14_1(X1)|X3=X4|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|v7_struct_0(X1)|esk14_1(X1)!=k4_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_28856, plain, (epred1_5(X1,X2,X3,X4,k12_matrix_1(X1,X5))|X3=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|esk12_1(X1)!=k5_struct_0(X1)|esk12_1(X1)!=k4_struct_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X2,X5)|~r1_xxreal_0(X3,X5)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_28471, negated_conjecture, (epred1_5(esk1_0,X1,esk3_0,X2,k12_matrix_1(esk1_0,esk2_0))|k3_matrix_1(k5_numbers,esk4_0,np__1,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,np__1,X1)))).#trainneg
cnf(i_0_27936, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),X1)|~v7_ordinal1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_28906, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_28978, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_28871, negated_conjecture, (esk12_1(esk1_0)=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1)|esk12_1(esk1_0)!=k4_struct_0(esk1_0))).#trainneg
cnf(i_0_1766, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k6_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28621, negated_conjecture, (np__2=esk3_0|np__2=esk6_0|r1_xxreal_0(esk6_0,esk3_0)|~v1_xxreal_0(esk3_0)|~v1_xxreal_0(np__2)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_29002, negated_conjecture, (esk3_0=esk6_0|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(esk6_0,X2)|~r1_xxreal_0(esk3_0,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1702, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,X2,X3)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k10_funct_7(X2,X3,X4)))),k12_finseq_1(X1,k7_partfun1(X1,k10_funct_7(X2,X3,X4),np__1))),k2_rfinseq(X1,X4,k10_funct_7(X2,X3,X4)))=k2_finseq_7(X1,k10_funct_7(X2,X3,X4),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k10_funct_7(X2,X3,X4),X1)|~r1_xxreal_0(X4,k3_finseq_1(k10_funct_7(X2,X3,X4)))|~v7_ordinal1(X4)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(X4,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1703, plain, (k8_finseq_1(X1,k8_finseq_1(X1,k8_finseq_1(X1,k12_finseq_1(X1,k7_partfun1(X1,k10_funct_7(X2,X3,np__1),X4)),k17_finseq_1(X1,k7_nat_d(X4,np__2),k2_rfinseq(X1,np__1,k10_funct_7(X2,X3,np__1)))),k12_finseq_1(X1,k7_partfun1(X1,X2,X3))),k2_rfinseq(X1,X4,k10_funct_7(X2,X3,np__1)))=k2_finseq_7(X1,k10_funct_7(X2,X3,np__1),np__1,X4)|v1_xboole_0(X1)|r1_xxreal_0(X4,np__1)|~m2_finseq_1(k10_funct_7(X2,X3,np__1),X1)|~r1_xxreal_0(X4,k3_finseq_1(k10_funct_7(X2,X3,np__1)))|~v7_ordinal1(X4)|~m1_finseq_1(X2,X1)|~r1_xxreal_0(np__1,k3_finseq_1(X2))|~r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_29156, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),X1,X2)=k5_struct_0(esk1_0)|k1_funct_1(k1_funct_1(esk4_0,esk3_0),X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29158, plain, (X1=np__1|epred1_5(X2,X3,X1,X4,k12_matrix_1(X2,X5))|X1=X3|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|esk14_1(X2)!=k5_struct_0(X2)|esk14_1(X2)!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X3,X5)|~r1_xxreal_0(X1,X5)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29157, plain, (epred1_5(X1,X2,X3,X4,k12_matrix_1(X1,X5))|X3=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|esk14_1(X1)!=k5_struct_0(X1)|esk14_1(X1)!=k4_struct_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X2,X5)|~r1_xxreal_0(X3,X5)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_28206, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,X1),X2,X3)=k5_struct_0(esk1_0)|X2=X3|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28454, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3)=k5_struct_0(esk1_0)|X2=X3|k3_matrix14(esk1_0,X1,X2)!=k1_funct_1(esk4_0,np__1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,esk2_0)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X3),u1_struct_0(esk1_0))|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_28339, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X4),X4,X5)=k4_struct_0(esk1_0)|np__1=X5|X2=X3|~r1_xxreal_0(X4,esk2_0)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X5,esk2_0)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_28014, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1)=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,X1),u1_struct_0(esk1_0))|k1_funct_1(k1_funct_1(esk4_0,np__1),X1)!=k4_struct_0(esk1_0)|~l2_struct_0(esk1_0))).#trainneg
cnf(i_0_29239, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29240, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18320, negated_conjecture, (r1_xxreal_0(esk3_0,esk4_0)|r1_xxreal_0(esk4_0,esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_29172, negated_conjecture, (esk14_1(esk1_0)=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk3_0)!=k1_funct_1(esk4_0,np__1)|esk14_1(esk1_0)!=k4_struct_0(esk1_0))).#trainneg
cnf(i_0_29327, negated_conjecture, (np__2=esk6_0|np__2=esk3_0|r1_xxreal_0(esk6_0,esk3_0)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_29335, negated_conjecture, (esk3_0=esk6_0|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_29006, negated_conjecture, (esk2_0=esk6_0|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(esk6_0,X2)|~r1_xxreal_0(esk2_0,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1763, plain, (k7_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4))),X5)=k9_finseq_1(X6)|k9_xtuple_0(k7_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4))),X5))!=k2_finseq_1(np__1)|k1_funct_1(k7_finseq_1(X1,k7_finseq_1(X2,X3)),np__1)!=X6|~v1_finseq_1(X5)|~v1_finseq_1(X4)|~v1_finseq_1(k7_finseq_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(k7_finseq_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(X5)|~v1_relat_1(X4)|~v1_relat_1(k7_finseq_1(X2,X3))|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,k7_finseq_1(X3,X4)))))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3))))|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1765, plain, (k7_finseq_1(X1,k7_finseq_1(k7_finseq_1(X2,X3),X4))=k9_finseq_1(X5)|k9_xtuple_0(k7_finseq_1(X1,k7_finseq_1(k7_finseq_1(X2,X3),X4)))!=k2_finseq_1(np__1)|k1_funct_1(k7_finseq_1(X1,X2),np__1)!=X5|~v1_finseq_1(X4)|~v1_finseq_1(X3)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,k7_finseq_1(X2,X3))))|~r1_xxreal_0(np__1,k3_finseq_1(k7_finseq_1(X1,X2)))|~v1_finseq_1(k7_finseq_1(X2,X3))|~v1_funct_1(k7_finseq_1(X2,X3))|~v1_relat_1(k7_finseq_1(X2,X3)))).#trainneg
cnf(i_0_29199, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X5)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k5_struct_0(X2)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29716, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X4)=k5_struct_0(X1)|np__1=X4|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_29715, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X1,X2),X3,X3)=k5_struct_0(esk1_0)|np__1=X3|X3=X2|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29500, plain, (X1=X2|X3=X2|epred1_5(esk1_0,X3,X1,X2,k12_matrix_1(esk1_0,X4))|X1=X3|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_29714, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X4)=k4_struct_0(X1)|np__1=X4|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_29497, plain, (epred1_5(esk1_0,X1,X2,X3,k12_matrix_1(esk1_0,X4))|X2=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_27953, negated_conjecture, (k2_nat_1(esk2_0,X1)=esk2_0|v2_struct_0(X2)|~v1_finseq_1(k12_matrix_1(X2,X1))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~r1_xxreal_0(k2_nat_1(esk2_0,X1),np__1)|~v7_ordinal1(k2_nat_1(esk2_0,X1))|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_29484, plain, (X1=np__1|epred1_5(esk1_0,X1,X2,X3,k4_matrix14(esk1_0,esk2_0,X2))|k1_funct_1(k1_funct_1(esk4_0,esk3_0),X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_29483, plain, (X1=np__1|epred1_5(esk1_0,X2,X1,X3,k4_matrix14(esk1_0,esk2_0,X1))|k1_funct_1(k1_funct_1(esk4_0,esk3_0),X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18517, negated_conjecture, (r1_xxreal_0(np__1,esk4_0)|r1_xxreal_0(esk4_0,esk6_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_29608, negated_conjecture, (np__2=esk3_0|np__2=esk6_0|r1_xxreal_0(esk6_0,esk3_0)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(esk6_0,np__2))).#trainneg
cnf(i_0_29650, negated_conjecture, (esk2_0=esk6_0|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27929, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X6),X7,X7)|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,X6)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_28190, negated_conjecture, (k2_nat_1(esk2_0,X1)=np__1|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v1_xxreal_0(np__1)|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~r1_xxreal_0(k2_nat_1(esk2_0,X1),np__1)|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_28748, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,np__1),X5)|~m2_finseq_1(k3_matrix14(X2,X3,np__1),u1_struct_0(X2))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27954, negated_conjecture, (k2_nat_1(esk2_0,X1)=esk2_0|v2_struct_0(X2)|~v1_xxreal_0(k2_nat_1(esk2_0,X1))|~v1_finseq_1(k12_matrix_1(X2,X1))|~v1_funct_1(k12_matrix_1(X2,X1))|~v1_relat_1(k12_matrix_1(X2,X1))|~r1_xxreal_0(k2_nat_1(esk2_0,X1),esk2_0)|~v7_ordinal1(X1)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_27470, negated_conjecture, (k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)|X1=esk3_0|np__1=X1|~m2_finseq_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k1_funct_1(k7_partfun1(k13_finseq_1(u1_struct_0(esk1_0)),esk4_0,X1),X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_29602, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__2,esk3_0)|~r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18664, negated_conjecture, (r1_xxreal_0(np__1,esk4_0)|r1_xxreal_0(esk4_0,esk3_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1771, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)),X3)=k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1772, plain, (k2_xcmplx_0(X2,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)))|~r1_xxreal_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)),k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_29829, plain, (k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X2,X3),X4,X5)=k5_struct_0(X1)|np__1=X4|X4=X5|X4=X3|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_29738, plain, (X1=np__1|epred1_5(esk1_0,X1,X1,X2,k4_matrix14(esk1_0,X3,X4))|X1=X4|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_29773, negated_conjecture, (X1=np__1|epred1_5(esk1_0,X1,X2,X3,k4_matrix14(esk1_0,esk2_0,X2))|k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)))).#trainneg
cnf(i_0_19004, negated_conjecture, (r1_xxreal_0(np__1,esk4_0)|r1_xxreal_0(esk4_0,esk2_0)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1775, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_29721, plain, (epred1_5(X1,X2,X2,X3,k4_matrix14(X1,X4,X5))|np__1=X2|X2=X5|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X4)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_29830, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X3,X4),X5,X6)|np__1=X5|X5=X6|X5=X4|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,X1,X2)!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X6,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29774, negated_conjecture, (X1=np__1|epred1_5(esk1_0,X1,X2,X3,k4_matrix14(esk1_0,esk2_0,X2))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_29778, negated_conjecture, (X1=np__1|epred1_5(esk1_0,X2,X1,X3,k4_matrix14(esk1_0,esk2_0,X1))|k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk3_0,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29777, negated_conjecture, (X1=np__1|epred1_5(esk1_0,X2,X1,X3,k4_matrix14(esk1_0,esk2_0,X1))|k3_matrix_1(k5_numbers,esk4_0,esk3_0,X2)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(X2,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,esk3_0,X2)))).#trainneg
cnf(i_0_29828, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X1,X2),X3,X4)=k5_struct_0(esk1_0)|np__1=X3|X3=X4|X3=X2|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X4,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29722, plain, (X1=np__1|epred1_5(X2,X1,X1,X3,k4_matrix14(X2,X4,X5))|X1=X5|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_21087, negated_conjecture, (v3_xxreal_0(esk10_1(k5_numbers))|v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,esk10_1(k5_numbers)))).#trainneg
cnf(i_0_1777, plain, (k2_xcmplx_0(k4_xcmplx_0(X3),k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_29039, negated_conjecture, (esk6_0=np__1|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(esk6_0,X2)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_30176, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X2,X3),X4,X5)=k4_struct_0(esk1_0)|np__1=X4|X4=X3|X4=X5|np__1=X1|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X5,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0))).#trainneg
cnf(i_0_30189, negated_conjecture, (epred1_5(esk1_0,X1,esk3_0,X2,esk4_0)|k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1)!=k5_struct_0(esk1_0)|~m2_finseq_1(esk4_0,k13_finseq_1(k5_numbers))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(k3_matrix_1(k5_numbers,esk4_0,esk3_0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_30037, plain, (epred1_5(X1,X2,X3,X4,k4_matrix14(X1,X5,X6))|np__1=X3|X3=X6|X3=X2|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k4_struct_0(X1)!=k5_struct_0(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X5)|~r1_xxreal_0(X2,X5)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_28907, negated_conjecture, (r1_xxreal_0(X1,esk6_0)|~v1_xcmplx_0(X1)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_21088, negated_conjecture, (v3_xxreal_0(esk10_1(k5_numbers))|v2_xxreal_0(esk3_0)|r1_xxreal_0(esk3_0,esk10_1(k5_numbers)))).#trainneg
cnf(i_0_28521, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),esk3_0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1782, plain, (k2_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X2),X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5786, plain, (v7_ordinal1(k7_partfun1(k5_numbers,X1,X2))|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_21089, negated_conjecture, (v3_xxreal_0(esk10_1(k5_numbers))|v2_xxreal_0(esk2_0)|r1_xxreal_0(esk2_0,esk10_1(k5_numbers)))).#trainneg
cnf(i_0_1783, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=X1|~v1_xxreal_0(k2_xcmplx_0(k4_xcmplx_0(X2),X1))|~r1_xxreal_0(k2_xcmplx_0(k4_xcmplx_0(X2),X1),X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3884, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),esk6_0)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3885, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk6_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_30477, negated_conjecture, (esk6_0=np__1|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_30488, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,X1,X2),X3,X4)=k4_struct_0(esk1_0)|np__1=X3|X3=X4|X3=X2|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X4,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29867, negated_conjecture, (k1_funct_1(k3_matrix14(esk1_0,esk2_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)|np__1=X1|X1=esk3_0|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k1_funct_1(k3_matrix14(esk1_0,esk2_0,X1),X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3409, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(esk2_0)|~v1_xxreal_0(np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(np__2,esk6_0))).#trainneg
cnf(i_0_25273, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_21654, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),np__0))).#trainneg
cnf(i_0_10629, negated_conjecture, (~v1_xboole_0(X1)|X1!=k13_finseq_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1786, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11617, plain, (v1_finset_1(np__0)|~v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4910, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk3_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),np__0)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_28651, negated_conjecture, (X1=esk3_0|v1_xboole_0(u1_struct_0(X2))|m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|v6_struct_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(esk3_0,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1787, plain, (k2_xcmplx_0(X1,np__1)=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_30800, plain, (epred1_5(esk1_0,X1,X2,X3,k4_matrix14(esk1_0,X4,X5))|np__1=X2|X2=X5|X2=X1|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X4)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_30799, negated_conjecture, (k1_funct_1(k1_funct_1(k4_matrix14(esk1_0,X1,X2),X3),X4)=X5|np__1=X3|X3=X2|X3=X4|X5!=k4_struct_0(esk1_0)|~m2_finseq_1(k4_matrix14(esk1_0,X1,X2),k13_finseq_1(u1_struct_0(esk1_0)))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,X1)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X5,u1_struct_0(esk1_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_30848, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X1)=k5_struct_0(esk1_0)|X1=esk3_0|np__1=X1|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_30849, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)=k4_struct_0(esk1_0)|X1=esk3_0|np__1=X1|X1=X2|~m2_finseq_1(k3_matrix14(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_30850, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(np__2)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_30976, negated_conjecture, (v1_finset_1(np__0)|~v1_xboole_0(esk2_0)|~v1_finset_1(esk4_0))).#trainneg
cnf(i_0_30979, plain, (v1_finset_1(np__0)|~v1_xboole_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_30792, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),esk2_0)|~r1_xxreal_0(k4_xcmplx_0(np__2),esk6_0)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_28655, negated_conjecture, (X1=esk2_0|v1_xboole_0(u1_struct_0(X2))|m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|v6_struct_0(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(esk2_0,X3)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29866, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)=k1_funct_1(k1_funct_1(esk4_0,X1),X2)|np__1=X1|X1=esk3_0|~m2_finseq_1(k1_funct_1(esk4_0,X1),u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k1_funct_1(k1_funct_1(esk4_0,X1),X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_29849, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k1_funct_1(k3_matrix14(X2,X3,np__1),X5)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29831, plain, (k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)=k3_matrix_1(u1_struct_0(X1),k4_matrix14(X1,X5,X6),X7,X8)|np__1=X7|X7=X8|X7=X6|v1_xboole_0(u1_struct_0(X1))|v6_struct_0(X1)|k3_matrix_1(u1_struct_0(X1),k12_matrix_1(X1,X2),X3,X4)!=k4_struct_0(X1)|~r1_xxreal_0(np__1,X8)|~r1_xxreal_0(np__1,X7)|~r1_xxreal_0(X7,X5)|~r1_xxreal_0(X8,X5)|~v7_ordinal1(X8)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_27447, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X5,X6)|np__1=X5|X5=X4|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|X5=X7|X1!=k1_funct_1(k1_funct_1(k4_matrix14(X2,X3,X7),X5),X6)|~m2_finseq_1(k1_funct_1(k4_matrix14(X2,X3,X7),X5),u1_struct_0(X2))|~m2_finseq_1(k12_matrix_1(X2,X3),k13_finseq_1(u1_struct_0(X2)))|~r1_xxreal_0(X5,k3_finseq_1(k12_matrix_1(X2,X3)))|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X6,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X6)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2)|~v7_ordinal1(X7)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2328, negated_conjecture, (esk3_0=X1|~v1_xxreal_0(X1)|~v1_xxreal_0(esk3_0)|~r1_xxreal_0(esk3_0,X1)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3057, plain, (r1_xxreal_0(np__0,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__0,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21726, plain, (v3_xxreal_0(np__2)|v2_xxreal_0(esk10_1(k5_numbers))|r1_xxreal_0(esk10_1(k5_numbers),np__2))).#trainneg
cnf(i_0_31053, negated_conjecture, (esk3_0=esk6_0|k4_struct_0(esk1_0)!=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_31082, plain, (v1_finset_1(np__0))).#trainneg
cnf(i_0_31203, negated_conjecture, (r1_xxreal_0(np__0,k2_xcmplx_0(np__1,esk2_0)))).#trainneg
cnf(i_0_1808, plain, (r1_xxreal_0(k6_xcmplx_0(X1,X2),k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1810, plain, (r1_xxreal_0(k1_xreal_0(X1,X2),k1_xreal_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31070, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X1)=k5_struct_0(esk1_0)|np__1=X1|X1=esk3_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_31073, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k12_matrix_1(esk1_0,esk2_0),X1,X2)=k4_struct_0(esk1_0)|np__1=X1|X1=esk3_0|X1=X2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,esk2_0)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31055, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X2),esk3_0,X3)|esk3_0=X2|esk3_0=X3|X1!=k4_struct_0(esk1_0)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,esk2_0)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31149, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,X5)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|np__1=X5|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X5,X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_31148, plain, (X1=k3_matrix_1(u1_struct_0(X2),k4_matrix14(X2,X3,X4),X4,np__1)|v1_xboole_0(u1_struct_0(X2))|v6_struct_0(X2)|v2_struct_0(X2)|X1!=k5_struct_0(X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v5_group_1(X2)|~v3_group_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v33_algstr_0(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_31040, negated_conjecture, (k1_funct_1(k1_funct_1(esk4_0,np__1),X1)=k4_struct_0(esk1_0)|esk3_0=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,esk2_0)|~v7_ordinal1(X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31075, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(X1,np__2))).#trainneg
cnf(i_0_31223, plain, (r1_xxreal_0(np__0,k2_xcmplx_0(np__1,np__2))|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2412, negated_conjecture, (v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,esk3_0)|~v2_xxreal_0(esk3_0)|~v1_xxreal_0(esk3_0))).#trainneg
cnf(i_0_18065, plain, (k2_finseq_1(X2)!=k2_finseq_1(np__1)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1812, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31439, negated_conjecture, (r1_xxreal_0(X1,esk2_0)|~r1_xxreal_0(np__2,esk6_0)|~r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_31251, plain, (v7_ordinal1(k1_card_1(np__0)))).#trainneg
cnf(i_0_25, negated_conjecture, (esk6_0!=np__1)).#trainneg
cnf(i_0_9934, negated_conjecture, (m2_finseq_1(k1_funct_1(esk4_0,esk6_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4401, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),esk6_0)=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_11961, plain, (epred1_5(esk1_0,esk6_0,esk6_0,X1,k12_matrix_1(esk1_0,esk2_0)))).#trainneg
cnf(i_0_32, negated_conjecture, (k4_matrix14(esk1_0,esk2_0,esk6_0)=esk4_0)).#trainneg
cnf(i_0_11418, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),k4_matrix14(esk1_0,esk2_0,X1),esk6_0,esk6_0)=k5_struct_0(esk1_0)|esk6_0=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5832, negated_conjecture, (v1_funct_1(k1_funct_1(esk4_0,esk6_0)))).#trainneg
cnf(i_0_10768, plain, (epred1_5(esk1_0,np__1,esk6_0,X1,esk4_0))).#trainneg
cnf(i_0_31587, plain, (v7_ordinal1(k3_finseq_1(np__0)))).#trainneg
cnf(i_0_21656, negated_conjecture, (v3_xxreal_0(np__0)|v2_xxreal_0(esk6_0))).#trainneg
cnf(i_0_28334, negated_conjecture, (k4_struct_0(esk1_0)=k5_struct_0(esk1_0)|k1_funct_1(esk4_0,esk6_0)!=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_2987, negated_conjecture, (~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_21600, negated_conjecture, (~r1_xxreal_0(esk6_0,np__0))).#trainneg
cnf(i_0_31128, negated_conjecture, (esk2_0=np__1|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v6_struct_0(X1)|~r1_xxreal_0(esk2_0,X2)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3937, negated_conjecture, (k1_funct_1(k12_matrix_1(esk1_0,esk2_0),np__1)=k1_funct_1(esk4_0,esk6_0))).#trainneg
cnf(i_0_5845, negated_conjecture, (v1_finseq_1(k1_funct_1(esk4_0,esk6_0)))).#trainneg
cnf(i_0_31771, plain, (epred1_5(esk1_0,esk6_0,esk6_0,X1,k4_matrix14(esk1_0,esk2_0,X2))|esk6_0=X2|~v7_ordinal1(X2))).#trainneg
cnf(i_0_31760, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,esk6_0,np__1)=k5_struct_0(esk1_0))).#trainneg
cnf(i_0_31630, negated_conjecture, (k3_matrix14(esk1_0,esk2_0,esk6_0)=k1_funct_1(esk4_0,np__1))).#trainneg
cnf(i_0_10907, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk4_0,np__1,esk6_0)=k5_struct_0(esk1_0))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 225
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 378
# Removed in clause preprocessing      : 43
# Initial clauses in saturation        : 335
# Processed clauses                    : 4737
# ...of these trivial                  : 97
# ...subsumed                          : 1910
# ...remaining for further processing  : 2730
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 418
# Backward-rewritten                   : 597
# Generated clauses                    : 31589
# ...of the previous two non-trivial   : 29946
# Contextual simplify-reflections      : 587
# Paramodulations                      : 31479
# Factorizations                       : 1
# Equation resolutions                 : 98
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1704
#    Positive orientable unit clauses  : 151
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 23
#    Non-unit-clauses                  : 1530
# Current number of unprocessed clauses: 24585
# ...number of literals in the above   : 281062
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3758
# Proof object given clauses           : 266
# Proof search given clauses           : 2730
# Clause-clause subsumption calls (NU) : 4351956
# Rec. Clause-clause subsumption calls : 133912
# Non-unit clause-clause subsumptions  : 2640
# Unit Clause-clause subsumption calls : 18639
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 110
# BW rewrite match successes           : 31
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1328031

# -------------------------------------------------
# User time                : 58.137 s
# System time              : 1.213 s
# Total time               : 59.350 s
# Maximum resident set size: 522580 pages
58.13user 1.31system 1:00.43elapsed 98%CPU (0avgtext+0avgdata 523376maxresident)k
0inputs+1648outputs (0major+438027minor)pagefaults 0swaps
