2020-08-31 10:32:55.351382: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_37
2020-08-31 10:32:55.443751: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-08-31 10:32:55.551954: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-08-31 10:32:55.879975: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-08-31 10:32:57.313401: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1962026 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_37' loaded (query=256; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.530 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_k1_seq_1)).
fof(redefinition_k8_nat_1, axiom, ![X1, X2, X3]:((((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))&v7_ordinal1(X3))=>k8_nat_1(X1,X2,X3)=k1_funct_1(X2,X3)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_k8_nat_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', cc1_relset_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', cc8_ordinal1)).
fof(free_g1_metric_1, axiom, ![X1, X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k2_zfmisc_1(X1,X1),k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))=>![X3, X4]:(g1_metric_1(X1,X2)=g1_metric_1(X3,X4)=>(X1=X3&X2=X4))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', free_g1_metric_1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_k5_numbers)).
fof(dt_k7_metric_1, axiom, ((v1_funct_1(k7_metric_1)&v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))&m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_k7_metric_1)).
fof(d13_metric_1, axiom, k8_metric_1=g1_metric_1(k1_numbers,k7_metric_1), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d13_metric_1)).
fof(abstractness_v1_metric_1, axiom, ![X1]:(l1_metric_1(X1)=>(v1_metric_1(X1)=>X1=g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', abstractness_v1_metric_1)).
fof(t4_topmetr3, conjecture, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(k8_metric_1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1)))))=>(r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),X1,X2)=>(((v2_comseq_2(X1)=>v2_tbsp_1(X2,k8_metric_1))&(v2_tbsp_1(X2,k8_metric_1)=>v2_comseq_2(X1)))&(v2_comseq_2(X1)=>k2_seq_2(X1)=k1_tbsp_1(k8_metric_1,X2)))))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', t4_topmetr3)).
fof(d7_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_comseq_2(X1)=>![X2]:(v1_xreal_0(X2)=>(X2=k1_seq_2(X1)<=>![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>?[X5]:((m2_subset_1(X5,k1_numbers,k5_numbers)&r1_xxreal_0(X4,X5))&r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X1,X5),X2)))))))))))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d7_seq_2)).
fof(d12_metric_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))=>(X1=k7_metric_1<=>![X2]:(m1_subset_1(X2,k1_numbers)=>![X3]:(m1_subset_1(X3,k1_numbers)=>k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k18_complex1(k9_real_1(X2,X3)))))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d12_metric_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', t1_subset)).
fof(d1_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)<=>r2_hidden(X1,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d1_xreal_0)).
fof(d1_metric_1, axiom, ![X1]:(l1_metric_1(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_metric_1(X1,X2,X3)=k1_metric_1(u1_struct_0(X1),u1_struct_0(X1),u1_metric_1(X1),X2,X3)))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d1_metric_1)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', cc18_valued_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', reflexivity_r1_tarski)).
fof(dt_k8_metric_1, axiom, (v1_metric_1(k8_metric_1)&l1_metric_1(k8_metric_1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_k8_metric_1)).
fof(d2_tbsp_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_metric_1(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>(v2_tbsp_1(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&![X4]:(m1_subset_1(X4,k1_numbers)=>~((~(r1_xxreal_0(X4,k6_numbers))&![X5]:(m2_subset_1(X5,k1_numbers,k5_numbers)=>?[X6]:((m2_subset_1(X6,k1_numbers,k5_numbers)&r1_xxreal_0(X5,X6))&r1_xxreal_0(X4,k2_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,X6),X3))))))))))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d2_tbsp_1)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', fc3_membered)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_m2_subset_1)).
fof(fc63_valued_0, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>v1_xreal_0(k1_funct_1(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', fc63_valued_0)).
fof(redefinition_k4_metric_1, axiom, ![X1, X2, X3]:((((v8_metric_1(X1)&l1_metric_1(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k4_metric_1(X1,X2,X3)=k2_metric_1(X1,X2,X3)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_k4_metric_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', fc1_numbers)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', fc2_struct_0)).
fof(dt_l1_metric_1, axiom, ![X1]:(l1_metric_1(X1)=>l1_struct_0(X1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_l1_metric_1)).
fof(redefinition_r1_funct_2, axiom, ![X1, X2, X3, X4, X5, X6]:((((((((~(v1_xboole_0(X2))&~(v1_xboole_0(X4)))&v1_funct_1(X5))&v1_funct_2(X5,X1,X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&v1_funct_2(X6,X3,X4))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(r1_funct_2(X1,X2,X3,X4,X5,X6)<=>X5=X6)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_r1_funct_2)).
fof(fc5_metric_1, axiom, (~(v2_struct_0(k8_metric_1))&v1_metric_1(k8_metric_1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', fc5_metric_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', fc6_ordinal1)).
fof(commutativity_k4_metric_1, axiom, ![X1, X2, X3]:((((v8_metric_1(X1)&l1_metric_1(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', commutativity_k4_metric_1)).
fof(d3_tbsp_1, axiom, ![X1]:((((((~(v2_struct_0(X1))&v6_metric_1(X1))&v7_metric_1(X1))&v8_metric_1(X1))&v9_metric_1(X1))&l1_metric_1(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>(v2_tbsp_1(X2,X1)=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k1_tbsp_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_numbers)=>~((~(r1_xxreal_0(X4,k6_numbers))&![X5]:(m2_subset_1(X5,k1_numbers,k5_numbers)=>?[X6]:((m2_subset_1(X6,k1_numbers,k5_numbers)&r1_xxreal_0(X5,X6))&r1_xxreal_0(X4,k4_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,X6),X3)))))))))))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d3_tbsp_1)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_k5_numbers)).
fof(dt_k8_nat_1, axiom, ![X1, X2, X3]:((((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))&v7_ordinal1(X3))=>m1_subset_1(k8_nat_1(X1,X2,X3),X1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_k8_nat_1)).
fof(redefinition_k1_metric_1, axiom, ![X1, X2, X3, X4, X5]:((((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers))))&m1_subset_1(X4,X1))&m1_subset_1(X5,X2))=>k1_metric_1(X1,X2,X3,X4,X5)=k1_binop_1(X3,X4,X5)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_k1_metric_1)).
fof(dt_u1_metric_1, axiom, ![X1]:(l1_metric_1(X1)=>((v1_funct_1(u1_metric_1(X1))&v1_funct_2(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers))&m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers))))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_u1_metric_1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', cc1_xreal_0)).
fof(dt_k2_tbsp_1, axiom, ![X1, X2, X3]:((((((~(v2_struct_0(X1))&l1_metric_1(X1))&v1_funct_1(X2))&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))&m1_subset_1(X3,k5_numbers))=>m1_subset_1(k2_tbsp_1(X1,X2,X3),u1_struct_0(X1))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_k2_tbsp_1)).
fof(redefinition_k2_tbsp_1, axiom, ![X1, X2, X3]:((((((~(v2_struct_0(X1))&l1_metric_1(X1))&v1_funct_1(X2))&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))&m1_subset_1(X3,k5_numbers))=>k2_tbsp_1(X1,X2,X3)=k1_funct_1(X2,X3)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_k2_tbsp_1)).
fof(redefinition_k2_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>k2_seq_2(X1)=k1_seq_2(X1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', redefinition_k2_seq_2)).
fof(dt_k1_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>v1_xreal_0(k1_seq_2(X1))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', dt_k1_seq_2)).
fof(d6_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_comseq_2(X1)<=>?[X2]:(v1_xreal_0(X2)&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>?[X5]:((m2_subset_1(X5,k1_numbers,k5_numbers)&r1_xxreal_0(X4,X5))&r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X1,X5),X2))))))))))), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', d6_seq_2)).
fof(fc6_metric_1, axiom, ((((v1_metric_1(k8_metric_1)&v6_metric_1(k8_metric_1))&v7_metric_1(k8_metric_1))&v8_metric_1(k8_metric_1))&v9_metric_1(k8_metric_1)), file('/scratch/mptp/hardpred/preds__0.005//t4_topmetr3', fc6_metric_1)).
fof(c_0_43, plain, ![X92, X93]:(~v1_relat_1(X92)|~v1_funct_1(X92)|~v3_valued_0(X92)|k1_seq_1(X92,X93)=k1_funct_1(X92,X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
fof(c_0_44, plain, ![X47, X48, X49]:(~v1_funct_1(X48)|~v1_funct_2(X48,k5_numbers,X47)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X47)))|~v7_ordinal1(X49)|k8_nat_1(X47,X48,X49)=k1_funct_1(X48,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_nat_1])])).
cnf(c_0_45, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_46, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_47, plain, ![X111, X112, X113]:(~m1_subset_1(X113,k1_zfmisc_1(k2_zfmisc_1(X111,X112)))|v1_relat_1(X113)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_48, plain, ![X140]:(~m1_subset_1(X140,k4_ordinal1)|v7_ordinal1(X140)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_49, plain, ![X53, X54, X55, X56]:((X53=X55|g1_metric_1(X53,X54)!=g1_metric_1(X55,X56)|(~v1_funct_1(X54)|~v1_funct_2(X54,k2_zfmisc_1(X53,X53),k1_numbers)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X53,X53),k1_numbers)))))&(X54=X56|g1_metric_1(X53,X54)!=g1_metric_1(X55,X56)|(~v1_funct_1(X54)|~v1_funct_2(X54,k2_zfmisc_1(X53,X53),k1_numbers)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X53,X53),k1_numbers)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_metric_1])])])])).
cnf(c_0_50, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_55, plain, (X1=X2|g1_metric_1(X3,X1)!=g1_metric_1(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, plain, (m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k7_metric_1])).
cnf(c_0_57, plain, (k8_metric_1=g1_metric_1(k1_numbers,k7_metric_1)), inference(split_conjunct,[status(thm)],[d13_metric_1])).
cnf(c_0_58, plain, (v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)), inference(split_conjunct,[status(thm)],[dt_k7_metric_1])).
cnf(c_0_59, plain, (v1_funct_1(k7_metric_1)), inference(split_conjunct,[status(thm)],[dt_k7_metric_1])).
fof(c_0_60, plain, ![X99]:(~l1_metric_1(X99)|(~v1_metric_1(X99)|X99=g1_metric_1(u1_struct_0(X99),u1_metric_1(X99)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_metric_1])])).
cnf(c_0_61, plain, (X1=X2|g1_metric_1(X1,X3)!=g1_metric_1(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_62, plain, (k8_nat_1(X3,X1,X2)=k1_seq_1(X1,X2)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51])])).
cnf(c_0_63, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_54])).
fof(c_0_66, negated_conjecture, ~(![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(k8_metric_1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1)))))=>(r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),X1,X2)=>(((v2_comseq_2(X1)=>v2_tbsp_1(X2,k8_metric_1))&(v2_tbsp_1(X2,k8_metric_1)=>v2_comseq_2(X1)))&(v2_comseq_2(X1)=>k2_seq_2(X1)=k1_tbsp_1(k8_metric_1,X2))))))), inference(assume_negation,[status(cth)],[t4_topmetr3])).
fof(c_0_67, plain, ![X39, X40, X41, X43, X45]:(((m2_subset_1(esk12_3(X39,X40,X41),k1_numbers,k5_numbers)|r1_xxreal_0(X41,k6_numbers)|~v1_xreal_0(X41)|X40!=k1_seq_2(X39)|~v1_xreal_0(X40)|~v2_comseq_2(X39)|(~v1_funct_1(X39)|~v1_funct_2(X39,k5_numbers,k1_numbers)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~m2_subset_1(X43,k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X39,X40,X41),X43)|~r1_xxreal_0(X41,k18_complex1(k9_real_1(k1_seq_1(X39,X43),X40)))|r1_xxreal_0(X41,k6_numbers)|~v1_xreal_0(X41)|X40!=k1_seq_2(X39)|~v1_xreal_0(X40)|~v2_comseq_2(X39)|(~v1_funct_1(X39)|~v1_funct_2(X39,k5_numbers,k1_numbers)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&((v1_xreal_0(esk13_2(X39,X40))|X40=k1_seq_2(X39)|~v1_xreal_0(X40)|~v2_comseq_2(X39)|(~v1_funct_1(X39)|~v1_funct_2(X39,k5_numbers,k1_numbers)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((~r1_xxreal_0(esk13_2(X39,X40),k6_numbers)|X40=k1_seq_2(X39)|~v1_xreal_0(X40)|~v2_comseq_2(X39)|(~v1_funct_1(X39)|~v1_funct_2(X39,k5_numbers,k1_numbers)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(((m2_subset_1(esk14_3(X39,X40,X45),k1_numbers,k5_numbers)|~m2_subset_1(X45,k1_numbers,k5_numbers)|X40=k1_seq_2(X39)|~v1_xreal_0(X40)|~v2_comseq_2(X39)|(~v1_funct_1(X39)|~v1_funct_2(X39,k5_numbers,k1_numbers)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(r1_xxreal_0(X45,esk14_3(X39,X40,X45))|~m2_subset_1(X45,k1_numbers,k5_numbers)|X40=k1_seq_2(X39)|~v1_xreal_0(X40)|~v2_comseq_2(X39)|(~v1_funct_1(X39)|~v1_funct_2(X39,k5_numbers,k1_numbers)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(r1_xxreal_0(esk13_2(X39,X40),k18_complex1(k9_real_1(k1_seq_1(X39,esk14_3(X39,X40,X45)),X40)))|~m2_subset_1(X45,k1_numbers,k5_numbers)|X40=k1_seq_2(X39)|~v1_xreal_0(X40)|~v2_comseq_2(X39)|(~v1_funct_1(X39)|~v1_funct_2(X39,k5_numbers,k1_numbers)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_seq_2])])])])])])).
fof(c_0_68, plain, ![X34, X35, X36]:((X34!=k7_metric_1|(~m1_subset_1(X35,k1_numbers)|(~m1_subset_1(X36,k1_numbers)|k1_metric_1(k1_numbers,k1_numbers,X34,X35,X36)=k18_complex1(k9_real_1(X35,X36))))|(~v1_funct_1(X34)|~v1_funct_2(X34,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))))&((m1_subset_1(esk10_1(X34),k1_numbers)|X34=k7_metric_1|(~v1_funct_1(X34)|~v1_funct_2(X34,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))))&((m1_subset_1(esk11_1(X34),k1_numbers)|X34=k7_metric_1|(~v1_funct_1(X34)|~v1_funct_2(X34,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))))&(k1_metric_1(k1_numbers,k1_numbers,X34,esk10_1(X34),esk11_1(X34))!=k18_complex1(k9_real_1(esk10_1(X34),esk11_1(X34)))|X34=k7_metric_1|(~v1_funct_1(X34)|~v1_funct_2(X34,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_metric_1])])])])])).
fof(c_0_69, plain, ![X152, X153]:(~r2_hidden(X152,X153)|m1_subset_1(X152,X153)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_70, plain, ![X86]:((~v1_xreal_0(X86)|r2_hidden(X86,k1_numbers))&(~r2_hidden(X86,k1_numbers)|v1_xreal_0(X86))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_xreal_0])])).
fof(c_0_71, plain, ![X66, X67, X68]:(~l1_metric_1(X66)|(~m1_subset_1(X67,u1_struct_0(X66))|(~m1_subset_1(X68,u1_struct_0(X66))|k2_metric_1(X66,X67,X68)=k1_metric_1(u1_struct_0(X66),u1_struct_0(X66),u1_metric_1(X66),X67,X68)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_metric_1])])])).
cnf(c_0_72, plain, (X1=X2|g1_metric_1(X3,X1)!=g1_metric_1(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),k1_numbers)))), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_73, plain, (m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_74, plain, (k8_metric_1=g1_metric_1(k1_numbers,k7_metric_1)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_75, plain, (v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_76, plain, (v1_funct_1(k7_metric_1)), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_77, plain, (X1=g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))|~l1_metric_1(X1)|~v1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_78, plain, (X1=X2|g1_metric_1(X1,X3)!=g1_metric_1(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_79, plain, (k8_nat_1(X1,X2,X3)=k1_seq_1(X2,X3)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)), inference(csr,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_80, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_64, c_0_65])).
fof(c_0_81, plain, ![X129, X130, X131]:(~v3_membered(X130)|(~m1_subset_1(X131,k1_zfmisc_1(k2_zfmisc_1(X129,X130)))|v3_valued_0(X131))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
fof(c_0_82, negated_conjecture, (((v1_funct_1(esk1_0)&v1_funct_2(esk1_0,k5_numbers,k1_numbers))&m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))&(((v1_funct_1(esk2_0)&v1_funct_2(esk2_0,k5_numbers,u1_struct_0(k8_metric_1)))&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1)))))&(r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),esk1_0,esk2_0)&((((v2_comseq_2(esk1_0)|(v2_tbsp_1(esk2_0,k8_metric_1)|v2_comseq_2(esk1_0)))&(k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|(v2_tbsp_1(esk2_0,k8_metric_1)|v2_comseq_2(esk1_0))))&((v2_comseq_2(esk1_0)|(~v2_comseq_2(esk1_0)|v2_comseq_2(esk1_0)))&(k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|(~v2_comseq_2(esk1_0)|v2_comseq_2(esk1_0)))))&(((v2_comseq_2(esk1_0)|(v2_tbsp_1(esk2_0,k8_metric_1)|~v2_tbsp_1(esk2_0,k8_metric_1)))&(k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|(v2_tbsp_1(esk2_0,k8_metric_1)|~v2_tbsp_1(esk2_0,k8_metric_1))))&((v2_comseq_2(esk1_0)|(~v2_comseq_2(esk1_0)|~v2_tbsp_1(esk2_0,k8_metric_1)))&(k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|(~v2_comseq_2(esk1_0)|~v2_tbsp_1(esk2_0,k8_metric_1))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_66])])])])).
fof(c_0_83, plain, ![X119, X120]:((~m1_subset_1(X119,k1_zfmisc_1(X120))|r1_tarski(X119,X120))&(~r1_tarski(X119,X120)|m1_subset_1(X119,k1_zfmisc_1(X120)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_84, plain, ![X150]:r1_tarski(X150,X150), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_85, plain, (r1_xxreal_0(X4,k6_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X2,X3,X4),X1)|~r1_xxreal_0(X4,k18_complex1(k9_real_1(k1_seq_1(X2,X1),X3)))|~v1_xreal_0(X4)|X3!=k1_seq_2(X2)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_86, plain, (k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k18_complex1(k9_real_1(X2,X3))|X1!=k7_metric_1|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_87, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_88, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_89, plain, (k2_metric_1(X1,X2,X3)=k1_metric_1(u1_struct_0(X1),u1_struct_0(X1),u1_metric_1(X1),X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_90, plain, (k7_metric_1=X1|g1_metric_1(X2,X1)!=k8_metric_1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74]), c_0_75]), c_0_76])])])).
cnf(c_0_91, plain, (X1=g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))|~l1_metric_1(X1)|~v1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_92, plain, (k1_numbers=X1|g1_metric_1(X1,X2)!=k8_metric_1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_73]), c_0_74]), c_0_75]), c_0_76])])])).
cnf(c_0_93, plain, (v1_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[dt_k8_metric_1])).
cnf(c_0_94, plain, (l1_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[dt_k8_metric_1])).
fof(c_0_95, plain, ![X18, X19, X21, X23, X24, X26]:(((m1_subset_1(esk6_2(X18,X19),u1_struct_0(X18))|~v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18)))&((m2_subset_1(esk7_3(X18,X19,X21),k1_numbers,k5_numbers)|r1_xxreal_0(X21,k6_numbers)|~m1_subset_1(X21,k1_numbers)|~v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18)))&(~m2_subset_1(X23,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X18,X19,X21),X23)|~r1_xxreal_0(X21,k2_metric_1(X18,k8_nat_1(u1_struct_0(X18),X19,X23),esk6_2(X18,X19)))|r1_xxreal_0(X21,k6_numbers)|~m1_subset_1(X21,k1_numbers)|~v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18)))))&((m1_subset_1(esk8_3(X18,X19,X24),k1_numbers)|~m1_subset_1(X24,u1_struct_0(X18))|v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18)))&((~r1_xxreal_0(esk8_3(X18,X19,X24),k6_numbers)|~m1_subset_1(X24,u1_struct_0(X18))|v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18)))&(((m2_subset_1(esk9_4(X18,X19,X24,X26),k1_numbers,k5_numbers)|~m2_subset_1(X26,k1_numbers,k5_numbers)|~m1_subset_1(X24,u1_struct_0(X18))|v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18)))&(r1_xxreal_0(X26,esk9_4(X18,X19,X24,X26))|~m2_subset_1(X26,k1_numbers,k5_numbers)|~m1_subset_1(X24,u1_struct_0(X18))|v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18))))&(r1_xxreal_0(esk8_3(X18,X19,X24),k2_metric_1(X18,k8_nat_1(u1_struct_0(X18),X19,esk9_4(X18,X19,X24,X26)),X24))|~m2_subset_1(X26,k1_numbers,k5_numbers)|~m1_subset_1(X24,u1_struct_0(X18))|v2_tbsp_1(X19,X18)|(~v1_funct_1(X19)|~v1_funct_2(X19,k5_numbers,u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X18)))))|(v2_struct_0(X18)|~l1_metric_1(X18))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_tbsp_1])])])])])])).
cnf(c_0_96, plain, (k8_nat_1(X1,X2,X3)=k1_seq_1(X2,X3)|~v3_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80])])).
cnf(c_0_97, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_98, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_99, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_100, plain, ![X146, X147, X148]:((~m2_subset_1(X148,X146,X147)|m1_subset_1(X148,X147)|(v1_xboole_0(X146)|v1_xboole_0(X147)|~m1_subset_1(X147,k1_zfmisc_1(X146))))&(~m1_subset_1(X148,X147)|m2_subset_1(X148,X146,X147)|(v1_xboole_0(X146)|v1_xboole_0(X147)|~m1_subset_1(X147,k1_zfmisc_1(X146))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_101, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_102, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_103, plain, ![X163, X164]:(~v1_relat_1(X163)|~v1_funct_1(X163)|~v3_valued_0(X163)|v1_xreal_0(k1_funct_1(X163,X164))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc63_valued_0])])).
fof(c_0_104, plain, ![X94, X95, X96]:(~v8_metric_1(X94)|~l1_metric_1(X94)|~m1_subset_1(X95,u1_struct_0(X94))|~m1_subset_1(X96,u1_struct_0(X94))|k4_metric_1(X94,X95,X96)=k2_metric_1(X94,X95,X96)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_metric_1])])).
cnf(c_0_105, plain, (r1_xxreal_0(X4,k6_numbers)|X3!=k1_seq_2(X2)|~v1_funct_1(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X2,X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X4,k18_complex1(k9_real_1(k1_seq_1(X2,X1),X3)))), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_106, plain, (k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k18_complex1(k9_real_1(X2,X3))|X1!=k7_metric_1|~v1_funct_1(X1)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_107, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_108, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_109, plain, (k2_metric_1(X1,X2,X3)=k1_metric_1(u1_struct_0(X1),u1_struct_0(X1),u1_metric_1(X1),X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_110, plain, (u1_metric_1(X1)=k7_metric_1|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
cnf(c_0_111, plain, (u1_struct_0(X1)=k1_numbers|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_91])])).
cnf(c_0_112, plain, (v1_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_113, plain, (l1_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_114, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_tbsp_1(X2,X1)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_115, plain, (k1_funct_1(X2,X3)=k1_seq_1(X2,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(X3,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_96])])).
cnf(c_0_116, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_117, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_118, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_119, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_120, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_121, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_122, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_123, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_124, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
fof(c_0_125, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
fof(c_0_126, plain, ![X161]:(v2_struct_0(X161)|~l1_struct_0(X161)|~v1_xboole_0(u1_struct_0(X161))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_127, plain, ![X104]:(~l1_metric_1(X104)|l1_struct_0(X104)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_metric_1])])).
fof(c_0_128, plain, ![X28, X29, X30, X31, X32, X33]:((~r1_funct_2(X28,X29,X30,X31,X32,X33)|X32=X33|(v1_xboole_0(X29)|v1_xboole_0(X31)|~v1_funct_1(X32)|~v1_funct_2(X32,X28,X29)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|~v1_funct_1(X33)|~v1_funct_2(X33,X30,X31)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))))&(X32!=X33|r1_funct_2(X28,X29,X30,X31,X32,X33)|(v1_xboole_0(X29)|v1_xboole_0(X31)|~v1_funct_1(X32)|~v1_funct_2(X32,X28,X29)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|~v1_funct_1(X33)|~v1_funct_2(X33,X30,X31)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r1_funct_2])])])])).
cnf(c_0_129, plain, (r1_xxreal_0(X4,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X2,X3,X4),X1)|~r1_xxreal_0(X4,k2_metric_1(X2,k8_nat_1(u1_struct_0(X2),X3,X1),esk6_2(X2,X3)))|~m1_subset_1(X4,k1_numbers)|~v2_tbsp_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~l1_metric_1(X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_130, plain, (k4_metric_1(X1,X2,X3)=k2_metric_1(X1,X2,X3)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_131, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(k1_numbers,k1_numbers,X5,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X5!=k7_metric_1|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v1_funct_2(X5,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106])])).
cnf(c_0_132, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108])])).
cnf(c_0_133, plain, (k1_metric_1(u1_struct_0(X1),u1_struct_0(X1),k7_metric_1,X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_134, plain, (u1_struct_0(k8_metric_1)=k1_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113])])])).
cnf(c_0_135, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_136, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_115, c_0_80])).
cnf(c_0_137, negated_conjecture, (v3_valued_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])])])).
cnf(c_0_138, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_139, negated_conjecture, (v1_funct_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_140, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_141, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123])])).
cnf(c_0_142, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_143, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_144, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_145, plain, (l1_struct_0(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_146, plain, (X5=X6|v1_xboole_0(X2)|v1_xboole_0(X4)|~r1_funct_2(X1,X2,X3,X4,X5,X6)|~v1_funct_1(X5)|~v1_funct_2(X5,X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_147, negated_conjecture, (r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_148, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1))))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_149, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,u1_struct_0(k8_metric_1))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_150, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_82])).
fof(c_0_151, plain, (~v2_struct_0(k8_metric_1)&v1_metric_1(k8_metric_1)), inference(fof_simplification,[status(thm)],[fc5_metric_1])).
fof(c_0_152, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_153, plain, (v2_struct_0(X2)|r1_xxreal_0(X4,k6_numbers)|~v1_funct_1(X3)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X4,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X2,X3,X4),X1)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~r1_xxreal_0(X4,k2_metric_1(X2,k8_nat_1(u1_struct_0(X2),X3,X1),esk6_2(X2,X3)))), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_154, plain, (k2_metric_1(X1,X2,X3)=k4_metric_1(X1,X2,X3)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_155, plain, (m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
fof(c_0_156, plain, ![X100, X101, X102]:(~v8_metric_1(X100)|~l1_metric_1(X100)|~m1_subset_1(X101,u1_struct_0(X100))|~m1_subset_1(X102,u1_struct_0(X100))|k4_metric_1(X100,X101,X102)=k4_metric_1(X100,X102,X101)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k4_metric_1])])).
cnf(c_0_157, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|X4!=k7_metric_1|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(k1_numbers,k1_numbers,X4,k1_seq_1(X3,X5),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X5)|~v2_comseq_2(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X5),k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)), inference(csr,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_158, plain, (k1_metric_1(k1_numbers,k1_numbers,k7_metric_1,X1,X2)=k2_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_112]), c_0_113])])])).
cnf(c_0_159, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k1_funct_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,X2,X3,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_51])])).
cnf(c_0_160, negated_conjecture, (k1_funct_1(esk1_0,X1)=k1_seq_1(esk1_0,X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_117]), c_0_137]), c_0_138]), c_0_139])])])).
cnf(c_0_161, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X1,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141])])).
cnf(c_0_162, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_142])])).
cnf(c_0_163, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_164, negated_conjecture, (v2_comseq_2(esk1_0)|v2_tbsp_1(esk2_0,k8_metric_1)|v2_comseq_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_165, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_166, plain, (l1_struct_0(X1)|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_167, plain, (X5=X6|v1_xboole_0(X4)|v1_xboole_0(X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_funct_2(X6,X3,X4)|~v1_funct_2(X5,X1,X2)|~r1_funct_2(X1,X2,X3,X4,X5,X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_168, negated_conjecture, (r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_169, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1))))), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_170, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,u1_struct_0(k8_metric_1))), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_171, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_172, plain, (~v2_struct_0(k8_metric_1)), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_173, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_152])).
cnf(c_0_174, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v8_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_175, plain, (v2_struct_0(X1)|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_176, plain, (k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_156])).
fof(c_0_177, plain, ![X9, X10, X11, X12, X14, X16]:(((m2_subset_1(esk3_4(X9,X10,X11,X12),k1_numbers,k5_numbers)|r1_xxreal_0(X12,k6_numbers)|~m1_subset_1(X12,k1_numbers)|X11!=k1_tbsp_1(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~v2_tbsp_1(X10,X9)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,u1_struct_0(X9))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X9)))))|(v2_struct_0(X9)|~v6_metric_1(X9)|~v7_metric_1(X9)|~v8_metric_1(X9)|~v9_metric_1(X9)|~l1_metric_1(X9)))&(~m2_subset_1(X14,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X9,X10,X11,X12),X14)|~r1_xxreal_0(X12,k4_metric_1(X9,k8_nat_1(u1_struct_0(X9),X10,X14),X11))|r1_xxreal_0(X12,k6_numbers)|~m1_subset_1(X12,k1_numbers)|X11!=k1_tbsp_1(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~v2_tbsp_1(X10,X9)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,u1_struct_0(X9))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X9)))))|(v2_struct_0(X9)|~v6_metric_1(X9)|~v7_metric_1(X9)|~v8_metric_1(X9)|~v9_metric_1(X9)|~l1_metric_1(X9))))&((m1_subset_1(esk4_3(X9,X10,X11),k1_numbers)|X11=k1_tbsp_1(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~v2_tbsp_1(X10,X9)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,u1_struct_0(X9))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X9)))))|(v2_struct_0(X9)|~v6_metric_1(X9)|~v7_metric_1(X9)|~v8_metric_1(X9)|~v9_metric_1(X9)|~l1_metric_1(X9)))&((~r1_xxreal_0(esk4_3(X9,X10,X11),k6_numbers)|X11=k1_tbsp_1(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~v2_tbsp_1(X10,X9)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,u1_struct_0(X9))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X9)))))|(v2_struct_0(X9)|~v6_metric_1(X9)|~v7_metric_1(X9)|~v8_metric_1(X9)|~v9_metric_1(X9)|~l1_metric_1(X9)))&(((m2_subset_1(esk5_4(X9,X10,X11,X16),k1_numbers,k5_numbers)|~m2_subset_1(X16,k1_numbers,k5_numbers)|X11=k1_tbsp_1(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~v2_tbsp_1(X10,X9)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,u1_struct_0(X9))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X9)))))|(v2_struct_0(X9)|~v6_metric_1(X9)|~v7_metric_1(X9)|~v8_metric_1(X9)|~v9_metric_1(X9)|~l1_metric_1(X9)))&(r1_xxreal_0(X16,esk5_4(X9,X10,X11,X16))|~m2_subset_1(X16,k1_numbers,k5_numbers)|X11=k1_tbsp_1(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~v2_tbsp_1(X10,X9)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,u1_struct_0(X9))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X9)))))|(v2_struct_0(X9)|~v6_metric_1(X9)|~v7_metric_1(X9)|~v8_metric_1(X9)|~v9_metric_1(X9)|~l1_metric_1(X9))))&(r1_xxreal_0(esk4_3(X9,X10,X11),k4_metric_1(X9,k8_nat_1(u1_struct_0(X9),X10,esk5_4(X9,X10,X11,X16)),X11))|~m2_subset_1(X16,k1_numbers,k5_numbers)|X11=k1_tbsp_1(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~v2_tbsp_1(X10,X9)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,u1_struct_0(X9))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X9)))))|(v2_struct_0(X9)|~v6_metric_1(X9)|~v7_metric_1(X9)|~v8_metric_1(X9)|~v9_metric_1(X9)|~l1_metric_1(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_tbsp_1])])])])])])).
cnf(c_0_178, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_73]), c_0_75]), c_0_76])])])).
cnf(c_0_179, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k2_metric_1(X1,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),X2))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,X2,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk9_4(X1,esk1_0,X2,X3),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_139])])])).
cnf(c_0_180, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_181, plain, (m2_subset_1(k1_funct_1(X1,X2),k1_numbers,k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162]), c_0_163])])).
cnf(c_0_182, negated_conjecture, (v1_relat_1(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_117])])).
cnf(c_0_183, negated_conjecture, (v2_comseq_2(esk1_0)|v2_tbsp_1(esk2_0,k8_metric_1)), inference(cn,[status(thm)],[c_0_164])).
cnf(c_0_184, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166])])).
cnf(c_0_185, negated_conjecture, (esk2_0=esk1_0|v1_xboole_0(u1_struct_0(k8_metric_1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_169]), c_0_117]), c_0_170]), c_0_138]), c_0_171]), c_0_139])]), c_0_163])])).
cnf(c_0_186, plain, (~v2_struct_0(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_187, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_188, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_189, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)), inference(csr,[status(thm)],[c_0_174, c_0_175])).
cnf(c_0_190, plain, (k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_191, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k4_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,X3,X4)),X3))|X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
cnf(c_0_192, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_178, c_0_132])).
cnf(c_0_193, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k2_metric_1(X1,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),X2))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk9_4(X1,esk1_0,X2,X3),k5_numbers)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_179, c_0_80])).
cnf(c_0_194, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_195, negated_conjecture, (m2_subset_1(k1_seq_1(esk1_0,X1),k1_numbers,k1_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_160]), c_0_137]), c_0_182]), c_0_139])])])).
cnf(c_0_196, negated_conjecture, (v2_comseq_2(esk1_0)|v2_tbsp_1(esk2_0,k8_metric_1)), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_197, negated_conjecture, (esk2_0=esk1_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_185]), c_0_113])]), c_0_186])])).
cnf(c_0_198, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_187])).
cnf(c_0_199, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_188, c_0_65])).
cnf(c_0_200, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190])])).
cnf(c_0_201, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|r1_xxreal_0(esk4_3(X1,X2,X3),k4_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,X3,X4)),X3))|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_202, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),esk9_4(k8_metric_1,esk1_0,X1,X2))|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193]), c_0_117]), c_0_138]), c_0_139]), c_0_113]), c_0_134]), c_0_117]), c_0_134]), c_0_134]), c_0_138])]), c_0_186])])).
cnf(c_0_203, negated_conjecture, (m1_subset_1(k1_seq_1(esk1_0,X1),k1_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_195]), c_0_141])]), c_0_163])])).
cnf(c_0_204, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|v2_comseq_2(esk1_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_196, c_0_197])])).
cnf(c_0_205, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_198]), c_0_199]), c_0_163])])).
cnf(c_0_206, plain, (r1_xxreal_0(X1,esk9_4(X2,X3,X4,X1))|v2_tbsp_1(X3,X2)|v2_struct_0(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~l1_metric_1(X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_207, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)), inference(csr,[status(thm)],[c_0_200, c_0_175])).
cnf(c_0_208, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,X1,k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4))))|v2_struct_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4)),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_190])])).
cnf(c_0_209, plain, (r1_xxreal_0(esk13_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk14_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_210, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),esk9_4(k8_metric_1,esk1_0,X1,X2))|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_202, c_0_132]), c_0_203]), c_0_204]), c_0_205])).
cnf(c_0_211, plain, (v2_struct_0(X2)|v2_tbsp_1(X3,X2)|r1_xxreal_0(X1,esk9_4(X2,X3,X4,X1))|~v1_funct_1(X3)|~l1_metric_1(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_212, plain, (m2_subset_1(esk9_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|v2_tbsp_1(X2,X1)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_213, plain, (r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),esk5_4(X1,X2,esk6_2(X1,X2),X3))|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),X3)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v9_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208])])).
cnf(c_0_214, plain, (r1_xxreal_0(X1,esk5_4(X2,X3,X4,X1))|X4=k1_tbsp_1(X2,X3)|v2_struct_0(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~v2_tbsp_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)), inference(split_conjunct,[status(thm)],[c_0_177])).
fof(c_0_215, plain, ![X75, X76, X77]:(~v1_funct_1(X76)|~v1_funct_2(X76,k5_numbers,X75)|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X75)))|~v7_ordinal1(X77)|m1_subset_1(k8_nat_1(X75,X76,X77),X75)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_nat_1])])).
cnf(c_0_216, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(esk13_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk14_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_217, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1))),k5_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_113]), c_0_134]), c_0_117]), c_0_134]), c_0_134]), c_0_138]), c_0_139])]), c_0_186])])).
cnf(c_0_218, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|m2_subset_1(esk9_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_219, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),X3),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),esk5_4(X1,X2,esk6_2(X1,X2),X3))|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),X3)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(csr,[status(thm)],[c_0_213, c_0_175])).
cnf(c_0_220, plain, (X4=k1_tbsp_1(X2,X3)|v2_struct_0(X2)|r1_xxreal_0(X1,esk5_4(X2,X3,X4,X1))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_221, plain, (m2_subset_1(esk7_3(X1,X2,X3),k1_numbers,k5_numbers)|r1_xxreal_0(X3,k6_numbers)|v2_struct_0(X1)|~m1_subset_1(X3,k1_numbers)|~v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_222, plain, (m1_subset_1(k8_nat_1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_215])).
cnf(c_0_223, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k1_metric_1(k1_numbers,k1_numbers,X4,k1_seq_1(X2,esk14_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X4!=k7_metric_1|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X3)),k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_106])])).
cnf(c_0_224, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1))),k5_numbers)), inference(csr,[status(thm)],[c_0_217, c_0_132])).
cnf(c_0_225, plain, (m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_218]), c_0_198])]), c_0_163]), c_0_199])])).
cnf(c_0_226, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220])])).
cnf(c_0_227, plain, (v2_struct_0(X1)|r1_xxreal_0(X3,k6_numbers)|m2_subset_1(esk7_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X2)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_221])).
cnf(c_0_228, negated_conjecture, (k8_nat_1(X1,esk1_0,X2)=k1_seq_1(esk1_0,X2)|~v7_ordinal1(X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk1_0,k5_numbers,X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_160]), c_0_139])])])).
cnf(c_0_229, plain, (m1_subset_1(k8_nat_1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_230, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k1_metric_1(k1_numbers,k1_numbers,X3,k1_seq_1(X2,esk14_3(X2,X1,X4)),X1))|X3!=k7_metric_1|~v1_xreal_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_223, c_0_132])).
cnf(c_0_231, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225]), c_0_113]), c_0_134]), c_0_117]), c_0_134]), c_0_134]), c_0_138]), c_0_139])]), c_0_186])])).
cnf(c_0_232, plain, (m2_subset_1(esk12_3(X1,X2,X3),k1_numbers,k5_numbers)|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|X2!=k1_seq_2(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_233, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_226, c_0_175]), c_0_227])).
cnf(c_0_234, negated_conjecture, (k8_nat_1(X1,esk1_0,X2)=k1_seq_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X1)), inference(csr,[status(thm)],[c_0_228, c_0_80])).
cnf(c_0_235, plain, (m1_subset_1(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_51])])).
cnf(c_0_236, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_96])])).
cnf(c_0_237, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k2_metric_1(k8_metric_1,k1_seq_1(X2,esk14_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_158]), c_0_73]), c_0_75]), c_0_76])])])).
cnf(c_0_238, plain, (v1_xreal_0(esk13_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_239, plain, (m2_subset_1(esk14_3(X1,X2,X3),k1_numbers,k5_numbers)|X2=k1_seq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_240, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_231, c_0_132])).
cnf(c_0_241, plain, (r1_xxreal_0(X3,k6_numbers)|m2_subset_1(esk12_3(X1,X2,X3),k1_numbers,k5_numbers)|X2!=k1_seq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_232])).
cnf(c_0_242, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))))),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234]), c_0_139])])])).
cnf(c_0_243, plain, (m1_subset_1(k1_seq_1(X1,X2),X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_50])])).
fof(c_0_244, plain, ![X87, X88, X89, X90, X91]:(~v1_funct_1(X89)|~m1_subset_1(X89,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X87,X88),k1_numbers)))|~m1_subset_1(X90,X87)|~m1_subset_1(X91,X88)|k1_metric_1(X87,X88,X89,X90,X91)=k1_binop_1(X89,X90,X91)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_metric_1])])).
cnf(c_0_245, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(X3,X4),esk6_2(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_236, c_0_205])).
cnf(c_0_246, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k2_metric_1(k8_metric_1,k1_seq_1(X2,esk14_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(csr,[status(thm)],[c_0_237, c_0_132])).
cnf(c_0_247, plain, (X2=k1_seq_2(X1)|v1_xreal_0(esk13_2(X1,X2))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_248, plain, (X2=k1_seq_2(X1)|m2_subset_1(esk14_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_249, plain, (X2=k1_seq_2(X1)|~r1_xxreal_0(esk13_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_250, plain, (v2_tbsp_1(X2,X1)|v2_struct_0(X1)|~r1_xxreal_0(esk8_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_251, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~v2_comseq_2(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241]), c_0_117]), c_0_138]), c_0_139])])])).
cnf(c_0_252, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))))),u1_struct_0(X1))|~m1_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k5_numbers)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_242, c_0_205])).
cnf(c_0_253, plain, (m1_subset_1(k1_seq_1(X1,X2),X3)|~v3_valued_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_243, c_0_63])).
cnf(c_0_254, plain, (m2_subset_1(esk5_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
cnf(c_0_255, plain, (k1_metric_1(X2,X3,X1,X4,X5)=k1_binop_1(X1,X4,X5)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),k1_numbers)))|~m1_subset_1(X4,X2)|~m1_subset_1(X5,X3)), inference(split_conjunct,[status(thm)],[c_0_244])).
fof(c_0_256, plain, ![X103]:(((v1_funct_1(u1_metric_1(X103))|~l1_metric_1(X103))&(v1_funct_2(u1_metric_1(X103),k2_zfmisc_1(u1_struct_0(X103),u1_struct_0(X103)),k1_numbers)|~l1_metric_1(X103)))&(m1_subset_1(u1_metric_1(X103),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X103),u1_struct_0(X103)),k1_numbers)))|~l1_metric_1(X103))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_metric_1])])])).
cnf(c_0_257, plain, (r1_xxreal_0(esk13_2(X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),esk14_3(X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk13_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk14_3(X1,esk6_2(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(k1_seq_1(X1,esk14_3(X1,esk6_2(k8_metric_1,X1),X2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_113]), c_0_134]), c_0_134])]), c_0_186])])).
cnf(c_0_258, plain, (m1_subset_1(esk13_2(X1,X2),k1_numbers)|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_247])])).
cnf(c_0_259, plain, (m1_subset_1(esk14_3(X1,X2,X3),k5_numbers)|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_248]), c_0_198])]), c_0_163]), c_0_199])])).
cnf(c_0_260, plain, (X2=k1_seq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk13_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_261, plain, (r1_xxreal_0(X1,esk14_3(X2,X3,X1))|X3=k1_seq_2(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_262, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~r1_xxreal_0(esk8_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_250])).
cnf(c_0_263, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_251, c_0_204])).
fof(c_0_264, plain, ![X149]:(~m1_subset_1(X149,k1_numbers)|v1_xreal_0(X149)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_265, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k5_numbers)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v7_ordinal1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_137]), c_0_139])])])).
cnf(c_0_266, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|m2_subset_1(esk5_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_254])).
cnf(c_0_267, plain, (k1_metric_1(X2,X3,X1,X4,X5)=k1_binop_1(X1,X4,X5)|~v1_funct_1(X1)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),k1_numbers)))), inference(evalgc,[status(thm)],[c_0_255])).
cnf(c_0_268, plain, (v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_256])).
cnf(c_0_269, plain, (m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_256])).
cnf(c_0_270, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),esk14_3(X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(k1_seq_1(X1,esk14_3(X1,esk6_2(k8_metric_1,X1),X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_257, c_0_258]), c_0_259]), c_0_260])).
cnf(c_0_271, plain, (X3=k1_seq_2(X2)|r1_xxreal_0(X1,esk14_3(X2,X3,X1))|~v1_funct_1(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_261])).
cnf(c_0_272, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(X1,k1_numbers)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_263]), c_0_113]), c_0_134]), c_0_117]), c_0_134]), c_0_134]), c_0_138]), c_0_139])]), c_0_186])])).
cnf(c_0_273, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_264])).
cnf(c_0_274, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k5_numbers)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_265, c_0_80])).
cnf(c_0_275, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),k5_numbers)|X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_266]), c_0_198])]), c_0_163]), c_0_199])])).
cnf(c_0_276, plain, (k1_binop_1(u1_metric_1(X1),X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~v1_funct_1(u1_metric_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_267])])).
cnf(c_0_277, plain, (v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_278, plain, (m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~l1_metric_1(X1)), inference(evalgc,[status(thm)],[c_0_269])).
fof(c_0_279, plain, ![X108, X109, X110]:(v2_struct_0(X108)|~l1_metric_1(X108)|~v1_funct_1(X109)|~v1_funct_2(X109,k5_numbers,u1_struct_0(X108))|~m1_subset_1(X109,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X108))))|~m1_subset_1(X110,k5_numbers)|m1_subset_1(k2_tbsp_1(X108,X109,X110),u1_struct_0(X108))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tbsp_1])])])).
fof(c_0_280, plain, ![X50, X51, X52]:(v2_struct_0(X50)|~l1_metric_1(X50)|~v1_funct_1(X51)|~v1_funct_2(X51,k5_numbers,u1_struct_0(X50))|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X50))))|~m1_subset_1(X52,k5_numbers)|k2_tbsp_1(X50,X51,X52)=k1_funct_1(X51,X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k2_tbsp_1])])])).
cnf(c_0_281, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(k1_seq_1(X1,esk14_3(X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_271])])).
cnf(c_0_282, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_272, c_0_132])).
cnf(c_0_283, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_273])).
cnf(c_0_284, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_numbers)|v2_tbsp_1(X2,X1)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_285, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
cnf(c_0_286, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275]), c_0_139])])])).
cnf(c_0_287, plain, (k1_binop_1(u1_metric_1(X1),X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_276, c_0_277]), c_0_278])).
cnf(c_0_288, plain, (v2_struct_0(X1)|m1_subset_1(k2_tbsp_1(X1,X2,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_289, plain, (v2_struct_0(X1)|k2_tbsp_1(X1,X2,X3)=k1_funct_1(X2,X3)|~l1_metric_1(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_280])).
cnf(c_0_290, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk14_3(X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1)))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_253])])).
cnf(c_0_291, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(X1)|~m1_subset_1(esk8_3(k8_metric_1,esk1_0,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283])])).
cnf(c_0_292, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|m1_subset_1(esk8_3(X1,X2,X3),k1_numbers)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_284])).
fof(c_0_293, plain, ![X124]:(~v1_funct_1(X124)|~v1_funct_2(X124,k5_numbers,k1_numbers)|~m1_subset_1(X124,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k2_seq_2(X124)=k1_seq_2(X124)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_seq_2])])).
cnf(c_0_294, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_285])).
cnf(c_0_295, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_227]), c_0_139])])])).
cnf(c_0_296, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
cnf(c_0_297, plain, (k1_binop_1(k7_metric_1,X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_110])])).
cnf(c_0_298, plain, (v2_struct_0(X1)|m1_subset_1(k2_tbsp_1(X1,X2,X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_288])).
cnf(c_0_299, plain, (k2_tbsp_1(X1,X2,X3)=k1_funct_1(X2,X3)|v2_struct_0(X1)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_289])).
cnf(c_0_300, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk14_3(X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1)))),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_80])])).
cnf(c_0_301, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_292]), c_0_113]), c_0_134]), c_0_117]), c_0_134]), c_0_134]), c_0_138]), c_0_139])]), c_0_186])])).
fof(c_0_302, plain, ![X132]:(~v1_funct_1(X132)|~v1_funct_2(X132,k5_numbers,k1_numbers)|~m1_subset_1(X132,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|v1_xreal_0(k1_seq_2(X132))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_2])])).
cnf(c_0_303, negated_conjecture, (k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|~v2_comseq_2(esk1_0)|~v2_tbsp_1(esk2_0,k8_metric_1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_304, plain, (k2_seq_2(X1)=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_293])).
cnf(c_0_305, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_295]), c_0_139])])])).
cnf(c_0_306, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_307, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_297])])).
cnf(c_0_308, plain, (v2_struct_0(X1)|m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299])])).
cnf(c_0_309, plain, (k1_metric_1(X1,X2,X3,X4,X5)=k1_metric_1(X6,X7,X3,X4,X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X6,X7),k1_numbers)))|~m1_subset_1(X5,X7)|~m1_subset_1(X4,X6)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_267])])).
cnf(c_0_310, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_300, c_0_259])).
cnf(c_0_311, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_301, c_0_132])).
cnf(c_0_312, plain, (v1_xreal_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_302])).
cnf(c_0_313, negated_conjecture, (k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|~v2_comseq_2(esk1_0)|~v2_tbsp_1(esk2_0,k8_metric_1)), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_314, plain, (k1_seq_2(X1)=k2_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_304])).
cnf(c_0_315, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_139])])])).
cnf(c_0_316, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)), inference(csr,[status(thm)],[c_0_307, c_0_175])).
cnf(c_0_317, plain, (v2_struct_0(X1)|m1_subset_1(k1_seq_1(X2,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_50])])).
fof(c_0_318, plain, ![X57, X59, X61, X62, X64]:(((v1_xreal_0(esk15_1(X57))|~v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((m2_subset_1(esk16_2(X57,X59),k1_numbers,k5_numbers)|r1_xxreal_0(X59,k6_numbers)|~v1_xreal_0(X59)|~v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~m2_subset_1(X61,k1_numbers,k5_numbers)|~r1_xxreal_0(esk16_2(X57,X59),X61)|~r1_xxreal_0(X59,k18_complex1(k9_real_1(k1_seq_1(X57,X61),esk15_1(X57))))|r1_xxreal_0(X59,k6_numbers)|~v1_xreal_0(X59)|~v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v1_xreal_0(esk17_2(X57,X62))|~v1_xreal_0(X62)|v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((~r1_xxreal_0(esk17_2(X57,X62),k6_numbers)|~v1_xreal_0(X62)|v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(((m2_subset_1(esk18_3(X57,X62,X64),k1_numbers,k5_numbers)|~m2_subset_1(X64,k1_numbers,k5_numbers)|~v1_xreal_0(X62)|v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(r1_xxreal_0(X64,esk18_3(X57,X62,X64))|~m2_subset_1(X64,k1_numbers,k5_numbers)|~v1_xreal_0(X62)|v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(r1_xxreal_0(esk17_2(X57,X62),k18_complex1(k9_real_1(k1_seq_1(X57,esk18_3(X57,X62,X64)),X62)))|~m2_subset_1(X64,k1_numbers,k5_numbers)|~v1_xreal_0(X62)|v2_comseq_2(X57)|(~v1_funct_1(X57)|~v1_funct_2(X57,k5_numbers,k1_numbers)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_seq_2])])])])])])).
cnf(c_0_319, plain, (k1_metric_1(X4,X5,X1,X2,X3)=k18_complex1(k9_real_1(X2,X3))|X1!=k7_metric_1|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X5),k1_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_309])])).
cnf(c_0_320, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|r1_xxreal_0(esk13_2(X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk13_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_227]), c_0_113]), c_0_134]), c_0_134])]), c_0_186])])).
cnf(c_0_321, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|~v1_xreal_0(k1_seq_2(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_311])])).
cnf(c_0_322, plain, (v1_xreal_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_312])).
cnf(c_0_323, plain, (v9_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_324, plain, (v8_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_325, plain, (v7_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_326, plain, (v6_metric_1(k8_metric_1)), inference(split_conjunct,[status(thm)],[fc6_metric_1])).
cnf(c_0_327, negated_conjecture, (k1_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|~v2_tbsp_1(esk2_0,k8_metric_1)|~v2_comseq_2(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_314]), c_0_117]), c_0_138]), c_0_139])])])).
cnf(c_0_328, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_175]), c_0_139])])])).
cnf(c_0_329, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(k1_numbers,X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_134]), c_0_112]), c_0_113])]), c_0_186])])).
cnf(c_0_330, plain, (v2_struct_0(X1)|m1_subset_1(k1_seq_1(X2,X3),u1_struct_0(X1))|~v3_valued_0(X2)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(csr,[status(thm)],[c_0_317, c_0_63])).
cnf(c_0_331, plain, (r1_xxreal_0(esk17_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_318])).
cnf(c_0_332, plain, (k18_complex1(k9_real_1(X4,X5))=k1_binop_1(X3,X4,X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X1)|~v1_funct_1(X3)|X3!=k7_metric_1|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X5,k1_numbers)|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X3,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_319])])).
cnf(c_0_333, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_320, c_0_258]), c_0_260])).
cnf(c_0_334, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_322]), c_0_117]), c_0_138]), c_0_139])])])).
cnf(c_0_335, plain, (v9_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_323])).
cnf(c_0_336, plain, (v8_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_324])).
cnf(c_0_337, plain, (v7_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_325])).
cnf(c_0_338, plain, (v6_metric_1(k8_metric_1)), inference(evalgc,[status(thm)],[c_0_326])).
cnf(c_0_339, negated_conjecture, (k1_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_327, c_0_197]), c_0_197])])).
cnf(c_0_340, negated_conjecture, (v2_struct_0(X1)|m1_subset_1(k1_tbsp_1(X1,esk1_0),u1_struct_0(X1))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_328]), c_0_139])])])).
cnf(c_0_341, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_96])])).
cnf(c_0_342, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_134]), c_0_113])]), c_0_186])])).
cnf(c_0_343, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk17_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_331])).
cnf(c_0_344, plain, (k18_complex1(k9_real_1(X1,X2))=k1_binop_1(k7_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_73]), c_0_73]), c_0_75]), c_0_76])])])).
cnf(c_0_345, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~v2_comseq_2(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_328]), c_0_137]), c_0_334]), c_0_117]), c_0_138]), c_0_139]), c_0_113]), c_0_335]), c_0_336]), c_0_337]), c_0_338]), c_0_134]), c_0_117]), c_0_134]), c_0_138])]), c_0_186])])).
cnf(c_0_346, negated_conjecture, (k1_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_comseq_2(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_339, c_0_334])])])).
cnf(c_0_347, negated_conjecture, (m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),k1_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_134]), c_0_113]), c_0_335]), c_0_336]), c_0_337]), c_0_338]), c_0_117]), c_0_138])]), c_0_186])])).
cnf(c_0_348, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(X2,X3),esk6_2(k8_metric_1,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_341, c_0_342]), c_0_205])).
cnf(c_0_349, plain, (r1_xxreal_0(esk17_2(X1,X2),k1_binop_1(k7_metric_1,k1_seq_1(X1,esk18_3(X1,X2,X3)),X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_344])])).
cnf(c_0_350, negated_conjecture, (~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~v2_comseq_2(esk1_0)), inference(csr,[status(thm)],[c_0_345, c_0_346])).
cnf(c_0_351, negated_conjecture, (m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_347, c_0_334])])).
cnf(c_0_352, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_328]), c_0_137]), c_0_334]), c_0_117]), c_0_138]), c_0_139]), c_0_113]), c_0_335]), c_0_336]), c_0_337]), c_0_338]), c_0_134]), c_0_117]), c_0_134]), c_0_138])]), c_0_186])])).
cnf(c_0_353, plain, (r1_xxreal_0(esk17_2(X1,X2),k1_binop_1(k7_metric_1,k1_seq_1(X1,esk18_3(X1,X2,X3)),X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_349, c_0_132])).
cnf(c_0_354, negated_conjecture, (~v2_comseq_2(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_283]), c_0_351])])])).
cnf(c_0_355, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1),k5_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_353]), c_0_117]), c_0_138]), c_0_139])]), c_0_354])])).
cnf(c_0_356, plain, (r1_xxreal_0(X1,esk18_3(X2,X3,X1))|v2_comseq_2(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X3)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_318])).
cnf(c_0_357, plain, (m2_subset_1(esk18_3(X1,X2,X3),k1_numbers,k5_numbers)|v2_comseq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_318])).
cnf(c_0_358, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1))|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1),k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)), inference(csr,[status(thm)],[c_0_355, c_0_203])).
cnf(c_0_359, plain, (v2_comseq_2(X2)|r1_xxreal_0(X1,esk18_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_356])).
cnf(c_0_360, plain, (v2_comseq_2(X1)|m2_subset_1(esk18_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_357])).
cnf(c_0_361, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_359]), c_0_117]), c_0_138]), c_0_139])]), c_0_354])])).
cnf(c_0_362, plain, (m1_subset_1(esk18_3(X1,X2,X3),k5_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_360]), c_0_198])]), c_0_163]), c_0_199])])).
cnf(c_0_363, plain, (v2_comseq_2(X1)|~r1_xxreal_0(esk17_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_318])).
cnf(c_0_364, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_362]), c_0_117]), c_0_138]), c_0_139])]), c_0_354])])).
cnf(c_0_365, plain, (v1_xreal_0(esk17_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_318])).
cnf(c_0_366, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk17_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_363])).
cnf(c_0_367, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_227]), c_0_113]), c_0_334]), c_0_134]), c_0_117]), c_0_134]), c_0_138]), c_0_139])]), c_0_186])])).
cnf(c_0_368, plain, (v2_comseq_2(X1)|v1_xreal_0(esk17_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_365])).
cnf(c_0_369, negated_conjecture, (~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_367]), c_0_117]), c_0_138]), c_0_139])]), c_0_354])])).
cnf(c_0_370, plain, (m1_subset_1(esk17_2(X1,X2),k1_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_368])])).
cnf(c_0_371, negated_conjecture, (~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_370]), c_0_117]), c_0_138]), c_0_139])]), c_0_354])])).
cnf(c_0_372, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_283]), c_0_351])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 373
# Proof object clause steps            : 293
# Proof object formula steps           : 80
# Proof object conjectures             : 76
# Proof object clause conjectures      : 73
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 81
# Proof object initial formulas used   : 43
# Proof object generating inferences   : 92
# Proof object simplifying inferences  : 295
# Training examples: 175 positive, 768 negative
# Training: Positive examples begin
cnf(c_0_65, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_76, plain, (v1_funct_1(k7_metric_1))).# trainpos
cnf(c_0_139, negated_conjecture, (v1_funct_1(esk1_0))).# trainpos
cnf(c_0_171, negated_conjecture, (v1_funct_1(esk2_0))).# trainpos
cnf(c_0_338, plain, (v6_metric_1(k8_metric_1))).# trainpos
cnf(c_0_337, plain, (v7_metric_1(k8_metric_1))).# trainpos
cnf(c_0_336, plain, (v8_metric_1(k8_metric_1))).# trainpos
cnf(c_0_335, plain, (v9_metric_1(k8_metric_1))).# trainpos
cnf(c_0_113, plain, (l1_metric_1(k8_metric_1))).# trainpos
cnf(c_0_112, plain, (v1_metric_1(k8_metric_1))).# trainpos
cnf(c_0_118, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_186, plain, (~v2_struct_0(k8_metric_1))).# trainpos
cnf(c_0_163, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_188, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_74, plain, (k8_metric_1=g1_metric_1(k1_numbers,k7_metric_1))).# trainpos
cnf(c_0_123, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_198, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_138, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers))).# trainpos
cnf(c_0_166, plain, (l1_struct_0(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_170, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,u1_struct_0(k8_metric_1)))).# trainpos
cnf(c_0_196, negated_conjecture, (v2_comseq_2(esk1_0)|v2_tbsp_1(esk2_0,k8_metric_1))).# trainpos
cnf(c_0_277, plain, (v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_108, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_283, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_64, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_117, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_75, plain, (v1_funct_2(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))).# trainpos
cnf(c_0_107, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_169, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1)))))).# trainpos
cnf(c_0_122, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_73, plain, (m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).# trainpos
cnf(c_0_168, negated_conjecture, (r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),esk1_0,esk2_0))).# trainpos
cnf(c_0_165, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_63, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_91, plain, (X1=g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))|~l1_metric_1(X1)|~v1_metric_1(X1))).# trainpos
cnf(c_0_313, negated_conjecture, (k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|~v2_comseq_2(esk1_0)|~v2_tbsp_1(esk2_0,k8_metric_1))).# trainpos
cnf(c_0_116, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_142, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_50, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_278, plain, (m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~l1_metric_1(X1))).# trainpos
cnf(c_0_314, plain, (k1_seq_2(X1)=k2_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_322, plain, (v1_xreal_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_140, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_194, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_190, plain, (k4_metric_1(X1,X2,X3)=k4_metric_1(X1,X3,X2)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_154, plain, (k2_metric_1(X1,X2,X3)=k4_metric_1(X1,X2,X3)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_51, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).# trainpos
cnf(c_0_229, plain, (m1_subset_1(k8_nat_1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).# trainpos
cnf(c_0_368, plain, (v2_comseq_2(X1)|v1_xreal_0(esk17_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_109, plain, (k2_metric_1(X1,X2,X3)=k1_metric_1(u1_struct_0(X1),u1_struct_0(X1),u1_metric_1(X1),X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_366, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk17_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_72, plain, (X1=X2|g1_metric_1(X3,X1)!=g1_metric_1(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),k1_numbers))))).# trainpos
cnf(c_0_78, plain, (X1=X2|g1_metric_1(X1,X3)!=g1_metric_1(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))).# trainpos
cnf(c_0_247, plain, (X2=k1_seq_2(X1)|v1_xreal_0(esk13_2(X1,X2))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_260, plain, (X2=k1_seq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk13_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_267, plain, (k1_metric_1(X2,X3,X1,X4,X5)=k1_binop_1(X1,X4,X5)|~v1_funct_1(X1)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),k1_numbers))))).# trainpos
cnf(c_0_359, plain, (v2_comseq_2(X2)|r1_xxreal_0(X1,esk18_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_175, plain, (v2_struct_0(X1)|m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_299, plain, (k2_tbsp_1(X1,X2,X3)=k1_funct_1(X2,X3)|v2_struct_0(X1)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_360, plain, (v2_comseq_2(X1)|m2_subset_1(esk18_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_298, plain, (v2_struct_0(X1)|m1_subset_1(k2_tbsp_1(X1,X2,X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_271, plain, (X3=k1_seq_2(X2)|r1_xxreal_0(X1,esk14_3(X2,X3,X1))|~v1_funct_1(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_248, plain, (X2=k1_seq_2(X1)|m2_subset_1(esk14_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_241, plain, (r1_xxreal_0(X3,k6_numbers)|m2_subset_1(esk12_3(X1,X2,X3),k1_numbers,k5_numbers)|X2!=k1_seq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_292, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|m1_subset_1(esk8_3(X1,X2,X3),k1_numbers)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_262, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~r1_xxreal_0(esk8_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_106, plain, (k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k18_complex1(k9_real_1(X2,X3))|X1!=k7_metric_1|~v1_funct_1(X1)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).# trainpos
cnf(c_0_343, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk17_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_227, plain, (v2_struct_0(X1)|r1_xxreal_0(X3,k6_numbers)|m2_subset_1(esk7_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X2)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_216, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(esk13_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk14_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_211, plain, (v2_struct_0(X2)|v2_tbsp_1(X3,X2)|r1_xxreal_0(X1,esk9_4(X2,X3,X4,X1))|~v1_funct_1(X3)|~l1_metric_1(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).# trainpos
cnf(c_0_218, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|m2_subset_1(esk9_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_167, plain, (X5=X6|v1_xboole_0(X4)|v1_xboole_0(X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~v1_funct_2(X6,X3,X4)|~v1_funct_2(X5,X1,X2)|~r1_funct_2(X1,X2,X3,X4,X5,X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_105, plain, (r1_xxreal_0(X4,k6_numbers)|X3!=k1_seq_2(X2)|~v1_funct_1(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X2,X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X4,k18_complex1(k9_real_1(k1_seq_1(X2,X1),X3))))).# trainpos
cnf(c_0_306, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_294, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_135, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|~v1_funct_1(X2)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_153, plain, (v2_struct_0(X2)|r1_xxreal_0(X4,k6_numbers)|~v1_funct_1(X3)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X4,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X2,X3,X4),X1)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~r1_xxreal_0(X4,k2_metric_1(X2,k8_nat_1(u1_struct_0(X2),X3,X1),esk6_2(X2,X3))))).# trainpos
cnf(c_0_220, plain, (X4=k1_tbsp_1(X2,X3)|v2_struct_0(X2)|r1_xxreal_0(X1,esk5_4(X2,X3,X4,X1))|~v1_funct_1(X3)|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).# trainpos
cnf(c_0_266, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|m2_subset_1(esk5_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_201, plain, (X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|r1_xxreal_0(esk4_3(X1,X2,X3),k4_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,X3,X4)),X3))|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).# trainpos
cnf(c_0_327, negated_conjecture, (k1_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk2_0)|~v2_tbsp_1(esk2_0,k8_metric_1)|~v2_comseq_2(esk1_0))).# trainpos
cnf(c_0_184, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_metric_1(X1))).# trainpos
cnf(c_0_185, negated_conjecture, (esk2_0=esk1_0|v1_xboole_0(u1_struct_0(k8_metric_1)))).# trainpos
cnf(c_0_205, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_259, plain, (m1_subset_1(esk14_3(X1,X2,X3),k5_numbers)|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_132, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_92, plain, (k1_numbers=X1|g1_metric_1(X1,X2)!=k8_metric_1)).# trainpos
cnf(c_0_90, plain, (k7_metric_1=X1|g1_metric_1(X2,X1)!=k8_metric_1)).# trainpos
cnf(c_0_182, negated_conjecture, (v1_relat_1(esk1_0))).# trainpos
cnf(c_0_309, plain, (k1_metric_1(X1,X2,X3,X4,X5)=k1_metric_1(X6,X7,X3,X4,X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X6,X7),k1_numbers)))|~m1_subset_1(X5,X7)|~m1_subset_1(X4,X6))).# trainpos
cnf(c_0_208, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,X1,k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4))))|v2_struct_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4)),u1_struct_0(X2)))).# trainpos
cnf(c_0_275, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),k5_numbers)|X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).# trainpos
cnf(c_0_62, plain, (k8_nat_1(X3,X1,X2)=k1_seq_1(X1,X2)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3))).# trainpos
cnf(c_0_174, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v8_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2)))).# trainpos
cnf(c_0_276, plain, (k1_binop_1(u1_metric_1(X1),X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~v1_funct_1(u1_metric_1(X1)))).# trainpos
cnf(c_0_159, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k1_funct_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,X2,X3,X4)))).# trainpos
cnf(c_0_223, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k1_metric_1(k1_numbers,k1_numbers,X4,k1_seq_1(X2,esk14_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X4!=k7_metric_1|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X3)),k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X4))).# trainpos
cnf(c_0_225, plain, (m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).# trainpos
cnf(c_0_137, negated_conjecture, (v3_valued_0(esk1_0))).# trainpos
cnf(c_0_141, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_362, plain, (m1_subset_1(esk18_3(X1,X2,X3),k5_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_235, plain, (m1_subset_1(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).# trainpos
cnf(c_0_131, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(k1_numbers,k1_numbers,X5,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X5!=k7_metric_1|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v1_funct_2(X5,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X5))).# trainpos
cnf(c_0_308, plain, (v2_struct_0(X1)|m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).# trainpos
cnf(c_0_197, negated_conjecture, (esk2_0=esk1_0)).# trainpos
cnf(c_0_111, plain, (u1_struct_0(X1)=k1_numbers|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_258, plain, (m1_subset_1(esk13_2(X1,X2),k1_numbers)|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_161, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X1,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_319, plain, (k1_metric_1(X4,X5,X1,X2,X3)=k18_complex1(k9_real_1(X2,X3))|X1!=k7_metric_1|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X5),k1_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X4))).# trainpos
cnf(c_0_96, plain, (k8_nat_1(X1,X2,X3)=k1_seq_1(X2,X3)|~v3_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).# trainpos
cnf(c_0_110, plain, (u1_metric_1(X1)=k7_metric_1|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1))).# trainpos
cnf(c_0_243, plain, (m1_subset_1(k1_seq_1(X1,X2),X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_200, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2)))).# trainpos
cnf(c_0_162, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_317, plain, (v2_struct_0(X1)|m1_subset_1(k1_seq_1(X2,X3),u1_struct_0(X1))|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_204, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|v2_comseq_2(esk1_0))).# trainpos
cnf(c_0_134, plain, (u1_struct_0(k8_metric_1)=k1_numbers)).# trainpos
cnf(c_0_339, negated_conjecture, (k1_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0))).# trainpos
cnf(c_0_181, plain, (m2_subset_1(k1_funct_1(X1,X2),k1_numbers,k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_133, plain, (k1_metric_1(u1_struct_0(X1),u1_struct_0(X1),k7_metric_1,X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1))).# trainpos
cnf(c_0_213, plain, (r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),esk5_4(X1,X2,esk6_2(X1,X2),X3))|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),X3)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v9_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).# trainpos
cnf(c_0_236, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).# trainpos
cnf(c_0_115, plain, (k1_funct_1(X2,X3)=k1_seq_1(X2,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(X3,k5_numbers))).# trainpos
cnf(c_0_158, plain, (k1_metric_1(k1_numbers,k1_numbers,k7_metric_1,X1,X2)=k2_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_226, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).# trainpos
cnf(c_0_160, negated_conjecture, (k1_funct_1(esk1_0,X1)=k1_seq_1(esk1_0,X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_228, negated_conjecture, (k8_nat_1(X1,esk1_0,X2)=k1_seq_1(esk1_0,X2)|~v7_ordinal1(X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk1_0,k5_numbers,X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_179, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k2_metric_1(X1,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),X2))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,X2,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk9_4(X1,esk1_0,X2,X3),k5_numbers))).# trainpos
cnf(c_0_195, negated_conjecture, (m2_subset_1(k1_seq_1(esk1_0,X1),k1_numbers,k1_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_178, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_242, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))))),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k5_numbers))).# trainpos
cnf(c_0_203, negated_conjecture, (m1_subset_1(k1_seq_1(esk1_0,X1),k1_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_202, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),esk9_4(k8_metric_1,esk1_0,X1,X2))|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_265, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)))),k5_numbers)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v7_ordinal1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))))))).# trainpos
cnf(c_0_217, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1))),k5_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_286, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1)))).# trainpos
cnf(c_0_231, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_297, plain, (k1_binop_1(k7_metric_1,X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1))).# trainpos
cnf(c_0_237, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k2_metric_1(k8_metric_1,k1_seq_1(X2,esk14_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_342, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_332, plain, (k18_complex1(k9_real_1(X4,X5))=k1_binop_1(X3,X4,X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X1)|~v1_funct_1(X3)|X3!=k7_metric_1|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X5,k1_numbers)|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X3,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))).# trainpos
cnf(c_0_344, plain, (k18_complex1(k9_real_1(X1,X2))=k1_binop_1(k7_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_251, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~v2_comseq_2(esk1_0))).# trainpos
cnf(c_0_295, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).# trainpos
cnf(c_0_307, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2)))).# trainpos
cnf(c_0_272, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(X1,k1_numbers)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_305, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers))).# trainpos
cnf(c_0_329, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(k1_numbers,X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).# trainpos
cnf(c_0_291, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(X1)|~m1_subset_1(esk8_3(k8_metric_1,esk1_0,X1),k1_numbers))).# trainpos
cnf(c_0_315, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).# trainpos
cnf(c_0_328, negated_conjecture, (esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).# trainpos
cnf(c_0_301, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_349, plain, (r1_xxreal_0(esk17_2(X1,X2),k1_binop_1(k7_metric_1,k1_seq_1(X1,esk18_3(X1,X2,X3)),X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers))).# trainpos
cnf(c_0_257, plain, (r1_xxreal_0(esk13_2(X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),esk14_3(X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk13_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk14_3(X1,esk6_2(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(k1_seq_1(X1,esk14_3(X1,esk6_2(k8_metric_1,X1),X2)),k1_numbers))).# trainpos
cnf(c_0_340, negated_conjecture, (v2_struct_0(X1)|m1_subset_1(k1_tbsp_1(X1,esk1_0),u1_struct_0(X1))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).# trainpos
cnf(c_0_321, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1)|~v1_xreal_0(k1_seq_2(esk1_0)))).# trainpos
cnf(c_0_281, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(k1_seq_1(X1,esk14_3(X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_334, negated_conjecture, (v2_tbsp_1(esk1_0,k8_metric_1))).# trainpos
cnf(c_0_290, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk14_3(X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))))))).# trainpos
cnf(c_0_346, negated_conjecture, (k1_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_comseq_2(esk1_0))).# trainpos
cnf(c_0_347, negated_conjecture, (m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),k1_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1))).# trainpos
cnf(c_0_300, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk14_3(X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk13_2(X1,esk6_2(k8_metric_1,X1)))),k5_numbers))).# trainpos
cnf(c_0_320, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|r1_xxreal_0(esk13_2(X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk13_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers))).# trainpos
cnf(c_0_345, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~v2_comseq_2(esk1_0))).# trainpos
cnf(c_0_354, negated_conjecture, (~v2_comseq_2(esk1_0))).# trainpos
cnf(c_0_341, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_352, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_355, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1),k5_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1)),k1_numbers))).# trainpos
cnf(c_0_361, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).# trainpos
cnf(c_0_364, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).# trainpos
cnf(c_0_367, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).# trainpos
cnf(c_0_369, negated_conjecture, (~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).# trainpos
cnf(c_0_370, plain, (m1_subset_1(esk17_2(X1,X2),k1_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_371, negated_conjecture, (~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_80, plain, (k1_xboole_0=k6_numbers)).#trainneg
cnf(i_0_102, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_142, plain, (v7_ordinal1(esk20_0))).#trainneg
cnf(i_0_91, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_117, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_121, plain, (v3_topmetr(k8_metric_1))).#trainneg
cnf(i_0_143, plain, (~v1_xboole_0(esk20_0))).#trainneg
cnf(i_0_128, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_137, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_101, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_109, plain, (v1_xboole_0(esk19_2(X1,X2)))).#trainneg
cnf(i_0_108, plain, (v1_relat_1(esk19_2(X1,X2)))).#trainneg
cnf(i_0_113, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_135, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_107, plain, (v4_relat_1(esk19_2(X1,X2),X1))).#trainneg
cnf(i_0_106, plain, (v5_relat_1(esk19_2(X1,X2),X2))).#trainneg
cnf(i_0_105, plain, (k16_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_125, plain, (v1_xreal_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_123, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_124, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_134, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_116, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_115, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_71, plain, (v1_xreal_0(X1)|~r2_hidden(X1,k1_numbers))).#trainneg
cnf(i_0_131, plain, (k18_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_100, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_110, plain, (m1_subset_1(esk19_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_139, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_89, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_88, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_103, plain, (k6_xcmplx_0(X1,X2)=k9_real_1(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_84, plain, (v1_funct_2(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)|~l1_metric_1(X1))).#trainneg
cnf(i_0_136, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_133, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_120, plain, (m1_subset_1(k2_metric_1(X1,X2,X3),k1_numbers)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_118, plain, (m1_subset_1(k2_seq_2(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_63, plain, (v1_xreal_0(esk15_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_114, plain, (m1_subset_1(k4_metric_1(X1,X2,X3),k1_numbers)|~v8_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_78, plain, (l1_metric_1(g1_metric_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))).#trainneg
cnf(i_0_79, plain, (v1_metric_1(g1_metric_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers))))).#trainneg
cnf(i_0_35, plain, (X1=k7_metric_1|m1_subset_1(esk10_1(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_34, plain, (X1=k7_metric_1|m1_subset_1(esk11_1(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_98, plain, (m1_subset_1(k1_metric_1(X2,X3,X1,X4,X5),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),k1_numbers))))).#trainneg
cnf(i_0_62, plain, (r1_xxreal_0(X2,k6_numbers)|m2_subset_1(esk16_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_33, plain, (X1=k7_metric_1|k1_metric_1(k1_numbers,k1_numbers,X1,esk10_1(X1),esk11_1(X1))!=k18_complex1(k9_real_1(esk10_1(X1),esk11_1(X1)))|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_69, plain, (v2_struct_0(X1)|m1_subset_1(k1_tbsp_1(X1,X2),u1_struct_0(X1))|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_70, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_funct_2(X4,X1,X6,X2,X3,X3)|~v1_funct_1(X5)|~v1_funct_1(X3)|~v1_funct_2(X5,X6,X2)|~v1_funct_2(X3,X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X1))))).#trainneg
cnf(i_0_61, plain, (r1_xxreal_0(X3,k6_numbers)|~v1_funct_1(X2)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~r1_xxreal_0(esk16_2(X2,X3),X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X2,X1),esk15_1(X2)))))).#trainneg
cnf(i_0_31, plain, (v1_xboole_0(X6)|v1_xboole_0(X4)|r1_funct_2(X3,X4,X5,X6,X1,X2)|X1!=X2|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X5,X6)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_67, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_funct_2(X4,X1,X6,X2,X5,X3)|~v1_funct_1(X5)|~v1_funct_1(X3)|~v1_funct_2(X5,X6,X2)|~v1_funct_2(X3,X4,X1)|~r1_funct_2(X4,X1,X6,X2,X3,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X1))))).#trainneg
cnf(i_0_22, plain, (v2_struct_0(X1)|r1_xxreal_0(X4,k6_numbers)|m2_subset_1(esk3_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|X3!=k1_tbsp_1(X1,X2)|~v1_funct_1(X2)|~v6_metric_1(X1)|~v7_metric_1(X1)|~v8_metric_1(X1)|~v9_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_21, plain, (v2_struct_0(X2)|r1_xxreal_0(X5,k6_numbers)|X4!=k1_tbsp_1(X2,X3)|~v1_funct_1(X3)|~v6_metric_1(X2)|~v7_metric_1(X2)|~v8_metric_1(X2)|~v9_metric_1(X2)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X5,k1_numbers)|~m1_subset_1(X4,u1_struct_0(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~r1_xxreal_0(esk3_4(X2,X3,X4,X5),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~r1_xxreal_0(X5,k4_metric_1(X2,k8_nat_1(u1_struct_0(X2),X3,X1),X4)))).#trainneg
cnf(i_0_248, plain, (r1_xxreal_0(esk13_2(X1,esk15_1(X1)),k6_numbers)|esk15_1(X1)=k1_seq_2(X1)|~v1_xreal_0(esk13_2(X1,esk15_1(X1)))|~m2_subset_1(esk14_3(X1,esk15_1(X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk16_2(X1,esk13_2(X1,esk15_1(X1))),esk14_3(X1,esk15_1(X1),X2))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk15_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_205, plain, (k8_nat_1(X1,X2,X3)=k8_nat_1(X4,X2,X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_267, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,X3,k8_nat_1(u1_struct_0(X2),X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2)))).#trainneg
cnf(i_0_160, negated_conjecture, (v1_relat_1(esk2_0))).#trainneg
cnf(i_0_204, plain, (v1_xreal_0(k8_nat_1(X3,X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3))).#trainneg
cnf(i_0_220, plain, (u1_struct_0(X1)=X2|g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))!=g1_metric_1(X2,X3)|~v1_funct_2(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_266, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_funct_1(X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_180, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_258, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_funct_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_257, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1)))).#trainneg
cnf(i_0_169, negated_conjecture, (v5_relat_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_255, negated_conjecture, (v1_xboole_0(u1_struct_0(k8_metric_1))|r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),esk2_0,esk1_0))).#trainneg
cnf(i_0_168, negated_conjecture, (v5_relat_1(esk2_0,u1_struct_0(k8_metric_1)))).#trainneg
cnf(i_0_263, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,k1_funct_1(X3,esk5_4(X2,X3,X1,X4)),X1))|v2_struct_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v7_ordinal1(esk5_4(X2,X3,X1,X4)))).#trainneg
cnf(i_0_227, plain, (m1_subset_1(esk16_2(X1,X2),k5_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_213, plain, (v1_xboole_0(k8_nat_1(k1_zfmisc_1(X1),X2,X3))|v1_xboole_0(X1)|m2_subset_1(X4,X1,k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~m1_subset_1(X4,k8_nat_1(k1_zfmisc_1(X1),X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(X1))))|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_162, plain, (v1_relat_1(k7_metric_1))).#trainneg
cnf(i_0_193, plain, (m1_subset_1(k1_seq_2(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_198, plain, (l1_metric_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))|~v1_funct_2(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_188, negated_conjecture, (v1_xboole_0(esk2_0)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1)))|m2_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(k8_metric_1)),esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_244, plain, (v1_xboole_0(X1)|r1_funct_2(X2,X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_241, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(k8_metric_1))|r1_funct_2(X2,X1,k5_numbers,u1_struct_0(k8_metric_1),X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_31, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_funct_2(X3,X1,X4,X2,X5,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X5,X4,X2)|~v1_funct_2(X5,X3,X1)|~v1_funct_1(X5))).#trainneg
cnf(i_0_222, plain, (X1=X2|g1_metric_1(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X3,X4))!=g1_metric_1(X2,X5)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X3,X4),k2_zfmisc_1(X1,X1),k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X3,X4))|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))))|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_216, plain, (u1_metric_1(X1)=X2|g1_metric_1(u1_struct_0(X1),u1_metric_1(X1))!=g1_metric_1(X3,X2)|~v1_funct_2(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_242, negated_conjecture, (v1_xboole_0(X1)|r1_funct_2(X2,X1,k5_numbers,k1_numbers,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_170, plain, (v5_relat_1(k7_metric_1,k1_numbers))).#trainneg
cnf(i_0_246, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(k1_numbers,k1_numbers,X4,k1_seq_1(X3,X2),esk15_1(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X4!=k7_metric_1|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(esk15_1(X3),k1_numbers)|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X4))).#trainneg
cnf(i_0_171, plain, (k9_real_1(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_226, plain, (m1_subset_1(k1_binop_1(X3,X4,X5),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_183, plain, (v5_relat_1(u1_metric_1(X1),k1_numbers)|~l1_metric_1(X1))).#trainneg
cnf(i_0_144, plain, (esk19_2(X1,X2)=k6_numbers)).#trainneg
cnf(i_0_184, plain, (v1_relat_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_237, plain, (r1_xxreal_0(esk17_2(X1,X2),k1_metric_1(k1_numbers,k1_numbers,X4,k1_seq_1(X1,esk18_3(X1,X2,X3)),X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|X4!=k7_metric_1|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X4))).#trainneg
cnf(i_0_189, negated_conjecture, (v1_xboole_0(esk1_0)|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|m2_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_262, plain, (m1_subset_1(esk3_4(X1,X2,X3,X4),k5_numbers)|r1_xxreal_0(X4,k6_numbers)|v2_struct_0(X1)|X3!=k1_tbsp_1(X1,X2)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_218, plain, (k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3)=X4|g1_metric_1(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3))!=g1_metric_1(X5,X4)|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3),k2_zfmisc_1(X1,X1),k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))))|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_238, plain, (m1_subset_1(esk7_3(X1,X2,X3),k5_numbers)|r1_xxreal_0(X3,k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_175, plain, (v3_valued_0(k7_metric_1))).#trainneg
cnf(i_0_231, plain, (m1_subset_1(esk12_3(X1,X2,X3),k5_numbers)|r1_xxreal_0(X3,k6_numbers)|X2!=k1_seq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_149, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_233, plain, (X1=k7_metric_1|k1_metric_1(k1_numbers,k1_numbers,X1,esk10_1(X1),esk11_1(X1))!=k1_metric_1(k1_numbers,k1_numbers,X2,esk10_1(X1),esk11_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X1)|X2!=k7_metric_1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(esk11_1(X1),k1_numbers)|~m1_subset_1(esk10_1(X1),k1_numbers)|~v1_funct_2(X2,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_243, plain, (v1_xboole_0(X1)|r1_funct_2(X2,X1,k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),k1_numbers,X4,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(u1_metric_1(X3),k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),k1_numbers)|~v1_funct_2(X4,X2,X1)|~v1_funct_1(u1_metric_1(X3))|~v1_funct_1(X4)|~l1_metric_1(X3))).#trainneg
cnf(i_0_206, plain, (l1_metric_1(g1_metric_1(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3)))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3),k2_zfmisc_1(X1,X1),k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))))|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_181, plain, (v3_valued_0(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_145, plain, (v1_xreal_0(k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_173, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(u1_struct_0(k8_metric_1)))).#trainneg
cnf(i_0_210, plain, (v5_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X2)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_165, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers))).#trainneg
cnf(i_0_201, plain, (v1_metric_1(g1_metric_1(u1_struct_0(X1),u1_metric_1(X1)))|~v1_funct_2(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_232, plain, (k18_complex1(k1_metric_1(k1_numbers,k1_numbers,X3,X1,X2))=k1_metric_1(k1_numbers,k1_numbers,X3,X1,X2)|~v1_xcmplx_0(k9_real_1(X1,X2))|X3!=k7_metric_1|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_191, plain, (v1_xboole_0(k7_metric_1)|v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))|m2_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers),k7_metric_1)|~m1_subset_1(X1,k7_metric_1))).#trainneg
cnf(i_0_164, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_211, plain, (v1_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_235, plain, (m1_subset_1(k18_complex1(k9_real_1(X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X1)|X1!=k7_metric_1|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))).#trainneg
cnf(i_0_245, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_funct_2(X3,X1,X4,X2,X5,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X4,X2)),X6,X7),X4,X2)|~v1_funct_2(X5,X3,X1)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X4,X2)),X6,X7))|~v1_funct_1(X5)|~v7_ordinal1(X7)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))|~v1_funct_2(X6,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~v1_funct_1(X6))).#trainneg
cnf(i_0_185, plain, (v4_relat_1(u1_metric_1(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_metric_1(X1))).#trainneg
cnf(i_0_190, plain, (v1_xboole_0(u1_metric_1(X1))|v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers))|m2_subset_1(X2,k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers),u1_metric_1(X1))|~m1_subset_1(X2,u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_212, plain, (v4_relat_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_207, plain, (v1_metric_1(g1_metric_1(X1,k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3)))|~v1_funct_2(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3),k2_zfmisc_1(X1,X1),k1_numbers)|~v1_funct_1(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)),X2,X3))|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))))|~v1_funct_2(X2,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_208, plain, (v3_valued_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v3_membered(X2)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_146, plain, (v1_xreal_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_166, plain, (v4_relat_1(k7_metric_1,k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_178, negated_conjecture, (v1_valued_0(esk1_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_177, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(u1_struct_0(k8_metric_1)))).#trainneg
cnf(i_0_147, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_209, plain, (v1_valued_0(k8_nat_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_membered(X2)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))|~v1_funct_2(X3,k5_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_179, plain, (v1_valued_0(k7_metric_1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_148, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_182, plain, (v1_valued_0(u1_metric_1(X1))|~v1_membered(k1_numbers)|~l1_metric_1(X1))).#trainneg
cnf(i_0_274, plain, (m1_subset_1(esk15_1(X1),k1_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_282, plain, (esk15_1(X1)=k1_seq_2(X1)|r1_xxreal_0(esk13_2(X1,esk15_1(X1)),k6_numbers)|~v1_xreal_0(esk13_2(X1,esk15_1(X1)))|~m2_subset_1(esk16_2(X1,esk13_2(X1,esk15_1(X1))),k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk15_1(X1)))).#trainneg
cnf(i_0_308, plain, (u1_struct_0(X1)=X2|X1!=g1_metric_1(X2,X3)|~l1_metric_1(X1)|~v1_metric_1(X1))).#trainneg
cnf(i_0_285, plain, (k8_nat_1(X1,X2,X3)=k8_nat_1(X4,X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_345, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v7_ordinal1(esk9_4(X1,X2,X3,X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~v1_funct_2(X2,k5_numbers,X5))).#trainneg
cnf(i_0_304, plain, (m1_subset_1(k8_nat_1(X1,X2,X3),k1_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_344, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v7_ordinal1(esk9_4(X1,X2,X3,X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_310, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X6,X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~v1_funct_2(X4,k5_numbers,X6))).#trainneg
cnf(i_0_319, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(X5,X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_286, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,X3,k1_funct_1(X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_403, plain, (m1_subset_1(k8_nat_1(X4,X1,X2),X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X1,k5_numbers,X4))).#trainneg
cnf(i_0_290, plain, (k1_metric_1(X4,X5,u1_metric_1(X1),X2,X3)=k2_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X5),k1_numbers)))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X4)|~v1_funct_1(u1_metric_1(X1)))).#trainneg
cnf(i_0_359, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_309, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_321, plain, (r1_xxreal_0(esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5))!=k1_tbsp_1(X1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X3,X4),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_520, plain, (m2_subset_1(esk13_2(X1,X2),k1_numbers,k1_numbers)|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1, negated_conjecture, (k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0))).#trainneg
cnf(i_0_465, plain, (esk15_1(X1)=k1_seq_2(X1)|~v1_xreal_0(esk15_1(X1))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk13_2(X1,esk15_1(X1))))).#trainneg
cnf(i_0_528, plain, (m2_subset_1(esk5_4(X1,X2,X3,X4),k5_numbers,k5_numbers)|X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_524, plain, (m2_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers,k5_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_536, plain, (m2_subset_1(k1_seq_1(X1,X2),k1_numbers,k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_518, plain, (m2_subset_1(esk15_1(X1),k1_numbers,k1_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_529, plain, (m2_subset_1(esk4_3(X1,X2,X3),k1_numbers,k1_numbers)|X3=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_542, plain, (v1_xboole_0(X1)|m2_subset_1(k8_nat_1(X1,X2,X3),X1,X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_641, plain, (m2_subset_1(k8_nat_1(X1,X2,X3),k1_numbers,k1_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_545, plain, (m2_subset_1(k2_seq_2(X1),k1_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_525, plain, (m2_subset_1(esk8_3(X1,X2,X3),k1_numbers,k1_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_527, plain, (v1_xboole_0(u1_struct_0(X1))|m2_subset_1(esk6_2(X1,X2),u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_571, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers))).#trainneg
cnf(i_0_470, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X1,k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers)|~v1_funct_2(X2,k5_numbers,X5))).#trainneg
cnf(i_0_472, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(X5,X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_543, plain, (m2_subset_1(k6_numbers,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_471, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers)|~v1_funct_2(X2,k5_numbers,X5))).#trainneg
cnf(i_0_626, plain, (v1_xboole_0(X1)|m2_subset_1(k1_seq_1(X2,X3),X1,X1)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_519, plain, (m2_subset_1(esk14_3(X1,X2,X3),k5_numbers,k5_numbers)|X2=k1_seq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_476, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X6,X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,X6))).#trainneg
cnf(i_0_544, plain, (v1_xboole_0(u1_struct_0(X1))|m2_subset_1(k1_tbsp_1(X1,X2),u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_516, plain, (m2_subset_1(esk18_3(X1,X2,X3),k5_numbers,k5_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_466, plain, (u1_struct_0(g1_metric_1(X1,X2))=X1|~v1_metric_1(g1_metric_1(X1,X2))|~l1_metric_1(g1_metric_1(X1,X2)))).#trainneg
cnf(i_0_687, plain, (v1_xboole_0(X1)|m2_subset_1(k8_nat_1(X2,X3,X4),X1,X1)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_541, plain, (m2_subset_1(k4_metric_1(X1,X2,X3),k1_numbers,k1_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_474, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,k8_nat_1(X5,X3,esk5_4(X2,X3,X1,X4)),X1))|v2_struct_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(esk5_4(X2,X3,X1,X4),k5_numbers)|~v1_funct_2(X3,k5_numbers,X5))).#trainneg
cnf(i_0_575, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,k1_seq_1(X3,esk5_4(X2,X3,X1,X4)),X1))|v2_struct_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(esk5_4(X2,X3,X1,X4),k5_numbers))).#trainneg
cnf(i_0_535, plain, (v1_xboole_0(X1)|m2_subset_1(k1_funct_1(X2,X3),X1,X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_577, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_521, plain, (m2_subset_1(esk12_3(X1,X2,X3),k5_numbers,k5_numbers)|r1_xxreal_0(X3,k6_numbers)|X2!=k1_seq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_477, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,X3,k8_nat_1(X6,X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X6,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,X6))).#trainneg
cnf(i_0_537, plain, (m2_subset_1(k1_seq_2(X1),k1_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_473, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X5,X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X5,X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_9, negated_conjecture, (r1_funct_2(k5_numbers,k1_numbers,k5_numbers,u1_struct_0(k8_metric_1),esk1_0,esk1_0))).#trainneg
cnf(i_0_555, plain, (k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k1_metric_1(X4,X5,X6,X2,X3)|X1!=k7_metric_1|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X1)|X6!=k7_metric_1|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X5),k1_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X4)|~v1_funct_2(X6,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X6))).#trainneg
cnf(i_0_651, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,X3,k8_nat_1(X6,X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X4,k5_numbers,X6)|~v1_funct_1(X4)|~m1_subset_1(k8_nat_1(X6,X4,X5),u1_struct_0(X2)))).#trainneg
cnf(i_0_661, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X4,X5,X3),esk6_2(X2,X5)))|~r1_xxreal_0(esk7_3(X2,X5,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X5,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X5,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X5,k5_numbers,X4)|~v1_funct_1(X5)|~v8_metric_1(X2)|~m1_subset_1(esk6_2(X2,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X4,X5,X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_629, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k8_nat_1(X5,X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X5,X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_596, plain, (v1_funct_2(k7_metric_1,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)|~l1_metric_1(X1)|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_628, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k1_seq_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_713, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(X6,X3,X4),X5))!=k1_tbsp_1(X1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(X6,X3,X4),X5)),esk8_3(X1,X2,k8_nat_1(X6,X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(X6,X3,X4),X5)),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X6,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X6))).#trainneg
cnf(i_0_578, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,X3,k1_seq_1(X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_627, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k1_funct_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_331, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_funct_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_110, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_725, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k4_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X2,esk5_4(k8_metric_1,X2,X1,X3)),X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_825, plain, (m2_subset_1(k1_tbsp_1(k8_metric_1,X1),k1_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_783, plain, (m2_subset_1(esk6_2(k8_metric_1,X1),k1_numbers,k1_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_738, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k2_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_730, plain, (g1_metric_1(k1_numbers,u1_metric_1(k8_metric_1))=k8_metric_1)).#trainneg
cnf(i_0_737, plain, (k1_metric_1(k1_numbers,k1_numbers,u1_metric_1(k8_metric_1),X1,X2)=k2_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_753, plain, (esk15_1(X1)=k1_seq_2(X1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk15_1(X1)))).#trainneg
cnf(i_0_728, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k4_metric_1(k8_metric_1,X1,k8_nat_1(k1_numbers,X2,esk5_4(k8_metric_1,X2,X1,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(k8_nat_1(k1_numbers,X2,esk5_4(k8_metric_1,X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_733, plain, (m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_723, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_765, plain, (m2_subset_1(k1_funct_1(X2,X3),k1_numbers,k1_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_839, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(u1_struct_0(X1),X3,X4)=k1_tbsp_1(X1,X2)|k8_nat_1(X5,X2,esk5_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X6))!=k1_tbsp_1(X1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k8_nat_1(X5,X2,esk5_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X6)),esk4_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X3,X4),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(X5,X2,esk5_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X6)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X5)|~v1_funct_1(X2))).#trainneg
cnf(i_0_747, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_724, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X3,X4),X2))|~r1_xxreal_0(esk3_4(k8_metric_1,X3,X2,X1),X4)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_810, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2))).#trainneg
cnf(i_0_726, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,esk6_2(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_877, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(X3,X4,X5)=k1_tbsp_1(X1,X2)|k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6))!=k1_tbsp_1(X1,X4)|~r1_xxreal_0(esk3_4(X1,X4,k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),esk4_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_732, plain, (m1_subset_1(u1_metric_1(k8_metric_1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))))).#trainneg
cnf(i_0_774, plain, (k2_metric_1(X1,X2,X3)=k1_binop_1(k7_metric_1,X2,X3)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_889, plain, (r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|~r1_xxreal_0(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),esk5_4(X1,X2,esk6_2(X1,X2),X3))|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),X3)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~m1_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v9_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_731, plain, (v1_funct_2(u1_metric_1(k8_metric_1),k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers))).#trainneg
cnf(i_0_727, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_954, plain, (r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|~v7_ordinal1(esk5_4(X1,X2,esk6_2(X1,X2),X3))|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),esk5_4(X1,X2,esk6_2(X1,X2),X3))|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v9_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_734, plain, (m1_subset_1(k1_tbsp_1(k8_metric_1,X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_849, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,X1,k8_nat_1(X4,X3,esk5_4(X2,X3,X1,X5))))|v2_struct_0(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~m1_subset_1(k8_nat_1(X4,X3,esk5_4(X2,X3,X1,X5)),u1_struct_0(X2)))).#trainneg
cnf(i_0_937, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(X3,X4,X5)=k1_tbsp_1(X1,X2)|k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6))!=k1_tbsp_1(X1,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X4,k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),esk4_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_777, plain, (k2_metric_1(X1,X2,X3)=k1_metric_1(X4,X5,k7_metric_1,X2,X3)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X5),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X4)|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_851, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(u1_struct_0(X1),X3,X4)=k1_tbsp_1(X1,X2)|k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5))!=k1_tbsp_1(X1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),esk4_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X3,X4),u1_struct_0(X1))|~m1_subset_1(k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~v3_valued_0(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_729, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,X2,k8_nat_1(k1_numbers,X3,X4)))|~r1_xxreal_0(esk3_4(k8_metric_1,X3,X2,X1),X4)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X4),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_766, plain, (m2_subset_1(k1_seq_1(X2,X3),k1_numbers,k1_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_801, plain, (r1_xxreal_0(esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k8_nat_1(X5,X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X6))!=k1_tbsp_1(X1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k8_nat_1(X5,X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X6)),esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X3,X4),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(X5,X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X6)),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4)),k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X5)|~v1_funct_1(X2))).#trainneg
cnf(i_0_933, plain, (r1_xxreal_0(esk8_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X6))!=k1_tbsp_1(X1,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X4,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),esk8_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1035, plain, (u1_metric_1(k8_metric_1)=k7_metric_1)).#trainneg
cnf(i_0_1057, plain, (k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k2_metric_1(k8_metric_1,X2,X3)|X1!=k7_metric_1|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1221, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~v7_ordinal1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))))|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1208, plain, (k1_binop_1(k7_metric_1,X1,X2)=k2_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1122, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(k8_metric_1,X2,k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1059, plain, (k1_metric_1(X3,X4,X5,X1,X2)=k2_metric_1(k8_metric_1,X1,X2)|X5!=k7_metric_1|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X4),k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X3)|~v1_funct_2(X5,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X5))).#trainneg
cnf(i_0_1089, negated_conjecture, (m2_subset_1(k1_funct_1(esk1_0,X1),k1_numbers,k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1083, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1119, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(k8_metric_1,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_1030, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k2_metric_1(k8_metric_1,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_1084, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(X4,X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_1031, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k2_metric_1(k8_metric_1,k8_nat_1(X4,X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,X4))).#trainneg
cnf(i_0_1210, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))))),u1_struct_0(X1))|~m1_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1225, plain, (r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|~r1_xxreal_0(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),esk5_4(X1,X2,esk6_2(X1,X2),X4))|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X2,esk5_4(X1,X2,esk6_2(X1,X2),X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~m1_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),X4),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X3)|~v1_funct_1(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v9_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1159, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,esk6_2(k8_metric_1,X3),k1_seq_1(X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1237, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,X1,k1_seq_1(X3,esk5_4(X2,X3,X1,X5))))|v2_struct_0(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(k1_seq_1(X3,esk5_4(X2,X3,X1,X5)),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(esk5_4(X2,X3,X1,X5),k5_numbers))).#trainneg
cnf(i_0_1017, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2)),k1_numbers)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1022, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k4_metric_1(k8_metric_1,k8_nat_1(X4,X2,esk5_4(k8_metric_1,X2,X1,X3)),X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(esk5_4(k8_metric_1,X2,X1,X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_1165, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2)),k1_numbers)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1120, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(k8_metric_1,k8_nat_1(X4,X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(X4,X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,X4))).#trainneg
cnf(i_0_1019, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~v7_ordinal1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2))|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1148, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k1_seq_1(X2,esk9_4(X1,X2,X3,X5))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk9_4(X1,X2,X3,X5),k5_numbers))).#trainneg
cnf(i_0_1212, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1079, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k4_metric_1(k8_metric_1,X1,k1_seq_1(X2,esk5_4(k8_metric_1,X2,X1,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(k1_seq_1(X2,esk5_4(k8_metric_1,X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk5_4(k8_metric_1,X2,X1,X3),k5_numbers))).#trainneg
cnf(i_0_1195, negated_conjecture, (X1=k1_tbsp_1(X2,esk1_0)|r1_xxreal_0(esk4_3(X2,esk1_0,X1),k4_metric_1(X2,k1_seq_1(esk1_0,esk5_4(X2,esk1_0,X1,X3)),X1))|v2_struct_0(X2)|~v7_ordinal1(esk5_4(X2,esk1_0,X1,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(esk5_4(X2,esk1_0,X1,X3),k5_numbers))).#trainneg
cnf(i_0_1151, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1154, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(k1_seq_1(X2,esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k5_numbers))).#trainneg
cnf(i_0_1080, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k4_metric_1(k8_metric_1,X1,k8_nat_1(X4,X2,esk5_4(k8_metric_1,X2,X1,X3))))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(k8_nat_1(X4,X2,esk5_4(k8_metric_1,X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(esk5_4(k8_metric_1,X2,X1,X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_1021, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k4_metric_1(k8_metric_1,k1_seq_1(X2,esk5_4(k8_metric_1,X2,X1,X3)),X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk5_4(k8_metric_1,X2,X1,X3),k5_numbers))).#trainneg
cnf(i_0_1385, negated_conjecture, (m2_subset_1(k1_seq_1(esk1_0,X1),k1_numbers,k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1603, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~v3_valued_0(X2)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m2_subset_1(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1440, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k2_metric_1(k8_metric_1,k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),X1))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers))).#trainneg
cnf(i_0_1559, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v7_ordinal1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))))|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1439, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k4_metric_1(k8_metric_1,X1,k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2))))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers))).#trainneg
cnf(i_0_1441, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k4_metric_1(k8_metric_1,k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),X1))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers))).#trainneg
cnf(i_0_1515, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~m2_subset_1(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k5_numbers)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1386, negated_conjecture, (m2_subset_1(k8_nat_1(X2,esk1_0,X1),k1_numbers,k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk1_0,k5_numbers,X2))).#trainneg
cnf(i_0_1398, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(k8_metric_1,X2,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1505, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(k8_nat_1(X3,X2,esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))))),u1_struct_0(X1))|~m1_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X3)|~v1_funct_1(X2)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1597, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(k1_seq_1(X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k5_numbers))).#trainneg
cnf(i_0_1363, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(k8_metric_1,X2,k8_nat_1(X4,X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(X4,X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,X4))).#trainneg
cnf(i_0_1575, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2)),k1_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1404, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k4_metric_1(X1,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),X2))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk9_4(X1,esk1_0,X2,X3),k5_numbers))).#trainneg
cnf(i_0_1125, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X5,X3,X4),X2))|~r1_xxreal_0(esk3_4(k8_metric_1,X3,X2,X1),X4)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X5))).#trainneg
cnf(i_0_1508, negated_conjecture, (X1=k1_tbsp_1(X2,esk1_0)|r1_xxreal_0(esk4_3(X2,esk1_0,X1),k4_metric_1(X2,X1,k1_seq_1(esk1_0,esk5_4(X2,esk1_0,X1,X3))))|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(X2,esk1_0,X1,X3)),u1_struct_0(X2))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2)))).#trainneg
cnf(i_0_1496, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1560, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k4_metric_1(X1,X2,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1600, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))))))).#trainneg
cnf(i_0_1124, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk3_4(k8_metric_1,X3,X2,X1),X4)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_1055, plain, (k1_metric_1(X3,X4,k7_metric_1,X1,X2)=k2_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X4),k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1598, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(k8_nat_1(X2,X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k5_numbers)|~v1_funct_2(X1,k5_numbers,X2))).#trainneg
cnf(i_0_1160, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,esk6_2(k8_metric_1,X3),k8_nat_1(X4,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(X4,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_1202, negated_conjecture, (v1_xboole_0(X1)|m2_subset_1(k1_seq_1(esk1_0,X2),X1,X1)|~v7_ordinal1(X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk1_0,k5_numbers,X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1213, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X4,X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(X4,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_876, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(X3,X4,X5)=k1_tbsp_1(X1,X2)|k8_nat_1(X6,X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X7))!=k1_tbsp_1(X1,X4)|~r1_xxreal_0(esk3_4(X1,X4,k8_nat_1(X6,X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),esk4_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(X6,X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X6)|~v1_funct_1(X2))).#trainneg
cnf(i_0_854, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,X1,k1_seq_1(X3,esk5_4(X2,X3,X1,X4))))|v2_struct_0(X2)|~v3_valued_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_seq_1(X3,esk5_4(X2,X3,X1,X4)),u1_struct_0(X2)))).#trainneg
cnf(i_0_1069, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_1(X3,X2),k1_seq_2(X3))))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_776, plain, (k2_metric_1(X1,X2,X3)=k18_complex1(k9_real_1(X2,X3))|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1))).#trainneg
cnf(i_0_1521, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X3))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(X2,X1,esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1196, negated_conjecture, (v1_xreal_0(k1_seq_1(esk1_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1571, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1348, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~m2_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk5_4(X1,X2,esk6_2(X1,X2),esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2)))),k5_numbers))).#trainneg
cnf(i_0_1443, negated_conjecture, (m1_subset_1(k1_seq_1(esk1_0,X2),X1)|~v7_ordinal1(X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk1_0,k5_numbers,X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1914, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|r1_xxreal_0(esk4_3(X1,X2,esk6_2(X1,X2)),k6_numbers)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m2_subset_1(esk7_3(X1,X2,esk4_3(X1,X2,esk6_2(X1,X2))),k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_1889, plain, (k18_complex1(k9_real_1(X1,X2))=k2_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1634, negated_conjecture, (m1_subset_1(k1_seq_1(esk1_0,X1),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1759, negated_conjecture, (r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|~r1_xxreal_0(esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))),esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),X2))|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),X2),k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v9_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),X2)),u1_struct_0(X1))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_1672, negated_conjecture, (m1_subset_1(k8_nat_1(X1,esk1_0,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk1_0,k5_numbers,X1))).#trainneg
cnf(i_0_1893, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,X2,k1_seq_1(X3,X4)))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1900, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k6_numbers)|v2_tbsp_1(X1,k8_metric_1)|X2!=k1_seq_2(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X1,X2,esk8_3(k8_metric_1,X1,X2)),esk9_4(k8_metric_1,X1,X2,X3))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1785, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~m2_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)))),k5_numbers))).#trainneg
cnf(i_0_872, plain, (r1_xxreal_0(esk8_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k8_nat_1(X6,X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X7))!=k1_tbsp_1(X1,X4)|~r1_xxreal_0(esk3_4(X1,X4,k8_nat_1(X6,X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),esk8_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(X6,X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X6)|~v1_funct_1(X2))).#trainneg
cnf(i_0_892, plain, (k1_metric_1(k1_numbers,k1_numbers,X4,X2,X3)=k2_metric_1(X1,X2,X3)|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k7_metric_1|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X4))).#trainneg
cnf(i_0_1764, negated_conjecture, (r1_xxreal_0(esk4_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(u1_struct_0(X1),X2,X3)=k1_tbsp_1(X1,esk1_0)|k1_seq_1(esk1_0,esk5_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4))!=k1_tbsp_1(X1,X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X2,k1_seq_1(esk1_0,esk5_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4)),esk4_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3))),X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_741, plain, (v2_tbsp_1(X1,k8_metric_1)|k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,k8_nat_1(k1_numbers,X2,X3),X4))!=k1_tbsp_1(k8_metric_1,X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(k8_metric_1,X2,k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,k8_nat_1(k1_numbers,X2,X3),X4)),esk8_3(k8_metric_1,X1,k8_nat_1(k1_numbers,X2,X3))),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,k8_nat_1(k1_numbers,X2,X3),X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_863, plain, (r1_xxreal_0(esk4_3(X1,X2,k1_seq_1(X3,X4)),k6_numbers)|v2_struct_0(X1)|k1_seq_1(X3,X4)=k1_tbsp_1(X1,X2)|k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k1_seq_1(X3,X4),X5))!=k1_tbsp_1(X1,X3)|~v3_valued_0(X3)|~r1_xxreal_0(esk3_4(X1,X3,k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k1_seq_1(X3,X4),X5)),esk4_3(X1,X2,k1_seq_1(X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k1_seq_1(X3,X4),X5)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k1_seq_1(X3,X4)),k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_1(X3,X4),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1236, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k4_metric_1(X2,X1,k1_funct_1(X3,esk5_4(X2,X3,X1,X5))))|v2_struct_0(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(k1_funct_1(X3,esk5_4(X2,X3,X1,X5)),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v7_ordinal1(esk5_4(X2,X3,X1,X5)))).#trainneg
cnf(i_0_1553, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(k8_metric_1,k1_funct_1(X1,esk9_4(k8_metric_1,X1,X2,X4)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(X1,esk9_4(k8_metric_1,X1,X2,X4)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X4),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~v7_ordinal1(esk9_4(k8_metric_1,X1,X2,X4)))).#trainneg
cnf(i_0_936, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(X3,X4,X5)=k1_tbsp_1(X1,X2)|k8_nat_1(X6,X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X7))!=k1_tbsp_1(X1,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X4,k8_nat_1(X6,X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),esk4_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(X6,X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X6)|~v1_funct_1(X2))).#trainneg
cnf(i_0_806, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,k1_funct_1(X2,esk9_4(X1,X2,X3,X5)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_funct_1(X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,X2,X3,X5)))).#trainneg
cnf(i_0_1206, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,esk1_0)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,X3,k1_seq_1(esk1_0,X4)))|~r1_xxreal_0(esk3_4(X2,esk1_0,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_1418, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,esk1_0)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(esk1_0,X4),X3))|~r1_xxreal_0(esk3_4(X2,esk1_0,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_1413, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(esk1_0,X3),esk6_2(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_1487, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k2_metric_1(k8_metric_1,k1_funct_1(X1,esk9_4(k8_metric_1,X1,X2,X4)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X4),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~v7_ordinal1(esk9_4(k8_metric_1,X1,X2,X4)))).#trainneg
cnf(i_0_1147, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k1_funct_1(X2,esk9_4(X1,X2,X3,X5))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_funct_1(X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,X2,X3,X5)))).#trainneg
cnf(i_0_987, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k1_seq_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_324, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1414, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,esk1_0),k1_seq_1(esk1_0,X3)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(esk1_0,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_708, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~v3_valued_0(X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,X3,k1_seq_1(X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(k1_seq_1(X4,X5),u1_struct_0(X2)))).#trainneg
cnf(i_0_540, plain, (m2_subset_1(k2_metric_1(X1,X2,X3),k1_numbers,k1_numbers)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1384, negated_conjecture, (m1_subset_1(k1_funct_1(esk1_0,X1),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_875, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(X3,X4,X5)=k1_tbsp_1(X1,X2)|k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6))!=k1_tbsp_1(X1,X4)|~r1_xxreal_0(esk3_4(X1,X4,k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),esk4_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~v3_valued_0(X2)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_932, plain, (r1_xxreal_0(esk8_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k8_nat_1(X6,X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X7))!=k1_tbsp_1(X1,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X4,k8_nat_1(X6,X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),esk8_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(X6,X2,esk9_4(X1,X2,k8_nat_1(X3,X4,X5),X7)),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~m2_subset_1(X7,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X6)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1988, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_1938, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk4_3(X1,X2,esk6_2(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_1973, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(esk1_0,X1,esk8_3(k8_metric_1,esk1_0,X1)),esk9_4(k8_metric_1,esk1_0,X1,X2))|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2146, negated_conjecture, (r1_xxreal_0(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,esk1_0)=k1_tbsp_1(X1,esk1_0)|~v7_ordinal1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),X2))|~m2_subset_1(esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X1,esk1_0,esk4_3(X1,esk1_0,esk6_2(X1,esk1_0))),esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),X2))|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk4_3(X1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v9_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(X1,esk1_0,esk6_2(X1,esk1_0),X2)),u1_struct_0(X1))|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_2124, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k2_metric_1(k8_metric_1,k1_funct_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),X1))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2168, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k1_seq_1(X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers))).#trainneg
cnf(i_0_2044, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k4_metric_1(k8_metric_1,k1_funct_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),X1))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1987, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k6_numbers)|v2_tbsp_1(X1,k8_metric_1)|X2!=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(esk12_3(X1,X2,esk8_3(k8_metric_1,X1,X2)),k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,esk12_3(X1,X2,esk8_3(k8_metric_1,X1,X2)))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,esk12_3(X1,X2,esk8_3(k8_metric_1,X1,X2))),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2140, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(X3,X4),esk6_2(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~m1_subset_1(k1_seq_1(X3,X4),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_1955, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|~m2_subset_1(esk5_4(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)))),k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers)|~v7_ordinal1(esk5_4(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0))))))).#trainneg
cnf(i_0_2132, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k4_metric_1(X1,X2,k1_funct_1(esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,X2,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1979, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,X2),k6_numbers)|X2=k1_tbsp_1(k8_metric_1,X1)|X2!=k1_seq_2(X1)|~v1_xreal_0(esk4_3(k8_metric_1,X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(esk5_4(k8_metric_1,X1,X2,X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X1,X2,esk4_3(k8_metric_1,X1,X2)),esk5_4(k8_metric_1,X1,X2,X3))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk5_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk5_4(k8_metric_1,X1,X2,X3),k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2167, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k1_funct_1(X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_funct_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,X2,X3,X4)))).#trainneg
cnf(i_0_1981, negated_conjecture, (r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,X1),k6_numbers)|X1=k1_tbsp_1(k8_metric_1,esk1_0)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk4_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk5_4(k8_metric_1,esk1_0,X1,X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1)),esk5_4(k8_metric_1,esk1_0,X1,X2))|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(k8_metric_1,esk1_0,X1,X2)),k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_318, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_595, plain, (g1_metric_1(u1_struct_0(X1),k7_metric_1)=X1|~v1_metric_1(X1)|~l1_metric_1(X1)|X1!=k8_metric_1)).#trainneg
cnf(i_0_720, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k1_seq_1(X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5))!=k1_tbsp_1(X1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k1_seq_1(X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5),k5_numbers))).#trainneg
cnf(i_0_988, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,esk6_2(X2,X4),k8_nat_1(X5,X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_572, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers))).#trainneg
cnf(i_0_608, plain, (u1_struct_0(X1)=X2|g1_metric_1(u1_struct_0(X1),k7_metric_1)!=g1_metric_1(X2,X3)|~l1_metric_1(X1)|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_1061, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X2),esk15_1(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(esk15_1(X3),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_719, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|k1_funct_1(X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5))!=k1_tbsp_1(X1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k1_funct_1(X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),esk8_3(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(k1_funct_1(X2,esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,X2,k8_nat_1(u1_struct_0(X1),X3,X4),X5)))).#trainneg
cnf(i_0_1203, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(esk1_0,X3),esk6_2(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_706, plain, (k1_metric_1(X1,X2,k7_metric_1,X4,X5)=k2_metric_1(X3,X4,X5)|~l1_metric_1(X3)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X5,u1_struct_0(X3))|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X1)|X3!=k8_metric_1|~v1_metric_1(X3))).#trainneg
cnf(i_0_642, plain, (m1_subset_1(k1_funct_1(X2,X3),k1_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1978, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,k1_seq_1(X2,X3)),k6_numbers)|k1_seq_1(X2,X3)=k1_tbsp_1(k8_metric_1,X1)|k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,k1_seq_1(X2,X3),X4))!=k1_seq_2(X2)|~v1_xreal_0(esk4_3(k8_metric_1,X1,k1_seq_1(X2,X3)))|~v1_xreal_0(k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,k1_seq_1(X2,X3),X4)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X2,k8_nat_1(k1_numbers,X1,esk5_4(k8_metric_1,X1,k1_seq_1(X2,X3),X4)),esk4_3(k8_metric_1,X1,k1_seq_1(X2,X3))),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1064, plain, (r1_xxreal_0(esk17_2(X1,X2),k2_metric_1(k8_metric_1,k1_seq_1(X1,esk18_3(X1,X2,X3)),X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1986, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,X2),k6_numbers)|X2=k1_tbsp_1(k8_metric_1,X1)|X2!=k1_seq_2(X1)|~v1_xreal_0(esk4_3(k8_metric_1,X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(esk5_4(k8_metric_1,X1,X2,X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X1,X2,esk4_3(k8_metric_1,X1,X2)),esk5_4(k8_metric_1,X1,X2,X3))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk5_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_935, plain, (r1_xxreal_0(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k6_numbers)|v2_struct_0(X1)|k8_nat_1(X3,X4,X5)=k1_tbsp_1(X1,X2)|k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6))!=k1_tbsp_1(X1,X4)|~v7_ordinal1(X5)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X4,k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),esk4_3(X1,X2,k8_nat_1(X3,X4,X5))),X5)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k1_seq_1(X2,esk5_4(X1,X2,k8_nat_1(X3,X4,X5),X6)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k8_nat_1(X3,X4,X5)),k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~v3_valued_0(X2)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X3,X4,X5),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_574, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_1265, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,X2,k1_seq_1(X3,X4)))|~r1_xxreal_0(esk3_4(k8_metric_1,X3,X2,X1),X4)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_2235, plain, (esk6_2(X1,X2)=k1_tbsp_1(X1,X2)|v2_struct_0(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2496, plain, (r1_xxreal_0(esk17_2(X1,esk6_2(k8_metric_1,X1)),k6_numbers)|v2_comseq_2(X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk17_2(X1,esk6_2(k8_metric_1,X1))),esk18_3(X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk17_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk18_3(X1,esk6_2(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,esk6_2(k8_metric_1,X1),X2)),k1_numbers))).#trainneg
cnf(i_0_2494, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0))),esk18_3(esk1_0,esk6_2(k8_metric_1,esk1_0),X1))|~v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk18_3(esk1_0,esk6_2(k8_metric_1,esk1_0),X1),k5_numbers)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,esk6_2(k8_metric_1,esk1_0),X1)),k1_numbers))).#trainneg
cnf(i_0_2327, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k1_seq_1(X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v7_ordinal1(esk9_4(X1,X2,X3,X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2277, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k4_metric_1(k8_metric_1,k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),X1))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2520, plain, (r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,X1,esk6_2(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_2500, plain, (r1_xxreal_0(esk17_2(X1,X2),k4_metric_1(k8_metric_1,k1_seq_1(X1,esk18_3(X1,X2,X3)),X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2280, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k4_metric_1(k8_metric_1,X1,k1_funct_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2))))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_funct_1(esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),k1_numbers))).#trainneg
cnf(i_0_2380, plain, (k1_binop_1(k7_metric_1,X2,X3)=k4_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_2233, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk7_3(k8_metric_1,X1,esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1))),k5_numbers))).#trainneg
cnf(i_0_2234, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|r1_xxreal_0(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2278, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k4_metric_1(k8_metric_1,k8_nat_1(X3,esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),X1))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk9_4(k8_metric_1,esk1_0,X1,X2))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_2242, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k2_metric_1(k8_metric_1,k8_nat_1(X3,esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X2)),X1))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X2),k5_numbers)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk9_4(k8_metric_1,esk1_0,X1,X2))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_2375, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1)))).#trainneg
cnf(i_0_2328, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X1,X3,k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v7_ordinal1(esk9_4(X1,X2,X3,X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~v1_funct_2(X2,k5_numbers,X5))).#trainneg
cnf(i_0_2286, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k6_numbers)|v2_tbsp_1(X1,k8_metric_1)|X2!=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,X2))|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,esk12_3(X1,X2,esk8_3(k8_metric_1,X1,X2))),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2372, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1)))).#trainneg
cnf(i_0_2412, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,esk15_1(X1)),k6_numbers)|v2_tbsp_1(X1,k8_metric_1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,esk15_1(X1)))|~m2_subset_1(esk9_4(k8_metric_1,X1,esk15_1(X1),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk16_2(X1,esk8_3(k8_metric_1,X1,esk15_1(X1))),esk9_4(k8_metric_1,X1,esk15_1(X1),X2))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,esk15_1(X1),X2)),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,X1,esk15_1(X1),X2),k5_numbers)|~m1_subset_1(esk15_1(X1),k1_numbers))).#trainneg
cnf(i_0_2295, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|~m2_subset_1(esk5_4(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)))),k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)))),k5_numbers))).#trainneg
cnf(i_0_2329, negated_conjecture, (X1=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,X1),k6_numbers)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk4_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk5_4(k8_metric_1,esk1_0,X1,esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1))),k1_numbers,k5_numbers)|~m2_subset_1(esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(k8_metric_1,esk1_0,X1,esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1)))),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2310, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k4_metric_1(X1,X2,k8_nat_1(X4,esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,X2,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(esk1_0,k5_numbers,X4))).#trainneg
cnf(i_0_2597, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk18_3(esk1_0,esk6_2(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)))),k5_numbers)|~m1_subset_1(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_2583, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2585, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2738, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_2566, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2595, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1))|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1),k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1))).#trainneg
cnf(i_0_2842, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,esk15_1(X1)),k6_numbers)|v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,esk15_1(X1)))|~m2_subset_1(esk16_2(X1,esk8_3(k8_metric_1,X1,esk15_1(X1))),k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(esk9_4(k8_metric_1,X1,esk15_1(X1),esk16_2(X1,esk8_3(k8_metric_1,X1,esk15_1(X1)))),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk15_1(X1),k1_numbers))).#trainneg
cnf(i_0_2695, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2)))).#trainneg
cnf(i_0_2713, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4)),X1))|v2_struct_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4)),u1_struct_0(X2)))).#trainneg
cnf(i_0_2626, plain, (k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k1_binop_1(k7_metric_1,X2,X3)|X1!=k7_metric_1|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2748, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk4_3(k8_metric_1,X1,esk6_2(k8_metric_1,X1)),k1_numbers))).#trainneg
cnf(i_0_2563, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2733, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k1_binop_1(k7_metric_1,X1,k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4))))|v2_struct_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X3,esk5_4(X2,X3,X1,X4)),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2))).#trainneg
cnf(i_0_2574, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(X3,X4),k1_tbsp_1(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2582, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X4,X5,X3),k1_tbsp_1(X2,X5)))|~r1_xxreal_0(esk7_3(X2,X5,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X5,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X5,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X5,k5_numbers,X4)|~v1_funct_1(X5)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2843, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk4_3(k8_metric_1,esk1_0,esk6_2(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk6_2(k8_metric_1,esk1_0),k1_numbers))).#trainneg
cnf(i_0_2581, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(X4,X5,X3),k1_tbsp_1(X2,X5)))|~r1_xxreal_0(esk7_3(X2,X5,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X5,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X5,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X5,k5_numbers,X4)|~v1_funct_1(X5)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2628, plain, (k1_metric_1(X1,X2,X3,X4,X5)=k1_binop_1(k7_metric_1,X4,X5)|X3!=k7_metric_1|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X5,k1_numbers)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X3,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2823, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k6_numbers)|v2_tbsp_1(X1,k8_metric_1)|X2!=k1_seq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,X2))|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(esk12_3(X1,X2,esk8_3(k8_metric_1,X1,X2)),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2580, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X3),k1_seq_1(X3,X4)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2577, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(X3,X4),k1_tbsp_1(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2564, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2575, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X3),k1_seq_1(X3,X4)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2744, plain, (k4_metric_1(X3,X1,X2)=k2_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|X3!=k8_metric_1|~v1_metric_1(X3)|~l1_metric_1(X3)|~v8_metric_1(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_2923, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_2878, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_2882, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2933, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2927, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,esk6_2(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2926, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,esk15_1(X1)),k6_numbers)|v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,esk15_1(X1)))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(esk16_2(X1,esk8_3(k8_metric_1,X1,esk15_1(X1))),k1_numbers,k5_numbers)|~m1_subset_1(esk15_1(X1),k1_numbers))).#trainneg
cnf(i_0_2883, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(X5,X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_2794, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),k1_numbers))).#trainneg
cnf(i_0_2980, plain, (esk6_2(k8_metric_1,X1)=k1_tbsp_1(k8_metric_1,X1)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_3093, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(k1_numbers,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2588, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2558, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,X2),k1_seq_1(X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2915, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2724, plain, (k4_metric_1(X1,X2,X3)=k1_binop_1(k7_metric_1,X3,X2)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_2911, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(k8_metric_1,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_2832, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers))).#trainneg
cnf(i_0_2586, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X4),k1_funct_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2584, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_funct_1(X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2579, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_seq_1(X3,X4),k1_tbsp_1(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2846, negated_conjecture, (X1=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,X1),k6_numbers)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk4_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk5_4(k8_metric_1,esk1_0,X1,esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1))),k1_numbers,k5_numbers)|~m2_subset_1(esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0))).#trainneg
cnf(i_0_2587, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_funct_1(X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3499, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(esk6_2(k8_metric_1,esk1_0),k1_numbers))).#trainneg
cnf(i_0_3217, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0))),k5_numbers))).#trainneg
cnf(i_0_3254, plain, (v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk15_1(X1),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,esk15_1(X1)))|~v2_comseq_2(X1))).#trainneg
cnf(i_0_3215, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_3224, plain, (r1_xxreal_0(esk17_2(X1,k1_tbsp_1(k8_metric_1,X1)),k6_numbers)|v2_comseq_2(X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk17_2(X1,k1_tbsp_1(k8_metric_1,X1))),esk18_3(X1,k1_tbsp_1(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk17_2(X1,k1_tbsp_1(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk18_3(X1,k1_tbsp_1(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,k1_tbsp_1(k8_metric_1,X1),X2)),k1_numbers))).#trainneg
cnf(i_0_3230, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3450, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(k8_metric_1,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3238, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers))).#trainneg
cnf(i_0_3250, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,esk6_2(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(esk6_2(k8_metric_1,X3),k1_numbers))).#trainneg
cnf(i_0_2309, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k4_metric_1(X1,X2,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,X2,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3179, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X2),X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers))).#trainneg
cnf(i_0_2411, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,esk15_1(esk1_0)),k6_numbers)|v2_tbsp_1(esk1_0,k8_metric_1)|~v1_xreal_0(esk8_3(k8_metric_1,esk1_0,esk15_1(esk1_0)))|~m2_subset_1(esk9_4(k8_metric_1,esk1_0,esk15_1(esk1_0),X1),k1_numbers,k5_numbers)|~r1_xxreal_0(esk16_2(esk1_0,esk8_3(k8_metric_1,esk1_0,esk15_1(esk1_0))),esk9_4(k8_metric_1,esk1_0,esk15_1(esk1_0),X1))|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(k8_metric_1,esk1_0,esk15_1(esk1_0),X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,esk15_1(esk1_0),X1),k5_numbers)|~m1_subset_1(esk15_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_2594, plain, (v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk17_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,esk6_2(k8_metric_1,X1),esk7_3(k8_metric_1,X1,esk17_2(X1,esk6_2(k8_metric_1,X1))))),k1_numbers)|~m1_subset_1(esk17_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2935, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers))).#trainneg
cnf(i_0_3225, plain, (r1_xxreal_0(esk13_2(X1,k1_tbsp_1(k8_metric_1,X1)),k6_numbers)|k1_seq_2(X1)=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk7_3(k8_metric_1,X1,esk13_2(X1,k1_tbsp_1(k8_metric_1,X1))),esk14_3(X1,k1_tbsp_1(k8_metric_1,X1),X2))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk13_2(X1,k1_tbsp_1(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(esk14_3(X1,k1_tbsp_1(k8_metric_1,X1),X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_2(X1)|~m1_subset_1(k1_seq_1(X1,esk14_3(X1,k1_tbsp_1(k8_metric_1,X1),X2)),k1_numbers))).#trainneg
cnf(i_0_3218, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1))).#trainneg
cnf(i_0_2661, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,X1),k6_numbers)|v2_comseq_2(esk1_0)|X1!=k1_tbsp_1(k8_metric_1,esk1_0)|~r1_xxreal_0(esk3_4(k8_metric_1,esk1_0,X1,esk17_2(esk1_0,X1)),esk18_3(esk1_0,X1,X2))|~v2_tbsp_1(esk1_0,k8_metric_1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk17_2(esk1_0,X1),k1_numbers)|~m1_subset_1(esk18_3(esk1_0,X1,X2),k5_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,X1,X2)),k1_numbers))).#trainneg
cnf(i_0_2735, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,X3,k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_3304, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,X3),k1_funct_1(X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3244, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,esk6_2(k8_metric_1,X3),k1_seq_1(X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2493, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k4_metric_1(k8_metric_1,k1_seq_1(X2,esk14_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2837, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(k8_metric_1,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),k1_numbers))).#trainneg
cnf(i_0_2718, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k1_binop_1(k7_metric_1,k8_nat_1(X4,X3,esk5_4(X2,X3,X1,X5)),X1))|v2_struct_0(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k8_nat_1(X4,X3,esk5_4(X2,X3,X1,X5)),u1_struct_0(X2)))).#trainneg
cnf(i_0_2734, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,X3,k8_nat_1(u1_struct_0(X2),X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2))).#trainneg
cnf(i_0_3510, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_tbsp_1(esk1_0,k8_metric_1))).#trainneg
cnf(i_0_3685, negated_conjecture, (v2_comseq_2(esk1_0)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,esk6_2(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_3514, plain, (v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~v1_xreal_0(esk8_3(k8_metric_1,X1,k1_seq_2(X1)))|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3670, plain, (v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk17_2(X1,esk6_2(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk17_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3681, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3605, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3536, negated_conjecture, (m2_subset_1(k1_tbsp_1(X1,esk1_0),u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_3649, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X2),X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3516, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_comseq_2(esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k5_numbers))).#trainneg
cnf(i_0_3300, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_funct_1(X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2896, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X4),k8_nat_1(X5,X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X5,X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_2824, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2370, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k1_binop_1(k7_metric_1,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),X2))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk9_4(X1,esk1_0,X2,X3),k5_numbers)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1)))).#trainneg
cnf(i_0_3609, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,esk6_2(k8_metric_1,X3),k1_seq_1(X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2591, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X4),k8_nat_1(X5,X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X4,k5_numbers,X5)|~v1_funct_1(X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_2774, negated_conjecture, (r1_xxreal_0(esk8_3(k8_metric_1,esk1_0,X1),k4_metric_1(k8_metric_1,X1,k8_nat_1(X2,esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X3))))|v2_tbsp_1(esk1_0,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(k8_metric_1,esk1_0,X1,X3),k5_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,X2)|~m1_subset_1(k8_nat_1(X2,esk1_0,esk9_4(k8_metric_1,esk1_0,X1,X3)),k1_numbers))).#trainneg
cnf(i_0_3727, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,X2),k1_seq_1(X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3305, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_funct_1(X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3193, plain, (k4_metric_1(X1,X2,X3)=k2_metric_1(k8_metric_1,X3,X2)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_3858, plain, (v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,X1))|~m2_subset_1(esk7_3(k8_metric_1,X1,esk17_2(X1,k1_tbsp_1(k8_metric_1,X1))),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk17_2(X1,k1_tbsp_1(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3852, negated_conjecture, (v2_comseq_2(esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1))).#trainneg
cnf(i_0_3942, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk14_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))))),k1_numbers))).#trainneg
cnf(i_0_4003, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,X2),k1_seq_1(X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3861, plain, (v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk17_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk7_3(k8_metric_1,X1,esk17_2(X1,esk6_2(k8_metric_1,X1))),k5_numbers))).#trainneg
cnf(i_0_3957, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3988, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers))).#trainneg
cnf(i_0_4065, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4089, plain, (k4_metric_1(X3,X2,X1)=k4_metric_1(k8_metric_1,X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|X3!=k8_metric_1|~v1_metric_1(X3)|~l1_metric_1(X3)|~v8_metric_1(X3)|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X3)))).#trainneg
cnf(i_0_3953, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_seq_2(esk1_0)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,esk6_2(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0)|~v7_ordinal1(esk14_3(esk1_0,esk6_2(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,esk6_2(k8_metric_1,esk1_0))))))).#trainneg
cnf(i_0_3950, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_seq_2(esk1_0)|~v1_xreal_0(esk6_2(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,esk6_2(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~v2_tbsp_1(esk1_0,k8_metric_1)|~v2_comseq_2(esk1_0)|~m1_subset_1(esk14_3(esk1_0,esk6_2(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,esk6_2(k8_metric_1,esk1_0)))),k5_numbers))).#trainneg
cnf(i_0_3510, negated_conjecture, (esk6_2(k8_metric_1,esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0))).#trainneg
cnf(i_0_4191, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~v2_comseq_2(esk1_0)|~m1_subset_1(esk14_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k5_numbers))).#trainneg
cnf(i_0_4190, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~v2_comseq_2(esk1_0)|~v7_ordinal1(esk14_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))))))).#trainneg
cnf(i_0_4192, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~v2_comseq_2(esk1_0))).#trainneg
cnf(i_0_1, negated_conjecture, (k2_seq_2(esk1_0)!=k1_tbsp_1(k8_metric_1,esk1_0)|~v2_comseq_2(esk1_0))).#trainneg
cnf(i_0_4200, plain, (v2_comseq_2(X1)|~v3_valued_0(X1)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,X1))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk7_3(k8_metric_1,X1,esk17_2(X1,k1_tbsp_1(k8_metric_1,X1))),k5_numbers)|~m1_subset_1(esk17_2(X1,k1_tbsp_1(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4374, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers))).#trainneg
cnf(i_0_4189, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk17_2(X1,k1_tbsp_1(k8_metric_1,X1)),k6_numbers)|~v3_valued_0(X1)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,X1))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk17_2(X1,k1_tbsp_1(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4216, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2846, negated_conjecture, (X1=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,X1),k6_numbers)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk4_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk5_4(k8_metric_1,esk1_0,X1,esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1))),k1_numbers,k5_numbers)|~m2_subset_1(esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~v2_comseq_2(esk1_0))).#trainneg
cnf(i_0_4202, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk17_2(X1,esk6_2(k8_metric_1,X1)),k6_numbers)|~v3_valued_0(X1)|~v1_xreal_0(esk6_2(k8_metric_1,X1))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(esk17_2(X1,esk6_2(k8_metric_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4348, plain, (r1_xxreal_0(esk17_2(X1,X2),k4_metric_1(k8_metric_1,X2,k1_seq_1(X1,esk18_3(X1,X2,X3))))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3296, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,X2),k1_seq_1(X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2378, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1)))).#trainneg
cnf(i_0_4356, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,X3,k8_nat_1(u1_struct_0(X2),X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2727, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,X3,k1_seq_1(X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v3_valued_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_2503, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k6_numbers)|X1!=k1_seq_2(X2)|~v3_valued_0(X2)|~v1_xreal_0(esk4_3(k8_metric_1,X2,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk5_4(k8_metric_1,X2,X1,esk12_3(X2,X1,esk4_3(k8_metric_1,X2,X1))),k1_numbers,k5_numbers)|~m2_subset_1(esk12_3(X2,X1,esk4_3(k8_metric_1,X2,X1)),k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk5_4(k8_metric_1,X2,X1,esk12_3(X2,X1,esk4_3(k8_metric_1,X2,X1)))),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4180, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(k8_metric_1,k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),k1_numbers))).#trainneg
cnf(i_0_4146, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,X3,k8_nat_1(u1_struct_0(X2),X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_2312, plain, (X1=k1_tbsp_1(k8_metric_1,X2)|r1_xxreal_0(esk4_3(k8_metric_1,X2,X1),k6_numbers)|X1!=k1_seq_2(X2)|~v3_valued_0(X2)|~v1_xreal_0(esk4_3(k8_metric_1,X2,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(X2,X1,esk4_3(k8_metric_1,X2,X1)),k1_numbers,k5_numbers)|~v2_tbsp_1(X2,k8_metric_1)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk5_4(k8_metric_1,X2,X1,esk12_3(X2,X1,esk4_3(k8_metric_1,X2,X1)))),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk5_4(k8_metric_1,X2,X1,esk12_3(X2,X1,esk4_3(k8_metric_1,X2,X1))),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2904, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k1_binop_1(k7_metric_1,X2,k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2787, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k1_binop_1(k7_metric_1,k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4455, negated_conjecture, (m2_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),k1_numbers,k1_numbers))).#trainneg
cnf(i_0_4506, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~v2_comseq_2(esk1_0)|~m1_subset_1(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k5_numbers))).#trainneg
cnf(i_0_4465, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(k1_numbers,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4459, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4456, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4473, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_funct_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4462, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4507, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~v2_comseq_2(esk1_0)|~m1_subset_1(esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_4502, negated_conjecture, (k1_seq_2(esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk13_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk14_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1))|~v2_comseq_2(esk1_0)|~m1_subset_1(k1_seq_1(esk1_0,esk14_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1)),k1_numbers))).#trainneg
cnf(i_0_4467, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_funct_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4470, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(k1_numbers,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4474, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_funct_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4468, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(k1_numbers,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4719, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k1_binop_1(k7_metric_1,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),X2))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_4464, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X3,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_4491, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(k1_numbers,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4478, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(X3,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_4463, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(X3,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_4485, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4475, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(k1_numbers,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4469, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(k1_numbers,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4486, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4461, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4460, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4457, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4546, negated_conjecture, (X1=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,X1),k6_numbers)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk4_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_3(esk1_0,X1,esk4_3(k8_metric_1,esk1_0,X1)),k1_numbers,k5_numbers)|~v2_comseq_2(esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4705, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k1_binop_1(k7_metric_1,X2,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_3295, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X2,X3),k1_tbsp_1(k8_metric_1,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2916, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(X4,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_2410, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X3,X2),esk15_1(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(esk15_1(X3),k1_numbers))).#trainneg
cnf(i_0_2409, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X3,X2),k1_seq_2(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2559, plain, (r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,X2),k8_nat_1(X3,X2,X4)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X4)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(k8_nat_1(X3,X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3967, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X4,X3,X5),k1_tbsp_1(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X4,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~m1_subset_1(k1_tbsp_1(X2,X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_3613, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2905, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k1_binop_1(k7_metric_1,X2,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk9_4(X1,esk1_0,X2,X3),k5_numbers))).#trainneg
cnf(i_0_3856, plain, (v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk8_3(k8_metric_1,X1,k1_seq_2(X1)),k1_numbers))).#trainneg
cnf(i_0_3460, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4669, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(k8_metric_1,k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4816, plain, (esk6_2(k8_metric_1,X1)=k1_seq_2(X1)|~v3_valued_0(X1)|~v2_tbsp_1(X1,k8_metric_1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk6_2(k8_metric_1,X1),k1_numbers))).#trainneg
cnf(i_0_5047, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4732, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(X3,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_5036, plain, (v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_2(X1),k1_numbers))).#trainneg
cnf(i_0_4988, plain, (r1_xxreal_0(esk8_3(X1,X2,k1_seq_2(X2)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v1_xreal_0(esk8_3(X1,X2,k1_seq_2(X2)))|~m2_subset_1(esk9_4(X1,X2,k1_seq_2(X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk16_2(X2,esk8_3(X1,X2,k1_seq_2(X2))),esk9_4(X1,X2,k1_seq_2(X2),X3))|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_4870, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2))|~m2_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2)),k1_numbers)|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_2306, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|k1_funct_1(esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4))!=k1_tbsp_1(X1,X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X2,k1_funct_1(esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4)),esk8_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3))),X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_funct_1(esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4)),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_2714, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X2),X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2)))).#trainneg
cnf(i_0_2447, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(X5,X6,k7_metric_1,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),k1_numbers)))|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X4),X6)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),X5))).#trainneg
cnf(i_0_3696, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,X1),k6_numbers)|v2_comseq_2(esk1_0)|X1!=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(X1)|~m2_subset_1(esk3_4(k8_metric_1,esk1_0,X1,esk17_2(esk1_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(esk18_3(esk1_0,X1,esk3_4(k8_metric_1,esk1_0,X1,esk17_2(esk1_0,X1))),k5_numbers)|~m1_subset_1(esk17_2(esk1_0,X1),k1_numbers))).#trainneg
cnf(i_0_2469, plain, (k2_metric_1(X5,X3,X4)=k18_complex1(k9_real_1(X3,X4))|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X1)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(X4,u1_struct_0(X5))|~m1_subset_1(X3,u1_struct_0(X5)))).#trainneg
cnf(i_0_2796, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X5,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(X3,u1_struct_0(X5))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5)))).#trainneg
cnf(i_0_2930, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,esk6_2(X2,X4),k1_seq_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2665, plain, (r1_xxreal_0(esk17_2(X1,X2),k6_numbers)|v2_comseq_2(X1)|X2!=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~r1_xxreal_0(esk3_4(k8_metric_1,X1,X2,esk17_2(X1,X2)),esk18_3(X1,X2,X3))|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(esk17_2(X1,X2),k1_numbers)|~m1_subset_1(esk18_3(X1,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers))).#trainneg
cnf(i_0_2621, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(X3,X2),esk15_1(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(esk15_1(X3),k1_numbers)|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers))).#trainneg
cnf(i_0_2795, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X5,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(X3,u1_struct_0(X5))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5)))).#trainneg
cnf(i_0_4934, negated_conjecture, (X1=k1_tbsp_1(k8_metric_1,esk1_0)|r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,X1),k6_numbers)|X1!=k1_seq_2(esk1_0)|~v1_xreal_0(esk4_3(k8_metric_1,esk1_0,X1))|~v1_xreal_0(X1)|~v2_comseq_2(esk1_0))).#trainneg
cnf(i_0_2729, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k1_binop_1(k7_metric_1,X2,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_2624, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(X3,X4),X2))|~r1_xxreal_0(esk12_3(X3,X2,X1),X4)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers))).#trainneg
cnf(i_0_4981, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X3,X2),k1_seq_2(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(k1_seq_2(X3),k1_numbers))).#trainneg
cnf(i_0_5102, plain, (k1_seq_2(X1)=k1_tbsp_1(k8_metric_1,X1)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v2_comseq_2(X1))).#trainneg
cnf(i_0_5184, plain, (v2_tbsp_1(k7_metric_1,k8_metric_1)|~v2_comseq_2(k7_metric_1)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k7_metric_1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_5251, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,X1),k6_numbers)|v2_comseq_2(esk1_0)|X1!=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(X1)|~m2_subset_1(esk3_4(k8_metric_1,esk1_0,X1,esk17_2(esk1_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(esk17_2(esk1_0,X1),k1_numbers))).#trainneg
cnf(i_0_5252, plain, (k2_metric_1(X1,X2,X3)=k18_complex1(k9_real_1(X2,X3))|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5113, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X3,X2),k1_seq_2(X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v2_comseq_2(X3))).#trainneg
cnf(i_0_5294, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(X3,X2),k1_seq_2(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_5180, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2),k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2)),k1_numbers)|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_5187, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v7_ordinal1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))))|~m2_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))))),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5378, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_2(X3),k1_seq_1(X3,X2)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(k1_seq_2(X3),k1_numbers))).#trainneg
cnf(i_0_3765, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X3!=k1_seq_2(X2)|~v1_xreal_0(esk8_3(X1,X2,X3))|~v1_xreal_0(X3)|~m2_subset_1(esk9_4(X1,X2,X3,X4),k1_numbers,k5_numbers)|~r1_xxreal_0(esk12_3(X2,X3,esk8_3(X1,X2,X3)),esk9_4(X1,X2,X3,X4))|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3283, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(k8_metric_1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2427, negated_conjecture, (v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|k1_seq_1(esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4))!=k1_tbsp_1(X1,X2)|~v7_ordinal1(esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X2,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4)),esk8_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3))),X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X4),k5_numbers))).#trainneg
cnf(i_0_3281, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(k8_metric_1,k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(X5,X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(esk9_4(X1,X2,X3,X4),k5_numbers)|~v1_funct_2(X2,k5_numbers,X5))).#trainneg
cnf(i_0_3963, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X3),k1_funct_1(X3,X5)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2789, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,k1_funct_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_funct_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,X2,X3,X4)))).#trainneg
cnf(i_0_5282, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v3_valued_0(X3)|~v1_metric_1(X2)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(X2,X3),k1_seq_1(X3,X4)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_4563, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,X1),k6_numbers)|v2_comseq_2(esk1_0)|X1!=k1_tbsp_1(k8_metric_1,esk1_0)|~r1_xxreal_0(esk3_4(k8_metric_1,esk1_0,X1,esk17_2(esk1_0,X1)),esk18_3(esk1_0,X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk17_2(esk1_0,X1),k1_numbers)|~m1_subset_1(esk18_3(esk1_0,X1,X2),k5_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,X1,X2)),k1_numbers))).#trainneg
cnf(i_0_5347, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2))|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2),k5_numbers)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2)),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_4941, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k2_metric_1(k8_metric_1,X2,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers))).#trainneg
cnf(i_0_2451, plain, (k1_metric_1(X4,X5,k7_metric_1,X2,X3)=k4_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X5),k1_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_5393, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,X1),k6_numbers)|X1!=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(X1)|~m1_subset_1(esk17_2(esk1_0,X1),k1_numbers)|~m1_subset_1(esk3_4(k8_metric_1,esk1_0,X1,esk17_2(esk1_0,X1)),k5_numbers))).#trainneg
cnf(i_0_5541, plain, (r1_xxreal_0(esk8_3(X1,X2,k1_seq_2(X2)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v3_valued_0(X2)|~m2_subset_1(esk9_4(X1,X2,k1_seq_2(X2),X3),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),esk9_4(X1,X2,k1_seq_2(X2),X3))|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk8_3(X1,X2,k1_seq_2(X2)),k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(k1_numbers,X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(k1_numbers,X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5434, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5394, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,X1),k6_numbers)|X1!=k1_tbsp_1(k8_metric_1,esk1_0)|~v1_xreal_0(X1)|~m1_subset_1(esk17_2(esk1_0,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5435, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(X4,X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_5437, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(k1_numbers,X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers))).#trainneg
cnf(i_0_5565, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_binop_1(k7_metric_1,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v7_ordinal1(esk9_4(X1,X2,X3,X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5427, plain, (k18_complex1(k9_real_1(X2,X3))=k4_metric_1(X1,X2,X3)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5420, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers))).#trainneg
cnf(i_0_5455, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))))),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))))),k1_numbers)|~m1_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k5_numbers)|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5419, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers))).#trainneg
cnf(i_0_5458, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v7_ordinal1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))))|~m2_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_funct_1(esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))))),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5459, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v7_ordinal1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))))|~m2_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))))),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k5_numbers))).#trainneg
cnf(i_0_5535, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~m2_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2))|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X2)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5463, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v7_ordinal1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))))|~m2_subset_1(esk9_4(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5599, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(X5,X6,k7_metric_1,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),k1_numbers)))|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),X6)|~m1_subset_1(esk6_2(X2,X4),X5))).#trainneg
cnf(i_0_5635, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(X5,X6,k7_metric_1,k1_tbsp_1(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),k1_numbers)))|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),X6)|~m1_subset_1(k1_tbsp_1(X2,X4),X5))).#trainneg
cnf(i_0_2912, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X5,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5))|~m1_subset_1(X3,u1_struct_0(X5)))).#trainneg
cnf(i_0_5702, negated_conjecture, (~v1_xreal_0(X1)|X1!=k1_tbsp_1(k8_metric_1,esk1_0)|~m1_subset_1(esk17_2(esk1_0,X1),k1_numbers))).#trainneg
cnf(i_0_5835, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(k1_numbers,X3,X2),k1_tbsp_1(k8_metric_1,X3))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(k1_tbsp_1(k8_metric_1,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_5715, plain, (r1_xxreal_0(esk8_3(X1,X2,k1_seq_2(X2)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v3_valued_0(X2)|~r1_xxreal_0(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),esk9_4(X1,X2,k1_seq_2(X2),X3))|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk8_3(X1,X2,k1_seq_2(X2)),k1_numbers)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),k1_numbers)|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~m1_subset_1(k1_seq_2(X2),k1_numbers))).#trainneg
cnf(i_0_5722, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_2(X2),k8_nat_1(k1_numbers,X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_2(X2),k1_numbers))).#trainneg
cnf(i_0_5839, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(k1_numbers,X3,X2),esk6_2(k8_metric_1,X3))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_5748, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers))).#trainneg
cnf(i_0_5784, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers))).#trainneg
cnf(i_0_5701, plain, (r1_xxreal_0(esk8_3(X1,X2,k1_seq_2(X2)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v3_valued_0(X2)|~r1_xxreal_0(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),esk9_4(X1,X2,k1_seq_2(X2),X3))|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk8_3(X1,X2,k1_seq_2(X2)),k1_numbers)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),u1_struct_0(X1))|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5717, plain, (r1_xxreal_0(esk8_3(X1,X2,k1_seq_2(X2)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v3_valued_0(X2)|~r1_xxreal_0(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),esk9_4(X1,X2,k1_seq_2(X2),X4))|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(esk8_3(X1,X2,k1_seq_2(X2)),k1_numbers)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),X4),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X3)|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(X3,X2,esk9_4(X1,X2,k1_seq_2(X2),X4)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(X3,X2,esk9_4(X1,X2,k1_seq_2(X2),X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5878, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_5834, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5788, plain, (k18_complex1(k9_real_1(X2,X3))=k4_metric_1(X1,X3,X2)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5836, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(k1_numbers,X3,X2),k1_seq_2(X3))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(k1_seq_2(X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v2_comseq_2(X3))).#trainneg
cnf(i_0_5752, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(X4,X3),esk6_2(X2,X4))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers))).#trainneg
cnf(i_0_5849, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X5,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X5))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X5)))).#trainneg
cnf(i_0_5846, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X5,k8_nat_1(u1_struct_0(X2),X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X5))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X5)))).#trainneg
cnf(i_0_5132, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_2(X2),k1_seq_1(X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v2_comseq_2(X2))).#trainneg
cnf(i_0_6239, plain, (r1_xxreal_0(esk8_3(X1,X2,k1_seq_2(X2)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~v3_valued_0(X2)|~r1_xxreal_0(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),esk9_4(X1,X2,k1_seq_2(X2),X3))|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk8_3(X1,X2,k1_seq_2(X2)),k1_numbers)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),X3),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),X3)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_6037, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk7_3(k8_metric_1,esk1_0,esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k5_numbers))).#trainneg
cnf(i_0_6038, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_6012, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(k1_numbers,esk1_0,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(k1_numbers,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6006, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(esk6_2(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(esk6_2(k8_metric_1,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6156, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(k1_numbers,X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_2(X2),k1_numbers))).#trainneg
cnf(i_0_6013, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(k1_tbsp_1(k8_metric_1,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6014, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_2(X3),k8_nat_1(k1_numbers,X3,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(k1_seq_2(X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v2_comseq_2(X3))).#trainneg
cnf(i_0_5979, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_2(X2),k1_seq_1(X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3888, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(X2,X4),k8_nat_1(X5,X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X5,X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X5,X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_3444, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k2_metric_1(k8_metric_1,X2,k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X1,esk9_4(k8_metric_1,X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3557, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,esk1_0),k8_nat_1(X4,esk1_0,X3)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v1_funct_2(esk1_0,k5_numbers,X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_5841, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(X4,X3),esk6_2(X2,X4))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3548, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,esk1_0),k8_nat_1(u1_struct_0(X2),esk1_0,X3)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),esk1_0,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_4066, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(X4,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_4783, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(X3,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(X3,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_6009, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(esk6_2(X2,X4),k1_seq_1(X4,X3))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2913, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X5,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5))|~m1_subset_1(X3,u1_struct_0(X5)))).#trainneg
cnf(i_0_2706, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k1_binop_1(k7_metric_1,k1_seq_1(X3,esk5_4(X2,X3,X1,X4)),X1))|v2_struct_0(X2)|~v3_valued_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_seq_1(X3,esk5_4(X2,X3,X1,X4)),u1_struct_0(X2)))).#trainneg
cnf(i_0_5508, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(k8_metric_1,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6242, negated_conjecture, (v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~l1_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(esk8_3(X1,esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_6302, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_2(X2),k8_nat_1(k1_numbers,X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_2(X2),k1_numbers))).#trainneg
cnf(i_0_6240, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6391, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(X3,X2),k1_tbsp_1(k8_metric_1,X3))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X3,X2),k1_numbers)|~m1_subset_1(k1_tbsp_1(k8_metric_1,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6390, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k1_funct_1(esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6301, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_2(X2),k8_nat_1(k1_numbers,X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_2(X2),k1_numbers))).#trainneg
cnf(i_0_6255, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(esk6_2(k8_metric_1,X3),k1_seq_1(X3,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4675, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,X3,k1_seq_1(X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X5),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_3543, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,esk1_0),k1_seq_1(esk1_0,X3)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3986, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X5,k1_tbsp_1(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X5))|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X5)))).#trainneg
cnf(i_0_3552, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,esk1_0),k1_funct_1(esk1_0,X3)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_5693, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(k8_nat_1(k1_numbers,X2,esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))))),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(k1_numbers,X2,esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))))),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5552, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_tbsp_1(X2,X3),k8_nat_1(X6,X3,X4)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X5)|~v1_funct_1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~v1_funct_2(X3,k5_numbers,X6))).#trainneg
cnf(i_0_5753, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(X4,X3),k1_tbsp_1(X2,X4))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers))).#trainneg
cnf(i_0_6243, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k1_seq_1(esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5670, plain, (k4_metric_1(X4,X2,X3)=k2_metric_1(X1,X2,X3)|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_numbers)))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4)))).#trainneg
cnf(i_0_3335, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(k1_numbers,k1_numbers,X4,k8_nat_1(k1_numbers,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X4!=k7_metric_1|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(k1_tbsp_1(k8_metric_1,X3),k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X4))).#trainneg
cnf(i_0_2839, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X5,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(X3,u1_struct_0(X5))|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5)))).#trainneg
cnf(i_0_4892, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(X3,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(X3,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_5339, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X4,X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X5))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|X4!=k8_metric_1|~v1_metric_1(X1)|~v1_metric_1(X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X4))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6636, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v7_ordinal1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1))|~m2_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1),k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers)|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),X1)),k1_numbers))).#trainneg
cnf(i_0_6553, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2)))),k5_numbers))).#trainneg
cnf(i_0_6536, negated_conjecture, (v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_tbsp_1(k8_metric_1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_6564, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_1(esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6684, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))))),u1_struct_0(X1))|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))))),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2)))),k5_numbers))).#trainneg
cnf(i_0_3546, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(u1_struct_0(X2),esk1_0,X3),k1_tbsp_1(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_3327, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_seq_1(X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2465, plain, (r1_xxreal_0(esk17_2(X1,X2),k2_metric_1(X4,k1_seq_1(X1,esk18_3(X1,X2,X3)),X2))|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,X3)),u1_struct_0(X4))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5843, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X5,X4,X3),esk6_2(X2,X4))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X5,X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X5,X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_4643, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X5,X3,k1_seq_1(X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5))|~m1_subset_1(X3,u1_struct_0(X5)))).#trainneg
cnf(i_0_5747, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(esk6_2(X2,X3),k1_seq_1(X3,X4))))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_seq_1(X3,X4),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~m1_subset_1(esk6_2(X2,X3),k1_numbers))).#trainneg
cnf(i_0_6575, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_2(X2),k1_seq_1(X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5449, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X4,k1_seq_1(X3,X2),k1_seq_2(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~m1_subset_1(k1_seq_2(X3),u1_struct_0(X4))|~m1_subset_1(k1_seq_1(X3,X2),u1_struct_0(X4)))).#trainneg
cnf(i_0_5867, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X5,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X5))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X5)))).#trainneg
cnf(i_0_5559, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~r1_xxreal_0(X1,k4_metric_1(X2,k1_funct_1(X3,X4),k1_tbsp_1(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X5)|~v1_funct_1(X3)|~m1_subset_1(k1_tbsp_1(X2,X3),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X3,X4),u1_struct_0(X2)))).#trainneg
cnf(i_0_4839, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(X3,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(X3,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_5997, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_1(X3,X2),esk6_2(k8_metric_1,X3))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5870, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X5,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X5))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X5)))).#trainneg
cnf(i_0_5749, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(esk6_2(X2,X4),k1_funct_1(X4,X3))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers))).#trainneg
cnf(i_0_5547, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(k8_metric_1,k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),u1_struct_0(X1))|~m1_subset_1(k8_nat_1(X4,X2,esk9_4(X1,X2,X3,X5)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6837, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k6_numbers)|~v7_ordinal1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))))|~v1_xreal_0(k1_tbsp_1(k8_metric_1,esk1_0))|~m2_subset_1(esk18_3(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0),esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)))),k1_numbers,k5_numbers)|~m2_subset_1(esk7_3(k8_metric_1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0))),k1_numbers,k5_numbers)|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(k8_metric_1,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_6855, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))))),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2)))),k5_numbers))).#trainneg
cnf(i_0_6965, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k1_seq_1(esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7079, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,esk1_0),k1_funct_1(esk1_0,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k1_funct_1(esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6838, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~l1_metric_1(X1)|~v8_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k5_numbers))).#trainneg
cnf(i_0_6864, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(u1_struct_0(X2),esk1_0,X3),k1_tbsp_1(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(k1_tbsp_1(X2,esk1_0),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),esk1_0,X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_7080, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,X3),k1_funct_1(X3,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(X3,X2),k1_numbers)|~m1_subset_1(k1_tbsp_1(k8_metric_1,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_7014, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X6,X3,X4),k1_tbsp_1(X2,X3)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X5)|~v1_funct_1(X3)|~v7_ordinal1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~v1_funct_2(X3,k5_numbers,X6))).#trainneg
cnf(i_0_7038, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_1(X2,X3),k1_tbsp_1(k8_metric_1,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6966, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,X2),k1_seq_1(X2,X3))))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_tbsp_1(X2,k8_metric_1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(k8_metric_1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3675, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,esk6_2(X2,X4),k1_seq_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5403, plain, (k1_metric_1(k1_numbers,k1_numbers,X1,X2,X3)=k2_metric_1(X4,X2,X3)|X1!=k7_metric_1|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X1)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4)))).#trainneg
cnf(i_0_3958, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X4,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_5774, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X4,X5,X3),k1_tbsp_1(X2,X5))))|~r1_xxreal_0(esk7_3(X2,X5,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X5,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X5,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X5,k5_numbers,X4)|~v1_funct_1(X5)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_tbsp_1(X2,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X4,X5,X3),u1_struct_0(X2))|~m1_subset_1(k1_tbsp_1(X2,X5),k1_numbers)|~m1_subset_1(k8_nat_1(X4,X5,X3),k1_numbers))).#trainneg
cnf(i_0_6898, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|~r1_xxreal_0(esk7_3(X1,esk1_0,esk17_2(esk1_0,esk6_2(X1,esk1_0))),esk18_3(esk1_0,esk6_2(X1,esk1_0),X2))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk18_3(esk1_0,esk6_2(X1,esk1_0),X2),k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(esk6_2(X1,esk1_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,esk6_2(X1,esk1_0),X2)),u1_struct_0(X1))|~m1_subset_1(k1_seq_1(esk1_0,esk18_3(esk1_0,esk6_2(X1,esk1_0),X2)),k1_numbers)|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1)))).#trainneg
cnf(i_0_5779, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X4,X5,X3),esk6_2(X2,X5))))|~r1_xxreal_0(esk7_3(X2,X5,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(X5,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X5,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X5,k5_numbers,X4)|~v1_funct_1(X5)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(esk6_2(X2,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X4,X5,X3),u1_struct_0(X2))|~m1_subset_1(esk6_2(X2,X5),k1_numbers)|~m1_subset_1(k8_nat_1(X4,X5,X3),k1_numbers))).#trainneg
cnf(i_0_3544, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(X4,esk1_0,X3),k1_tbsp_1(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v1_funct_2(esk1_0,k5_numbers,X4)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_5864, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X5,X4,X3),k1_tbsp_1(X2,X4))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X5,X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X5,X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_3545, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X2,k8_nat_1(X4,esk1_0,X3),k1_tbsp_1(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v1_funct_2(esk1_0,k5_numbers,X4)|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_7340, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(esk6_2(X1,esk1_0))|~m2_subset_1(esk7_3(X1,esk1_0,esk17_2(esk1_0,esk6_2(X1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk18_3(esk1_0,esk6_2(X1,esk1_0),esk7_3(X1,esk1_0,esk17_2(esk1_0,esk6_2(X1,esk1_0)))),k5_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_7140, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk9_4(X1,X2,k1_seq_2(X2),esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2)))),k5_numbers))).#trainneg
cnf(i_0_7333, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(k1_tbsp_1(X1,esk1_0))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(esk7_3(X1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0))),esk18_3(esk1_0,k1_tbsp_1(X1,esk1_0),X2))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(X1,esk1_0),X2),k5_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k1_numbers)|~m1_subset_1(k1_tbsp_1(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_4604, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,X3,k1_seq_1(X4,X5)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X5),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_7143, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(esk1_0,X2),k1_numbers))).#trainneg
cnf(i_0_5293, plain, (v2_comseq_2(X1)|X2!=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(esk3_4(k8_metric_1,X1,X2,esk17_2(X1,X2)),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,X2,esk3_4(k8_metric_1,X1,X2,esk17_2(X1,X2)))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk17_2(X1,X2),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3257, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_funct_1(X4,X5),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_3602, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(X2,X4),k1_seq_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4672, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,X2,k8_nat_1(k1_numbers,X3,X4)))|~r1_xxreal_0(esk3_4(k8_metric_1,X3,X2,X1),X4)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X4),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2865, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3)),k6_numbers)|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|k8_nat_1(X4,esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X5))!=k1_tbsp_1(X1,X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X2,k8_nat_1(X4,esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X5)),esk8_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3))),X3)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(X4,esk1_0,esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X5)),u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3)),k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v7_ordinal1(esk9_4(X1,esk1_0,k8_nat_1(u1_struct_0(X1),X2,X3),X5))|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,X4))).#trainneg
cnf(i_0_4644, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(X5,X3,k1_seq_1(X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5))|~m1_subset_1(X3,u1_struct_0(X5)))).#trainneg
cnf(i_0_4803, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(X3,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(X3,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_7153, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_funct_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_funct_1(esk1_0,X2),k1_numbers))).#trainneg
cnf(i_0_6867, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(u1_struct_0(X2),esk1_0,X3),k1_tbsp_1(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_tbsp_1(X2,esk1_0),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),esk1_0,X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_5448, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X4,k1_seq_1(X3,X2),k1_seq_2(X3)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(k1_seq_2(X3),u1_struct_0(X4))|~m1_subset_1(k1_seq_1(X3,X2),u1_struct_0(X4)))).#trainneg
cnf(i_0_6073, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k18_complex1(k9_real_1(X3,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_3971, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(X2,X3),k8_nat_1(X4,X3,X5)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X4,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,X4)|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_tbsp_1(X2,X3),u1_struct_0(X2)))).#trainneg
cnf(i_0_6285, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,X3),k1_seq_1(X3,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_7280, plain, (k2_metric_1(X4,X1,X2)=k4_metric_1(X3,X1,X2)|X3!=k8_metric_1|~v1_metric_1(X3)|~l1_metric_1(X3)|~v8_metric_1(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X4)))).#trainneg
cnf(i_0_7470, plain, (v2_comseq_2(X1)|X2!=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(esk3_4(k8_metric_1,X1,X2,esk17_2(X1,X2)),k1_numbers,k5_numbers)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk17_2(X1,X2),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7424, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(esk6_2(X1,esk1_0))|~m2_subset_1(esk7_3(X1,esk1_0,esk17_2(esk1_0,esk6_2(X1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_7427, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(k1_tbsp_1(X1,esk1_0))|~m2_subset_1(esk7_3(X1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk18_3(esk1_0,k1_tbsp_1(X1,esk1_0),esk7_3(X1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)))),k5_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_7425, plain, (v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~l1_metric_1(X1)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(esk7_3(k8_metric_1,X2,esk8_3(X1,X2,k1_seq_2(X2))),k5_numbers))).#trainneg
cnf(i_0_7560, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(X3,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_3428, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k1_binop_1(k7_metric_1,k1_funct_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),X2))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,X2,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k1_funct_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1)))).#trainneg
cnf(i_0_3887, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(X2,X4),k1_seq_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3445, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k2_metric_1(k8_metric_1,X2,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1))|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk9_4(X1,esk1_0,X2,X3),k5_numbers))).#trainneg
cnf(i_0_6026, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(X2,X4),k1_seq_1(X4,X3))))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2838, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k4_metric_1(X5,k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(X3,u1_struct_0(X5))|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X5)))).#trainneg
cnf(i_0_4819, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k4_metric_1(X4,X2,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4)))).#trainneg
cnf(i_0_3106, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X4,X3,X2),k1_tbsp_1(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(X4,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_3253, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(k1_numbers,k1_numbers,X4,esk6_2(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X4!=k7_metric_1|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~m1_subset_1(esk6_2(k8_metric_1,X3),k1_numbers)|~v1_funct_2(X4,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X4))).#trainneg
cnf(i_0_3990, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X5,k8_nat_1(u1_struct_0(X2),X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~m1_subset_1(k1_tbsp_1(X2,X4),u1_struct_0(X5))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X5)))).#trainneg
cnf(i_0_4943, plain, (r1_xxreal_0(esk8_3(k8_metric_1,X1,X2),k2_metric_1(X4,X2,k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3))))|v2_tbsp_1(X1,k8_metric_1)|~v3_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_4(k8_metric_1,X1,X2,X3),k5_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~m1_subset_1(k1_seq_1(X1,esk9_4(k8_metric_1,X1,X2,X3)),u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4)))).#trainneg
cnf(i_0_5043, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_funct_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_6859, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k8_nat_1(X4,esk1_0,X3),k1_tbsp_1(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X4))).#trainneg
cnf(i_0_6565, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X3,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_7144, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_seq_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(esk1_0,X2),k1_numbers))).#trainneg
cnf(i_0_7824, plain, (v2_comseq_2(X1)|X2!=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~v1_xreal_0(X2)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk17_2(X1,X2),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk3_4(k8_metric_1,X1,X2,esk17_2(X1,X2)),k5_numbers))).#trainneg
cnf(i_0_7833, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(k1_tbsp_1(X1,esk1_0))|~m2_subset_1(esk7_3(X1,esk1_0,esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0))),k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_7832, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(esk6_2(X1,esk1_0))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk6_2(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_7825, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk17_2(X1,X2),k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X1)|~v3_valued_0(X1)|~v1_xreal_0(X2)|~v2_tbsp_1(X1,k8_metric_1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk17_2(X1,X2),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8026, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|X3!=k8_metric_1|~v1_metric_1(X2)|~v1_metric_1(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X3,k8_nat_1(u1_struct_0(X2),X5,X4),k1_tbsp_1(X2,X5)))|~r1_xxreal_0(esk7_3(X2,X5,X1),X4)|~l1_metric_1(X2)|~l1_metric_1(X3)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X5,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X5,X4),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X5,X4),u1_struct_0(X3))|~m1_subset_1(k1_tbsp_1(X2,X5),u1_struct_0(X3))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X5,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X5)|~v8_metric_1(X3))).#trainneg
cnf(i_0_8064, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,esk6_2(X2,X4),k1_funct_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2731, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k1_binop_1(k7_metric_1,X2,k1_funct_1(esk1_0,esk9_4(X1,esk1_0,X2,X3))))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~v7_ordinal1(esk9_4(X1,esk1_0,X2,X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k1_funct_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1)))).#trainneg
cnf(i_0_7154, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k1_funct_1(esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_funct_1(esk1_0,X2),k1_numbers))).#trainneg
cnf(i_0_6244, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(X3,esk1_0,X2))))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(X3,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_5438, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X4,k1_seq_2(X2),k8_nat_1(k1_numbers,X2,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),u1_struct_0(X4))|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X4))|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers))).#trainneg
cnf(i_0_4642, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k2_metric_1(k8_metric_1,X3,k1_seq_1(X2,esk9_4(X1,X2,X3,X4))))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|X1!=k8_metric_1|~v3_valued_0(X2)|~v1_metric_1(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_seq_1(X2,esk9_4(X1,X2,X3,X4)),k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_4092, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X4,k8_nat_1(k1_numbers,X3,X2),esk6_2(k8_metric_1,X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(esk6_2(k8_metric_1,X3),u1_struct_0(X4))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),u1_struct_0(X4))|~m1_subset_1(esk6_2(k8_metric_1,X3),k1_numbers))).#trainneg
cnf(i_0_2446, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),k1_metric_1(X5,X6,k7_metric_1,k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X3))|v2_struct_0(X1)|v2_tbsp_1(X2,X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|X1!=k8_metric_1|~v1_metric_1(X1)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X5,X6),k1_numbers)))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X3,X6)|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk9_4(X1,X2,X3,X4)),X5))).#trainneg
cnf(i_0_3248, plain, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(X4,X5,k7_metric_1,esk6_2(k8_metric_1,X3),k8_nat_1(k1_numbers,X3,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~m1_subset_1(k7_metric_1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X4,X5),k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X2),X5)|~m1_subset_1(esk6_2(k8_metric_1,X3),X4))).#trainneg
cnf(i_0_6296, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_seq_2(X2),k1_seq_1(X2,X3))))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3426, negated_conjecture, (r1_xxreal_0(esk8_3(X1,esk1_0,X2),k1_binop_1(k7_metric_1,k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),X2))|v2_struct_0(X1)|v2_tbsp_1(esk1_0,X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~l1_metric_1(X1)|~v8_metric_1(X1)|~m1_subset_1(k1_seq_1(esk1_0,esk9_4(X1,esk1_0,X2,X3)),u1_struct_0(X1))|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|X1!=k8_metric_1|~v1_metric_1(X1))).#trainneg
cnf(i_0_5445, plain, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X4,k1_seq_2(X3),k1_seq_1(X3,X2)))|~r1_xxreal_0(esk16_2(X3,X1),X2)|~v2_comseq_2(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X2),k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(k1_seq_1(X3,X2),u1_struct_0(X4))|~m1_subset_1(k1_seq_2(X3),u1_struct_0(X4)))).#trainneg
cnf(i_0_4107, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X5,esk6_2(X2,X4),k8_nat_1(u1_struct_0(X2),X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X5!=k8_metric_1|~v1_metric_1(X5)|~l1_metric_1(X5)|~v8_metric_1(X5)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X3),u1_struct_0(X5))|~m1_subset_1(esk6_2(X2,X4),u1_struct_0(X5)))).#trainneg
cnf(i_0_8119, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(k1_tbsp_1(X1,esk1_0))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k1_numbers)|~m1_subset_1(k1_tbsp_1(X1,esk1_0),u1_struct_0(X1))|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1))).#trainneg
cnf(i_0_8121, negated_conjecture, (r1_xxreal_0(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|X1!=k8_metric_1|~v1_metric_1(X1)|~v1_xreal_0(esk6_2(X1,esk1_0))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_8175, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(k8_metric_1,esk1_0),k8_nat_1(X3,esk1_0,X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_6115, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(u1_struct_0(X2),X4,X5),X3)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers))).#trainneg
cnf(i_0_4532, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_funct_1(X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4881, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X3,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(k8_nat_1(X3,esk1_0,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3))).#trainneg
cnf(i_0_5439, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X4,k8_nat_1(k1_numbers,X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|X4!=k8_metric_1|~v1_metric_1(X4)|~l1_metric_1(X4)|~v8_metric_1(X4)|~m1_subset_1(k1_seq_2(X2),u1_struct_0(X4))|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),u1_struct_0(X4))|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~m1_subset_1(k8_nat_1(k1_numbers,X2,X3),k1_numbers))).#trainneg
cnf(i_0_5183, plain, (v2_tbsp_1(u1_metric_1(X1),k8_metric_1)|~v2_comseq_2(u1_metric_1(X1))|~m1_subset_1(u1_metric_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(u1_metric_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(u1_metric_1(X1))|~l1_metric_1(X1))).#trainneg
cnf(i_0_5044, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_seq_1(X4,X3),esk6_2(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X4,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X4,X3),k1_numbers)|~m1_subset_1(esk6_2(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v3_valued_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5123, plain, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_funct_1(X3,X2),k1_seq_2(X3)))|~r1_xxreal_0(esk7_3(k8_metric_1,X3,X1),X2)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v2_comseq_2(X3))).#trainneg
cnf(i_0_2740, plain, (X1=k1_tbsp_1(X2,X3)|r1_xxreal_0(esk4_3(X2,X3,X1),k1_binop_1(k7_metric_1,X1,k1_seq_1(X3,esk5_4(X2,X3,X1,X4))))|v2_struct_0(X2)|~v3_valued_0(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(k1_seq_1(X3,esk5_4(X2,X3,X1,X4)),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2))).#trainneg
cnf(i_0_5720, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X4,X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(X4,X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_5782, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v3_valued_0(X3)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(X2,X3),k1_seq_1(X3,X4))))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3)|X2!=k8_metric_1|~v1_metric_1(X2)|~m1_subset_1(k1_seq_1(X3,X4),u1_struct_0(X2))|~m1_subset_1(k1_tbsp_1(X2,X3),u1_struct_0(X2))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_3757, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk13_2(X2,X1),k4_metric_1(k8_metric_1,X1,k1_seq_1(X2,esk14_3(X2,X1,X3))))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_2(X2)|~m1_subset_1(k1_seq_1(X2,esk14_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3540, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X2,k1_seq_1(esk1_0,X3),k1_tbsp_1(X2,esk1_0)))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_517, plain, (m2_subset_1(esk16_2(X1,X2),k5_numbers,k5_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3785, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_tbsp_1(k8_metric_1,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,X2,k8_nat_1(k1_numbers,X3,X4)))|~r1_xxreal_0(esk3_4(k8_metric_1,X3,X2,X1),X4)|~v2_tbsp_1(X3,k8_metric_1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k8_nat_1(k1_numbers,X3,X4),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3907, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X6,X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X6,X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X6,X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,X6))).#trainneg
cnf(i_0_8487, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_seq_1(X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8488, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k8_nat_1(X4,X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_8381, negated_conjecture, (v2_struct_0(X1)|~v1_xreal_0(k1_tbsp_1(X1,esk1_0))|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,k1_tbsp_1(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_8382, negated_conjecture, (v2_struct_0(X1)|~v1_xreal_0(esk6_2(X1,esk1_0))|X1!=k8_metric_1|~v1_metric_1(X1)|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk17_2(esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_8532, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v3_valued_0(X3)|~v1_metric_1(X2)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(k8_metric_1,k1_tbsp_1(X2,X3),k1_seq_1(X3,X4)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_8422, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_tbsp_1(X2,X4),k1_funct_1(X4,X3)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_8473, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v3_valued_0(X3)|~v1_metric_1(X2)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,esk6_2(X2,X3),k1_seq_1(X3,X4)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~m1_subset_1(esk6_2(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_8533, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v3_valued_0(X3)|~v1_metric_1(X2)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k1_tbsp_1(X2,X3),k1_seq_1(X3,X4)))|~r1_xxreal_0(esk7_3(X2,X3,X1),X4)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_8477, plain, (r1_xxreal_0(esk17_2(X1,esk6_2(X2,X1)),k6_numbers)|v2_struct_0(X2)|v2_comseq_2(X1)|X2!=k8_metric_1|~v3_valued_0(X1)|~v1_metric_1(X2)|~r1_xxreal_0(esk7_3(X2,X1,esk17_2(X1,esk6_2(X2,X1))),esk18_3(X1,esk6_2(X2,X1),X3))|~l1_metric_1(X2)|~v2_tbsp_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_seq_1(X1,esk18_3(X1,esk6_2(X2,X1),X3)),k1_numbers)|~m1_subset_1(esk6_2(X2,X1),k1_numbers)|~m1_subset_1(esk17_2(X1,esk6_2(X2,X1)),k1_numbers)|~m1_subset_1(esk18_3(X1,esk6_2(X2,X1),X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X1)|~v1_xreal_0(esk6_2(X2,X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_8489, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_valued_0(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k1_funct_1(X2,X3),k1_seq_2(X2)))|~r1_xxreal_0(esk7_3(k8_metric_1,X2,X1),X3)|~v2_comseq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_2(X2),k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_8421, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(k8_metric_1,k8_nat_1(X5,X4,X3),k1_tbsp_1(X2,X4)))|~r1_xxreal_0(esk7_3(X2,X4,X1),X3)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X5,X4,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,X4),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X5)))|~v1_funct_2(X4,k5_numbers,X5))).#trainneg
cnf(i_0_8484, negated_conjecture, (r1_xxreal_0(esk4_3(k8_metric_1,esk1_0,esk6_2(X1,esk1_0)),k6_numbers)|v2_struct_0(X1)|esk6_2(X1,esk1_0)=k1_tbsp_1(k8_metric_1,esk1_0)|X1!=k8_metric_1|~v1_metric_1(X1)|~r1_xxreal_0(esk7_3(X1,esk1_0,esk4_3(k8_metric_1,esk1_0,esk6_2(X1,esk1_0))),esk5_4(k8_metric_1,esk1_0,esk6_2(X1,esk1_0),X2))|~l1_metric_1(X1)|~v2_tbsp_1(esk1_0,X1)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_1(esk1_0,esk5_4(k8_metric_1,esk1_0,esk6_2(X1,esk1_0),X2)),k1_numbers)|~m1_subset_1(esk6_2(X1,esk1_0),k1_numbers)|~m1_subset_1(esk4_3(k8_metric_1,esk1_0,esk6_2(X1,esk1_0)),k1_numbers)|~m1_subset_1(esk5_4(k8_metric_1,esk1_0,esk6_2(X1,esk1_0),X2),k5_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8411, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_metric_1(k1_numbers,k1_numbers,X6,k8_nat_1(u1_struct_0(X2),X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X6!=k7_metric_1|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)))|~v1_funct_2(X6,k2_zfmisc_1(k1_numbers,k1_numbers),k1_numbers)|~v1_funct_1(X6))).#trainneg
cnf(i_0_8412, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k2_metric_1(X6,k8_nat_1(u1_struct_0(X2),X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X6!=k8_metric_1|~v1_metric_1(X6)|~l1_metric_1(X6)|~m1_subset_1(X3,u1_struct_0(X6))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X6)))).#trainneg
cnf(i_0_8408, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k8_nat_1(X6,X4,X5),X3)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(X6,X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(X6,X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X6)))|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,X6))).#trainneg
cnf(i_0_8406, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_funct_1(X4,X5),X3)))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k1_funct_1(X4,X5),u1_struct_0(X2))|~m1_subset_1(k1_funct_1(X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_6015, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k1_tbsp_1(X2,esk1_0),k8_nat_1(u1_struct_0(X2),esk1_0,X3))))|~r1_xxreal_0(esk7_3(X2,esk1_0,X1),X3)|~l1_metric_1(X2)|~v8_metric_1(X2)|~v2_tbsp_1(esk1_0,X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),esk1_0,X3),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),esk1_0,X3),k1_numbers)|~m1_subset_1(k1_tbsp_1(X2,esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,u1_struct_0(X2))|~v9_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2))).#trainneg
cnf(i_0_8097, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k1_binop_1(k7_metric_1,k8_nat_1(X3,esk1_0,X2),k1_tbsp_1(k8_metric_1,esk1_0)))|~r1_xxreal_0(esk7_3(k8_metric_1,esk1_0,X1),X2)|~m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(esk1_0,k5_numbers,X3)|~m1_subset_1(k8_nat_1(X3,esk1_0,X2),k1_numbers))).#trainneg
cnf(i_0_710, plain, (r1_xxreal_0(esk4_3(X1,X2,k1_seq_1(X3,X4)),k6_numbers)|v2_struct_0(X1)|k1_seq_1(X3,X4)=k1_tbsp_1(X1,X2)|k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k1_seq_1(X3,X4),X5))!=k1_tbsp_1(X1,X3)|~v3_valued_0(X3)|~v7_ordinal1(X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_4(X1,X3,k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k1_seq_1(X3,X4),X5)),esk4_3(X1,X2,k1_seq_1(X3,X4))),X4)|~l1_metric_1(X1)|~v9_metric_1(X1)|~v8_metric_1(X1)|~v7_metric_1(X1)|~v6_metric_1(X1)|~v2_tbsp_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k8_nat_1(u1_struct_0(X1),X2,esk5_4(X1,X2,k1_seq_1(X3,X4),X5)),u1_struct_0(X1))|~m1_subset_1(esk4_3(X1,X2,k1_seq_1(X3,X4)),k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X3)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~v2_tbsp_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k1_seq_1(X3,X4),u1_struct_0(X1))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_8415, plain, (r1_xxreal_0(X1,k6_numbers)|v2_struct_0(X2)|X3!=k1_tbsp_1(X2,X4)|X2!=k8_metric_1|~v1_metric_1(X2)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k4_metric_1(X6,k8_nat_1(u1_struct_0(X2),X4,X5),X3))|~r1_xxreal_0(esk3_4(X2,X4,X3,X1),X5)|~l1_metric_1(X2)|~v9_metric_1(X2)|~v8_metric_1(X2)|~v7_metric_1(X2)|~v6_metric_1(X2)|~v2_tbsp_1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X2))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X4)|X6!=k8_metric_1|~v1_metric_1(X6)|~l1_metric_1(X6)|~v8_metric_1(X6)|~m1_subset_1(X3,u1_struct_0(X6))|~m1_subset_1(k8_nat_1(u1_struct_0(X2),X4,X5),u1_struct_0(X6)))).#trainneg
cnf(i_0_8878, plain, (m2_subset_1(esk17_2(X1,X2),k1_numbers,k1_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 74
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 143
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 137
# Processed clauses                    : 2225
# ...of these trivial                  : 98
# ...subsumed                          : 1184
# ...remaining for further processing  : 943
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 161
# Backward-rewritten                   : 45
# Generated clauses                    : 8820
# ...of the previous two non-trivial   : 7935
# Contextual simplify-reflections      : 463
# Paramodulations                      : 8811
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 736
#    Positive orientable unit clauses  : 46
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 684
# Current number of unprocessed clauses: 5609
# ...number of literals in the above   : 103894
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1149
# Proof object given clauses           : 175
# Proof search given clauses           : 943
# Clause-clause subsumption calls (NU) : 758283
# Rec. Clause-clause subsumption calls : 5352
# Non-unit clause-clause subsumptions  : 1783
# Unit Clause-clause subsumption calls : 2400
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 38
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 661392

# -------------------------------------------------
# User time                : 72.666 s
# System time              : 1.877 s
# Total time               : 74.542 s
# Maximum resident set size: 504276 pages
72.66user 2.01system 1:16.60elapsed 97%CPU (0avgtext+0avgdata 504276maxresident)k
0inputs+1064outputs (0major+583174minor)pagefaults 0swaps
