# local_data = 0x2dfb830
#TFS# Evaluating axioms..
# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
#TF#SERVER: Lgb model filter skipped 23 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=0, query=323, total=327).
#TF#SERVER: Sent 70078 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1296 bytes expected).
#TF#SERVER: Received 1296 bytes from the TF server.
# Preprocessing time       : 0.764 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 56 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=335, query=286, total=625).
#TF#SERVER: Sent 169123 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2488 bytes expected).
#TF#SERVER: Received 2488 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 186 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=565, query=431, total=1000).
#TF#SERVER: Sent 376337 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3988 bytes expected).
#TF#SERVER: Received 3988 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 353 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=706, query=679, total=1389).
#TF#SERVER: Sent 649570 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5544 bytes expected).
#TF#SERVER: Received 5544 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 278 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=824, query=756, total=1584).
#TF#SERVER: Sent 735190 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6324 bytes expected).
#TF#SERVER: Received 6324 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 298 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=958, query=740, total=1702).
#TF#SERVER: Sent 857785 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6796 bytes expected).
#TF#SERVER: Received 6796 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 389 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1079, query=639, total=1722).
#TF#SERVER: Sent 911349 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6876 bytes expected).
#TF#SERVER: Received 6876 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 451 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1195, query=574, total=1773).
#TF#SERVER: Sent 922075 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7080 bytes expected).
#TF#SERVER: Received 7080 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 535 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1288, query=527, total=1819).
#TF#SERVER: Sent 915563 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7264 bytes expected).
#TF#SERVER: Received 7264 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 372 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1352, query=657, total=2013).
#TF#SERVER: Sent 1084301 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8040 bytes expected).
#TF#SERVER: Received 8040 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 385 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1450, query=685, total=2139).
#TF#SERVER: Sent 1162906 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8544 bytes expected).
#TF#SERVER: Received 8544 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 479 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1489, query=643, total=2136).
#TF#SERVER: Sent 1158121 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8532 bytes expected).
#TF#SERVER: Received 8532 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 423 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=640, total=2180).
#TF#SERVER: Sent 1160957 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8708 bytes expected).
#TF#SERVER: Received 8708 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 346 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=751, total=2291).
#TF#SERVER: Sent 1252660 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9152 bytes expected).
#TF#SERVER: Received 9152 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 443 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=607, total=2147).
#TF#SERVER: Sent 1081973 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8576 bytes expected).
#TF#SERVER: Received 8576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 410 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=674, total=2214).
#TF#SERVER: Sent 1179802 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8844 bytes expected).
#TF#SERVER: Received 8844 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 373 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=731, total=2271).
#TF#SERVER: Sent 1221399 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9072 bytes expected).
#TF#SERVER: Received 9072 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 357 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=699, total=2239).
#TF#SERVER: Sent 1173760 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8944 bytes expected).
#TF#SERVER: Received 8944 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 465 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=565, total=2105).
#TF#SERVER: Sent 1102402 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8408 bytes expected).
#TF#SERVER: Received 8408 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 421 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=676, total=2216).
#TF#SERVER: Sent 1249230 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8852 bytes expected).
#TF#SERVER: Received 8852 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 281 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=764, total=2304).
#TF#SERVER: Sent 1358283 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9204 bytes expected).
#TF#SERVER: Received 9204 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 380 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=702, total=2242).
#TF#SERVER: Sent 1283079 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8956 bytes expected).
#TF#SERVER: Received 8956 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 383 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=703, total=2243).
#TF#SERVER: Sent 1217106 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8960 bytes expected).
#TF#SERVER: Received 8960 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 255 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=803, total=2343).
#TF#SERVER: Sent 1220792 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9360 bytes expected).
#TF#SERVER: Received 9360 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 202 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=827, total=2367).
#TF#SERVER: Sent 1228157 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9456 bytes expected).
#TF#SERVER: Received 9456 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 391 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=767, total=2307).
#TF#SERVER: Sent 1267382 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9216 bytes expected).
#TF#SERVER: Received 9216 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 289 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=748, total=2288).
#TF#SERVER: Sent 1308329 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9140 bytes expected).
#TF#SERVER: Received 9140 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 288 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=758, total=2298).
#TF#SERVER: Sent 1459316 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9180 bytes expected).
#TF#SERVER: Received 9180 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 449 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=645, total=2185).
#TF#SERVER: Sent 1273324 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8728 bytes expected).
#TF#SERVER: Received 8728 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 482 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=689, total=2229).
#TF#SERVER: Sent 1312614 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8904 bytes expected).
#TF#SERVER: Received 8904 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 395 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=673, total=2213).
#TF#SERVER: Sent 1309601 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8840 bytes expected).
#TF#SERVER: Received 8840 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 159 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=871, total=2411).
#TF#SERVER: Sent 1431625 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9632 bytes expected).
#TF#SERVER: Received 9632 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 239 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=850, total=2390).
#TF#SERVER: Sent 1356649 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9548 bytes expected).
#TF#SERVER: Received 9548 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 392 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=918, total=2458).
#TF#SERVER: Sent 1460585 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9820 bytes expected).
#TF#SERVER: Received 9820 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 206 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=820, total=2360).
#TF#SERVER: Sent 1309297 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9428 bytes expected).
#TF#SERVER: Received 9428 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 204 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=911, total=2451).
#TF#SERVER: Sent 1392713 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9792 bytes expected).
#TF#SERVER: Received 9792 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 309 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=769, total=2309).
#TF#SERVER: Sent 1307034 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9224 bytes expected).
#TF#SERVER: Received 9224 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 388 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=859, total=2399).
#TF#SERVER: Sent 1412955 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9584 bytes expected).
#TF#SERVER: Received 9584 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 338 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=773, total=2313).
#TF#SERVER: Sent 1364140 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9240 bytes expected).
#TF#SERVER: Received 9240 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 215 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=852, total=2392).
#TF#SERVER: Sent 1454034 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9556 bytes expected).
#TF#SERVER: Received 9556 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 359 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=725, total=2265).
#TF#SERVER: Sent 1336123 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9048 bytes expected).
#TF#SERVER: Received 9048 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 366 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=697, total=2237).
#TF#SERVER: Sent 1283407 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8936 bytes expected).
#TF#SERVER: Received 8936 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 423 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=639, total=2179).
#TF#SERVER: Sent 1260665 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8704 bytes expected).
#TF#SERVER: Received 8704 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 431 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=840, total=2380).
#TF#SERVER: Sent 1400632 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9508 bytes expected).
#TF#SERVER: Received 9508 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 372 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=652, total=2192).
#TF#SERVER: Sent 1221637 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8756 bytes expected).
#TF#SERVER: Received 8756 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 365 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=715, total=2255).
#TF#SERVER: Sent 1273171 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9008 bytes expected).
#TF#SERVER: Received 9008 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 363 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=711, total=2251).
#TF#SERVER: Sent 1209573 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8992 bytes expected).
#TF#SERVER: Received 8992 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 340 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=776, total=2316).
#TF#SERVER: Sent 1358516 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9252 bytes expected).
#TF#SERVER: Received 9252 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 239 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=798, total=2338).
#TF#SERVER: Sent 1299212 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9340 bytes expected).
#TF#SERVER: Received 9340 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 283 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=762, total=2302).
#TF#SERVER: Sent 1377665 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9196 bytes expected).
#TF#SERVER: Received 9196 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 239 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=857, total=2397).
#TF#SERVER: Sent 1295660 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9576 bytes expected).
#TF#SERVER: Received 9576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 393 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=659, total=2199).
#TF#SERVER: Sent 1218854 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8784 bytes expected).
#TF#SERVER: Received 8784 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 498 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=615, total=2155).
#TF#SERVER: Sent 1202722 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8608 bytes expected).
#TF#SERVER: Received 8608 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 443 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=653, total=2193).
#TF#SERVER: Sent 1291471 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8760 bytes expected).
#TF#SERVER: Received 8760 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 392 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=734, total=2274).
#TF#SERVER: Sent 1284219 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9084 bytes expected).
#TF#SERVER: Received 9084 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 433 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=680, total=2220).
#TF#SERVER: Sent 1231913 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8868 bytes expected).
#TF#SERVER: Received 8868 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 644 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=758, total=2298).
#TF#SERVER: Sent 1321264 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9180 bytes expected).
#TF#SERVER: Received 9180 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 386 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=696, total=2236).
#TF#SERVER: Sent 1271481 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8932 bytes expected).
#TF#SERVER: Received 8932 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 309 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=724, total=2264).
#TF#SERVER: Sent 1285863 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9044 bytes expected).
#TF#SERVER: Received 9044 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 433 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=610, total=2150).
#TF#SERVER: Sent 1191588 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8588 bytes expected).
#TF#SERVER: Received 8588 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 294 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=754, total=2294).
#TF#SERVER: Sent 1303925 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9164 bytes expected).
#TF#SERVER: Received 9164 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 464 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=635, total=2175).
#TF#SERVER: Sent 1213546 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8688 bytes expected).
#TF#SERVER: Received 8688 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 465 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=721, total=2261).
#TF#SERVER: Sent 1239704 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9032 bytes expected).
#TF#SERVER: Received 9032 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 425 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=602, total=2142).
#TF#SERVER: Sent 1187708 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8556 bytes expected).
#TF#SERVER: Received 8556 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 285 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=742, total=2282).
#TF#SERVER: Sent 1231030 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9116 bytes expected).
#TF#SERVER: Received 9116 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 565 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=675, total=2215).
#TF#SERVER: Sent 1199117 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8848 bytes expected).
#TF#SERVER: Received 8848 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 455 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=855, total=2395).
#TF#SERVER: Sent 1427285 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9568 bytes expected).
#TF#SERVER: Received 9568 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 471 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=558, total=2098).
#TF#SERVER: Sent 1180847 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8380 bytes expected).
#TF#SERVER: Received 8380 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 476 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=616, total=2156).
#TF#SERVER: Sent 1262027 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8612 bytes expected).
#TF#SERVER: Received 8612 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 520 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=539, total=2079).
#TF#SERVER: Sent 1155901 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8304 bytes expected).
#TF#SERVER: Received 8304 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 429 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=596, total=2136).
#TF#SERVER: Sent 1135824 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8532 bytes expected).
#TF#SERVER: Received 8532 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 511 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=555, total=2095).
#TF#SERVER: Sent 1162335 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8368 bytes expected).
#TF#SERVER: Received 8368 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 335 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=741, total=2281).
#TF#SERVER: Sent 1175306 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9112 bytes expected).
#TF#SERVER: Received 9112 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 569 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=523, total=2063).
#TF#SERVER: Sent 1130895 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8240 bytes expected).
#TF#SERVER: Received 8240 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 619 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=424, total=1964).
#TF#SERVER: Sent 1067833 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7844 bytes expected).
#TF#SERVER: Received 7844 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 591 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=573, total=2113).
#TF#SERVER: Sent 1120664 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8440 bytes expected).
#TF#SERVER: Received 8440 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 505 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=661, total=2201).
#TF#SERVER: Sent 1312554 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8792 bytes expected).
#TF#SERVER: Received 8792 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 479 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=702, total=2242).
#TF#SERVER: Sent 1352763 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8956 bytes expected).
#TF#SERVER: Received 8956 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 491 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=555, total=2095).
#TF#SERVER: Sent 1248759 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8368 bytes expected).
#TF#SERVER: Received 8368 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 511 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=594, total=2134).
#TF#SERVER: Sent 1242462 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8524 bytes expected).
#TF#SERVER: Received 8524 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 486 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=676, total=2216).
#TF#SERVER: Sent 1228447 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8852 bytes expected).
#TF#SERVER: Received 8852 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 562 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=566, total=2106).
#TF#SERVER: Sent 1096400 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8412 bytes expected).
#TF#SERVER: Received 8412 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 470 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=614, total=2154).
#TF#SERVER: Sent 1194374 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8604 bytes expected).
#TF#SERVER: Received 8604 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 441 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=676, total=2216).
#TF#SERVER: Sent 1269714 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8852 bytes expected).
#TF#SERVER: Received 8852 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 371 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=772, total=2312).
#TF#SERVER: Sent 1331609 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9236 bytes expected).
#TF#SERVER: Received 9236 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 431 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=1098, total=2638).
#TF#SERVER: Sent 1567696 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10540 bytes expected).
#TF#SERVER: Received 10540 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 492 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=960, total=2500).
#TF#SERVER: Sent 1426198 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9988 bytes expected).
#TF#SERVER: Received 9988 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 436 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=845, total=2385).
#TF#SERVER: Sent 1391942 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9528 bytes expected).
#TF#SERVER: Received 9528 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 413 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=750, total=2290).
#TF#SERVER: Sent 1293770 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9148 bytes expected).
#TF#SERVER: Received 9148 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 521 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=655, total=2195).
#TF#SERVER: Sent 1212551 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8768 bytes expected).
#TF#SERVER: Received 8768 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 654 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=527, total=2067).
#TF#SERVER: Sent 1153408 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8256 bytes expected).
#TF#SERVER: Received 8256 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 860 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=465, total=2005).
#TF#SERVER: Sent 1101332 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8008 bytes expected).
#TF#SERVER: Received 8008 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 545 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=679, total=2219).
#TF#SERVER: Sent 1263579 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8864 bytes expected).
#TF#SERVER: Received 8864 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 283 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=743, total=2283).
#TF#SERVER: Sent 1433478 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9120 bytes expected).
#TF#SERVER: Received 9120 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 733 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=700, total=2240).
#TF#SERVER: Sent 1399012 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8948 bytes expected).
#TF#SERVER: Received 8948 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 558 clauses
#TF#SERVER: Sending data to the TF server (conj=4, context=1536, query=960, total=2500).
#TF#SERVER: Sent 1487846 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9988 bytes expected).
#TF#SERVER: Received 9988 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_fcont_1, axiom, ![X1]:((v1_xreal_0(X1)=>![X2]:((v1_xreal_0(X2)=>![X3]:((((v1_funct_1(X3)&v1_funct_2(X3,k1_numbers,k1_numbers))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k1_fcont_1(X1,X2)<=>![X4]:((v1_xreal_0(X4)=>k1_seq_1(X3,X4)=k2_xcmplx_0(k3_xcmplx_0(X1,X4),X2)))))))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', d4_fcont_1)).
fof(dt_k1_fcont_1, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>((v1_funct_1(k1_fcont_1(X1,X2))&v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers))&m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', dt_k1_fcont_1)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:((((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', redefinition_k1_seq_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc1_relset_1)).
fof(cc18_valued_0, axiom, ![X1, X2]:((v3_membered(X2)=>![X3]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc18_valued_0)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc3_membered)).
fof(spc7_arithm, axiom, ![X1, X2, X3]:((((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc7_arithm)).
fof(spc2_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc2_arithm)).
fof(fc61_valued_0, axiom, ![X1, X2]:((((v1_relat_1(X1)&v1_funct_1(X1))&v1_valued_0(X1))=>v1_xcmplx_0(k1_funct_1(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc61_valued_0)).
fof(commutativity_k3_xcmplx_0, axiom, ![X1, X2]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', commutativity_k3_xcmplx_0)).
fof(cc6_valued_0, axiom, ![X1]:(((v1_relat_1(X1)&v3_valued_0(X1))=>(v1_relat_1(X1)&v1_valued_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc6_valued_0)).
fof(dt_k4_xcmplx_0, axiom, ![X1]:((v1_xcmplx_0(X1)=>v1_xcmplx_0(k4_xcmplx_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', dt_k4_xcmplx_0)).
fof(cc14_membered, axiom, ![X1]:((v3_membered(X1)=>![X2]:((m1_subset_1(X2,X1)=>v1_xreal_0(X2))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc14_membered)).
fof(fc3_xreal_0, axiom, ![X1]:((v1_xreal_0(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_xreal_0(k4_xcmplx_0(X1))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc3_xreal_0)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:((v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', involutiveness_k4_xcmplx_0)).
fof(t3_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t3_arithm)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc1_numerals)).
fof(cc3_xreal_0, axiom, ![X1]:((v1_xreal_0(X1)=>v1_xcmplx_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc3_xreal_0)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc0_numerals)).
fof(rqRealAdd__k2_xcmplx_0__rm1_r0_rm1, axiom, k2_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', rqRealAdd__k2_xcmplx_0__rm1_r0_rm1)).
fof(fc6_xreal_0, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k3_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc6_xreal_0)).
fof(t6_boole, axiom, ![X1]:((v1_xboole_0(X1)=>X1=k1_xboole_0)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t6_boole)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', redefinition_k6_numbers)).
fof(t5_xcmplx_1, axiom, ![X1]:((v1_xcmplx_0(X1)=>![X2]:((v1_xcmplx_0(X2)=>![X3]:((v1_xcmplx_0(X3)=>(k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X3,X1)=>(X1=k6_numbers|X2=X3)))))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t5_xcmplx_1)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc0_boole)).
fof(t50_fcont_1, conjecture, ![X1]:((v1_xreal_0(X1)=>![X2]:((v1_xreal_0(X2)=>(X1!=k6_numbers=>v2_funct_1(k1_fcont_1(X1,X2))))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t50_fcont_1)).
fof(spc1_arithm, axiom, ![X1, X2]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc1_arithm)).
fof(spc8_arithm, axiom, ![X1, X2]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc8_arithm)).
fof(spc6_arithm, axiom, ![X1, X2, X3]:((((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc6_arithm)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', commutativity_k2_xcmplx_0)).
fof(spc9_arithm, axiom, ![X1, X2]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc9_arithm)).
fof(fc7_xreal_0, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k6_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc7_xreal_0)).
fof(cc12_membered, axiom, ![X1]:((v1_membered(X1)=>![X2]:((m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc12_membered)).
fof(cc21_membered, axiom, ![X1]:((v3_membered(X1)=>![X2]:((m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc21_membered)).
fof(cc5_membered, axiom, ![X1]:((v3_membered(X1)=>v1_membered(X1))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc5_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', dt_k5_numbers)).
fof(spc5_arithm, axiom, ![X1, X2, X3]:((((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', spc5_arithm)).
fof(t2_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k3_xcmplx_0(X1,k6_numbers)=k6_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t2_arithm)).
fof(fc16_xreal_0, axiom, ![X1]:(((~(v3_xxreal_0(X1))&v1_xreal_0(X1))=>(v1_xcmplx_0(k4_xcmplx_0(X1))&~(v2_xxreal_0(k4_xcmplx_0(X1)))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc16_xreal_0)).
fof(rqRealMult__k3_xcmplx_0__r0_r1_r0, axiom, k3_xcmplx_0(np__0,np__1)=np__0, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', rqRealMult__k3_xcmplx_0__r0_r1_r0)).
fof(cc3_nat_1, axiom, ![X1]:((v7_ordinal1(X1)=>(v7_ordinal1(X1)&~(v3_xxreal_0(X1))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc3_nat_1)).
fof(cc17_membered, axiom, ![X1]:((v6_membered(X1)=>![X2]:((m1_subset_1(X2,X1)=>v7_ordinal1(X2))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', redefinition_k5_numbers)).
fof(rqRealNeg__k4_xcmplx_0__r0_r0, axiom, k4_xcmplx_0(np__0)=np__0, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', rqRealNeg__k4_xcmplx_0__r0_r0)).
fof(t4_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k6_xcmplx_0(X1,k6_numbers)=X1)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t4_arithm)).
fof(rqRealAdd__k2_xcmplx_0__r1_rm1_r0, axiom, k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', rqRealAdd__k2_xcmplx_0__r1_rm1_r0)).
fof(rqRealAdd__k2_xcmplx_0__r0_r1_r1, axiom, k2_xcmplx_0(np__0,np__1)=np__1, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', rqRealAdd__k2_xcmplx_0__r0_r1_r1)).
fof(t1_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1)), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t1_arithm)).
fof(fc5_xreal_0, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k2_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', fc5_xreal_0)).
fof(rqRealAdd__k2_xcmplx_0__rm1_r1_r0, axiom, k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0, file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', rqRealAdd__k2_xcmplx_0__rm1_r1_r0)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:(((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', dt_k1_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:(((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', redefinition_k1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2)))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', cc2_relset_1)).
fof(d4_funct_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2, X3]:((((r2_hidden(X2,k9_xtuple_0(X1))&r2_hidden(X3,k9_xtuple_0(X1)))&k1_funct_1(X1,X2)=k1_funct_1(X1,X3))=>X2=X3))))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', d4_funct_1)).
fof(t1_subset, axiom, ![X1, X2]:((r2_hidden(X1,X2)=>m1_subset_1(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/t50_fcont_1', t1_subset)).
fof(c_0_57, plain, ![X309, X310, X311, X312]:(((X311!=k1_fcont_1(X309,X310)|(~v1_xreal_0(X312)|k1_seq_1(X311,X312)=k2_xcmplx_0(k3_xcmplx_0(X309,X312),X310))|(~v1_funct_1(X311)|~v1_funct_2(X311,k1_numbers,k1_numbers)|~m1_subset_1(X311,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~v1_xreal_0(X310)|~v1_xreal_0(X309))&((v1_xreal_0(esk3_3(X309,X310,X311))|X311=k1_fcont_1(X309,X310)|(~v1_funct_1(X311)|~v1_funct_2(X311,k1_numbers,k1_numbers)|~m1_subset_1(X311,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~v1_xreal_0(X310)|~v1_xreal_0(X309))&(k1_seq_1(X311,esk3_3(X309,X310,X311))!=k2_xcmplx_0(k3_xcmplx_0(X309,esk3_3(X309,X310,X311)),X310)|X311=k1_fcont_1(X309,X310)|(~v1_funct_1(X311)|~v1_funct_2(X311,k1_numbers,k1_numbers)|~m1_subset_1(X311,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~v1_xreal_0(X310)|~v1_xreal_0(X309))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_fcont_1])])])])])).
cnf(c_0_58, plain, (k1_seq_1(X1,X4)=k2_xcmplx_0(k3_xcmplx_0(X2,X4),X3)|X1!=k1_fcont_1(X2,X3)|~v1_xreal_0(X4)|~v1_funct_1(X1)|~v1_funct_2(X1,k1_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_59, plain, ![X322, X323]:((((v1_funct_1(k1_fcont_1(X322,X323))|(~v1_xreal_0(X322)|~v1_xreal_0(X323)))&(v1_funct_2(k1_fcont_1(X322,X323),k1_numbers,k1_numbers)|(~v1_xreal_0(X322)|~v1_xreal_0(X323))))&(m1_subset_1(k1_fcont_1(X322,X323),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_xreal_0(X322)|~v1_xreal_0(X323))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_fcont_1])])])).
fof(c_0_60, plain, ![X351, X352]:((~v1_relat_1(X351)|~v1_funct_1(X351)|~v3_valued_0(X351)|k1_seq_1(X351,X352)=k1_funct_1(X351,X352))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
cnf(c_0_61, plain, (k1_seq_1(X1,X4)=k2_xcmplx_0(k3_xcmplx_0(X2,X4),X3)|X1!=k1_fcont_1(X2,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k1_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_62, plain, (v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_63, plain, (v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_64, plain, (m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_65, plain, ![X379, X380, X381]:((~m1_subset_1(X381,k1_zfmisc_1(k2_zfmisc_1(X379,X380)))|v1_relat_1(X381))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_66, plain, ![X408, X409, X410]:((~v3_membered(X409)|(~m1_subset_1(X410,k1_zfmisc_1(k2_zfmisc_1(X408,X409)))|v3_valued_0(X410)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
cnf(c_0_67, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, plain, (k1_seq_1(k1_fcont_1(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2)|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_61])])).
cnf(c_0_69, plain, (v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_70, plain, (v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_71, plain, (m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_72, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_75, plain, ![X369, X370, X371]:((~v1_xcmplx_0(X369)|~v1_xcmplx_0(X370)|~v1_xcmplx_0(X371)|k3_xcmplx_0(k3_xcmplx_0(X369,X370),X371)=k3_xcmplx_0(X369,k3_xcmplx_0(X370,X371)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc7_arithm])])).
fof(c_0_76, plain, ![X350]:((~v1_xcmplx_0(X350)|k3_xcmplx_0(X350,k4_xcmplx_0(np__1))=k4_xcmplx_0(X350))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc2_arithm])])).
fof(c_0_77, plain, ![X494, X495]:((~v1_relat_1(X494)|~v1_funct_1(X494)|~v1_valued_0(X494)|v1_xcmplx_0(k1_funct_1(X494,X495)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc61_valued_0])])).
cnf(c_0_78, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_79, plain, (k1_seq_1(k1_fcont_1(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_71])).
cnf(c_0_80, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_81, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_82, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_83, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_84, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_85, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_86, plain, (k1_funct_1(k1_fcont_1(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2)|~v3_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79])])).
cnf(c_0_87, plain, (v1_relat_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_71])])).
cnf(c_0_88, plain, (v3_valued_0(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_71]), c_0_82])])])).
fof(c_0_89, plain, ![X325, X326]:((~v1_xcmplx_0(X325)|~v1_xcmplx_0(X326)|k3_xcmplx_0(X325,X326)=k3_xcmplx_0(X326,X325))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_xcmplx_0])])).
cnf(c_0_90, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_91, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_92, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_93, plain, (k1_funct_1(k1_fcont_1(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_86, c_0_69]), c_0_87]), c_0_88])).
fof(c_0_94, plain, ![X518]:(((v1_relat_1(X518)|(~v1_relat_1(X518)|~v3_valued_0(X518)))&(v1_valued_0(X518)|(~v1_relat_1(X518)|~v3_valued_0(X518))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_valued_0])])])).
cnf(c_0_95, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_96, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
fof(c_0_97, plain, ![X348]:((~v1_xcmplx_0(X348)|v1_xcmplx_0(k4_xcmplx_0(X348)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_xcmplx_0])])).
cnf(c_0_98, plain, (v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2))|~v1_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_99, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
fof(c_0_100, plain, ![X346, X347]:((~v3_membered(X346)|(~m1_subset_1(X347,X346)|v1_xreal_0(X347)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_101, plain, ![X328]:(((v1_xcmplx_0(k4_xcmplx_0(X328))|~v1_xreal_0(X328))&(v1_xreal_0(k4_xcmplx_0(X328))|~v1_xreal_0(X328)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xreal_0])])])).
fof(c_0_102, plain, ![X330]:((~v1_xcmplx_0(X330)|k4_xcmplx_0(k4_xcmplx_0(X330))=X330)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
cnf(c_0_103, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_104, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_91])])).
cnf(c_0_105, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_106, plain, (v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X3))|~v1_valued_0(k1_fcont_1(X1,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_98, c_0_69]), c_0_87])).
cnf(c_0_107, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_99])).
fof(c_0_108, plain, ![X337]:((~v1_xcmplx_0(X337)|k3_xcmplx_0(np__1,X337)=X337)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
cnf(c_0_109, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_110, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_111, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_112, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_113, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_114, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_115, plain, (v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X3))|~v3_valued_0(k1_fcont_1(X1,X3))|~v1_relat_1(k1_fcont_1(X1,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107])])).
cnf(c_0_116, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_117, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_118, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_110])).
fof(c_0_119, plain, ![X338]:((~v1_xreal_0(X338)|v1_xcmplx_0(X338))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_120, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_121, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_122, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
cnf(c_0_123, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_124, plain, (v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_115, c_0_87]), c_0_88])).
cnf(c_0_125, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_126, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_82])])])).
cnf(c_0_127, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_128, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_129, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_130, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_131, plain, (m1_subset_1(np__0,k1_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_132, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
cnf(c_0_133, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_91])])).
cnf(c_0_134, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_121])])).
cnf(c_0_135, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126])])])).
cnf(c_0_136, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_137, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1)), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__rm1_r0_rm1])).
cnf(c_0_138, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_126])])])).
cnf(c_0_139, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_140, plain, (m1_subset_1(np__0,k1_numbers)), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_141, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_142, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_134, c_0_114])).
cnf(c_0_143, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_144, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_145, plain, (v1_xreal_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_126])])])).
cnf(c_0_146, plain, (v1_xreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_140]), c_0_82])])])).
cnf(c_0_147, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_141])])).
cnf(c_0_148, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_142])])).
cnf(c_0_149, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_145]), c_0_146])])])).
cnf(c_0_150, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_147, c_0_114])).
cnf(c_0_151, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_149])])])).
cnf(c_0_152, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_150, c_0_149])])])).
fof(c_0_153, plain, ![X339, X340]:((~v1_xreal_0(X339)|~v1_xreal_0(X340)|v1_xreal_0(k3_xcmplx_0(X339,X340)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_xreal_0])])).
cnf(c_0_154, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152])])).
cnf(c_0_155, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_156, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_136])])).
cnf(c_0_157, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_158, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_103])])).
fof(c_0_159, plain, ![X324]:((~v1_xboole_0(X324)|X324=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_160, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157])])).
cnf(c_0_161, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3))=k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_158])])).
cnf(c_0_162, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_163, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_164, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_160, c_0_136])).
cnf(c_0_165, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_91])])).
fof(c_0_166, plain, ![X319, X320, X321]:((~v1_xcmplx_0(X319)|(~v1_xcmplx_0(X320)|(~v1_xcmplx_0(X321)|(k3_xcmplx_0(X320,X319)!=k3_xcmplx_0(X321,X319)|(X319=k6_numbers|X320=X321)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_xcmplx_1])])])).
cnf(c_0_167, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_168, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_169, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_170, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_121])])).
cnf(c_0_171, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_91])])).
cnf(c_0_172, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_165, c_0_114])).
cnf(c_0_173, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_136])])).
cnf(c_0_174, plain, (X1=k6_numbers|X2=X3|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_175, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_167, c_0_168])).
cnf(c_0_176, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_177, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_170, c_0_136]), c_0_139])).
cnf(c_0_178, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172])])).
cnf(c_0_179, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_157])])).
cnf(c_0_180, plain, (X2=X3|X1=k6_numbers|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_181, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176])])).
cnf(c_0_182, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_177])])).
cnf(c_0_183, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_178, c_0_149])])])).
cnf(c_0_184, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_179, c_0_136])).
cnf(c_0_185, plain, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_180, c_0_181])])).
cnf(c_0_186, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_149])])])).
cnf(c_0_187, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_136])])).
cnf(c_0_188, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184])])).
cnf(c_0_189, plain, (k4_xcmplx_0(X2)=X3|X1=np__0|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186])])).
cnf(c_0_190, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_157])])).
cnf(c_0_191, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_188, c_0_136])).
cnf(c_0_192, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X1,k4_xcmplx_0(X3))!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_189, c_0_114])).
cnf(c_0_193, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_190, c_0_120])).
cnf(c_0_194, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_121])])).
cnf(c_0_195, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_121])])).
cnf(c_0_196, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193])])).
cnf(c_0_197, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_194, c_0_136]), c_0_139])).
cnf(c_0_198, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_195, c_0_136]), c_0_139])).
cnf(c_0_199, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_196, c_0_136]), c_0_136])).
cnf(c_0_200, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_201, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200])])).
cnf(c_0_202, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_201, c_0_136])).
fof(c_0_203, negated_conjecture, ~(![X1]:((v1_xreal_0(X1)=>![X2]:((v1_xreal_0(X2)=>(X1!=k6_numbers=>v2_funct_1(k1_fcont_1(X1,X2)))))))), inference(assume_negation,[status(cth)],[t50_fcont_1])).
fof(c_0_204, plain, ![X331, X332]:((~v1_xcmplx_0(X331)|~v1_xcmplx_0(X332)|k2_xcmplx_0(X331,k4_xcmplx_0(X332))=k6_xcmplx_0(X331,X332))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc1_arithm])])).
fof(c_0_205, plain, ![X372, X373]:((~v1_xcmplx_0(X372)|~v1_xcmplx_0(X373)|k2_xcmplx_0(k4_xcmplx_0(X372),k4_xcmplx_0(X373))=k4_xcmplx_0(k2_xcmplx_0(X372,X373)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc8_arithm])])).
cnf(c_0_206, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_198])])).
fof(c_0_207, negated_conjecture, (v1_xreal_0(esk1_0)&(v1_xreal_0(esk2_0)&(esk1_0!=k6_numbers&~v2_funct_1(k1_fcont_1(esk1_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_203])])])).
fof(c_0_208, plain, ![X395, X396, X397]:((~v1_xcmplx_0(X395)|~v1_xcmplx_0(X396)|~v1_xcmplx_0(X397)|k2_xcmplx_0(k2_xcmplx_0(X395,X396),X397)=k2_xcmplx_0(X395,k2_xcmplx_0(X396,X397)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc6_arithm])])).
fof(c_0_209, plain, ![X333, X334]:((~v1_xcmplx_0(X333)|~v1_xcmplx_0(X334)|k2_xcmplx_0(X333,X334)=k2_xcmplx_0(X334,X333))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_210, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_211, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_212, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_206, c_0_120])).
cnf(c_0_213, negated_conjecture, (esk1_0!=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_214, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_208])).
cnf(c_0_215, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_209])).
fof(c_0_216, plain, ![X363, X364]:((~v1_xcmplx_0(X363)|~v1_xcmplx_0(X364)|k6_xcmplx_0(k4_xcmplx_0(X363),k4_xcmplx_0(X364))=k6_xcmplx_0(X364,X363))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc9_arithm])])).
cnf(c_0_217, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_218, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_219, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_200])])).
cnf(c_0_220, negated_conjecture, (v1_xreal_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_221, negated_conjecture, (esk1_0!=k6_numbers), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_222, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_223, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_224, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_216])).
cnf(c_0_225, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218])])).
cnf(c_0_226, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|X2=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_219])])).
cnf(c_0_227, negated_conjecture, (v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_220])).
cnf(c_0_228, negated_conjecture, (esk1_0!=np__0), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_221, c_0_181])])).
cnf(c_0_229, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_217])])).
cnf(c_0_230, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_217])])).
cnf(c_0_231, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_232, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_225, c_0_114])).
cnf(c_0_233, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227]), c_0_228])])).
cnf(c_0_234, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_229, c_0_114])).
cnf(c_0_235, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_230, c_0_114])).
cnf(c_0_236, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_121])])).
cnf(c_0_237, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232])])).
cnf(c_0_238, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233])])).
cnf(c_0_239, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235])])).
cnf(c_0_240, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_236, c_0_114])).
cnf(c_0_241, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_237, c_0_114])).
cnf(c_0_242, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_238, c_0_139])).
cnf(c_0_243, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_240])])).
cnf(c_0_244, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_242])])).
cnf(c_0_245, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)), inference(csr,[status(thm)],[c_0_243, c_0_114])).
cnf(c_0_246, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_244, c_0_114]), c_0_139])).
cnf(c_0_247, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_233])])).
cnf(c_0_248, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_233])])).
cnf(c_0_249, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246])])).
cnf(c_0_250, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_247, c_0_139])).
cnf(c_0_251, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_248, c_0_139])).
cnf(c_0_252, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_249, c_0_136])).
cnf(c_0_253, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_246])])).
cnf(c_0_254, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251])])).
cnf(c_0_255, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253])])).
cnf(c_0_256, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_121])])).
cnf(c_0_257, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255])])).
cnf(c_0_258, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_256, c_0_114])).
cnf(c_0_259, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_254])])).
cnf(c_0_260, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_233])])).
fof(c_0_261, plain, ![X335, X336]:((~v1_xreal_0(X335)|~v1_xreal_0(X336)|v1_xreal_0(k6_xcmplx_0(X335,X336)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_xreal_0])])).
cnf(c_0_262, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_258, c_0_149])])])).
cnf(c_0_263, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_259])])).
cnf(c_0_264, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_260, c_0_139])).
cnf(c_0_265, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_266, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_233])])).
cnf(c_0_267, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_193])])).
cnf(c_0_268, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264])])).
cnf(c_0_269, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_270, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_233])])).
cnf(c_0_271, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_266, c_0_120])).
cnf(c_0_272, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_267, c_0_136])).
cnf(c_0_273, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_268, c_0_120]), c_0_269]), c_0_136])).
cnf(c_0_274, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_270, c_0_139])).
cnf(c_0_275, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_233])])).
cnf(c_0_276, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_271])])).
cnf(c_0_277, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_272])])).
cnf(c_0_278, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_274])])).
cnf(c_0_279, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_275, c_0_120])).
cnf(c_0_280, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_276, c_0_139])).
cnf(c_0_281, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_277, c_0_120])).
cnf(c_0_282, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279])])).
cnf(c_0_283, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_233])])).
cnf(c_0_284, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_121])])).
fof(c_0_285, plain, ![X367, X368]:((~v1_membered(X367)|(~m1_subset_1(X368,X367)|v1_xcmplx_0(X368)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
fof(c_0_286, plain, ![X374, X375]:((~v3_membered(X374)|(~m1_subset_1(X375,k1_zfmisc_1(X374))|v3_membered(X375)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_287, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_282, c_0_120])).
cnf(c_0_288, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_283, c_0_120])).
cnf(c_0_289, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_121])])).
cnf(c_0_290, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_284, c_0_136]), c_0_139])).
cnf(c_0_291, plain, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_103])])).
cnf(c_0_292, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_285])).
cnf(c_0_293, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_294, plain, ![X384]:((~v3_membered(X384)|v1_membered(X384))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_295, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_286])).
cnf(c_0_296, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_297, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_250])])).
cnf(c_0_298, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_250])])).
cnf(c_0_299, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_289, c_0_114])).
cnf(c_0_300, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_149])])])).
cnf(c_0_301, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X2,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_290])])).
cnf(c_0_302, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X1,k4_xcmplx_0(X3))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_186])])).
fof(c_0_303, plain, ![X376, X377, X378]:((~v1_xcmplx_0(X376)|~v1_xcmplx_0(X377)|~v1_xcmplx_0(X378)|k3_xcmplx_0(k2_xcmplx_0(X376,X377),X378)=k2_xcmplx_0(k3_xcmplx_0(X376,X378),k3_xcmplx_0(X377,X378)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc5_arithm])])).
fof(c_0_304, plain, ![X327]:((~v1_xcmplx_0(X327)|k3_xcmplx_0(X327,k6_numbers)=k6_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_arithm])])).
cnf(c_0_305, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_292])).
cnf(c_0_306, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_293])).
cnf(c_0_307, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_294])).
cnf(c_0_308, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_295])).
cnf(c_0_309, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_310, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_311, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_297])])).
cnf(c_0_312, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299])])).
cnf(c_0_313, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_300]), c_0_126])])])).
cnf(c_0_314, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_301, c_0_120]), c_0_120])).
cnf(c_0_315, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X1,k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_302, c_0_114])).
cnf(c_0_316, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_303])).
cnf(c_0_317, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_318, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_90])])).
cnf(c_0_319, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306])])).
cnf(c_0_320, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_321, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_309]), c_0_82])])])).
cnf(c_0_322, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_323, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_311, c_0_114]), c_0_143])).
cnf(c_0_324, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_312, c_0_136]), c_0_114])).
cnf(c_0_325, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_246])])).
cnf(c_0_326, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_313, c_0_136])).
cnf(c_0_327, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_233])])).
cnf(c_0_328, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_152])])).
cnf(c_0_329, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_316])).
cnf(c_0_330, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_317])).
fof(c_0_331, plain, ![X1]:(((~v3_xxreal_0(X1)&v1_xreal_0(X1))=>(v1_xcmplx_0(k4_xcmplx_0(X1))&~v2_xxreal_0(k4_xcmplx_0(X1))))), inference(fof_simplification,[status(thm)],[fc16_xreal_0])).
cnf(c_0_332, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_125])])).
cnf(c_0_333, plain, (v1_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_320]), c_0_321])])])).
cnf(c_0_334, plain, (k3_xcmplx_0(np__0,np__1)=np__0), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r0_r1_r0])).
cnf(c_0_335, plain, (v1_xcmplx_0(np__0)|~v1_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_322])])).
fof(c_0_336, plain, ![X1]:((v7_ordinal1(X1)=>(v7_ordinal1(X1)&~v3_xxreal_0(X1)))), inference(fof_simplification,[status(thm)],[cc3_nat_1])).
fof(c_0_337, plain, ![X353, X354]:((~v6_membered(X353)|(~m1_subset_1(X354,X353)|v7_ordinal1(X354)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_338, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_339, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_340, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_299])])).
cnf(c_0_341, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324])])).
cnf(c_0_342, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_233])])).
cnf(c_0_343, negated_conjecture, (k3_xcmplx_0(X1,np__1)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_233])])).
cnf(c_0_344, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_327, c_0_120])).
cnf(c_0_345, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_233])])).
cnf(c_0_346, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,X2),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_329])])).
cnf(c_0_347, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_330, c_0_181]), c_0_181])])).
fof(c_0_348, plain, ![X382]:(((v1_xcmplx_0(k4_xcmplx_0(X382))|(v3_xxreal_0(X382)|~v1_xreal_0(X382)))&(~v2_xxreal_0(k4_xcmplx_0(X382))|(v3_xxreal_0(X382)|~v1_xreal_0(X382))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_331])])])).
cnf(c_0_349, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_332, c_0_333])])).
cnf(c_0_350, plain, (k3_xcmplx_0(np__0,np__1)=np__0), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_351, plain, (v1_xcmplx_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_320]), c_0_321])])])).
cnf(c_0_352, plain, (k4_xcmplx_0(np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__r0_r0])).
fof(c_0_353, plain, ![X413]:(((v7_ordinal1(X413)|~v7_ordinal1(X413))&(~v3_xxreal_0(X413)|~v7_ordinal1(X413)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_336])])])).
cnf(c_0_354, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_337])).
cnf(c_0_355, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_338])).
cnf(c_0_356, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_339])).
cnf(c_0_357, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_341])])).
fof(c_0_358, plain, ![X361]:((~v1_xcmplx_0(X361)|k6_xcmplx_0(X361,k6_numbers)=X361)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_arithm])])).
cnf(c_0_359, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_342, c_0_139])).
cnf(c_0_360, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_125])])).
cnf(c_0_361, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_344]), c_0_126])])])).
cnf(c_0_362, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_rm1_r0])).
cnf(c_0_363, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_345, c_0_139])).
cnf(c_0_364, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(X2,np__1),X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_125])])).
cnf(c_0_365, plain, (k3_xcmplx_0(np__0,X1)=np__0|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_347])])).
cnf(c_0_366, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r0_r1_r1])).
cnf(c_0_367, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_348])).
cnf(c_0_368, plain, (k4_xcmplx_0(np__1)=X2|X1=np__0|k3_xcmplx_0(X2,X1)!=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_133])])).
cnf(c_0_369, plain, (k3_xcmplx_0(X1,np__0)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_350]), c_0_351])])])).
cnf(c_0_370, plain, (k4_xcmplx_0(np__0)=np__0), inference(evalgc,[status(thm)],[c_0_352])).
cnf(c_0_371, plain, (~v3_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_353])).
cnf(c_0_372, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_354])).
cnf(c_0_373, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_355, c_0_356])).
cnf(c_0_374, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_253])])).
cnf(c_0_375, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_121])])).
cnf(c_0_376, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_358])).
cnf(c_0_377, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_233])])).
cnf(c_0_378, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_360, c_0_333])])])).
cnf(c_0_379, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_361, c_0_120])).
cnf(c_0_380, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0), inference(evalgc,[status(thm)],[c_0_362])).
cnf(c_0_381, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_200])])).
cnf(c_0_382, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(X2,np__1),X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_364, c_0_333])])).
cnf(c_0_383, plain, (k3_xcmplx_0(np__0,X1)=np__0|~v1_xcmplx_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_365, c_0_351])])).
cnf(c_0_384, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(evalgc,[status(thm)],[c_0_366])).
cnf(c_0_385, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[c_0_367])).
cnf(c_0_386, plain, (k4_xcmplx_0(np__1)=X1|X2=np__0|k3_xcmplx_0(X1,X2)!=k4_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_136]), c_0_145])])])).
cnf(c_0_387, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_300]), c_0_370]), c_0_149]), c_0_351])])])).
cnf(c_0_388, plain, (v2_xxreal_0(np__1)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_389, plain, (~v7_ordinal1(X1)|~v3_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_371])).
cnf(c_0_390, plain, (v7_ordinal1(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_322]), c_0_373])])])).
cnf(c_0_391, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_374, c_0_114])).
cnf(c_0_392, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_375, c_0_114])).
cnf(c_0_393, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_376])).
cnf(c_0_394, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_377, c_0_120])).
cnf(c_0_395, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_288])])).
cnf(c_0_396, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379]), c_0_380]), c_0_149])])])).
cnf(c_0_397, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_381, c_0_136])).
cnf(c_0_398, plain, (k3_xcmplx_0(np__1,X1)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_383]), c_0_384]), c_0_351])])])).
cnf(c_0_399, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_385, c_0_121])])).
cnf(c_0_400, plain, (k4_xcmplx_0(np__1)=np__0|np__0!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_387]), c_0_129]), c_0_351]), c_0_149])])])).
cnf(c_0_401, plain, (v2_xxreal_0(np__1)), inference(evalgc,[status(thm)],[c_0_388])).
cnf(c_0_402, plain, (~v3_xxreal_0(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_390])])).
cnf(c_0_403, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_325])])).
cnf(c_0_404, plain, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_217])])).
cnf(c_0_405, plain, (k6_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_393, c_0_181])])).
cnf(c_0_406, negated_conjecture, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_394])])).
cnf(c_0_407, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_395, c_0_136])).
cnf(c_0_408, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=k3_xcmplx_0(np__0,X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_396, c_0_136])).
cnf(c_0_409, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_200])])).
cnf(c_0_410, plain, (k3_xcmplx_0(X1,np__1)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_398]), c_0_333])])])).
cnf(c_0_411, plain, (np__0!=np__1), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_400]), c_0_401]), c_0_333]), c_0_146])]), c_0_402])])).
fof(c_0_412, plain, ![X329]:((~v1_xcmplx_0(X329)|k2_xcmplx_0(X329,k6_numbers)=X329)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_413, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_403, c_0_254])])).
cnf(c_0_414, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_405]), c_0_351])])])).
cnf(c_0_415, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407])])).
cnf(c_0_416, negated_conjecture, (k6_xcmplx_0(X1,X1)=k3_xcmplx_0(np__0,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_408])])).
cnf(c_0_417, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=k3_xcmplx_0(X2,np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_410]), c_0_126])]), c_0_411])])).
cnf(c_0_418, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_412])).
cnf(c_0_419, negated_conjecture, (k2_xcmplx_0(X2,k2_xcmplx_0(k4_xcmplx_0(X1),X3))=k2_xcmplx_0(k6_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_359])])).
cnf(c_0_420, negated_conjecture, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_233])])).
cnf(c_0_421, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_406])])).
fof(c_0_422, plain, ![X344, X345]:((~v1_xreal_0(X344)|~v1_xreal_0(X345)|v1_xreal_0(k2_xcmplx_0(X344,X345)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc5_xreal_0])])).
cnf(c_0_423, negated_conjecture, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=k2_xcmplx_0(k4_xcmplx_0(X1),X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_416])])).
cnf(c_0_424, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_410])])).
cnf(c_0_425, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_418])).
cnf(c_0_426, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_419, c_0_139])).
cnf(c_0_427, negated_conjecture, (k2_xcmplx_0(X1,np__0)=X1|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_420, c_0_139])).
cnf(c_0_428, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_421, c_0_136])).
cnf(c_0_429, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_422])).
cnf(c_0_430, negated_conjecture, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=k2_xcmplx_0(k4_xcmplx_0(X1),X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_423, c_0_120]), c_0_139])).
cnf(c_0_431, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_424, c_0_136])).
cnf(c_0_432, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_425, c_0_181])])).
cnf(c_0_433, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),np__0)=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_427]), c_0_146])])])).
cnf(c_0_434, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_428])])).
cnf(c_0_435, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_429])).
cnf(c_0_436, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_394])])).
cnf(c_0_437, negated_conjecture, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_430])])).
cnf(c_0_438, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,X1),k3_xcmplx_0(X3,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_103])])).
cnf(c_0_439, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__rm1_r1_r0])).
cnf(c_0_440, negated_conjecture, (X1=X2|X1!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_432])])).
cnf(c_0_441, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),np__0)=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_433, c_0_120])).
cnf(c_0_442, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_434, c_0_435]), c_0_435]), c_0_136])).
cnf(c_0_443, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_436, c_0_136])).
cnf(c_0_444, negated_conjecture, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_437, c_0_139])).
cnf(c_0_445, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(np__1),X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_91])])).
cnf(c_0_446, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0), inference(evalgc,[status(thm)],[c_0_439])).
cnf(c_0_447, negated_conjecture, (X1=k6_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_441])])).
cnf(c_0_448, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_442])])).
cnf(c_0_449, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_264])])).
cnf(c_0_450, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_444])])).
cnf(c_0_451, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_121])])).
cnf(c_0_452, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_445, c_0_125]), c_0_446]), c_0_333])])])).
cnf(c_0_453, negated_conjecture, (X1=k6_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_447, c_0_269])).
cnf(c_0_454, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_448, c_0_435])).
cnf(c_0_455, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_449, c_0_120]), c_0_136])).
cnf(c_0_456, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_450])])).
cnf(c_0_457, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_451, c_0_114])).
cnf(c_0_458, plain, (k6_xcmplx_0(X1,X1)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_452])])).
cnf(c_0_459, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_453, c_0_454])])).
cnf(c_0_460, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_455, c_0_456])])).
cnf(c_0_461, plain, (k4_xcmplx_0(k6_xcmplx_0(X2,X1))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_235])])).
cnf(c_0_462, plain, (k6_xcmplx_0(X1,X1)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_458, c_0_149])])])).
fof(c_0_463, plain, ![X404, X405]:((~v1_relat_1(X405)|~v4_relat_1(X405,X404)|m1_subset_1(k1_relset_1(X404,X405),k1_zfmisc_1(X404)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
fof(c_0_464, plain, ![X355, X356]:((~v1_relat_1(X356)|~v4_relat_1(X356,X355)|k1_relset_1(X355,X356)=k9_xtuple_0(X356))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_465, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_459, c_0_120]), c_0_435])).
cnf(c_0_466, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_450])])).
cnf(c_0_467, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X1))|~v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_460])])).
cnf(c_0_468, plain, (k4_xcmplx_0(k3_xcmplx_0(np__0,X1))=k6_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462])])).
cnf(c_0_469, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_463])).
cnf(c_0_470, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_464])).
fof(c_0_471, plain, ![X454, X455, X456]:(((v4_relat_1(X456,X454)|~m1_subset_1(X456,k1_zfmisc_1(k2_zfmisc_1(X454,X455))))&(v5_relat_1(X456,X455)|~m1_subset_1(X456,k1_zfmisc_1(k2_zfmisc_1(X454,X455)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_472, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X3),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_466])])).
cnf(c_0_473, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_467, c_0_269])).
cnf(c_0_474, plain, (k6_xcmplx_0(X1,X1)=np__0|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_383]), c_0_370])])).
fof(c_0_475, plain, ![X314, X315, X316]:(((~v2_funct_1(X314)|(~r2_hidden(X315,k9_xtuple_0(X314))|~r2_hidden(X316,k9_xtuple_0(X314))|k1_funct_1(X314,X315)!=k1_funct_1(X314,X316)|X315=X316)|(~v1_relat_1(X314)|~v1_funct_1(X314)))&((((r2_hidden(esk4_1(X314),k9_xtuple_0(X314))|v2_funct_1(X314)|(~v1_relat_1(X314)|~v1_funct_1(X314)))&(r2_hidden(esk5_1(X314),k9_xtuple_0(X314))|v2_funct_1(X314)|(~v1_relat_1(X314)|~v1_funct_1(X314))))&(k1_funct_1(X314,esk4_1(X314))=k1_funct_1(X314,esk5_1(X314))|v2_funct_1(X314)|(~v1_relat_1(X314)|~v1_funct_1(X314))))&(esk4_1(X314)!=esk5_1(X314)|v2_funct_1(X314)|(~v1_relat_1(X314)|~v1_funct_1(X314)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_funct_1])])])])])).
fof(c_0_476, plain, ![X385, X386]:((~r2_hidden(X385,X386)|m1_subset_1(X385,X386))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_477, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_469])).
cnf(c_0_478, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_470])).
cnf(c_0_479, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_471])).
cnf(c_0_480, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_200])])).
cnf(c_0_481, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_472, c_0_254])])).
cnf(c_0_482, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_474])])).
cnf(c_0_483, plain, (k1_funct_1(X1,esk4_1(X1))=k1_funct_1(X1,esk5_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_475])).
cnf(c_0_484, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_476])).
cnf(c_0_485, plain, (r2_hidden(esk5_1(X1),k9_xtuple_0(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_475])).
cnf(c_0_486, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_478])])).
cnf(c_0_487, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_479])).
cnf(c_0_488, plain, (r2_hidden(esk4_1(X1),k9_xtuple_0(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_475])).
cnf(c_0_489, plain, (k4_xcmplx_0(X3)=np__0|k4_xcmplx_0(X1)=X2|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_177])])).
cnf(c_0_490, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_186])])).
cnf(c_0_491, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_481, c_0_269]), c_0_136])).
cnf(c_0_492, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,np__0)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_482, c_0_136])).
cnf(c_0_493, plain, (k1_funct_1(X1,esk5_1(X1))=k1_funct_1(X1,esk4_1(X1))|v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_483])).
cnf(c_0_494, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_484])).
cnf(c_0_495, plain, (v2_funct_1(X1)|r2_hidden(esk5_1(X1),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_485])).
cnf(c_0_496, plain, (v3_membered(k9_xtuple_0(X1))|~v3_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_486])])).
cnf(c_0_497, plain, (v4_relat_1(k1_fcont_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_487, c_0_71])])).
cnf(c_0_498, plain, (v2_funct_1(X1)|r2_hidden(esk4_1(X1),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_488])).
cnf(c_0_499, plain, (k4_xcmplx_0(X1)=np__0|k4_xcmplx_0(X2)=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_489, c_0_120])).
cnf(c_0_500, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|X1=np__0|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_212])])).
cnf(c_0_501, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_490, c_0_136]), c_0_136])).
cnf(c_0_502, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_492])])).
cnf(c_0_503, plain, (k1_funct_1(k1_fcont_1(X1,X2),esk4_1(k1_fcont_1(X1,X2)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_493, c_0_93])])).
cnf(c_0_504, plain, (m1_subset_1(esk5_1(X1),k9_xtuple_0(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_495])])).
cnf(c_0_505, plain, (v3_membered(k9_xtuple_0(k1_fcont_1(X1,X2)))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_497]), c_0_82])])])).
cnf(c_0_506, plain, (m1_subset_1(esk4_1(X1),k9_xtuple_0(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_498])])).
cnf(c_0_507, plain, (k4_xcmplx_0(X2)=X3|X1=np__0|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_500])])).
cnf(c_0_508, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_501])])).
cnf(c_0_509, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_432])])])).
cnf(c_0_510, plain, (k1_funct_1(k1_fcont_1(X1,X2),esk4_1(k1_fcont_1(X1,X2)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_503, c_0_69]), c_0_87])).
cnf(c_0_511, plain, (v2_funct_1(X1)|v1_xreal_0(esk5_1(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_504])])).
cnf(c_0_512, plain, (v3_membered(k9_xtuple_0(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_505, c_0_87])).
cnf(c_0_513, plain, (v2_funct_1(X1)|v1_xreal_0(esk4_1(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_506])])).
cnf(c_0_514, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X1,k4_xcmplx_0(X3))!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_507, c_0_120])).
cnf(c_0_515, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_508, c_0_120]), c_0_120])).
cnf(c_0_516, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_509, c_0_136])).
cnf(c_0_517, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_510])])).
cnf(c_0_518, plain, (v2_funct_1(k1_fcont_1(X1,X2))|v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_512])])).
cnf(c_0_519, plain, (v2_funct_1(k1_fcont_1(X1,X2))|v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_512])])).
cnf(c_0_520, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X3,k4_xcmplx_0(X1))!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_514, c_0_515])])).
cnf(c_0_521, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2),X2)=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_516, c_0_517])])).
cnf(c_0_522, plain, (v2_funct_1(k1_fcont_1(X1,X2))|v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_518, c_0_69]), c_0_87])).
cnf(c_0_523, plain, (v2_funct_1(k1_fcont_1(X1,X2))|v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_519, c_0_69]), c_0_87])).
cnf(c_0_524, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(k4_xcmplx_0(X1),X3)!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_520, c_0_200])])).
cnf(c_0_525, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2),X2)=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_521, c_0_522]), c_0_523])).
cnf(c_0_526, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(k4_xcmplx_0(X1),X3)!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_524, c_0_120])).
cnf(c_0_527, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_516, c_0_525])])).
cnf(c_0_528, negated_conjecture, (X3=np__0|X1=X2|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_233])])).
cnf(c_0_529, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_527, c_0_157])])).
cnf(c_0_530, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_528, c_0_120])).
cnf(c_0_531, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_529, c_0_522])).
cnf(c_0_532, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_200])])).
cnf(c_0_533, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_531, c_0_157])])).
cnf(c_0_534, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_532, c_0_200])])).
cnf(c_0_535, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_533, c_0_523])).
cnf(c_0_536, negated_conjecture, (X2=esk5_1(k1_fcont_1(X1,X3))|X1=np__0|v2_funct_1(k1_fcont_1(X1,X3))|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X3)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X3)))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_534, c_0_535])])).
cnf(c_0_537, negated_conjecture, (X1=esk5_1(k1_fcont_1(X2,X3))|X2=np__0|v2_funct_1(k1_fcont_1(X2,X3))|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X3)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_536, c_0_522])).
cnf(c_0_538, plain, (v2_funct_1(X1)|esk4_1(X1)!=esk5_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_475])).
cnf(c_0_539, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=esk4_1(k1_fcont_1(X1,X2))|X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_537])])).
cnf(c_0_540, plain, (v2_funct_1(X1)|esk5_1(X1)!=esk4_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_538])).
cnf(c_0_541, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=esk4_1(k1_fcont_1(X1,X2))|X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_539, c_0_523])).
cnf(c_0_542, negated_conjecture, (~v2_funct_1(k1_fcont_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_543, negated_conjecture, (X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_540, c_0_541])])).
cnf(c_0_544, negated_conjecture, (v1_xreal_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_545, negated_conjecture, (~v2_funct_1(k1_fcont_1(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[c_0_542])).
cnf(c_0_546, negated_conjecture, (X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_543, c_0_69]), c_0_87])).
cnf(c_0_547, negated_conjecture, (v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_544])).
cnf(c_0_548, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_545, c_0_546]), c_0_227]), c_0_547])]), c_0_228]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 549
# Proof object clause steps            : 447
# Proof object formula steps           : 102
# Proof object conjectures             : 143
# Proof object clause conjectures      : 140
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 69
# Proof object initial formulas used   : 57
# Proof object generating inferences   : 198
# Proof object simplifying inferences  : 224
# Training examples: 279 positive, 2803 negative
# Training: Positive examples begin
cnf(c_0_221, negated_conjecture, (esk1_0!=k6_numbers)).# trainpos
cnf(c_0_545, negated_conjecture, (~v2_funct_1(k1_fcont_1(esk1_0,esk2_0)))).# trainpos
cnf(c_0_540, plain, (v2_funct_1(X1)|esk5_1(X1)!=esk4_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_547, negated_conjecture, (v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_498, plain, (v2_funct_1(X1)|r2_hidden(esk4_1(X1),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_227, negated_conjecture, (v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_495, plain, (v2_funct_1(X1)|r2_hidden(esk5_1(X1),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_493, plain, (k1_funct_1(X1,esk5_1(X1))=k1_funct_1(X1,esk4_1(X1))|v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_168, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_69, plain, (v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_356, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_71, plain, (m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_136, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_167, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_80, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_70, plain, (v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_355, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_180, plain, (X2=X3|X1=k6_numbers|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_176, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_370, plain, (k4_xcmplx_0(np__0)=np__0)).# trainpos
cnf(c_0_494, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_309, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_82, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_61, plain, (k1_seq_1(X1,X4)=k2_xcmplx_0(k3_xcmplx_0(X2,X4),X3)|X1!=k1_fcont_1(X2,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k1_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_330, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_117, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_322, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_125, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_425, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_140, plain, (m1_subset_1(np__0,k1_numbers))).# trainpos
cnf(c_0_306, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_118, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_350, plain, (k3_xcmplx_0(np__0,np__1)=np__0)).# trainpos
cnf(c_0_103, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_384, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).# trainpos
cnf(c_0_139, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_320, plain, (v1_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_223, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_380, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0)).# trainpos
cnf(c_0_401, plain, (v2_xxreal_0(np__1))).# trainpos
cnf(c_0_120, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_393, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_114, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_157, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_121, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_308, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_385, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_446, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0)).# trainpos
cnf(c_0_129, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_269, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_435, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_78, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_217, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_90, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_389, plain, (~v7_ordinal1(X1)|~v3_xxreal_0(X1))).# trainpos
cnf(c_0_81, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_91, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_222, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_329, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_218, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_231, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_107, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_144, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_487, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_478, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_372, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_305, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_92, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_477, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_181, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_68, plain, (k1_seq_1(k1_fcont_1(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2)|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_146, plain, (v1_xreal_0(np__0))).# trainpos
cnf(c_0_506, plain, (m1_subset_1(esk4_1(X1),k9_xtuple_0(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_504, plain, (m1_subset_1(esk5_1(X1),k9_xtuple_0(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_87, plain, (v1_relat_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_126, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_486, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).# trainpos
cnf(c_0_128, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_497, plain, (v4_relat_1(k1_fcont_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_321, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_318, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_390, plain, (v7_ordinal1(np__0))).# trainpos
cnf(c_0_158, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_88, plain, (v3_valued_0(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_360, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_438, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,X1),k3_xcmplx_0(X3,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_319, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).# trainpos
cnf(c_0_335, plain, (v1_xcmplx_0(np__0)|~v1_membered(k5_numbers))).# trainpos
cnf(c_0_346, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,X2),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_230, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_236, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_225, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_229, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_133, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_375, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_96, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_399, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_132, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_228, negated_conjecture, (esk1_0!=np__0)).# trainpos
cnf(c_0_402, plain, (~v3_xxreal_0(np__0))).# trainpos
cnf(c_0_185, plain, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_86, plain, (k1_funct_1(k1_fcont_1(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2)|~v3_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_432, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_347, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_351, plain, (v1_xcmplx_0(np__0))).# trainpos
cnf(c_0_405, plain, (k6_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_333, plain, (v1_xcmplx_0(np__1))).# trainpos
cnf(c_0_104, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_141, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_404, plain, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_138, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_451, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_364, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(X2,np__1),X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_171, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_237, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_496, plain, (v3_membered(k9_xtuple_0(X1))|~v3_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).# trainpos
cnf(c_0_513, plain, (v2_funct_1(X1)|v1_xreal_0(esk4_1(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_239, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_332, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_445, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(np__1),X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_511, plain, (v2_funct_1(X1)|v1_xreal_0(esk5_1(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_161, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3))=k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_291, plain, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_503, plain, (k1_funct_1(k1_fcont_1(X1,X2),esk4_1(k1_fcont_1(X1,X2)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_145, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_365, plain, (k3_xcmplx_0(np__0,X1)=np__0|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_505, plain, (v3_membered(k9_xtuple_0(k1_fcont_1(X1,X2)))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_243, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_256, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_368, plain, (k4_xcmplx_0(np__1)=X2|X1=np__0|k3_xcmplx_0(X2,X1)!=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_378, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_369, plain, (k3_xcmplx_0(X1,np__0)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_414, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_461, plain, (k4_xcmplx_0(k6_xcmplx_0(X2,X1))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_517, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_519, plain, (v2_funct_1(k1_fcont_1(X1,X2))|v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_518, plain, (v2_funct_1(k1_fcont_1(X1,X2))|v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_386, plain, (k4_xcmplx_0(np__1)=X1|X2=np__0|k3_xcmplx_0(X1,X2)!=k4_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_165, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_113, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_147, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_398, plain, (k3_xcmplx_0(np__1,X1)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_452, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_410, plain, (k3_xcmplx_0(X1,np__1)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_458, plain, (k6_xcmplx_0(X1,X1)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_98, plain, (v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2))|~v1_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_134, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_178, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_115, plain, (v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,X2),X3))|~v3_valued_0(k1_fcont_1(X1,X3))|~v1_relat_1(k1_fcont_1(X1,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_148, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_135, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_149, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_262, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_186, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_183, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_300, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_151, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_462, plain, (k6_xcmplx_0(X1,X1)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_152, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_154, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_173, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2))))).# trainpos
cnf(c_0_387, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__0)).# trainpos
cnf(c_0_302, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X1,k4_xcmplx_0(X3))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_468, plain, (k4_xcmplx_0(k3_xcmplx_0(np__0,X1))=k6_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_189, plain, (k4_xcmplx_0(X2)=X3|X1=np__0|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_474, plain, (k6_xcmplx_0(X1,X1)=np__0|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_328, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_179, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_400, plain, (k4_xcmplx_0(np__1)=np__0|np__0!=np__1)).# trainpos
cnf(c_0_156, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2))))).# trainpos
cnf(c_0_160, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_411, plain, (np__0!=np__1)).# trainpos
cnf(c_0_188, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_194, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_170, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_195, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_490, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_200, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_182, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_508, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_187, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_313, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_190, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_196, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_267, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).# trainpos
cnf(c_0_201, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_277, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).# trainpos
cnf(c_0_284, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_206, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_480, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_500, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|X1=np__0|~v1_xreal_0(X1))).# trainpos
cnf(c_0_219, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_301, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X2,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_489, plain, (k4_xcmplx_0(X3)=np__0|k4_xcmplx_0(X1)=X2|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_226, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|X2=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_507, plain, (k4_xcmplx_0(X2)=X3|X1=np__0|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_233, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_520, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X3,k4_xcmplx_0(X1))!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_343, negated_conjecture, (k3_xcmplx_0(X1,np__1)=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_247, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_270, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_420, negated_conjecture, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_248, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_254, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_238, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_345, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_327, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_381, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_409, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_417, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=k3_xcmplx_0(X2,np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_524, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(k4_xcmplx_0(X1),X3)!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_528, negated_conjecture, (X3=np__0|X1=X2|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_424, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_361, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_532, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_534, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_440, negated_conjecture, (X1=X2|X1!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_244, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_253, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_325, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_266, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_260, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_342, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_289, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_275, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_377, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_406, negated_conjecture, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_249, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_255, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_257, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3))).# trainpos
cnf(c_0_276, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_259, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3))).# trainpos
cnf(c_0_283, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3))).# trainpos
cnf(c_0_263, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3))).# trainpos
cnf(c_0_436, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_395, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_415, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3))).# trainpos
cnf(c_0_268, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_449, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_421, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_278, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_434, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_448, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_282, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_297, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_311, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_340, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_298, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_396, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_312, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_416, negated_conjecture, (k6_xcmplx_0(X1,X1)=k3_xcmplx_0(np__0,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_341, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_423, negated_conjecture, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=k2_xcmplx_0(k4_xcmplx_0(X1),X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_357, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_437, negated_conjecture, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_374, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_450, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_403, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_456, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_460, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_467, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X1))|~v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_482, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_413, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_419, negated_conjecture, (k2_xcmplx_0(X2,k2_xcmplx_0(k4_xcmplx_0(X1),X3))=k2_xcmplx_0(k6_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_466, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_433, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),np__0)=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_447, negated_conjecture, (X1=k6_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_459, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_472, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X3),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_481, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_502, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_509, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_521, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2),X2)=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_527, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_529, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_533, negated_conjecture, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_536, negated_conjecture, (X2=esk5_1(k1_fcont_1(X1,X3))|X1=np__0|v2_funct_1(k1_fcont_1(X1,X3))|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X3)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X3)))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_539, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=esk4_1(k1_fcont_1(X1,X2))|X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_543, negated_conjecture, (X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_116, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_77, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_53, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_82, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_133, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_66, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_318, plain, (~v1_xboole_0(k4_numbers))).#trainneg
cnf(i_0_112, plain, (~v1_xboole_0(esk9_0))).#trainneg
cnf(i_0_204, plain, (~v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_259, plain, (~v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_127, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_214, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_124, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_169, plain, (v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_288, plain, (~v1_xboole_0(esk23_0))).#trainneg
cnf(i_0_300, plain, (~v1_xboole_0(esk24_0))).#trainneg
cnf(i_0_355, plain, (~v1_xboole_0(esk26_0))).#trainneg
cnf(i_0_381, plain, (~v1_xboole_0(esk28_0))).#trainneg
cnf(i_0_405, plain, (~v1_xboole_0(esk30_0))).#trainneg
cnf(i_0_248, plain, (~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_247, plain, (~v4_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_46, plain, (v1_xreal_0(esk7_0))).#trainneg
cnf(i_0_126, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12, plain, (X2=X3|k1_funct_1(X1,X2)!=k1_funct_1(X1,X3)|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_80, plain, (v1_xboole_0(o_0_0_xboole_0))).#trainneg
cnf(i_0_190, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_371, plain, (v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_131, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_36, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_137, plain, (v1_xreal_0(esk10_0))).#trainneg
cnf(i_0_207, plain, (v1_xreal_0(esk14_0))).#trainneg
cnf(i_0_229, plain, (v1_xreal_0(esk17_0))).#trainneg
cnf(i_0_41, plain, (v1_funct_1(esk6_0))).#trainneg
cnf(i_0_200, plain, (v1_funct_1(esk13_0))).#trainneg
cnf(i_0_260, plain, (v1_funct_1(esk18_0))).#trainneg
cnf(i_0_296, plain, (v1_funct_1(esk24_0))).#trainneg
cnf(i_0_44, plain, (v1_relat_1(esk6_0))).#trainneg
cnf(i_0_203, plain, (v1_relat_1(esk13_0))).#trainneg
cnf(i_0_106, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_385, plain, (v3_xxreal_2(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_245, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_213, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_134, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_55, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_141, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_263, plain, (v1_relat_1(esk18_0))).#trainneg
cnf(i_0_299, plain, (v1_relat_1(esk24_0))).#trainneg
cnf(i_0_380, plain, (v1_relat_1(esk28_0))).#trainneg
cnf(i_0_48, plain, (v1_xcmplx_0(esk7_0))).#trainneg
cnf(i_0_140, plain, (v1_xcmplx_0(esk10_0))).#trainneg
cnf(i_0_232, plain, (v1_xcmplx_0(esk17_0))).#trainneg
cnf(i_0_40, plain, (v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_49, plain, (v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_109, plain, (v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_115, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_128, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_110, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_78, plain, (k3_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_178, plain, (v1_xboole_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_relat_1(X1))).#trainneg
cnf(i_0_173, plain, (k3_xcmplx_0(np__1,np__0)=np__0)).#trainneg
cnf(i_0_91, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k1_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_352, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_242, plain, (v5_xxreal_2(X1)|~v3_membered(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_152, plain, (v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_39, plain, (v1_valued_0(esk6_0))).#trainneg
cnf(i_0_197, plain, (v1_valued_0(esk13_0))).#trainneg
cnf(i_0_256, plain, (v1_valued_0(esk18_0))).#trainneg
cnf(i_0_292, plain, (v1_valued_0(esk24_0))).#trainneg
cnf(i_0_38, plain, (v2_valued_0(esk6_0))).#trainneg
cnf(i_0_196, plain, (v2_valued_0(esk13_0))).#trainneg
cnf(i_0_255, plain, (v2_valued_0(esk18_0))).#trainneg
cnf(i_0_291, plain, (v2_valued_0(esk24_0))).#trainneg
cnf(i_0_73, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_97, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_81, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_301, plain, (m1_subset_1(esk24_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_135, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_95, plain, (~v3_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_246, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_344, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_378, plain, (v4_xxreal_2(X1)|~v3_membered(X1)|~v2_xxreal_2(X1))).#trainneg
cnf(i_0_37, plain, (v3_valued_0(esk6_0))).#trainneg
cnf(i_0_195, plain, (v3_valued_0(esk13_0))).#trainneg
cnf(i_0_254, plain, (v3_valued_0(esk18_0))).#trainneg
cnf(i_0_290, plain, (v3_valued_0(esk24_0))).#trainneg
cnf(i_0_47, plain, (v1_xxreal_0(esk7_0))).#trainneg
cnf(i_0_108, plain, (v1_xxreal_0(esk8_0))).#trainneg
cnf(i_0_139, plain, (v1_xxreal_0(esk10_0))).#trainneg
cnf(i_0_224, plain, (v1_xxreal_0(esk16_0))).#trainneg
cnf(i_0_231, plain, (v1_xxreal_0(esk17_0))).#trainneg
cnf(i_0_96, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_205, plain, (m1_subset_1(esk13_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_107, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_176, plain, (k6_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_266, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_392, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_119, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_28, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_125, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_264, plain, (m1_subset_1(esk18_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_362, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_142, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_265, plain, (v1_xxreal_0(esk19_0))).#trainneg
cnf(i_0_359, plain, (v3_membered(esk26_0))).#trainneg
cnf(i_0_111, plain, (v7_ordinal1(esk9_0))).#trainneg
cnf(i_0_285, plain, (v7_ordinal1(esk22_0))).#trainneg
cnf(i_0_356, plain, (v6_membered(esk26_0))).#trainneg
cnf(i_0_230, plain, (v2_xxreal_0(esk17_0))).#trainneg
cnf(i_0_65, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_287, plain, (v3_ordinal1(esk23_0))).#trainneg
cnf(i_0_75, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_376, plain, (v1_relat_1(esk27_1(X1)))).#trainneg
cnf(i_0_45, plain, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_373, plain, (v1_funct_1(esk27_1(X1)))).#trainneg
cnf(i_0_63, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_361, plain, (v1_membered(esk26_0))).#trainneg
cnf(i_0_138, plain, (v3_xxreal_0(esk10_0))).#trainneg
cnf(i_0_223, plain, (v3_xxreal_0(esk16_0))).#trainneg
cnf(i_0_240, plain, (v5_membered(k4_numbers))).#trainneg
cnf(i_0_357, plain, (v5_membered(esk26_0))).#trainneg
cnf(i_0_335, plain, (v4_membered(k3_numbers))).#trainneg
cnf(i_0_358, plain, (v4_membered(esk26_0))).#trainneg
cnf(i_0_253, plain, (v1_fcont_1(esk18_0))).#trainneg
cnf(i_0_360, plain, (v2_membered(esk26_0))).#trainneg
cnf(i_0_293, plain, (v1_funct_2(esk24_0,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_402, plain, (v1_relat_1(esk29_1(X1)))).#trainneg
cnf(i_0_273, plain, (v1_relat_1(esk20_1(X1)))).#trainneg
cnf(i_0_399, plain, (v1_funct_1(esk29_1(X1)))).#trainneg
cnf(i_0_225, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_74, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_275, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_150, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k9_xtuple_0(X1)))).#trainneg
cnf(i_0_194, plain, (v4_valued_0(esk13_0))).#trainneg
cnf(i_0_193, plain, (v5_valued_0(esk13_0))).#trainneg
cnf(i_0_353, plain, (v5_xxreal_2(esk26_0))).#trainneg
cnf(i_0_295, plain, (v3_funct_1(esk24_0))).#trainneg
cnf(i_0_354, plain, (v1_xxreal_2(esk26_0))).#trainneg
cnf(i_0_404, plain, (v1_zfmisc_1(esk30_0))).#trainneg
cnf(i_0_43, plain, (v4_relat_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_202, plain, (v4_relat_1(esk13_0,k5_numbers))).#trainneg
cnf(i_0_262, plain, (v4_relat_1(esk18_0,k1_numbers))).#trainneg
cnf(i_0_226, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_340, plain, (v1_relat_1(esk25_1(X1)))).#trainneg
cnf(i_0_316, plain, (v3_relat_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_94, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_117, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_294, plain, (v1_partfun1(esk24_0,k5_numbers))).#trainneg
cnf(i_0_172, plain, (m1_subset_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_270, plain, (v1_funct_1(esk20_1(X1)))).#trainneg
cnf(i_0_201, plain, (v5_relat_1(esk13_0,k5_numbers))).#trainneg
cnf(i_0_319, plain, (v1_xboole_0(X1)|v1_xxreal_2(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_191, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_338, plain, (v1_funct_1(esk25_1(X1)))).#trainneg
cnf(i_0_397, plain, (v1_valued_0(esk29_1(X1)))).#trainneg
cnf(i_0_396, plain, (v2_valued_0(esk29_1(X1)))).#trainneg
cnf(i_0_298, plain, (v4_relat_1(esk24_0,k5_numbers))).#trainneg
cnf(i_0_42, plain, (v5_relat_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_261, plain, (v5_relat_1(esk18_0,k1_numbers))).#trainneg
cnf(i_0_297, plain, (v5_relat_1(esk24_0,k1_numbers))).#trainneg
cnf(i_0_199, plain, (v1_partfun1(esk13_0,k5_numbers))).#trainneg
cnf(i_0_258, plain, (v1_partfun1(esk18_0,k1_numbers))).#trainneg
cnf(i_0_144, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_158, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_72, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_289, plain, (m1_subset_1(esk23_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_187, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X2)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_395, plain, (v3_valued_0(esk29_1(X1)))).#trainneg
cnf(i_0_336, plain, (v4_valued_0(esk25_1(X1)))).#trainneg
cnf(i_0_394, plain, (v4_valued_0(esk29_1(X1)))).#trainneg
cnf(i_0_269, plain, (v3_funct_1(esk20_1(X1)))).#trainneg
cnf(i_0_198, plain, (v1_funct_2(esk13_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_257, plain, (v1_funct_2(esk18_0,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_244, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_372, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_181, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_389, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v3_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_129, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_120, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_180, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_130, plain, (v1_xboole_0(k2_zfmisc_1(X2,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_235, plain, (v1_xboole_0(X1)|v2_xxreal_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_222, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X2)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_375, plain, (v2_relat_1(esk27_1(X1)))).#trainneg
cnf(i_0_100, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_149, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_252, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_387, plain, (~v1_xboole_0(X1)|~v2_membered(X1)|~v1_xxreal_2(X1))).#trainneg
cnf(i_0_365, plain, (~v1_xboole_0(X1)|~v2_membered(X1)|~v2_xxreal_2(X1))).#trainneg
cnf(i_0_98, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_156, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_267, plain, (v6_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_136, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_153, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_154, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_327, plain, (v3_xxreal_0(X2)|v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_101, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_241, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_276, plain, (v1_fcont_1(X1)|~v1_funct_1(X1)|~v2_fcont_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_272, plain, (v4_relat_1(esk20_1(X1),k5_numbers))).#trainneg
cnf(i_0_400, plain, (v5_relat_1(esk29_1(X1),k3_numbers))).#trainneg
cnf(i_0_339, plain, (v4_relat_1(esk25_1(X1),X1))).#trainneg
cnf(i_0_183, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_151, plain, (v1_rat_1(X2)|~v4_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_160, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_212, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_331, plain, (v1_fcont_1(X1)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_145, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_250, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_368, plain, (v1_comseq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_177, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_184, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_374, plain, (v4_relat_1(esk27_1(X1),X1))).#trainneg
cnf(i_0_398, plain, (v1_partfun1(esk29_1(X1),X1))).#trainneg
cnf(i_0_348, plain, (v2_fcont_1(X1)|~v1_funct_1(X1)|~v3_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_206, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_271, plain, (v5_relat_1(esk20_1(X1),X1))).#trainneg
cnf(i_0_322, plain, (v1_comseq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v3_funct_1(X1))).#trainneg
cnf(i_0_329, plain, (v1_rfunct_2(X1)|~v1_funct_1(X1)|~v7_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_401, plain, (v4_relat_1(esk29_1(X1),X1))).#trainneg
cnf(i_0_337, plain, (v1_partfun1(esk25_1(X1),X1))).#trainneg
cnf(i_0_280, plain, (v1_funct_1(esk21_2(X1,X2)))).#trainneg
cnf(i_0_217, plain, (v1_relat_1(esk15_2(X1,X2)))).#trainneg
cnf(i_0_333, plain, (v1_fcont_1(X1)|~v1_funct_1(X1)|~v3_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_346, plain, (v2_fcont_1(X1)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_165, plain, (v1_seq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_164, plain, (v2_seq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_384, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6, plain, (X3=k1_fcont_1(X1,X2)|v1_xreal_0(esk3_3(X1,X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,k1_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_383, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_102, plain, (v1_fcont_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_382, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_351, plain, (v1_xboole_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_251, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_175, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_328, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_350, plain, (v1_xboole_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_179, plain, (v1_partfun1(X2,X1)|~v1_xboole_0(X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_182, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_208, plain, (v1_comseq_2(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_seq_2(X1)|~v2_seq_2(X1))).#trainneg
cnf(i_0_326, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_283, plain, (v1_relat_1(esk21_2(X1,X2)))).#trainneg
cnf(i_0_218, plain, (v1_xboole_0(esk15_2(X1,X2)))).#trainneg
cnf(i_0_216, plain, (v4_relat_1(esk15_2(X1,X2),X1))).#trainneg
cnf(i_0_282, plain, (v4_relat_1(esk21_2(X1,X2),X1))).#trainneg
cnf(i_0_233, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_268, plain, (v1_funct_2(esk20_1(X1),k5_numbers,X1))).#trainneg
cnf(i_0_274, plain, (m1_subset_1(esk20_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_377, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_161, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_403, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_249, plain, (v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k6_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_163, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_189, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_278, plain, (v2_xxreal_0(X2)|v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_174, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_345, plain, (v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k6_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_162, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k3_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_393, plain, (v4_relat_1(X3,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_123, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_215, plain, (v5_relat_1(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_281, plain, (v5_relat_1(esk21_2(X1,X2),X2))).#trainneg
cnf(i_0_367, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_279, plain, (v1_funct_2(esk21_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_219, plain, (m1_subset_1(esk15_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_284, plain, (m1_subset_1(esk21_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_887, plain, (~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1098, plain, (~v3_xxreal_0(esk22_0))).#trainneg
cnf(i_0_1099, plain, (~v3_xxreal_0(esk9_0))).#trainneg
cnf(i_0_1268, plain, (~v1_xboole_0(esk17_0))).#trainneg
cnf(i_0_1271, plain, (~v2_xxreal_0(esk8_0))).#trainneg
cnf(i_0_998, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_997, plain, (X1=esk6_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_901, plain, (v2_funct_1(X1)|~v1_xboole_0(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_875, negated_conjecture, (~v1_xboole_0(k1_fcont_1(esk1_0,esk2_0)))).#trainneg
cnf(i_0_903, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_879, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1272, plain, (~v2_xxreal_0(esk7_0))).#trainneg
cnf(i_0_1277, plain, (~v1_xboole_0(esk16_0))).#trainneg
cnf(i_0_1278, plain, (~v1_xboole_0(esk10_0))).#trainneg
cnf(i_0_1279, plain, (~v3_xxreal_0(esk8_0))).#trainneg
cnf(i_0_1280, plain, (~v3_xxreal_0(esk7_0))).#trainneg
cnf(i_0_1288, plain, (~v3_xxreal_0(esk17_0))).#trainneg
cnf(i_0_1289, plain, (~v2_xxreal_0(esk16_0))).#trainneg
cnf(i_0_1290, plain, (~v2_xxreal_0(esk10_0))).#trainneg
cnf(i_0_978, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1298, plain, (~v1_xboole_0(k1_fcont_1(X1,X2))|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_952, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1184, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_940, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_890, plain, (v1_funct_1(np__0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_990, plain, (v1_xreal_0(esk9_0))).#trainneg
cnf(i_0_991, plain, (v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_1030, plain, (~m1_subset_1(esk10_0,k5_numbers))).#trainneg
cnf(i_0_1031, plain, (~m1_subset_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_1267, plain, (~v2_xxreal_0(esk19_0)|~v1_xboole_0(esk19_0))).#trainneg
cnf(i_0_1275, plain, (~v3_xxreal_0(esk19_0)|~v1_xboole_0(esk19_0))).#trainneg
cnf(i_0_1274, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1432, plain, (v1_comseq_2(X1)|~v3_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_921, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1314, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_909, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1054, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),X3)|~v3_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_934, plain, (X1=k1_fcont_1(np__1,X2)|k1_seq_1(X1,esk3_3(np__1,X2,X1))!=k2_xcmplx_0(esk3_3(np__1,X2,X1),X2)|~v1_xcmplx_0(esk3_3(np__1,X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(np__1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1100, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_895, plain, (v1_funct_1(np__0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1032, plain, (v3_membered(k3_numbers))).#trainneg
cnf(i_0_1232, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_1287, plain, (~v3_xxreal_0(esk19_0)|~v2_xxreal_0(esk19_0))).#trainneg
cnf(i_0_1504, plain, (~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1059, plain, (v1_xboole_0(k1_funct_1(X1,X2))|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1082, plain, (X1=k1_fcont_1(k3_xcmplx_0(X2,X3),X4)|k1_seq_1(X1,esk3_3(k3_xcmplx_0(X2,X3),X4,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,esk3_3(k3_xcmplx_0(X2,X3),X4,X1))),X4)|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(X2,X3),X4,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,X3))|~v1_xreal_0(X4))).#trainneg
cnf(i_0_996, plain, (X1=esk7_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1320, plain, (v1_partfun1(k1_fcont_1(X1,X2),k1_numbers)|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1141, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),X2)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1262, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_898, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1371, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_1374, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_994, plain, (X1=esk11_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_912, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_958, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_959, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_961, plain, (v1_relat_1(np__0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_962, plain, (v1_relat_1(np__0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1147, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1071, plain, (k6_xcmplx_0(X1,np__0)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_941, plain, (X1=X2|v1_xboole_0(k9_xtuple_0(X3))|k1_funct_1(X3,X1)!=k1_funct_1(X3,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~m1_subset_1(X2,k9_xtuple_0(X3))|~v1_funct_1(X3)|~v2_funct_1(X3))).#trainneg
cnf(i_0_1179, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1029, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_1351, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1261, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_910, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k1_fcont_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_995, plain, (X1=esk8_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1248, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1251, plain, (v1_xcmplx_0(np__0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1328, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1329, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1315, plain, (~v4_relat_1(X3,X1)|~v1_xboole_0(X1)|~r2_hidden(X2,k1_relset_1(X1,X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1133, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1166, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1294, plain, (~v3_xxreal_0(X1)|~v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_948, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk24_0))).#trainneg
cnf(i_0_1331, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1332, plain, (v1_xxreal_0(np__0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1456, plain, (v3_valued_0(esk20_1(k1_numbers)))).#trainneg
cnf(i_0_1359, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1386, plain, (v1_rat_1(np__1)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_1387, plain, (v1_rat_1(np__1)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1344, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1519, plain, (v3_xxreal_0(k1_funct_1(X1,X2))|v2_xxreal_0(k1_funct_1(X1,X2))|v1_xboole_0(k1_funct_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1120, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1182, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1175, plain, (v5_relat_1(k1_fcont_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1038, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1039, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_977, plain, (v1_xboole_0(X2)|~r2_hidden(X2,X1)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_972, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~r2_hidden(X1,esk13_0))).#trainneg
cnf(i_0_1389, plain, (v1_rat_1(np__0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_1390, plain, (v1_rat_1(np__0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_1401, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1402, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1404, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_1405, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1416, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_989, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_1119, plain, (v1_funct_2(esk13_0,k5_numbers,X1)|~m1_subset_1(esk13_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1117, plain, (v1_funct_2(esk24_0,k5_numbers,X1)|~m1_subset_1(esk24_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1308, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1659, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(X2)|~m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1626, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1125, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1181, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1621, plain, (v1_xboole_0(esk20_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1712, plain, (v4_relat_1(esk24_0,X1)|~v4_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),X1))).#trainneg
cnf(i_0_1061, plain, (X1=k9_xtuple_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_975, plain, (v2_funct_1(X1)|~r2_hidden(k9_xtuple_0(X1),esk5_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_976, plain, (v2_funct_1(X1)|~r2_hidden(k9_xtuple_0(X1),esk4_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1742, plain, (v1_partfun1(esk20_1(k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1022, plain, (v3_membered(esk24_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1023, plain, (v3_membered(esk18_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1024, plain, (v3_membered(esk13_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1025, plain, (v3_membered(esk6_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1352, plain, (v2_membered(esk24_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1354, plain, (v2_membered(esk18_0)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1178, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1118, plain, (v1_funct_2(esk18_0,k1_numbers,X1)|~m1_subset_1(esk18_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1306, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1037, plain, (v3_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1311, plain, (v1_relat_1(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1305, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1723, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1067, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1176, plain, (v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1638, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_986, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,esk18_0))).#trainneg
cnf(i_0_1191, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk23_0))).#trainneg
cnf(i_0_1355, plain, (v2_membered(esk13_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1357, plain, (v2_membered(esk6_0)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1409, plain, (v1_membered(esk24_0)|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1411, plain, (v1_membered(esk18_0)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1412, plain, (v1_membered(esk13_0)|~v1_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1414, plain, (v1_membered(esk6_0)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1065, plain, (X1=k1_fcont_1(X2,k4_xcmplx_0(X3))|k1_seq_1(X1,esk3_3(X2,k4_xcmplx_0(X3),X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)),X3)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1168, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1672, plain, (v2_xxreal_0(k6_xcmplx_0(X2,X1))|v2_xxreal_0(k4_xcmplx_0(X2))|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1159, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1719, plain, (v4_relat_1(k1_relset_1(X1,X2),X3)|~v4_relat_1(X1,X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1169, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1531, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(k6_xcmplx_0(X2,X1))|~v2_xxreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1312, plain, (v3_membered(k1_relset_1(X1,X2))|~v3_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1596, plain, (v1_valued_0(k1_fcont_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1005, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,esk6_0))).#trainneg
cnf(i_0_1077, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk12_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1625, plain, (~v3_xxreal_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1722, plain, (~v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1036, plain, (v2_xxreal_0(k4_xcmplx_0(np__1))|~v3_xxreal_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1255, plain, (v3_xxreal_0(esk19_0)|v2_xxreal_0(esk19_0)|v1_xboole_0(esk19_0))).#trainneg
cnf(i_0_1637, plain, (~v4_valued_0(X1)|~v3_xxreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1518, plain, (~v3_xxreal_0(k1_funct_1(X1,X2))|~v2_xxreal_0(k1_funct_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1512, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xboole_0(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1663, plain, (v3_xxreal_0(k6_xcmplx_0(X2,X1))|v2_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_1581, plain, (v2_xxreal_0(k4_xcmplx_0(X2))|v2_xxreal_0(X1)|~v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1688, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1553, plain, (v1_xboole_0(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1358, plain, (v2_membered(k1_relset_1(X1,X2))|~v2_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1335, plain, (v3_xxreal_0(X2)|~v3_xxreal_0(X1)|~m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1016, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1027, plain, (v3_membered(k1_fcont_1(X1,X2))|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1360, plain, (v2_membered(k1_fcont_1(X1,X2))|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1417, plain, (v1_membered(k1_fcont_1(X1,X2))|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1597, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk20_1(X1))|~m1_subset_1(esk20_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1470, plain, (v2_xxreal_0(X2)|~v3_xxreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1486, plain, (v2_xxreal_0(X1)|~v3_xxreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1069, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__1))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1139, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1727, plain, (v7_ordinal1(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1160, plain, (v1_xreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1734, plain, (v7_ordinal1(k6_xcmplx_0(X1,X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1415, plain, (v1_membered(k1_relset_1(X1,X2))|~v1_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1598, plain, (v1_partfun1(esk20_1(X1),k5_numbers)|v1_xboole_0(X1)|~m1_subset_1(esk20_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_1538, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(k6_xcmplx_0(X2,X1))|~v3_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_1339, plain, (v3_xxreal_0(k6_xcmplx_0(X2,X1))|v3_xxreal_0(k4_xcmplx_0(X2))|~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1482, plain, (v3_xxreal_0(k6_xcmplx_0(X1,X2))|v2_xxreal_0(k4_xcmplx_0(X2))|~v3_xxreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1345, plain, (v1_funct_1(k1_seq_1(X1,X2))|~v4_funct_1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1428, plain, (v3_xxreal_0(k4_xcmplx_0(X2))|v3_xxreal_0(X1)|~v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1053, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1469, plain, (v1_valued_0(k1_fcont_1(X1,X2))|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1377, plain, (v3_ordinal1(k1_fcont_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1370, plain, (v3_ordinal1(k1_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1143, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1149, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__1)=k2_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1153, plain, (k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__0)=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1346, plain, (v1_relat_1(k1_seq_1(X1,X2))|~v4_funct_1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1347, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1498, plain, (v3_xxreal_0(k6_xcmplx_0(X1,X2))|v2_xxreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1092, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1348, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1376, plain, (v3_ordinal1(k1_seq_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1391, plain, (v1_rat_1(k1_seq_1(X1,X2))|~v4_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1406, plain, (v1_int_1(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1145, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(np__1)),k4_xcmplx_0(X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1155, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1643, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(k4_xcmplx_0(X2))|~v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1695, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(k4_xcmplx_0(X2))|~v3_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1089, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X4)=k3_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1124, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_53, plain, (o_0_0_xboole_0=np__0)).#trainneg
cnf(i_0_24, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_2007, plain, (~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_1810, plain, (esk6_0=np__0)).#trainneg
cnf(i_0_2098, plain, (v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,X3)|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X3)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1811, plain, (k9_xtuple_0(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1809, plain, (esk7_0=np__0)).#trainneg
cnf(i_0_1899, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1912, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1807, plain, (esk11_0=np__0)).#trainneg
cnf(i_0_1808, plain, (esk8_0=np__0)).#trainneg
cnf(i_0_1840, negated_conjecture, (~v1_xboole_0(k9_xtuple_0(k1_fcont_1(esk1_0,esk2_0)))|~v1_relat_1(k1_fcont_1(esk1_0,esk2_0))|~v1_funct_1(k1_fcont_1(esk1_0,esk2_0)))).#trainneg
cnf(i_0_1906, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1803, plain, (k1_fcont_1(X1,X2)=k1_fcont_1(X3,X4)|k2_xcmplx_0(k3_xcmplx_0(X1,esk3_3(X3,X4,k1_fcont_1(X1,X2))),X2)!=k2_xcmplx_0(k3_xcmplx_0(X3,esk3_3(X3,X4,k1_fcont_1(X1,X2))),X4)|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk3_3(X3,X4,k1_fcont_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2013, plain, (k1_funct_1(X1,X2)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3020, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k4_xcmplx_0(X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2658, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,X1)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_1856, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1805, plain, (m1_subset_1(k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2),k1_numbers)|~v3_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2012, plain, (~v1_xboole_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1911, plain, (v3_membered(k4_numbers))).#trainneg
cnf(i_0_2337, plain, (v2_xxreal_0(esk9_0))).#trainneg
cnf(i_0_2011, plain, (~m1_subset_1(k4_xcmplx_0(np__1),k5_numbers))).#trainneg
cnf(i_0_1900, plain, (~v4_relat_1(X3,X1)|~v1_xboole_0(X1)|~r2_hidden(X2,k9_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1812, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1902, plain, (X1=k1_fcont_1(np__1,X2)|k1_seq_1(X1,esk3_3(np__1,X2,X1))!=k2_xcmplx_0(X2,esk3_3(np__1,X2,X1))|~v1_xcmplx_0(esk3_3(np__1,X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2050, plain, (k1_fcont_1(X1,X2)=k1_fcont_1(k3_xcmplx_0(X3,X4),X5)|k2_xcmplx_0(k3_xcmplx_0(X1,esk3_3(k3_xcmplx_0(X3,X4),X5,k1_fcont_1(X1,X2))),X2)!=k2_xcmplx_0(k3_xcmplx_0(X3,k3_xcmplx_0(X4,esk3_3(k3_xcmplx_0(X3,X4),X5,k1_fcont_1(X1,X2)))),X5)|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(X3,X4),X5,k1_fcont_1(X1,X2)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk3_3(k3_xcmplx_0(X3,X4),X5,k1_fcont_1(X1,X2)))|~v1_xreal_0(k3_xcmplx_0(X3,X4))|~v1_xreal_0(X5)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1989, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2602, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2975, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1806, plain, (esk15_2(X1,X2)=np__0)).#trainneg
cnf(i_0_2359, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2526, plain, (v1_xxreal_0(np__0))).#trainneg
cnf(i_0_3199, plain, (v3_xxreal_0(np__1)|~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1938, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1947, plain, (k1_fcont_1(X1,X2)=k1_fcont_1(np__1,X3)|k2_xcmplx_0(k3_xcmplx_0(X1,esk3_3(np__1,X3,k1_fcont_1(X1,X2))),X2)!=k2_xcmplx_0(esk3_3(np__1,X3,k1_fcont_1(X1,X2)),X3)|~v1_xcmplx_0(esk3_3(np__1,X3,k1_fcont_1(X1,X2)))|~m1_subset_1(k1_fcont_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(k1_fcont_1(X1,X2),k1_numbers,k1_numbers)|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(esk3_3(np__1,X3,k1_fcont_1(X1,X2)))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2468, plain, (k6_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1892, plain, (v1_relat_1(k9_xtuple_0(X1))|~v4_relat_1(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2967, plain, (v1_xreal_0(k2_xcmplx_0(X1,esk22_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1813, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2883, plain, (esk20_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2336, plain, (v2_xxreal_0(esk22_0)|v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_3010, plain, (v1_xreal_0(k1_funct_1(esk13_0,X1)))).#trainneg
cnf(i_0_3342, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(np__0)|~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3197, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3391, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3392, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2051, plain, (X1=k1_fcont_1(k3_xcmplx_0(X2,X3),X4)|k1_funct_1(X1,esk3_3(k3_xcmplx_0(X2,X3),X4,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,esk3_3(k3_xcmplx_0(X2,X3),X4,X1))),X4)|~v3_valued_0(X1)|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(X2,X3),X4,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,X3))|~v1_xreal_0(X4))).#trainneg
cnf(i_0_2085, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1937, plain, (v1_xboole_0(k1_relset_1(X2,X3))|m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k1_relset_1(X2,X3)))).#trainneg
cnf(i_0_2305, plain, (k3_xcmplx_0(k2_xcmplx_0(X2,X2),X1)=k3_xcmplx_0(k2_xcmplx_0(X1,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3696, plain, (v1_xboole_0(k9_xtuple_0(X1))|v1_xboole_0(X2)|m1_subset_1(X3,X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~m1_subset_1(X3,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_1894, plain, (v1_membered(k9_xtuple_0(X1))|~v1_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2968, plain, (v1_xreal_0(k2_xcmplx_0(X1,esk9_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1823, plain, (v1_funct_1(esk4_1(X1))|v2_funct_1(X1)|~v4_funct_1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1847, plain, (v1_xreal_0(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2887, plain, (v1_xreal_0(esk20_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2969, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3196, plain, (~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3250, plain, (~v3_xxreal_0(k6_xcmplx_0(X2,X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3514, plain, (~v3_xxreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1951, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2088, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3021, plain, (X1=k1_fcont_1(X2,np__0)|k1_seq_1(X1,esk3_3(X2,np__0,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)),np__0)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)))|~v1_xcmplx_0(np__0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2469, plain, (k6_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3118, plain, (k4_xcmplx_0(k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2570, plain, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2362, plain, (v1_xboole_0(k1_relset_1(X2,X1))|~v4_relat_1(X1,X2)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~m1_subset_1(X3,k1_relset_1(X2,X1)))).#trainneg
cnf(i_0_1895, plain, (v2_membered(k9_xtuple_0(X1))|~v2_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2970, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1824, plain, (v1_relat_1(esk4_1(X1))|v2_funct_1(X1)|~v4_funct_1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3012, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,esk23_0))).#trainneg
cnf(i_0_1848, plain, (v1_xreal_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1849, plain, (v1_xreal_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2525, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,esk24_0))).#trainneg
cnf(i_0_3878, plain, (~v2_membered(k1_numbers)|~v3_xxreal_0(k1_seq_1(X1,X2))|~v2_xxreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1885, plain, (v1_relat_1(k9_xtuple_0(X1))|~v4_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1))).#trainneg
cnf(i_0_2025, plain, (X1=k1_fcont_1(k4_xcmplx_0(X2),X3)|k1_seq_1(X1,esk3_3(k4_xcmplx_0(X2),X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(k4_xcmplx_0(np__1),esk3_3(k4_xcmplx_0(X2),X3,X1))),X3)|~v1_xcmplx_0(esk3_3(k4_xcmplx_0(X2),X3,X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1965, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2923, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1825, plain, (v1_xcmplx_0(esk4_1(X1))|v2_funct_1(X1)|~v1_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2203, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,X1),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1826, plain, (v1_xxreal_0(esk4_1(X1))|v2_funct_1(X1)|~v2_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2015, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1861, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_1865, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_1882, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v3_ordinal1(k1_zfmisc_1(X2))|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1919, plain, (v2_xxreal_0(X2)|~v2_xxreal_0(X1)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1920, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(X2)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2517, plain, (k4_xcmplx_0(k2_xcmplx_0(np__0,X1))=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2288, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,X1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1)|~v1_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1948, plain, (X1=k1_fcont_1(np__1,X2)|k1_funct_1(X1,esk3_3(np__1,X2,X1))!=k2_xcmplx_0(esk3_3(np__1,X2,X1),X2)|~v3_valued_0(X1)|~v1_xcmplx_0(esk3_3(np__1,X2,X1))|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2062, plain, (X1=k1_fcont_1(k3_xcmplx_0(X2,X3),X4)|k1_seq_1(X1,esk3_3(k3_xcmplx_0(X2,X3),X4,X1))!=k2_xcmplx_0(k3_xcmplx_0(X3,k3_xcmplx_0(X2,esk3_3(k3_xcmplx_0(X2,X3),X4,X1))),X4)|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(X2,X3),X4,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,X3))|~v1_xreal_0(X4))).#trainneg
cnf(i_0_2925, plain, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1901, plain, (v4_relat_1(k9_xtuple_0(X1),X2)|~v4_relat_1(X3,X2)|~v4_relat_1(X1,X3)|~v1_relat_1(X3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2974, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1827, plain, (v7_ordinal1(esk4_1(X1))|v2_funct_1(X1)|~v6_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2343, plain, (v2_xxreal_0(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1922, plain, (v3_xxreal_0(X2)|~v2_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1923, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2003, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(X2)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2004, plain, (v3_xxreal_0(X2)|~v3_xxreal_0(X1)|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2005, plain, (v2_xxreal_0(X1)|~v3_xxreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2006, plain, (v2_xxreal_0(X2)|~v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1886, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v4_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2822, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3071, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3112, plain, (k4_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1834, plain, (v1_xcmplx_0(esk5_1(X1))|v2_funct_1(X1)|~v1_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1988, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2519, plain, (v2_xxreal_0(X2)|~v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v2_xxreal_0(X1)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1828, plain, (v3_ordinal1(esk4_1(X1))|v2_funct_1(X1)|~v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1887, plain, (v4_relat_1(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1891, plain, (v5_relat_1(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,k2_zfmisc_1(X3,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_2230, plain, (v1_funct_1(k1_funct_1(X1,X2))|~v4_funct_1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1888, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v4_relat_1(X1,k2_zfmisc_1(X3,X2))|~v1_xboole_0(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3339, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1272, plain, (~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_45, plain, (m1_subset_1(np__0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_921, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4316, plain, (v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6367, plain, (v1_xreal_0(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,esk24_0))).#trainneg
cnf(i_0_1357, plain, (v2_membered(np__0)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_37, plain, (v3_valued_0(np__0))).#trainneg
cnf(i_0_41, plain, (v1_funct_1(np__0))).#trainneg
cnf(i_0_44, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_7182, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_39, plain, (v1_valued_0(np__0))).#trainneg
cnf(i_0_38, plain, (v2_valued_0(np__0))).#trainneg
cnf(i_0_4311, plain, (m1_subset_1(esk4_1(X1),X2)|v2_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6211, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4310, plain, (m1_subset_1(esk5_1(X1),X2)|v2_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4421, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k4_xcmplx_0(X3))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4235, plain, (k3_xcmplx_0(X2,X3)=X4|X1=np__0|k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1))!=k3_xcmplx_0(X4,X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4616, plain, (v2_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1005, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_1414, plain, (v1_membered(np__0)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1025, plain, (v3_membered(np__0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_42, plain, (v5_relat_1(np__0,k1_numbers))).#trainneg
cnf(i_0_43, plain, (v4_relat_1(np__0,k1_numbers))).#trainneg
cnf(i_0_219, plain, (m1_subset_1(np__0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6545, plain, (v1_xcmplx_0(esk4_1(X1))|v2_funct_1(X1)|~v3_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4406, plain, (X1=k1_fcont_1(X2,k4_xcmplx_0(np__1))|k1_seq_1(X1,esk3_3(X2,k4_xcmplx_0(np__1),X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(np__1),X1)),np__1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(np__1),X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4312, plain, (v1_xboole_0(k9_xtuple_0(X3))|m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k9_xtuple_0(X3)))).#trainneg
cnf(i_0_4390, plain, (k1_fcont_1(X1,X2)=k1_fcont_1(X3,X4)|k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk3_3(X3,X4,k1_fcont_1(X1,X2))))!=k2_xcmplx_0(k3_xcmplx_0(X3,esk3_3(X3,X4,k1_fcont_1(X1,X2))),X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,esk3_3(X3,X4,k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1071, plain, (k6_xcmplx_0(X1,np__0)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5322, plain, (v3_funct_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5783, plain, (v3_membered(k9_xtuple_0(esk24_0)))).#trainneg
cnf(i_0_5786, plain, (v3_membered(k9_xtuple_0(esk18_0)))).#trainneg
cnf(i_0_5787, plain, (v3_membered(k9_xtuple_0(esk13_0)))).#trainneg
cnf(i_0_5270, plain, (v1_relat_1(k9_xtuple_0(esk24_0))|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_5273, plain, (v1_relat_1(k9_xtuple_0(esk18_0))|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_4277, plain, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6038, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5246, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6269, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5508, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1182, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_215, plain, (v5_relat_1(np__0,X2))).#trainneg
cnf(i_0_216, plain, (v4_relat_1(np__0,X1))).#trainneg
cnf(i_0_6873, plain, (k6_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_5274, plain, (v1_relat_1(k9_xtuple_0(esk13_0))|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_5795, plain, (v1_membered(k9_xtuple_0(esk24_0))|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_5798, plain, (v1_membered(k9_xtuple_0(esk18_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_5799, plain, (v1_membered(k9_xtuple_0(esk13_0))|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_6343, plain, (v2_membered(k9_xtuple_0(esk24_0))|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_7905, plain, (v1_xcmplx_0(esk5_1(X1))|v2_funct_1(X1)|~v3_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4488, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_6274, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5617, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1179, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5277, plain, (v1_xreal_0(k2_xcmplx_0(esk22_0,esk22_0)))).#trainneg
cnf(i_0_6346, plain, (v2_membered(k9_xtuple_0(esk18_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_6347, plain, (v2_membered(k9_xtuple_0(esk13_0))|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_5278, plain, (v1_xreal_0(k2_xcmplx_0(esk9_0,esk22_0)))).#trainneg
cnf(i_0_5280, plain, (v1_xreal_0(k2_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_5282, plain, (v1_xreal_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_6253, plain, (k2_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5947, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1141, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),X2)=k3_xcmplx_0(k2_xcmplx_0(X1,np__1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5655, plain, (k3_xcmplx_0(X1,k2_xcmplx_0(X2,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X1),X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1143, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__0)=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5802, plain, (v1_xreal_0(k2_xcmplx_0(esk22_0,esk9_0)))).#trainneg
cnf(i_0_5803, plain, (v1_xreal_0(k2_xcmplx_0(esk9_0,esk9_0)))).#trainneg
cnf(i_0_5805, plain, (v1_xreal_0(k2_xcmplx_0(np__1,esk9_0)))).#trainneg
cnf(i_0_5807, plain, (v1_xreal_0(k2_xcmplx_0(np__0,esk9_0)))).#trainneg
cnf(i_0_5816, plain, (v1_xreal_0(k2_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_5817, plain, (v1_xreal_0(k2_xcmplx_0(esk9_0,np__1)))).#trainneg
cnf(i_0_5819, plain, (v1_xreal_0(k2_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_6350, plain, (v1_xreal_0(k2_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_4276, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4608, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4236, plain, (k3_xcmplx_0(X2,X3)=X4|X1=np__0|k3_xcmplx_0(X3,k3_xcmplx_0(X2,X1))!=k3_xcmplx_0(X4,X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_4948, plain, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6501, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5924, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X3))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_5980, plain, (k3_xcmplx_0(X2,X3)=X4|X1=np__0|k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))!=k3_xcmplx_0(X4,X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4376, plain, (k1_fcont_1(X1,X2)=k1_fcont_1(X3,X4)|k2_xcmplx_0(k3_xcmplx_0(X1,esk3_3(X3,X4,k1_fcont_1(X1,X2))),X2)!=k2_xcmplx_0(X4,k3_xcmplx_0(X3,esk3_3(X3,X4,k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk3_3(X3,X4,k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7633, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1169, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6351, plain, (v1_xreal_0(k2_xcmplx_0(esk9_0,np__0)))).#trainneg
cnf(i_0_6779, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_7969, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k3_xcmplx_0(np__1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_7974, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k3_xcmplx_0(np__0,k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_4262, plain, (v6_membered(np__0)|~v4_relat_1(X1,k5_numbers)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5353, plain, (~v3_xxreal_0(X1)|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5340, plain, (~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6090, plain, (k3_xcmplx_0(X1,np__1)=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5514, plain, (k3_xcmplx_0(X1,X2)=np__0|X3=X4|k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))!=k3_xcmplx_0(X1,k3_xcmplx_0(X4,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4556, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6990, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(np__1,esk3_3(X2,X3,X1))),X3)|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_4268, plain, (m1_subset_1(np__0,k1_zfmisc_1(X2))|~v4_relat_1(X1,X2)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1166, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),np__0)=k4_xcmplx_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4967, plain, (v1_xreal_0(k2_xcmplx_0(esk22_0,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4968, plain, (v1_xreal_0(k2_xcmplx_0(esk9_0,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4970, plain, (v1_xreal_0(k2_xcmplx_0(np__1,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5356, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7284, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5829, plain, (~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6208, plain, (~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k2_xcmplx_0(np__0,X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6803, plain, (v2_xxreal_0(np__0)|~v2_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5107, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6163, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9381, plain, (v2_funct_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8243, plain, (v1_xreal_0(X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,esk24_0))).#trainneg
cnf(i_0_1153, plain, (k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__0)=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7182, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_8736, plain, (v3_funct_1(np__0))).#trainneg
cnf(i_0_11733, plain, (v6_membered(np__0))).#trainneg
cnf(i_0_1037, plain, (v3_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_9374, plain, (v1_relat_1(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_8741, plain, (v2_funct_1(esk24_0)|v1_xreal_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_8742, plain, (v2_funct_1(esk24_0)|v1_xreal_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_8743, plain, (v1_xcmplx_0(esk4_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_8136, plain, (X1=np__0|X2=X3|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8092, plain, (X1=np__0|X2=np__0|k3_xcmplx_0(X2,X1)!=np__0|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9549, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8108, plain, (X1=np__0|np__1=X2|k3_xcmplx_0(X1,X2)!=X1|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8276, plain, (v2_funct_1(X1)|v1_xreal_0(esk4_1(X1))|~v4_relat_1(X1,k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_12311, plain, (m1_subset_1(np__0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8630, plain, (v1_xcmplx_0(esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8278, plain, (v1_funct_1(esk4_1(X1))|v2_funct_1(X1)|~v4_funct_1(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1139, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9371, plain, (v1_xboole_0(k9_xtuple_0(np__0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9377, plain, (v1_membered(k9_xtuple_0(np__0))|~v1_membered(X1))).#trainneg
cnf(i_0_8745, plain, (v2_funct_1(esk18_0)|v1_xreal_0(esk4_1(esk18_0)))).#trainneg
cnf(i_0_8746, plain, (v2_funct_1(esk18_0)|v1_xreal_0(esk5_1(esk18_0)))).#trainneg
cnf(i_0_8747, plain, (v1_xcmplx_0(esk4_1(esk18_0))|v2_funct_1(esk18_0))).#trainneg
cnf(i_0_8749, plain, (v2_funct_1(esk13_0)|v1_xreal_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_8750, plain, (v2_funct_1(esk13_0)|v1_xreal_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_12301, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(np__1,esk3_3(X2,X3,X1))),X3)|~v3_valued_0(X1)|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_8255, plain, (v2_funct_1(X1)|v1_xreal_0(esk4_1(X1))|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9523, plain, (v1_xcmplx_0(esk5_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8230, plain, (X1=np__0|X2=np__1|k3_xcmplx_0(X2,X1)!=X1|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8256, plain, (m1_subset_1(esk4_1(X1),k1_numbers)|v2_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8453, plain, (k3_xcmplx_0(X1,X2)=X3|X4=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X4,X2))!=k3_xcmplx_0(X3,X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8302, plain, (X1=k1_fcont_1(X2,np__0)|k1_seq_1(X1,esk3_3(X2,np__0,X1))!=k3_xcmplx_0(X2,esk3_3(X2,np__0,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8279, plain, (v1_relat_1(esk4_1(X1))|v2_funct_1(X1)|~v4_funct_1(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9378, plain, (v3_membered(k9_xtuple_0(np__0))|~v3_membered(X1))).#trainneg
cnf(i_0_1092, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9380, plain, (v2_membered(k9_xtuple_0(np__0))|~v2_membered(X1))).#trainneg
cnf(i_0_9372, plain, (v4_relat_1(k9_xtuple_0(np__0),X1))).#trainneg
cnf(i_0_8751, plain, (v1_xcmplx_0(esk4_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_9520, plain, (v1_xcmplx_0(esk5_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_9521, plain, (v1_xcmplx_0(esk5_1(esk18_0))|v2_funct_1(esk18_0))).#trainneg
cnf(i_0_9522, plain, (v1_xcmplx_0(esk5_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_12296, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,k3_xcmplx_0(np__1,esk3_3(X2,X3,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(np__1,esk3_3(X2,X3,X1))))|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_8361, plain, (v2_funct_1(X1)|v1_xreal_0(esk5_1(X1))|~v4_relat_1(X1,k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8303, plain, (X1=k1_fcont_1(X2,np__0)|k1_funct_1(X1,esk3_3(X2,np__0,X1))!=k3_xcmplx_0(X2,esk3_3(X2,np__0,X1))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8678, plain, (k1_fcont_1(X1,np__0)=k1_fcont_1(X2,X3)|k3_xcmplx_0(X1,esk3_3(X2,X3,k1_fcont_1(X1,np__0)))!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,k1_fcont_1(X1,np__0))),X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,esk3_3(X2,X3,k1_fcont_1(X1,np__0))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12162, plain, (X1=np__0|X2=np__0|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8340, plain, (v2_funct_1(X1)|v1_xreal_0(esk5_1(X1))|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8286, plain, (v1_xboole_0(esk4_1(X1))|v1_xboole_0(X2)|v2_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~m1_subset_1(X2,esk4_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_12063, plain, (X1=np__0|X2=np__0|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7384, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_8280, plain, (v1_xcmplx_0(esk4_1(X1))|v2_funct_1(X1)|~v1_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1149, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__1)=k2_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9375, plain, (v5_relat_1(k9_xtuple_0(np__0),X1))).#trainneg
cnf(i_0_11791, plain, (~v3_xxreal_0(k3_xcmplx_0(X1,np__1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11862, plain, (~v3_xxreal_0(k3_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11452, plain, (k1_fcont_1(X1,np__0)=k1_fcont_1(X2,X3)|k3_xcmplx_0(X1,esk3_3(X2,X3,k1_fcont_1(X1,np__0)))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,esk3_3(X2,X3,k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk3_3(X2,X3,k1_fcont_1(X1,np__0))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7447, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8341, plain, (m1_subset_1(esk5_1(X1),k1_numbers)|v2_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8281, plain, (v1_xxreal_0(esk4_1(X1))|v2_funct_1(X1)|~v2_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8242, plain, (v1_xreal_0(X1)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,esk24_0))).#trainneg
cnf(i_0_1829, plain, (v1_rat_1(esk4_1(X1))|v2_funct_1(X1)|~v4_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_12076, plain, (~v3_xxreal_0(k3_xcmplx_0(X1,np__0))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12141, plain, (~v3_xxreal_0(k3_xcmplx_0(np__0,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9432, plain, (v1_xxreal_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_9953, plain, (v1_xxreal_0(esk4_1(esk18_0))|v2_funct_1(esk18_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_9955, plain, (v1_xxreal_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_8722, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8729, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8119, plain, (k3_xcmplx_0(X2,X3)=X4|X1=np__0|k3_xcmplx_0(X2,k3_xcmplx_0(X3,X1))!=k3_xcmplx_0(X1,X4)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8282, plain, (v7_ordinal1(esk4_1(X1))|v2_funct_1(X1)|~v6_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8703, plain, (v1_xboole_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1830, plain, (v1_int_1(esk4_1(X1))|v2_funct_1(X1)|~v5_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4972, plain, (v1_xreal_0(k2_xcmplx_0(np__0,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5287, plain, (v1_xreal_0(k2_xcmplx_0(X1,esk22_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5812, plain, (v1_xreal_0(k2_xcmplx_0(X1,esk9_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8730, plain, (~v3_xxreal_0(k2_xcmplx_0(X1,np__0))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4315, plain, (v3_membered(k9_xtuple_0(X1))|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7720, plain, (k3_xcmplx_0(X1,X2)=np__0|X3=X4|k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))!=k3_xcmplx_0(X2,k3_xcmplx_0(X1,X4))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4422, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,np__0))=k4_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4551, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4295, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(np__0,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8283, plain, (v3_ordinal1(esk4_1(X1))|v2_funct_1(X1)|~v3_ordinal1(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4446, plain, (v2_xxreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1832, plain, (v1_funct_1(esk5_1(X1))|v2_funct_1(X1)|~v4_funct_1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5826, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6360, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5326, plain, (v1_funct_2(np__0,k5_numbers,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8726, plain, (~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,np__0),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12437, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xboole_0(k6_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_12497, plain, (~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xboole_0(k6_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4504, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8732, plain, (k2_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6336, plain, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8214, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8284, plain, (v1_rat_1(esk4_1(X1))|v2_funct_1(X1)|~v4_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4647, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_1833, plain, (v1_relat_1(esk5_1(X1))|v2_funct_1(X1)|~v4_funct_1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5286, plain, (v1_xreal_0(k2_xcmplx_0(X1,esk22_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5811, plain, (v1_xreal_0(k2_xcmplx_0(X1,esk9_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_9325, plain, (~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k6_xcmplx_0(X1,np__0),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_12475, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k6_xcmplx_0(np__0,X1),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_8020, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(k2_xcmplx_0(np__0,X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5347, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11738, plain, (~v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4951, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,k4_xcmplx_0(X3)),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_7542, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(X1,np__0))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4529, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k4_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9232, plain, (k3_xcmplx_0(X2,X3)=X4|X1=np__0|k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X1,X4)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_12060, plain, (X1=np__0|np__0=X2|k3_xcmplx_0(X1,np__0)!=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5657, plain, (k2_xcmplx_0(X2,X2)=X3|X1=np__0|k3_xcmplx_0(k2_xcmplx_0(X1,X1),X2)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8285, plain, (v1_int_1(esk4_1(X1))|v2_funct_1(X1)|~v5_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6203, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1835, plain, (v1_xxreal_0(esk5_1(X1))|v2_funct_1(X1)|~v2_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7148, plain, (v1_xreal_0(k2_xcmplx_0(esk22_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7149, plain, (v1_xreal_0(k2_xcmplx_0(esk9_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7151, plain, (v1_xreal_0(k2_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11743, plain, (~v3_xxreal_0(k2_xcmplx_0(np__0,X1))|~v2_xxreal_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(np__0,X1)))).#trainneg
cnf(i_0_4615, plain, (~v4_relat_1(X1,X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1)|~r2_hidden(X3,np__0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6368, plain, (v1_funct_1(X1)|~v4_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,esk24_0))).#trainneg
cnf(i_0_6369, plain, (v1_relat_1(X1)|~v4_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,esk24_0))).#trainneg
cnf(i_0_6370, plain, (v1_xcmplx_0(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,esk24_0))).#trainneg
cnf(i_0_4324, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__0))=np__0|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12630, plain, (X1=k1_fcont_1(X2,k4_xcmplx_0(X3))|k1_seq_1(X1,esk3_3(X2,k4_xcmplx_0(X3),X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)),k6_xcmplx_0(X3,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_12929, plain, (v2_funct_1(np__0))).#trainneg
cnf(i_0_12881, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2)=k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17141, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k2_xcmplx_0(esk3_3(X2,X3,X1),np__0)),X3)|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_17143, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k2_xcmplx_0(esk3_3(X2,X3,X1),np__0)),X3)|~v3_valued_0(X1)|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_15257, plain, (v1_xcmplx_0(esk5_1(X1))|v2_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13232, plain, (v3_membered(np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_13233, plain, (v1_membered(np__0)|~v1_membered(X1))).#trainneg
cnf(i_0_1836, plain, (v7_ordinal1(esk5_1(X1))|v2_funct_1(X1)|~v6_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13015, plain, (v3_membered(np__0))).#trainneg
cnf(i_0_13234, plain, (v2_membered(np__0)|~v2_membered(X1))).#trainneg
cnf(i_0_13240, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,np__0))).#trainneg
cnf(i_0_14955, plain, (~v1_xboole_0(k2_xcmplx_0(esk9_0,np__0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_13335, plain, (v1_xboole_0(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_13340, plain, (v1_membered(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_13576, plain, (v3_membered(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_15258, plain, (v1_xcmplx_0(esk4_1(X1))|v2_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13380, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,k3_xcmplx_0(np__1,esk3_3(X2,X3,X1))))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(np__1,esk3_3(X2,X3,X1))))|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_13781, plain, (X1=k1_fcont_1(X2,np__0)|k1_funct_1(X1,esk3_3(X2,np__0,X1))!=k3_xcmplx_0(X2,esk3_3(X2,np__0,X1))|~v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_13563, plain, (X1=k1_fcont_1(X2,np__0)|k1_seq_1(X1,esk3_3(X2,np__0,X1))!=k3_xcmplx_0(X2,esk3_3(X2,np__0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17566, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k6_xcmplx_0(k4_xcmplx_0(X3),np__0))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_12871, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))),np__0)=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17148, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),np__0)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8363, plain, (v1_funct_1(esk5_1(X1))|v2_funct_1(X1)|~v4_funct_1(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13160, plain, (v1_funct_1(X1)|~v4_funct_1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1837, plain, (v3_ordinal1(esk5_1(X1))|v2_funct_1(X1)|~v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13719, plain, (v2_membered(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_14149, plain, (v3_valued_0(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_13161, plain, (v1_relat_1(X1)|~v4_funct_1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14954, plain, (v1_xboole_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(esk22_0,np__0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_14612, plain, (v2_funct_1(esk24_0)|~v2_membered(k5_numbers)|~v3_xxreal_0(esk4_1(esk24_0))|~v2_xxreal_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_14614, plain, (v2_funct_1(esk18_0)|~v2_membered(k1_numbers)|~v3_xxreal_0(esk4_1(esk18_0))|~v2_xxreal_0(esk4_1(esk18_0)))).#trainneg
cnf(i_0_14616, plain, (v2_funct_1(esk13_0)|~v2_membered(k5_numbers)|~v3_xxreal_0(esk4_1(esk13_0))|~v2_xxreal_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_13730, plain, (v1_relat_1(k9_xtuple_0(k9_xtuple_0(np__0))))).#trainneg
cnf(i_0_17152, plain, (k3_xcmplx_0(X1,k2_xcmplx_0(X2,np__0))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13792, plain, (k1_fcont_1(X1,np__0)=k1_fcont_1(X2,X3)|k3_xcmplx_0(X1,esk3_3(X2,X3,k1_fcont_1(X1,np__0)))!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,k1_fcont_1(X1,np__0))),X3)|~v1_xreal_0(k3_xcmplx_0(X1,esk3_3(X2,X3,k1_fcont_1(X1,np__0))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8364, plain, (v1_relat_1(esk5_1(X1))|v2_funct_1(X1)|~v4_funct_1(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13165, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1838, plain, (v1_rat_1(esk5_1(X1))|v2_funct_1(X1)|~v4_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13166, plain, (v1_rat_1(X1)|~v4_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13167, plain, (v1_int_1(X1)|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13239, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_14944, plain, (v7_ordinal1(esk4_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_14948, plain, (v7_ordinal1(esk4_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_15725, plain, (v3_ordinal1(esk4_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_15729, plain, (v3_ordinal1(esk4_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_13245, plain, (v1_funct_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_12864, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,k4_xcmplx_0(X2)))),k4_xcmplx_0(X2))=k6_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,k4_xcmplx_0(X2)))),X2)|v2_funct_1(k1_fcont_1(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,k4_xcmplx_0(X2)))))|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,k4_xcmplx_0(X2))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,k4_xcmplx_0(X2))))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16547, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_12912, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17561, plain, (X1=k1_fcont_1(X2,k4_xcmplx_0(X3))|k1_seq_1(X1,esk3_3(X2,k4_xcmplx_0(X3),X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)),k2_xcmplx_0(X3,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8365, plain, (v1_xcmplx_0(esk5_1(X1))|v2_funct_1(X1)|~v1_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13737, plain, (v2_funct_1(k9_xtuple_0(np__0))|~v1_xboole_0(X1)|~v1_funct_1(k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_1839, plain, (v1_int_1(esk5_1(X1))|v2_funct_1(X1)|~v5_membered(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_13727, plain, (v1_xboole_0(k9_xtuple_0(k9_xtuple_0(np__0)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_16304, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(np__0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13248, plain, (v1_funct_1(esk4_1(esk18_0))|v2_funct_1(esk18_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_13249, plain, (v1_funct_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_13567, plain, (v1_relat_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_13570, plain, (v1_relat_1(esk4_1(esk18_0))|v2_funct_1(esk18_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_13571, plain, (v1_relat_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_13858, plain, (v1_xboole_0(esk4_1(X1))|v2_funct_1(X1)|v2_funct_1(X2)|~v4_relat_1(X1,esk4_1(X2))|~v4_relat_1(X2,esk4_1(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_13207, plain, (k3_xcmplx_0(X1,X2)=np__1|X3=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))!=X3|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16990, plain, (k2_xcmplx_0(X1,X1)=X2|X3=np__0|k3_xcmplx_0(k2_xcmplx_0(X3,X3),X1)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X1)))).#trainneg
cnf(i_0_12914, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,k3_xcmplx_0(X1,X2)))),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(esk5_1(k1_fcont_1(X1,k3_xcmplx_0(X1,X2))),X2),X1)|v2_funct_1(k1_fcont_1(X1,k3_xcmplx_0(X1,X2)))|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,k3_xcmplx_0(X1,X2))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,k3_xcmplx_0(X1,X2))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,k3_xcmplx_0(X1,X2))))|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16102, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(np__0,X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_14249, plain, (k4_xcmplx_0(np__1)=X1|X2=np__0|k3_xcmplx_0(X2,X1)!=k4_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8366, plain, (v1_xxreal_0(esk5_1(X1))|v2_funct_1(X1)|~v2_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_15444, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,np__0)))).#trainneg
cnf(i_0_1941, plain, (X1=k1_fcont_1(X2,k4_xcmplx_0(X3))|k1_funct_1(X1,esk3_3(X2,k4_xcmplx_0(X3),X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)),X3)|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k4_xcmplx_0(X3),X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_15728, plain, (v3_ordinal1(esk4_1(esk18_0))|v2_funct_1(esk18_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_16294, plain, (v1_rat_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_16866, plain, (k3_xcmplx_0(X1,X2)=X3|X4=np__0|k3_xcmplx_0(k3_xcmplx_0(X4,X2),X1)!=k3_xcmplx_0(X4,X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X4,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13101, plain, (k2_xcmplx_0(X2,X2)=X3|X1=np__0|k3_xcmplx_0(k2_xcmplx_0(X1,X1),X2)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(k2_xcmplx_0(X2,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_13501, plain, (k3_xcmplx_0(X1,np__1)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13058, plain, (X2=k3_xcmplx_0(X3,np__1)|X1=np__0|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(k3_xcmplx_0(X3,np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_15719, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,np__0))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16641, plain, (k2_xcmplx_0(X1,np__1)=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8367, plain, (v7_ordinal1(esk5_1(X1))|v2_funct_1(X1)|~v6_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_15995, plain, (v1_xboole_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1944, plain, (X1=k1_fcont_1(np__1,k4_xcmplx_0(X2))|k1_seq_1(X1,esk3_3(np__1,k4_xcmplx_0(X2),X1))!=k6_xcmplx_0(esk3_3(np__1,k4_xcmplx_0(X2),X1),X2)|~v1_xcmplx_0(esk3_3(np__1,k4_xcmplx_0(X2),X1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_15915, plain, (~v3_xxreal_0(k2_xcmplx_0(X1,np__0))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16002, plain, (~v3_xxreal_0(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16297, plain, (v1_rat_1(esk4_1(esk18_0))|v2_funct_1(esk18_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_16298, plain, (v1_rat_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_17259, plain, (v1_int_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_17262, plain, (v1_int_1(esk4_1(esk18_0))|v2_funct_1(esk18_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_17263, plain, (v1_int_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_13094, plain, (k3_xcmplx_0(X1,X2)=np__0|X3=X4|k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))!=k3_xcmplx_0(k3_xcmplx_0(X1,X2),X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16035, plain, (k2_xcmplx_0(X2,k6_xcmplx_0(X1,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_13138, plain, (k3_xcmplx_0(X1,X2)=np__0|X3=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))!=np__0|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13109, plain, (k3_xcmplx_0(X2,X3)=X4|X1=np__0|k3_xcmplx_0(X3,k3_xcmplx_0(X2,X1))!=k3_xcmplx_0(X1,X4)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_13066, plain, (X2=k3_xcmplx_0(X3,X4)|X1=np__0|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X4,k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_13538, plain, (k3_xcmplx_0(X1,X2)=X3|X4=np__0|k3_xcmplx_0(k3_xcmplx_0(X4,X2),X1)!=k3_xcmplx_0(X3,X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X4,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12910, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X3))),X3))=k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X3))),k2_xcmplx_0(X1,X3))|v2_funct_1(k1_fcont_1(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X3))),X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X3))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(esk4_1(k1_fcont_1(X2,X3)))|~v1_xreal_0(esk5_1(k1_fcont_1(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8368, plain, (v3_ordinal1(esk5_1(X1))|v2_funct_1(X1)|~v3_ordinal1(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2019, plain, (X1=k1_fcont_1(np__0,X2)|k1_seq_1(X1,esk3_3(np__0,X2,X1))!=k2_xcmplx_0(k3_xcmplx_0(np__1,k3_xcmplx_0(np__0,esk3_3(np__0,X2,X1))),X2)|~v1_xcmplx_0(esk3_3(np__0,X2,X1))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17301, plain, (v1_xxreal_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_17302, plain, (v1_xxreal_0(esk5_1(esk18_0))|v2_funct_1(esk18_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_17303, plain, (v1_xxreal_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_14613, plain, (v3_xxreal_0(esk4_1(esk24_0))|v2_xxreal_0(esk4_1(esk24_0))|v1_xboole_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_14615, plain, (v3_xxreal_0(esk4_1(esk18_0))|v2_xxreal_0(esk4_1(esk18_0))|v1_xboole_0(esk4_1(esk18_0))|v2_funct_1(esk18_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_12911, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))=k2_xcmplx_0(X2,k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X1))),X1))|v2_funct_1(k1_fcont_1(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X3,X1)))|~v1_xreal_0(esk5_1(k1_fcont_1(X3,X1)))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_14926, plain, (k3_xcmplx_0(X1,X2)=X3|X4=np__0|k3_xcmplx_0(k3_xcmplx_0(X2,X4),X1)!=k3_xcmplx_0(X4,X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X2,X4))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13135, plain, (k3_xcmplx_0(X1,X2)=np__0|X3=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))!=np__0|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14487, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16545, plain, (k2_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1)),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8369, plain, (v1_rat_1(esk5_1(X1))|v2_funct_1(X1)|~v4_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8700, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2020, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(np__1,k3_xcmplx_0(X2,esk3_3(X2,X3,X1))),X3)|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_14617, plain, (v3_xxreal_0(esk4_1(esk13_0))|v2_xxreal_0(esk4_1(esk13_0))|v1_xboole_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_15317, plain, (k3_xcmplx_0(X1,X2)=np__0|X3=X4|k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))!=k3_xcmplx_0(X2,k3_xcmplx_0(X4,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X4,X1))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12890, plain, (k2_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16594, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,np__0),X1)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,np__0))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_15868, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13208, plain, (k3_xcmplx_0(X1,X2)=np__1|X3=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))!=X3|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15524, plain, (k4_xcmplx_0(k2_xcmplx_0(np__0,X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8370, plain, (v1_int_1(esk5_1(X1))|v2_funct_1(X1)|~v5_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8720, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2021, plain, (X1=k1_fcont_1(np__0,X2)|k1_seq_1(X1,esk3_3(np__0,X2,X1))!=k2_xcmplx_0(k3_xcmplx_0(np__0,k3_xcmplx_0(np__1,esk3_3(np__0,X2,X1))),X2)|~v1_xcmplx_0(esk3_3(np__0,X2,X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17793, plain, (v1_membered(np__0))).#trainneg
cnf(i_0_17805, plain, (v2_membered(np__0))).#trainneg
cnf(i_0_17814, plain, (~v1_xboole_0(k2_xcmplx_0(np__0,esk9_0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_17818, plain, (k9_xtuple_0(np__0)=np__0)).#trainneg
cnf(i_0_17881, plain, (X1=k1_fcont_1(X2,np__0)|k1_funct_1(X1,esk3_3(X2,np__0,X1))!=k3_xcmplx_0(X2,esk3_3(X2,np__0,X1))|~v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk3_3(X2,np__0,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17895, plain, (X1=k1_fcont_1(X2,np__0)|k1_seq_1(X1,esk3_3(X2,np__0,X1))!=k3_xcmplx_0(X2,esk3_3(X2,np__0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk3_3(X2,np__0,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17916, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k2_xcmplx_0(k4_xcmplx_0(X3),np__0))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17972, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17938, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17937, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k6_xcmplx_0(k4_xcmplx_0(X3),np__0))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17650, plain, (k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))=k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8277, plain, (v2_funct_1(X1)|v1_xreal_0(esk4_1(X1))|~v3_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_18302, plain, (v2_funct_1(esk24_0)|~v3_xxreal_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_2030, plain, (X1=k1_fcont_1(k3_xcmplx_0(X2,np__1),X3)|k1_seq_1(X1,esk3_3(k3_xcmplx_0(X2,np__1),X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(k3_xcmplx_0(X2,np__1),X3,X1)),X3)|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(X2,np__1),X3,X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,np__1))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_18313, plain, (v2_funct_1(esk13_0)|~v3_xxreal_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_19576, plain, (v1_xreal_0(k4_xcmplx_0(esk22_0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_18140, plain, (v1_xboole_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk22_0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_18864, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18371, plain, (k6_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k4_xcmplx_0(X2))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22662, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19623, plain, (X1=k1_fcont_1(X2,np__0)|k1_funct_1(X1,esk3_3(X2,np__0,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)),np__0)|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_18881, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_17773, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,k2_xcmplx_0(esk3_3(X2,X3,X1),np__0)))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,k2_xcmplx_0(esk3_3(X2,X3,X1),np__0)))|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_18843, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8362, plain, (v2_funct_1(X1)|v1_xreal_0(esk5_1(X1))|~v3_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_21424, plain, (v2_funct_1(esk24_0)|~v2_membered(k5_numbers)|~v3_xxreal_0(esk5_1(esk24_0))|~v2_xxreal_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_2032, plain, (X1=k1_fcont_1(k3_xcmplx_0(np__1,X2),X3)|k1_seq_1(X1,esk3_3(k3_xcmplx_0(np__1,X2),X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(k3_xcmplx_0(np__1,X2),X3,X1)),X3)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(k3_xcmplx_0(np__1,X2),X3,X1)))|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(np__1,X2),X3,X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(np__1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_19579, plain, (v1_xreal_0(k4_xcmplx_0(esk9_0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_21426, plain, (v2_funct_1(esk18_0)|~v2_membered(k1_numbers)|~v3_xxreal_0(esk5_1(esk18_0))|~v2_xxreal_0(esk5_1(esk18_0)))).#trainneg
cnf(i_0_21428, plain, (v2_funct_1(esk13_0)|~v2_membered(k5_numbers)|~v3_xxreal_0(esk5_1(esk13_0))|~v2_xxreal_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_19023, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(np__0))))).#trainneg
cnf(i_0_20162, plain, (v7_ordinal1(esk5_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_20166, plain, (v7_ordinal1(esk5_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_19853, plain, (k3_xcmplx_0(X1,np__1)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_17981, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0))))=k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,np__0))),X1)|v2_funct_1(k1_fcont_1(X2,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0))))|~v1_xcmplx_0(X1)|~v1_xreal_0(esk4_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17739, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k2_xcmplx_0(np__0,esk3_3(X2,X3,X1))),X3)|~v3_valued_0(X1)|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_21290, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19852, plain, (k2_xcmplx_0(X1,np__0)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17683, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19594, plain, (X1=k1_fcont_1(X2,np__0)|k1_funct_1(X1,esk3_3(X2,np__0,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)),np__0)|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,np__0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_18836, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18299, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22861, plain, (v1_funct_1(k4_xcmplx_0(X1))|~v4_funct_1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2035, plain, (X1=k1_fcont_1(k3_xcmplx_0(X2,X3),k4_xcmplx_0(X4))|k1_seq_1(X1,esk3_3(k3_xcmplx_0(X2,X3),k4_xcmplx_0(X4),X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,esk3_3(k3_xcmplx_0(X2,X3),k4_xcmplx_0(X4),X1))),X4)|~v1_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,esk3_3(k3_xcmplx_0(X2,X3),k4_xcmplx_0(X4),X1))))|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(X2,X3),k4_xcmplx_0(X4),X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4)))).#trainneg
cnf(i_0_20020, plain, (~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20139, plain, (~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22862, plain, (v1_relat_1(k4_xcmplx_0(X1))|~v4_funct_1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21378, plain, (v3_ordinal1(esk5_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_21382, plain, (v3_ordinal1(esk5_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_18125, plain, (v1_funct_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_18128, plain, (v1_funct_1(esk5_1(esk18_0))|v2_funct_1(esk18_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_18129, plain, (v1_funct_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_13053, plain, (X1=np__0|X2=np__0|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_18235, plain, (X2=k2_xcmplx_0(X3,np__0)|X1=np__0|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(k2_xcmplx_0(X3,np__0))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_21669, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|v2_funct_1(k1_fcont_1(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_12898, plain, (k6_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13096, plain, (X1=np__0|np__1=X2|k3_xcmplx_0(np__1,X1)!=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7147, plain, (v4_relat_1(k9_xtuple_0(X1),k1_numbers)|~v4_relat_1(X1,k1_fcont_1(X2,X3))|~v1_relat_1(k1_fcont_1(X2,X3))|~v1_relat_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_19572, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(np__0,X1)))).#trainneg
cnf(i_0_20027, plain, (~v3_xxreal_0(k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_18289, plain, (v1_relat_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_18292, plain, (v1_relat_1(esk5_1(esk18_0))|v2_funct_1(esk18_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_18293, plain, (v1_relat_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_13857, plain, (v1_xboole_0(esk4_1(X1))|v2_funct_1(X1)|v2_funct_1(X2)|~v4_relat_1(X1,esk5_1(X2))|~v4_relat_1(X2,esk4_1(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_14792, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17083, plain, (X1=np__0|X2=np__1|k3_xcmplx_0(X2,X1)!=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_18880, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_17078, plain, (X1=np__0|np__1=X2|k3_xcmplx_0(X1,X2)!=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15934, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,X1),np__0)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20128, plain, (~v3_xxreal_0(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21381, plain, (v3_ordinal1(esk5_1(esk18_0))|v2_funct_1(esk18_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_22217, plain, (v1_rat_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_22220, plain, (v1_rat_1(esk5_1(esk18_0))|v2_funct_1(esk18_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_22221, plain, (v1_rat_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_17009, plain, (k1_fcont_1(np__1,X1)=k1_fcont_1(X2,X3)|k2_xcmplx_0(k2_xcmplx_0(esk3_3(X2,X3,k1_fcont_1(np__1,X1)),np__0),X1)!=k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,k1_fcont_1(np__1,X1))),X3)|~v1_xcmplx_0(esk3_3(X2,X3,k1_fcont_1(np__1,X1)))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12860, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,k6_xcmplx_0(X2,X3)))),k6_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,k6_xcmplx_0(X2,X3)))),X3))|v2_funct_1(k1_fcont_1(X1,k6_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,k6_xcmplx_0(X2,X3)))))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,k6_xcmplx_0(X2,X3))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,k6_xcmplx_0(X2,X3))))|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13112, plain, (k3_xcmplx_0(X2,X3)=X4|X1=np__0|k3_xcmplx_0(X3,k3_xcmplx_0(X1,X2))!=k3_xcmplx_0(X1,X4)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_18269, plain, (k3_xcmplx_0(X1,k2_xcmplx_0(np__0,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12878, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,k2_xcmplx_0(X2,X3)))),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,k2_xcmplx_0(X2,X3)))),X3))|v2_funct_1(k1_fcont_1(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,k2_xcmplx_0(X2,X3)))))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,k2_xcmplx_0(X2,X3))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,k2_xcmplx_0(X2,X3))))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14227, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(np__1)|X3=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))!=k4_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13832, plain, (k3_xcmplx_0(X1,X2)=np__0|X3=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))!=k3_xcmplx_0(np__0,X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4253, plain, (v1_xxreal_0(k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2))|~v2_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19586, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22864, plain, (v1_xxreal_0(k4_xcmplx_0(X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22866, plain, (v3_ordinal1(k4_xcmplx_0(X1))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23101, plain, (v1_int_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_23104, plain, (v1_int_1(esk5_1(esk18_0))|v2_funct_1(esk18_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_23105, plain, (v1_int_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_18311, plain, (v1_xboole_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(esk4_1(esk24_0),np__0))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_18322, plain, (v1_xboole_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(esk4_1(esk13_0),np__0))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_19822, plain, (k2_xcmplx_0(X1,X1)=X2|X3=np__0|k3_xcmplx_0(k2_xcmplx_0(X3,X3),X1)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X1)))).#trainneg
cnf(i_0_17733, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,k2_xcmplx_0(esk3_3(X2,X3,X1),np__0)))|~v1_xcmplx_0(k3_xcmplx_0(X2,k2_xcmplx_0(esk3_3(X2,X3,X1),np__0)))|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_16548, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_18236, plain, (X2=k2_xcmplx_0(X3,np__0)|X1=np__0|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(k2_xcmplx_0(X3,np__0))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_17699, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k2_xcmplx_0(np__0,esk3_3(X2,X3,X1))),X3)|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_4254, plain, (v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2))|~v4_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2066, plain, (k3_xcmplx_0(X2,k3_xcmplx_0(X1,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_22867, plain, (v1_rat_1(k4_xcmplx_0(X1))|~v4_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22868, plain, (v1_int_1(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22875, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_20174, plain, (v1_xboole_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_20175, plain, (v1_xboole_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_18305, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk24_0),esk22_0)))).#trainneg
cnf(i_0_18306, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk24_0),esk9_0)))).#trainneg
cnf(i_0_18307, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk24_0),np__1)))).#trainneg
cnf(i_0_14230, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(np__1)|X3=np__0|k3_xcmplx_0(X1,k3_xcmplx_0(X3,X2))!=k4_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16115, plain, (k6_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2),X3)=k2_xcmplx_0(X2,k6_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13060, plain, (X2=k3_xcmplx_0(X3,np__1)|X1=np__0|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(k3_xcmplx_0(X3,np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_17091, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,np__0),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15939, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,np__0))=k6_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17680, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X3)))))=k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X3))),k2_xcmplx_0(X3,X1))|v2_funct_1(k1_fcont_1(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X3,k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X3)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X3))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X3))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2530, plain, (v1_xreal_0(k2_xcmplx_0(k3_xcmplx_0(X1,X3),X2))|~v3_valued_0(k1_fcont_1(X1,X2))|~v1_relat_1(k1_fcont_1(X1,X2))|~v1_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7153, plain, (v1_xreal_0(k2_xcmplx_0(np__0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_25986, plain, (~v1_xboole_0(k2_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_25979, plain, (~v1_xboole_0(k2_xcmplx_0(esk22_0,np__1))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_25980, plain, (~v1_xboole_0(k2_xcmplx_0(esk9_0,np__1))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_25994, plain, (~v1_xboole_0(k2_xcmplx_0(np__1,esk22_0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_25995, plain, (~v1_xboole_0(k2_xcmplx_0(np__1,esk9_0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_24578, plain, (v2_funct_1(esk24_0)|~v3_xxreal_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_24632, plain, (k3_xcmplx_0(X1,np__1)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23289, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23737, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23258, plain, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k6_xcmplx_0(np__0,X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_23735, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1)))))=k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))),X1)|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_25400, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_25718, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_29191, plain, (m1_subset_1(X1,k1_numbers)|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24590, plain, (v2_funct_1(esk13_0)|~v3_xxreal_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_2079, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3)),X4)=k3_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_24574, plain, (esk12_1(k1_zfmisc_1(np__0))=np__0)).#trainneg
cnf(i_0_29186, plain, (m1_subset_1(esk22_0,k1_numbers)|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_29187, plain, (m1_subset_1(esk9_0,k1_numbers)|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_25981, plain, (v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(esk5_1(esk24_0),np__1))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_25982, plain, (v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(esk5_1(esk13_0),np__1))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_25983, plain, (v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(esk4_1(esk24_0),np__1))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_25984, plain, (v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(esk4_1(esk13_0),np__1))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_25996, plain, (v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(np__1,esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_26487, plain, (k6_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23694, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30498, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,X1))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_25717, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_26547, plain, (np__1=X1|X2=np__0|k3_xcmplx_0(X2,np__1)!=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_24393, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_24631, plain, (k2_xcmplx_0(X1,np__0)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,np__0)))).#trainneg
cnf(i_0_30522, plain, (v1_xreal_0(k2_xcmplx_0(np__0,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_28297, plain, (v1_xreal_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2081, plain, (k3_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X1,X3)),X4)=k3_xcmplx_0(X3,k3_xcmplx_0(k3_xcmplx_0(X1,X2),X4))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_30502, plain, (v1_xreal_0(k2_xcmplx_0(np__1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30504, plain, (v1_xreal_0(k2_xcmplx_0(np__0,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25997, plain, (v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(np__1,esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_25998, plain, (v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(np__1,esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_25999, plain, (v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(np__1,esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_29169, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_24587, plain, (v1_xboole_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(esk5_1(esk24_0),np__0))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_24112, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_23686, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23412, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26376, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1)))))=k2_xcmplx_0(X1,k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_24360, plain, (k2_xcmplx_0(k4_xcmplx_0(X3),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25347, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24690, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))),k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0)))))=k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,np__0))),k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X1))),X1))|v2_funct_1(k1_fcont_1(X2,np__0))|v2_funct_1(k1_fcont_1(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))),k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xreal_0(esk4_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_27568, plain, (k3_xcmplx_0(X1,X2)=X3|X4=np__0|k3_xcmplx_0(X2,k3_xcmplx_0(X4,X1))!=k3_xcmplx_0(X3,X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2091, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X4)=k3_xcmplx_0(X3,k3_xcmplx_0(k3_xcmplx_0(X1,X2),X4))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24588, plain, (v1_xboole_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_24599, plain, (v1_xboole_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(esk5_1(esk13_0),np__0))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_24600, plain, (v1_xboole_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_23276, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k2_xcmplx_0(k4_xcmplx_0(X3),np__0))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_23914, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))),X3)=k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))),k2_xcmplx_0(X1,X3))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_26316, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X1))),X1))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|v2_funct_1(k1_fcont_1(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_26370, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))),k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0)))))=k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,np__0))),k2_xcmplx_0(X1,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|v2_funct_1(k1_fcont_1(X2,np__0))|v2_funct_1(k1_fcont_1(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))),k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xreal_0(esk4_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_30642, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_30158, plain, (k6_xcmplx_0(np__0,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))=k4_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24627, plain, (k3_xcmplx_0(np__1,X1)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(np__1,X1)))).#trainneg
cnf(i_0_23718, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))=k4_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25636, plain, (k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30356, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X1))),k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|v2_funct_1(k1_fcont_1(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2148, plain, (v3_xxreal_0(k6_xcmplx_0(X2,X1))|v3_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27222, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,np__1),np__0)=np__0)).#trainneg
cnf(i_0_18308, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk24_0),np__0)))).#trainneg
cnf(i_0_26344, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X1)))))|v2_funct_1(k1_fcont_1(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_24316, plain, (k2_xcmplx_0(X3,k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23450, plain, (k6_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))),X3)=k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))),k6_xcmplx_0(X1,X3))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_28485, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X1))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2149, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v3_xxreal_0(X1)|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8785, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_18316, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk13_0),esk22_0)))).#trainneg
cnf(i_0_18317, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk13_0),esk9_0)))).#trainneg
cnf(i_0_18318, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk13_0),np__1)))).#trainneg
cnf(i_0_18319, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk4_1(esk13_0),np__0)))).#trainneg
cnf(i_0_24581, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk24_0),esk22_0)))).#trainneg
cnf(i_0_27651, plain, (k3_xcmplx_0(X1,X2)=X3|X4=np__0|k3_xcmplx_0(k3_xcmplx_0(X4,X1),X2)!=k3_xcmplx_0(X4,X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X4,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_30005, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__0)=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24914, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,k2_xcmplx_0(np__0,esk3_3(X2,X3,X1))))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,k2_xcmplx_0(np__0,esk3_3(X2,X3,X1))))|~v1_xcmplx_0(esk3_3(X2,X3,X1))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_26101, plain, (X1=k2_xcmplx_0(X2,np__0)|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,np__0)))).#trainneg
cnf(i_0_28875, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_30544, plain, (v1_xreal_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8765, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2150, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24582, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk24_0),esk9_0)))).#trainneg
cnf(i_0_24583, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk24_0),np__1)))).#trainneg
cnf(i_0_24584, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk24_0),np__0)))).#trainneg
cnf(i_0_24593, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk13_0),esk22_0)))).#trainneg
cnf(i_0_24594, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk13_0),esk9_0)))).#trainneg
cnf(i_0_25261, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30752, plain, (k3_xcmplx_0(X1,np__1)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36427, plain, (k3_xcmplx_0(np__1,X1)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30829, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36368, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))=k4_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30709, plain, (k3_xcmplx_0(X1,np__1)=np__0|X2=np__0|k3_xcmplx_0(X1,X2)!=np__0|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_34118, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36471, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))=k4_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))),X1))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_30696, plain, (k3_xcmplx_0(X1,np__1)=np__1|X2=np__0|k3_xcmplx_0(X1,X2)!=X2|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39225, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|v1_xboole_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36311, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_2151, plain, (v3_xxreal_0(k3_xcmplx_0(X2,X1))|v3_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k3_xcmplx_0(X2,X1))|~v1_xreal_0(k3_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_31887, plain, (~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_31888, plain, (m1_subset_1(esk22_0,k1_numbers))).#trainneg
cnf(i_0_31889, plain, (m1_subset_1(esk9_0,k1_numbers))).#trainneg
cnf(i_0_35635, plain, (X1=k1_fcont_1(X2,X3)|k1_funct_1(X1,esk3_3(X2,X3,X1))!=k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)),k6_xcmplx_0(np__0,X3))|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_36480, plain, (k6_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1)))),X1)=k4_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))),X1))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_39540, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_31436, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_31243, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_35956, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_35963, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_34361, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k6_xcmplx_0(X2,X1))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_36982, plain, (v1_xcmplx_0(k2_xcmplx_0(np__0,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_39196, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2152, plain, (v3_xxreal_0(k3_xcmplx_0(X1,X2))|v3_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_39185, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_33731, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_35964, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_34699, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X3,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_35957, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_32142, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30815, plain, (k3_xcmplx_0(X1,np__1)=X2|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32728, plain, (X1=np__0|np__1=X2|k3_xcmplx_0(X1,np__1)!=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_39096, plain, (v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30854, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))),np__0)=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2163, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,X3),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_39184, plain, (~v2_xxreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36962, plain, (v1_xcmplx_0(k2_xcmplx_0(np__1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36964, plain, (v1_xcmplx_0(k2_xcmplx_0(np__0,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39169, plain, (v1_xreal_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32226, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39425, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36389, plain, (np__1=X1|X2=np__0|k3_xcmplx_0(np__1,X2)!=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_37860, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))),k6_xcmplx_0(X1,k4_xcmplx_0(X3)))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_38236, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31451, plain, (m1_subset_1(esk5_1(esk24_0),k1_numbers)|v2_funct_1(esk24_0)|~v1_xboole_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_31452, plain, (m1_subset_1(esk5_1(esk18_0),k1_numbers)|v2_funct_1(esk18_0)|~v1_xboole_0(esk5_1(esk18_0)))).#trainneg
cnf(i_0_31453, plain, (m1_subset_1(esk5_1(esk13_0),k1_numbers)|v2_funct_1(esk13_0)|~v1_xboole_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_31456, plain, (m1_subset_1(esk4_1(esk24_0),k1_numbers)|v2_funct_1(esk24_0)|~v1_xboole_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_31457, plain, (m1_subset_1(esk4_1(esk18_0),k1_numbers)|v2_funct_1(esk18_0)|~v1_xboole_0(esk4_1(esk18_0)))).#trainneg
cnf(i_0_38230, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k6_xcmplx_0(k2_xcmplx_0(X3,X1),X2)|~v1_xcmplx_0(k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_36935, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(X2,k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_33058, plain, (k2_xcmplx_0(X1,np__0)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,np__0)))).#trainneg
cnf(i_0_34320, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k4_xcmplx_0(X3),X2))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X3),X2))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39590, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32040, plain, (k6_xcmplx_0(np__0,k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1)))))=k4_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))),X1))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8766, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10470, plain, (k3_xcmplx_0(np__1,X1)=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8075, plain, (np__0=np__1|X1=np__1|k3_xcmplx_0(X1,np__1)!=np__1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_31458, plain, (m1_subset_1(esk4_1(esk13_0),k1_numbers)|v2_funct_1(esk13_0)|~v1_xboole_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_34321, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(k4_xcmplx_0(X3),X2))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X3),X2))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36353, plain, (k4_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))))=k6_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))),np__0)|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38150, plain, (k6_xcmplx_0(k2_xcmplx_0(X3,X1),X2)=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_43171, plain, (k3_xcmplx_0(X1,np__1)=X2|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40277, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40430, plain, (k2_xcmplx_0(X1,np__0)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39872, plain, (k3_xcmplx_0(np__1,X1)=X2|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40583, plain, (v2_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_40695, plain, (v1_funct_1(esk22_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_2179, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1)=k2_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_40696, plain, (v1_relat_1(esk22_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_40697, plain, (v1_xcmplx_0(esk22_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_40698, plain, (v1_xxreal_0(esk22_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_40700, plain, (v3_ordinal1(esk22_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_40701, plain, (v1_rat_1(esk22_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_40702, plain, (v1_int_1(esk22_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_40706, plain, (v1_funct_1(esk9_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_40707, plain, (v1_relat_1(esk9_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_39806, plain, (k3_xcmplx_0(np__1,X1)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39913, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__1|X1=np__0|v2_funct_1(k1_fcont_1(X1,np__0))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))!=X1|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(X1)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40432, plain, (X1=k2_xcmplx_0(X2,np__0)|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_43111, plain, (k3_xcmplx_0(X1,np__1)=np__1|X2=np__0|k3_xcmplx_0(X2,X1)!=X2|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40066, plain, (k4_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))))=k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45335, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_44571, plain, (X1=k4_xcmplx_0(X2)|v3_xxreal_0(X2)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_40708, plain, (v1_xcmplx_0(esk9_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_40709, plain, (v1_xxreal_0(esk9_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_40711, plain, (v3_ordinal1(esk9_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_40712, plain, (v1_rat_1(esk9_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_40713, plain, (v1_int_1(esk9_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_42239, plain, (v1_funct_1(k4_xcmplx_0(np__1))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_40452, plain, (v1_xcmplx_0(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_46005, plain, (k3_xcmplx_0(X1,np__1)=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_46356, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_44506, plain, (k2_xcmplx_0(k4_xcmplx_0(X3),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_41633, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_40578, plain, (v2_xxreal_0(X1)|v1_xboole_0(k2_xcmplx_0(X1,np__0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_42240, plain, (v1_relat_1(k4_xcmplx_0(np__1))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_2186, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=k3_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(np__1)),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_45704, plain, (~v3_xxreal_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43785, plain, (v1_xcmplx_0(k2_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_42242, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_42244, plain, (v3_ordinal1(k4_xcmplx_0(np__1))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_42245, plain, (v1_rat_1(k4_xcmplx_0(np__1))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_42246, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_44605, plain, (v1_xxreal_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v2_membered(k1_numbers)|~v1_xboole_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_42471, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_46140, plain, (k2_xcmplx_0(np__0,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_41872, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_40431, plain, (X1=k2_xcmplx_0(X2,np__0)|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_39920, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|X2=X1|v2_funct_1(k1_fcont_1(X1,np__0))|k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X1,np__0)))!=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39990, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|k3_xcmplx_0(X1,np__1)=X2|v2_funct_1(k1_fcont_1(X1,np__0))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))!=k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40232, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|k3_xcmplx_0(X1,np__1)=np__0|v2_funct_1(k1_fcont_1(X1,np__0))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))!=np__0|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45435, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_42977, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_46038, plain, (v1_xreal_0(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2191, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,X1),np__1)=k2_xcmplx_0(np__1,k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_43873, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43878, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8076, plain, (np__0=np__1|X1=np__0|k3_xcmplx_0(X1,np__1)!=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_44608, plain, (v1_rat_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v4_membered(k1_numbers)|~v1_xboole_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_44609, plain, (v1_int_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v5_membered(k1_numbers)|~v1_xboole_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_44627, plain, (v1_xxreal_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v2_membered(k1_numbers)|~v1_xboole_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_44630, plain, (v1_rat_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v4_membered(k1_numbers)|~v1_xboole_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_44631, plain, (v1_int_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v5_membered(k1_numbers)|~v1_xboole_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_43124, plain, (k3_xcmplx_0(X1,np__1)=np__0|X2=np__0|k3_xcmplx_0(X2,X1)!=np__0|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43874, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45424, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39760, plain, (k3_xcmplx_0(np__1,X1)=np__0|X2=np__0|k3_xcmplx_0(X1,X2)!=np__0|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40433, plain, (k2_xcmplx_0(X1,X1)=X2|X3=np__0|k3_xcmplx_0(k2_xcmplx_0(X3,X3),X1)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40125, plain, (k2_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k6_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))))|v2_funct_1(k1_fcont_1(X1,np__0))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k6_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45441, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_41994, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k6_xcmplx_0(X2,k4_xcmplx_0(X1)))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_39983, plain, (esk5_1(k1_fcont_1(X2,np__0))=np__0|k3_xcmplx_0(X1,np__1)=X2|v2_funct_1(k1_fcont_1(X2,np__0))|k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X2,np__0)))!=k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,np__0)))|~v1_xcmplx_0(esk5_1(k1_fcont_1(X2,np__0)))|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X2,np__0)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_43418, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))=k4_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40569, plain, (v2_xxreal_0(esk22_0)|v2_xxreal_0(X1)|v1_xboole_0(k2_xcmplx_0(X1,esk22_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2193, plain, (k2_xcmplx_0(X1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8107, plain, (np__0=np__1|np__1=X1|k3_xcmplx_0(np__1,X1)!=np__1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8113, plain, (np__0=np__1|np__0=X1|k3_xcmplx_0(np__1,X1)!=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_44638, plain, (v1_xxreal_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v2_membered(k1_numbers)|~v1_xboole_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_44641, plain, (v1_rat_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v4_membered(k1_numbers)|~v1_xboole_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_44642, plain, (v1_int_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v5_membered(k1_numbers)|~v1_xboole_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_46185, plain, (v1_xxreal_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v2_membered(k1_numbers)|~v1_xboole_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_46188, plain, (v1_rat_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v4_membered(k1_numbers)|~v1_xboole_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_45564, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40434, plain, (k2_xcmplx_0(X1,X1)=X2|X3=np__0|k3_xcmplx_0(k2_xcmplx_0(X3,X3),X1)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_44550, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39914, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|X1=np__0|v2_funct_1(k1_fcont_1(X1,np__0))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))!=k3_xcmplx_0(X1,np__0)|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(X1)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45445, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))),k2_xcmplx_0(X1,X1))=k2_xcmplx_0(X1,k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))),X1))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))),X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k2_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_45418, plain, (k2_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))),k2_xcmplx_0(X1,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))=k2_xcmplx_0(X1,k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|v2_funct_1(k1_fcont_1(X3,X1))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))),k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_45446, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X2))=k2_xcmplx_0(X2,k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X2,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39993, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|k3_xcmplx_0(np__1,X1)=X2|v2_funct_1(k1_fcont_1(X1,np__0))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))!=k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xcmplx_0(X2)|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40478, plain, (v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_6382, plain, (v1_relat_1(k9_xtuple_0(X1))|~v4_relat_1(X1,np__0)|~v1_xboole_0(X3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2195, plain, (k3_xcmplx_0(k2_xcmplx_0(np__0,X1),np__0)=k2_xcmplx_0(np__0,k3_xcmplx_0(np__0,X1))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_46189, plain, (v1_int_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v5_membered(k1_numbers)|~v1_xboole_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_37255, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k4_xcmplx_0(esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_38403, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k4_xcmplx_0(esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_45254, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40810, plain, (k4_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2))=k4_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_52058, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_46916, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3073, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_52477, plain, (k3_xcmplx_0(k2_xcmplx_0(np__1,esk5_1(k1_fcont_1(X1,X1))),X1)=k2_xcmplx_0(X1,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X1))))|v2_funct_1(k1_fcont_1(X1,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X1))))|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_53431, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_53021, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_46984, plain, (k2_xcmplx_0(X1,np__0)=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28875, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47517, plain, (v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_47518, plain, (v1_xxreal_0(esk22_0))).#trainneg
cnf(i_0_48208, plain, (v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_48209, plain, (v1_xxreal_0(esk9_0))).#trainneg
cnf(i_0_48925, plain, (v2_xxreal_0(esk22_0)|v1_xboole_0(k2_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_52185, plain, (v2_xxreal_0(esk22_0)|v1_xboole_0(k2_xcmplx_0(esk22_0,esk22_0)))).#trainneg
cnf(i_0_52194, plain, (v2_xxreal_0(esk22_0)|v1_xboole_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_39364, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k4_xcmplx_0(esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_48513, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14487, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_48129, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4556, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_52955, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k4_xcmplx_0(X1)))=k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k6_xcmplx_0(X2,X1))|v2_funct_1(k1_fcont_1(X3,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_5246, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47117, plain, (k3_xcmplx_0(np__1,X1)=X2|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_47206, plain, (v2_xxreal_0(X1)|v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_47201, plain, (v1_xboole_0(k2_xcmplx_0(np__0,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_54176, plain, (m1_subset_1(esk4_1(esk24_0),k1_numbers)|v2_funct_1(esk24_0)|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_54179, plain, (m1_subset_1(esk4_1(esk13_0),k1_numbers)|v2_funct_1(esk13_0)|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_24595, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk13_0),np__1)))).#trainneg
cnf(i_0_24596, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk5_1(esk13_0),np__0)))).#trainneg
cnf(i_0_29198, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(np__1,esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_4608, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_52014, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55159, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55175, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55161, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_58882, plain, (X1=k2_xcmplx_0(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_14954, plain, (v1_xboole_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_14805, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_25994, plain, (~v1_xboole_0(k2_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_25979, plain, (~v1_xboole_0(k2_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_25995, plain, (~v1_xboole_0(k2_xcmplx_0(np__1,esk9_0)))).#trainneg
cnf(i_0_25980, plain, (~v1_xboole_0(k2_xcmplx_0(esk9_0,np__1)))).#trainneg
cnf(i_0_14955, plain, (~v1_xboole_0(k2_xcmplx_0(esk9_0,np__0)))).#trainneg
cnf(i_0_17814, plain, (~v1_xboole_0(k2_xcmplx_0(np__0,esk9_0)))).#trainneg
cnf(i_0_18140, plain, (v1_xboole_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_19576, plain, (v1_xreal_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_56593, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_56230, plain, (k2_xcmplx_0(X1,np__0)=np__0|X2=np__0|k3_xcmplx_0(X1,X2)!=np__0|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_20040, plain, (v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14792, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_19579, plain, (v1_xreal_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_59306, plain, (v1_funct_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_59307, plain, (v1_relat_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_59309, plain, (v1_xxreal_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_59312, plain, (v1_rat_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_59313, plain, (v1_int_1(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_54585, plain, (k4_xcmplx_0(X1)=np__0|X2=k4_xcmplx_0(X3)|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_57397, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16309, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v4_relat_1(X1,np__0)|~v1_xboole_0(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_20130, plain, (v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7781, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59152, plain, (k2_xcmplx_0(np__0,X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_58006, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_56480, plain, (k2_xcmplx_0(esk22_0,np__0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_59317, plain, (v1_funct_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_59318, plain, (v1_relat_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_59320, plain, (v1_xxreal_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_59323, plain, (v1_rat_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_59324, plain, (v1_int_1(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_54502, plain, (k4_xcmplx_0(X1)=np__0|X2=k4_xcmplx_0(X3)|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_55136, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_60226, plain, (k2_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))),k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2)))))|v2_funct_1(k1_fcont_1(X3,X2))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2)))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk4_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X3,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4977, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_22248, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6629, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_56498, plain, (k2_xcmplx_0(esk22_0,esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_56512, plain, (k2_xcmplx_0(np__0,esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_58878, plain, (k2_xcmplx_0(X1,X2)=np__0|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_59314, plain, (v1_xboole_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~m1_subset_1(k1_numbers,esk4_1(esk24_0)))).#trainneg
cnf(i_0_59325, plain, (v1_xboole_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~m1_subset_1(k1_numbers,esk4_1(esk13_0)))).#trainneg
cnf(i_0_59132, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_60255, plain, (k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_60598, plain, (k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4551, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_63856, plain, (v1_xreal_0(k3_xcmplx_0(np__0,esk22_0))|~v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_63858, plain, (v1_xreal_0(k3_xcmplx_0(np__0,esk9_0))|~v1_xboole_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_61543, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(k4_xcmplx_0(X3),X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_63408, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X1,k4_xcmplx_0(X3))!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_60594, plain, (k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_62747, plain, (k4_xcmplx_0(X1)=np__0|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_38240, plain, (v3_xxreal_0(X1)|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4529, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k4_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_62956, plain, (esk22_0=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_64184, plain, (k4_xcmplx_0(esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_63855, plain, (v1_xreal_0(k3_xcmplx_0(np__0,esk22_0))|~m1_subset_1(k4_xcmplx_0(esk22_0),k5_numbers))).#trainneg
cnf(i_0_63857, plain, (v1_xreal_0(k3_xcmplx_0(np__0,esk9_0))|~m1_subset_1(k4_xcmplx_0(esk9_0),k5_numbers))).#trainneg
cnf(i_0_62428, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__0)=np__0)).#trainneg
cnf(i_0_61883, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_60628, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_61135, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_60621, plain, (esk5_1(k1_fcont_1(X1,np__0))=np__0|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63466, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X3,k4_xcmplx_0(X1))!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_63445, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_61600, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X1,k4_xcmplx_0(X2))!=k3_xcmplx_0(X1,k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_40582, plain, (v2_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4524, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_62809, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_56529, plain, (m1_subset_1(esk5_1(esk24_0),k1_numbers)|v2_funct_1(esk24_0)|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_59330, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k4_xcmplx_0(esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_61420, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_61148, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1)))).#trainneg
cnf(i_0_64556, plain, (X1=np__0|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X3,X2)|~v1_xboole_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_60977, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),k6_xcmplx_0(X2,X3))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63312, plain, (k3_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_68268, plain, (k1_funct_1(k1_fcont_1(X1,np__0),esk4_1(k1_fcont_1(X1,np__0)))=k1_funct_1(k1_fcont_1(X1,np__0),np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65534, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_67154, plain, (X1=k4_xcmplx_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_69894, plain, (np__0=np__1|k3_xcmplx_0(np__1,X1)!=np__1|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3071, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_66889, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65535, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68365, plain, (k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68565, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6339, plain, (v1_xboole_0(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_65560, plain, (v1_xboole_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2331, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(np__1),X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_69523, plain, (v1_funct_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_69524, plain, (v1_relat_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_69526, plain, (v1_xxreal_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_69529, plain, (v1_rat_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_69530, plain, (v1_int_1(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_66890, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65516, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X1))=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_65944, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_66639, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X3,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_65754, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4496, plain, (m1_subset_1(k2_xcmplx_0(np__0,X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_66972, plain, (np__0=X1|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2093, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_62969, plain, (k6_xcmplx_0(np__0,esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_69531, plain, (v1_xboole_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~m1_subset_1(k1_numbers,esk5_1(esk24_0)))).#trainneg
cnf(i_0_69535, plain, (m1_subset_1(esk5_1(esk13_0),k1_numbers)|v2_funct_1(esk13_0)|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_64097, plain, (k3_xcmplx_0(np__0,esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_67358, plain, (k6_xcmplx_0(esk22_0,np__0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_65515, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X1)))=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_72394, plain, (k3_xcmplx_0(np__0,X1)=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68362, plain, (k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65574, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(np__0,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_74654, plain, (np__0=X1|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71261, plain, (np__0=np__1|X1!=np__1|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2066, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_75012, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_75014, plain, (v2_xxreal_0(esk22_0)|~v2_xxreal_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_75019, plain, (v1_funct_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_75020, plain, (v1_relat_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_75022, plain, (v1_xxreal_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_75025, plain, (v1_rat_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_71163, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_65670, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k3_xcmplx_0(k4_xcmplx_0(X1),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68361, plain, (k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68013, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_71587, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65532, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X1)))=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68014, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68488, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(k4_xcmplx_0(X1),X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_72780, plain, (v1_xboole_0(k9_xtuple_0(X2))|~v4_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_71252, plain, (np__0=np__1|k3_xcmplx_0(X1,np__1)!=np__1|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2025, plain, (X1=k1_fcont_1(k4_xcmplx_0(X2),X3)|k1_seq_1(X1,esk3_3(k4_xcmplx_0(X2),X3,X1))!=k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(k4_xcmplx_0(np__1),esk3_3(k4_xcmplx_0(X2),X3,X1))),X3)|~v1_xcmplx_0(esk3_3(k4_xcmplx_0(X2),X3,X1))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_72698, plain, (np__1=X1|k3_xcmplx_0(np__1,X1)!=np__1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_72715, plain, (np__0=X1|k3_xcmplx_0(np__1,X1)!=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_75026, plain, (v1_int_1(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_75027, plain, (v1_xboole_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~m1_subset_1(k1_numbers,esk5_1(esk13_0)))).#trainneg
cnf(i_0_68648, plain, (k4_xcmplx_0(X1)=np__0|X2=k4_xcmplx_0(X3)|k3_xcmplx_0(X1,k4_xcmplx_0(X2))!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_66879, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k2_xcmplx_0(k3_xcmplx_0(X1,np__0),np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68703, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X1,k4_xcmplx_0(X2))!=k3_xcmplx_0(k4_xcmplx_0(X3),X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_65538, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_76340, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_72777, plain, (X1=k1_relset_1(X2,X3)|~v4_relat_1(X3,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_76624, plain, (v2_funct_1(esk24_0)|esk5_1(esk24_0)!=np__1|~v1_xboole_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_1988, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_76620, plain, (esk22_0!=np__1|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_76648, plain, (X1!=np__1|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_78457, plain, (k3_xcmplx_0(np__1,X1)!=np__1|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_76954, plain, (v2_xxreal_0(esk22_0)|v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_75118, plain, (k6_xcmplx_0(esk22_0,esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_76147, plain, (k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_76401, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_77574, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_79410, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_72818, plain, (v1_xboole_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_78662, plain, (np__1=X1|k3_xcmplx_0(X1,np__1)!=np__1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1069, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__1))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_78461, plain, (k2_xcmplx_0(X1,np__0)!=np__1|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_37256, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(np__0,esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_38400, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(np__1,esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_77575, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65753, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_66619, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_66495, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_72340, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68028, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_74586, plain, (m1_subset_1(k2_xcmplx_0(np__0,X1),k1_numbers)|~v1_xboole_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1145, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(np__1)),k4_xcmplx_0(X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_78462, plain, (k2_xcmplx_0(np__0,X1)!=np__1|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_78672, plain, (np__1=X1|k2_xcmplx_0(X1,np__0)!=np__1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_38404, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(np__0,esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_39361, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(np__1,esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_70767, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k4_xcmplx_0(k3_xcmplx_0(X2,X1))!=k3_xcmplx_0(X3,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_72381, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_80996, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_82538, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_81456, plain, (X1=k4_xcmplx_0(esk22_0)|v2_xxreal_0(esk22_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_81455, plain, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(esk22_0)!=np__1)).#trainneg
cnf(i_0_1116, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82541, plain, (~v1_xboole_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_81382, plain, (esk20_1(X1)!=np__1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_81390, plain, (k9_xtuple_0(X1)!=np__1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_82540, plain, (v1_xboole_0(esk22_0)|~v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_81391, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(esk22_0,esk22_0)!=np__1)).#trainneg
cnf(i_0_81392, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(esk22_0,np__0)!=np__1)).#trainneg
cnf(i_0_81393, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(np__0,esk22_0)!=np__1)).#trainneg
cnf(i_0_82542, plain, (v1_xboole_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|~v1_xboole_0(k4_xcmplx_0(esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_82376, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_83756, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_82525, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_66510, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1)))).#trainneg
cnf(i_0_78673, plain, (np__1=X1|k2_xcmplx_0(np__0,X1)!=np__1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1110, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81394, plain, (k2_xcmplx_0(np__0,X1)!=np__1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_81387, plain, (k4_xcmplx_0(X1)!=np__1|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81395, plain, (k2_xcmplx_0(X1,X2)!=np__1|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_82543, plain, (v1_xboole_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|~v1_xboole_0(k4_xcmplx_0(esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_82544, plain, (v1_xboole_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|~v1_xboole_0(k4_xcmplx_0(esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_82545, plain, (v1_xboole_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|~v1_xboole_0(k4_xcmplx_0(esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_79734, plain, (k3_xcmplx_0(np__0,k2_xcmplx_0(np__1,np__1))=np__0)).#trainneg
cnf(i_0_81310, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_66407, plain, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(k4_xcmplx_0(X3),k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_72166, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_75424, plain, (k6_xcmplx_0(X3,k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_70939, plain, (k3_xcmplx_0(X2,k4_xcmplx_0(X1))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_78708, plain, (np__0=X1|k3_xcmplx_0(X1,np__1)!=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1155, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_84966, plain, (k3_xcmplx_0(np__1,esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_39365, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(np__0,esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0)))).#trainneg
cnf(i_0_66886, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k4_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X1,np__0)))|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_66877, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k2_xcmplx_0(np__0,k3_xcmplx_0(X1,np__0))|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v4_relat_1(k1_fcont_1(X1,np__0),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0))))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_74947, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),np__0)=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68652, plain, (k4_xcmplx_0(X1)=np__0|X2=np__0|k3_xcmplx_0(X1,k4_xcmplx_0(X2))!=k3_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_66587, plain, (k4_xcmplx_0(X3)=np__0|k4_xcmplx_0(X1)=X2|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(k4_xcmplx_0(X3),X2)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_2200, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50665, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(np__0,esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_83329, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68266, plain, (k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|esk4_1(k1_fcont_1(X1,np__0))!=np__0|~v1_relat_1(k1_fcont_1(X1,np__0))|~v1_funct_1(k1_fcont_1(X1,np__0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_67606, plain, (k6_xcmplx_0(X3,k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_78718, plain, (np__0=X1|k2_xcmplx_0(X1,np__0)!=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2216, plain, (X1=X2|v1_xboole_0(k9_xtuple_0(X3))|k1_funct_1(X3,X1)!=k1_funct_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X2,k9_xtuple_0(X3))|~m1_subset_1(X1,k9_xtuple_0(X3))|~v1_funct_1(X3)|~v2_funct_1(X3))).#trainneg
cnf(i_0_78719, plain, (np__0=X1|k2_xcmplx_0(np__0,X1)!=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59327, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(np__1,esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0)))).#trainneg
cnf(i_0_66496, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_85542, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_85959, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_72772, plain, (v1_xreal_0(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2231, plain, (v1_relat_1(k1_funct_1(X1,X2))|~v4_funct_1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_85929, plain, (v2_xxreal_0(esk22_0)|esk22_0!=np__1)).#trainneg
cnf(i_0_87309, plain, (k2_xcmplx_0(np__1,np__1)=np__0|k3_xcmplx_0(np__1,k2_xcmplx_0(np__1,np__1))!=np__0)).#trainneg
cnf(i_0_87346, plain, (k2_xcmplx_0(np__1,np__1)=np__0|k2_xcmplx_0(k2_xcmplx_0(np__1,np__1),np__0)!=np__0)).#trainneg
cnf(i_0_88645, plain, (k4_xcmplx_0(np__1)=esk22_0|v2_xxreal_0(esk22_0)|k4_xcmplx_0(np__1)!=np__0)).#trainneg
cnf(i_0_88648, plain, (k3_xcmplx_0(esk22_0,np__1)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_88124, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,np__0))=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_89201, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X1,np__1)))=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_86491, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_86184, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_40614, plain, (v3_xxreal_0(X2)|v2_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2232, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_18304, plain, (v2_xxreal_0(esk4_1(esk24_0))|v1_xboole_0(esk4_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_90430, plain, (X1=esk5_1(X2)|v1_xboole_0(k9_xtuple_0(X3))|v2_funct_1(X2)|k1_funct_1(X3,X1)!=k1_funct_1(X3,esk5_1(X2))|~v4_relat_1(X2,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_relat_1(X2)|~m1_subset_1(X1,k9_xtuple_0(X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v2_funct_1(X3))).#trainneg
cnf(i_0_65797, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k4_xcmplx_0(k3_xcmplx_0(X1,X3))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_70715, plain, (k3_xcmplx_0(k2_xcmplx_0(np__0,X1),X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81308, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(X1))=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_68054, plain, (k6_xcmplx_0(X3,k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k3_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_81808, plain, (k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,np__0)))=k3_xcmplx_0(X1,np__0)|k3_xcmplx_0(X1,np__1)=X1|v2_funct_1(k1_fcont_1(X1,np__0))|~v1_xcmplx_0(k3_xcmplx_0(X1,np__0))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65529, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,k2_xcmplx_0(X1,X2)),X2)=np__0|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_91302, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_91938, plain, (v1_xreal_0(k9_xtuple_0(X2))|~v4_relat_1(X2,X1)|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2233, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_91947, plain, (k2_xcmplx_0(np__1,np__1)=np__0|k2_xcmplx_0(np__0,k2_xcmplx_0(np__1,np__1))!=np__0)).#trainneg
cnf(i_0_93660, plain, (esk4_1(X1)=esk5_1(X2)|v2_funct_1(X2)|v2_funct_1(X1)|k1_funct_1(X3,esk4_1(X1))!=k1_funct_1(X3,esk5_1(X2))|~v4_relat_1(X2,k9_xtuple_0(X3))|~v4_relat_1(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_funct_1(X3))).#trainneg
cnf(i_0_93658, plain, (esk5_1(X1)=esk5_1(X2)|v2_funct_1(X2)|v2_funct_1(X1)|k1_funct_1(X3,esk5_1(X1))!=k1_funct_1(X3,esk5_1(X2))|~v4_relat_1(X2,k9_xtuple_0(X3))|~v4_relat_1(X1,k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_funct_1(X3))).#trainneg
cnf(i_0_90944, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_93522, plain, (X1=k3_xcmplx_0(X2,X3)|v3_xxreal_0(X3)|v2_xxreal_0(X3)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_89099, plain, (k4_xcmplx_0(X1)=np__0|k3_xcmplx_0(X1,np__0)!=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2235, plain, (v3_ordinal1(k1_funct_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_86995, plain, (k6_xcmplx_0(np__0,X1)!=np__1|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_84962, plain, (v1_xreal_0(k3_xcmplx_0(np__1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_89339, plain, (v1_xreal_0(k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18315, plain, (v2_xxreal_0(esk4_1(esk13_0))|v1_xboole_0(esk4_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_24580, plain, (v2_xxreal_0(esk5_1(esk24_0))|v1_xboole_0(esk5_1(esk24_0))|v2_funct_1(esk24_0))).#trainneg
cnf(i_0_24592, plain, (v2_xxreal_0(esk5_1(esk13_0))|v1_xboole_0(esk5_1(esk13_0))|v2_funct_1(esk13_0))).#trainneg
cnf(i_0_29194, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk22_0,esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_91166, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_91156, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X3,X1))!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_91523, plain, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X2),X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40615, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2236, plain, (v1_rat_1(k1_funct_1(X1,X2))|~v4_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_87523, plain, (k4_xcmplx_0(k3_xcmplx_0(np__0,X1))=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_29196, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk9_0,esk4_1(esk24_0)))|~v1_xcmplx_0(esk4_1(esk24_0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_38396, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk22_0,esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_38398, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk9_0,esk4_1(esk13_0)))|~v1_xcmplx_0(esk4_1(esk13_0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_91151, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k4_xcmplx_0(k3_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_98073, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_98324, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_97842, plain, (k4_xcmplx_0(X3)=np__0|k4_xcmplx_0(X1)=X2|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_97843, plain, (k4_xcmplx_0(X3)=np__0|k4_xcmplx_0(X1)=X2|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_98050, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X1,X3))!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_97934, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X3,X1))!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_98472, plain, (X1=k3_xcmplx_0(X2,X3)|v3_xxreal_0(X3)|v3_xxreal_0(X2)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_97085, plain, (k4_xcmplx_0(X1)=np__0|k3_xcmplx_0(np__0,X1)!=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2237, plain, (v1_int_1(k1_funct_1(X1,X2))|~v5_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_38408, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk22_0,esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_39359, plain, (v2_funct_1(esk24_0)|v1_xreal_0(k2_xcmplx_0(esk9_0,esk5_1(esk24_0)))|~v1_xcmplx_0(esk5_1(esk24_0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_97829, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_99401, plain, (k4_xcmplx_0(X1)=np__0|X2=X3|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k4_xcmplx_0(k3_xcmplx_0(X3,X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_42135, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v2_xxreal_0(X2)|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2301, plain, (X1=k1_fcont_1(X2,k3_xcmplx_0(X3,X2))|k1_seq_1(X1,esk3_3(X2,k3_xcmplx_0(X3,X2),X1))!=k3_xcmplx_0(k2_xcmplx_0(esk3_3(X2,k3_xcmplx_0(X3,X2),X1),X3),X2)|~v1_xcmplx_0(esk3_3(X2,k3_xcmplx_0(X3,X2),X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X3,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_39369, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk22_0,esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0))|~v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_39371, plain, (v2_funct_1(esk13_0)|v1_xreal_0(k2_xcmplx_0(esk9_0,esk5_1(esk13_0)))|~v1_xcmplx_0(esk5_1(esk13_0))|~v1_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_67360, plain, (k4_xcmplx_0(k2_xcmplx_0(np__0,esk22_0))=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_97819, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X3,X1))!=k3_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_100756, plain, (k4_xcmplx_0(X1)=np__0|k4_xcmplx_0(X2)=np__0|k3_xcmplx_0(X2,X1)!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_96978, plain, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_99735, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k4_xcmplx_0(k3_xcmplx_0(X3,X1))!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_102839, plain, (X1=k3_xcmplx_0(X2,X3)|v3_xxreal_0(X2)|v2_xxreal_0(X3)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_99659, plain, (np__0=X1|k3_xcmplx_0(np__1,X1)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2302, plain, (X1=k1_fcont_1(X2,k3_xcmplx_0(X3,X2))|k1_funct_1(X1,esk3_3(X2,k3_xcmplx_0(X3,X2),X1))!=k3_xcmplx_0(k2_xcmplx_0(esk3_3(X2,k3_xcmplx_0(X3,X2),X1),X3),X2)|~v3_valued_0(X1)|~v1_xcmplx_0(esk3_3(X2,k3_xcmplx_0(X3,X2),X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X3,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_103358, plain, (v2_xxreal_0(esk22_0)|~v2_xxreal_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_103359, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_42136, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|v2_xxreal_0(X2)|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_101429, plain, (np__0=X1|k3_xcmplx_0(X1,np__1)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2306, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__1),np__1)=k2_xcmplx_0(k3_xcmplx_0(np__1,X1),np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_86993, plain, (k2_xcmplx_0(np__0,X1)!=np__1|~v1_xboole_0(k6_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_87009, plain, (k2_xcmplx_0(np__0,X1)!=np__1|~v1_xboole_0(k2_xcmplx_0(X1,np__0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_103369, plain, (v2_xxreal_0(esk22_0)|m1_subset_1(k2_xcmplx_0(np__0,esk22_0),k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_67361, plain, (k4_xcmplx_0(k2_xcmplx_0(esk22_0,np__0))=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_100388, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_99284, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_2308, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),X1)=k3_xcmplx_0(k2_xcmplx_0(X2,np__1),X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_106444, plain, (v2_xxreal_0(esk22_0)|~v2_xxreal_0(k2_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_106445, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k2_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_104331, plain, (np__1=X1|X2=np__0|k3_xcmplx_0(X2,X1)!=X2|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_105776, plain, (np__0=X1|X2=np__0|k3_xcmplx_0(X1,X2)!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_105887, plain, (X1=k3_xcmplx_0(X2,X3)|v3_xxreal_0(X3)|v3_xxreal_0(X2)|v2_xxreal_0(X3)|~v1_xboole_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_104330, plain, (np__1=X1|k3_xcmplx_0(np__1,X1)!=np__1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2310, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,np__0),np__0)=k2_xcmplx_0(k3_xcmplx_0(np__0,X1),np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_106401, plain, (v2_xxreal_0(esk22_0)|m1_subset_1(k2_xcmplx_0(np__0,esk22_0),k1_numbers))).#trainneg
cnf(i_0_106455, plain, (v2_xxreal_0(esk22_0)|m1_subset_1(k2_xcmplx_0(esk22_0,np__0),k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_102527, plain, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_105872, plain, (k4_xcmplx_0(X3)=np__0|k4_xcmplx_0(X1)=X2|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_97796, plain, (k4_xcmplx_0(X3)=np__0|k4_xcmplx_0(X1)=X2|k4_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X1,X3)))!=k3_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_104306, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=np__0|X2=np__0|k3_xcmplx_0(X1,X2)!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2316, plain, (k2_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(np__1),X1),k4_xcmplx_0(X2))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_65596, plain, (k4_xcmplx_0(k3_xcmplx_0(np__0,X1))=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_75145, plain, (k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_75146, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_98323, plain, (k3_xcmplx_0(k4_xcmplx_0(X2),X1)=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_110511, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|k4_xcmplx_0(X1)=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_108848, plain, (k3_xcmplx_0(X1,np__1)=X1|X1=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_108847, plain, (k3_xcmplx_0(np__1,X1)=X1|X1=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_108690, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=np__0|X1=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_109028, plain, (np__1=X1|k3_xcmplx_0(X1,np__1)!=np__1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_112041, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))!=np__1|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_110089, plain, (v2_xxreal_0(esk22_0)|v1_funct_1(k2_xcmplx_0(np__0,esk22_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_110090, plain, (v2_xxreal_0(esk22_0)|v1_relat_1(k2_xcmplx_0(np__0,esk22_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_110091, plain, (v2_xxreal_0(esk22_0)|v1_xcmplx_0(k2_xcmplx_0(np__0,esk22_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_110092, plain, (v2_xxreal_0(esk22_0)|v1_xxreal_0(k2_xcmplx_0(np__0,esk22_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_111090, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=np__0|X1!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_108691, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=np__0|X1=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_94659, plain, (k4_xcmplx_0(X1)=np__0|k4_xcmplx_0(X2)=np__0|k3_xcmplx_0(X2,X1)!=np__0|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_104431, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X3,X1)!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2323, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1)=k2_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_111012, plain, (k4_xcmplx_0(k4_xcmplx_0(esk22_0))=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_114959, plain, (k3_xcmplx_0(np__1,X1)=np__0|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_114698, plain, (v2_xxreal_0(esk22_0)|k6_xcmplx_0(np__0,esk22_0)!=np__1|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_114701, plain, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(k2_xcmplx_0(np__0,esk22_0))!=np__1|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_114704, plain, (v2_xxreal_0(esk22_0)|v1_xcmplx_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_110094, plain, (v3_ordinal1(k2_xcmplx_0(np__0,esk22_0))|v2_xxreal_0(esk22_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_110095, plain, (v1_rat_1(k2_xcmplx_0(np__0,esk22_0))|v2_xxreal_0(esk22_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_114714, plain, (k4_xcmplx_0(X1)=np__0|X1!=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_114388, plain, (np__1=X1|X2=np__0|k3_xcmplx_0(X2,X1)!=X2|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_113845, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X1))=np__0|X1=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_114313, plain, (X1=np__0|k3_xcmplx_0(X1,np__0)!=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_114391, plain, (np__1=X1|X2=np__0|k3_xcmplx_0(X1,X2)!=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_115127, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X1,k4_xcmplx_0(X3))!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_6784, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v4_relat_1(X1,X2)|~v1_xboole_0(k1_zfmisc_1(X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_114960, plain, (k3_xcmplx_0(X1,np__1)=np__0|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_114751, plain, (X1!=np__0|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_114752, plain, (X1!=np__0|~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_110096, plain, (v1_int_1(k2_xcmplx_0(np__0,esk22_0))|v2_xxreal_0(esk22_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_112031, plain, (v2_xxreal_0(esk22_0)|v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_116379, plain, (v2_xxreal_0(esk22_0)|m1_subset_1(k4_xcmplx_0(esk22_0),k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_110109, plain, (v2_xxreal_0(esk22_0)|m1_subset_1(k2_xcmplx_0(esk22_0,np__0),k1_numbers))).#trainneg
cnf(i_0_104110, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_37130, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|v2_xxreal_0(X2)|v2_xxreal_0(X1)|v1_xboole_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2340, plain, (v2_xxreal_0(k1_funct_1(X1,X2))|v1_xboole_0(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_81490, plain, (k2_xcmplx_0(k4_xcmplx_0(esk22_0),esk22_0)=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_81491, plain, (k2_xcmplx_0(esk22_0,k4_xcmplx_0(esk22_0))=np__0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_85553, plain, (k2_xcmplx_0(esk22_0,esk22_0)=k4_xcmplx_0(esk22_0)|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_114953, plain, (k2_xcmplx_0(X1,np__0)=np__0|X1!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39224, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|v2_xxreal_0(X2)|v2_xxreal_0(X1)|v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2341, plain, (v2_xxreal_0(k2_xcmplx_0(X1,X2))|v1_xboole_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_117802, plain, (esk9_0!=np__0)).#trainneg
cnf(i_0_117858, plain, (k4_xcmplx_0(np__1)!=np__0)).#trainneg
cnf(i_0_117798, plain, (v1_xboole_0(esk22_0)|esk22_0!=np__0)).#trainneg
cnf(i_0_120286, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0)))|~v7_ordinal1(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_119070, plain, (v3_ordinal1(k9_xtuple_0(esk24_0))|~v1_xboole_0(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_116843, negated_conjecture, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116723, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=np__0|k4_xcmplx_0(X1)=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116650, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116722, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=np__0|k4_xcmplx_0(X1)=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116654, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_117661, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X3,k4_xcmplx_0(X1))!=k3_xcmplx_0(X3,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_117974, plain, (X1!=np__0|~v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_120502, plain, (v1_xboole_0(X1)|X1!=np__0|~v7_ordinal1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2342, plain, (v2_xxreal_0(k3_xcmplx_0(X1,X2))|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_121109, plain, (esk22_0!=np__1|esk22_0!=np__0)).#trainneg
cnf(i_0_117975, plain, (X1!=np__0|~v3_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_121126, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k6_xcmplx_0(esk22_0,np__0))|~v7_ordinal1(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_121128, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k3_xcmplx_0(np__1,esk22_0))|~v7_ordinal1(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_121129, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k3_xcmplx_0(esk22_0,np__1))|~v7_ordinal1(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_119073, plain, (v3_ordinal1(k9_xtuple_0(esk18_0))|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_119074, plain, (v3_ordinal1(k9_xtuple_0(esk13_0))|~v1_xboole_0(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_116672, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116737, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_116766, negated_conjecture, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_117872, negated_conjecture, (np__0=X1|k4_xcmplx_0(X1)!=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2345, plain, (v1_xboole_0(k1_fcont_1(X2,X3))|m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_fcont_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_120267, plain, (v2_xxreal_0(esk22_0)|v7_ordinal1(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_124147, plain, (esk17_0!=np__0)).#trainneg
cnf(i_0_124331, plain, (esk10_0!=np__0)).#trainneg
cnf(i_0_124330, plain, (esk16_0!=np__0|~v1_xreal_0(esk16_0))).#trainneg
cnf(i_0_124874, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116761, negated_conjecture, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_116749, negated_conjecture, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_116753, negated_conjecture, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_125401, negated_conjecture, (k6_xcmplx_0(np__0,X1)=k4_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116751, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_124169, plain, (v3_xxreal_0(X1)|v1_xboole_0(X1)|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_118708, plain, (X1=np__0|k3_xcmplx_0(np__0,X1)!=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_124361, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_124366, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_125909, plain, (v2_xxreal_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__1))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_125910, plain, (v2_xxreal_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(np__1,k4_xcmplx_0(esk22_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_120293, plain, (v2_xxreal_0(esk22_0)|v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_126294, negated_conjecture, (k2_xcmplx_0(np__0,X1)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_126210, negated_conjecture, (X1=k1_fcont_1(X2,X3)|k1_seq_1(X1,esk3_3(X2,X3,X1))!=k2_xcmplx_0(X3,k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_xcmplx_0(X2,esk3_3(X2,X3,X1)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_126076, negated_conjecture, (k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_126183, negated_conjecture, (k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2)=k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116684, negated_conjecture, (k4_xcmplx_0(X1)=np__0|~v1_xboole_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_122072, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X1))=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_129530, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_116788, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_119021, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|X3=np__0|k3_xcmplx_0(X1,X3)!=k3_xcmplx_0(X3,X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2367, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_124174, plain, (v1_xboole_0(esk5_1(esk24_0))|v2_funct_1(esk24_0)|esk5_1(esk24_0)!=np__0)).#trainneg
cnf(i_0_124175, plain, (v1_xboole_0(esk5_1(esk13_0))|v2_funct_1(esk13_0)|esk5_1(esk13_0)!=np__0)).#trainneg
cnf(i_0_124176, plain, (v1_xboole_0(esk4_1(esk24_0))|v2_funct_1(esk24_0)|esk4_1(esk24_0)!=np__0)).#trainneg
cnf(i_0_124177, plain, (v1_xboole_0(esk4_1(esk13_0))|v2_funct_1(esk13_0)|esk4_1(esk13_0)!=np__0)).#trainneg
cnf(i_0_116599, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_129666, negated_conjecture, (k1_fcont_1(esk1_0,esk2_0)!=np__0|~v1_xreal_0(k1_fcont_1(esk1_0,esk2_0)))).#trainneg
cnf(i_0_129624, negated_conjecture, (k9_xtuple_0(k1_fcont_1(esk1_0,esk2_0))!=np__0|~v1_relat_1(k1_fcont_1(esk1_0,esk2_0))|~v1_funct_1(k1_fcont_1(esk1_0,esk2_0))|~v1_xreal_0(k9_xtuple_0(k1_fcont_1(esk1_0,esk2_0))))).#trainneg
cnf(i_0_126078, negated_conjecture, (k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_117681, plain, (np__0=X1|X2=np__0|k3_xcmplx_0(X2,np__0)!=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_116754, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_134268, plain, (v2_funct_1(esk24_0)|esk5_1(esk24_0)!=np__1|esk5_1(esk24_0)!=np__0)).#trainneg
cnf(i_0_2376, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_129612, plain, (k4_xcmplx_0(esk9_0)!=np__0)).#trainneg
cnf(i_0_129586, plain, (esk30_0!=np__0|~v1_xreal_0(esk30_0))).#trainneg
cnf(i_0_129587, plain, (esk28_0!=np__0|~v1_xreal_0(esk28_0))).#trainneg
cnf(i_0_129626, plain, (k2_xcmplx_0(esk22_0,np__1)!=np__0)).#trainneg
cnf(i_0_129628, plain, (k2_xcmplx_0(esk9_0,np__1)!=np__0)).#trainneg
cnf(i_0_129629, plain, (k2_xcmplx_0(esk9_0,np__0)!=np__0)).#trainneg
cnf(i_0_129638, plain, (k2_xcmplx_0(np__1,esk22_0)!=np__0)).#trainneg
cnf(i_0_129639, plain, (k2_xcmplx_0(np__1,esk9_0)!=np__0)).#trainneg
cnf(i_0_116600, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_137523, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_136063, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(X2,X1))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_136071, plain, (X1=np__0|k3_xcmplx_0(X1,np__0)!=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_120874, plain, (v2_xxreal_0(X1)|v1_xboole_0(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_134285, plain, (v2_funct_1(esk13_0)|esk5_1(esk13_0)!=np__1|esk5_1(esk13_0)!=np__0)).#trainneg
cnf(i_0_2377, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X4)=k2_xcmplx_0(X3,k2_xcmplx_0(k2_xcmplx_0(X1,X2),X4))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_129588, plain, (esk26_0!=np__0|~v1_xreal_0(esk26_0))).#trainneg
cnf(i_0_129589, plain, (esk24_0!=np__0|~v1_xreal_0(esk24_0))).#trainneg
cnf(i_0_129590, plain, (esk23_0!=np__0|~v1_xreal_0(esk23_0))).#trainneg
cnf(i_0_129644, plain, (k2_xcmplx_0(np__1,np__1)!=np__0)).#trainneg
cnf(i_0_129647, plain, (k2_xcmplx_0(np__0,esk9_0)!=np__0)).#trainneg
cnf(i_0_138063, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_138064, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_139366, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_134302, plain, (v2_funct_1(esk24_0)|esk4_1(esk24_0)!=np__1|esk4_1(esk24_0)!=np__0)).#trainneg
cnf(i_0_129592, plain, (esk18_0!=np__0|~v1_xreal_0(esk18_0))).#trainneg
cnf(i_0_129595, plain, (esk13_0!=np__0|~v1_xreal_0(esk13_0))).#trainneg
cnf(i_0_129598, plain, (k4_numbers!=np__0|~v1_xreal_0(k4_numbers))).#trainneg
cnf(i_0_139620, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_139928, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_140009, plain, (X1=np__0|k3_xcmplx_0(np__0,X1)!=k2_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_141461, negated_conjecture, (k3_xcmplx_0(np__0,X1)=k3_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_134319, plain, (v2_funct_1(esk13_0)|esk4_1(esk13_0)!=np__1|esk4_1(esk13_0)!=np__0)).#trainneg
cnf(i_0_2381, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_129610, plain, (k5_numbers!=np__0|~v1_xreal_0(k5_numbers))).#trainneg
cnf(i_0_129665, plain, (np__0!=k1_numbers|~v1_xreal_0(k1_numbers))).#trainneg
cnf(i_0_139007, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_136089, plain, (X1=np__0|k3_xcmplx_0(X1,np__0)!=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143428, negated_conjecture, (k6_xcmplx_0(X1,np__0)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143427, negated_conjecture, (k6_xcmplx_0(X1,np__0)=np__0|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_142372, negated_conjecture, (X1=np__0|X2=np__0|k3_xcmplx_0(X2,X1)!=np__0|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143425, negated_conjecture, (k6_xcmplx_0(X1,np__0)=k3_xcmplx_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_130062, negated_conjecture, (k2_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))),X1)=k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|v2_funct_1(k1_fcont_1(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk5_1(k1_fcont_1(X2,X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1))))|~v1_xreal_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X1)))))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_129611, plain, (v1_xboole_0(esk22_0)|k4_xcmplx_0(esk22_0)!=np__0)).#trainneg
cnf(i_0_2390, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(X3,k2_xcmplx_0(k2_xcmplx_0(X1,X2),X4))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_145324, plain, (k4_xcmplx_0(k2_xcmplx_0(np__1,np__1))!=np__0)).#trainneg
cnf(i_0_129675, plain, (X1!=np__1|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143555, plain, (k2_xcmplx_0(k2_xcmplx_0(np__1,np__1),np__0)!=np__0)).#trainneg
cnf(i_0_129627, plain, (v1_xboole_0(esk22_0)|k2_xcmplx_0(esk22_0,np__0)!=np__0)).#trainneg
cnf(i_0_129646, plain, (v1_xboole_0(esk22_0)|k2_xcmplx_0(np__0,esk22_0)!=np__0)).#trainneg
cnf(i_0_129776, plain, (v2_xxreal_0(esk22_0)|~v1_xboole_0(k6_xcmplx_0(np__1,esk22_0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_116665, negated_conjecture, (X1=np__0|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_142314, negated_conjecture, (k3_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_142693, negated_conjecture, (X1=np__0|X2=np__0|k3_xcmplx_0(X1,X2)!=np__0|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116703, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_131523, negated_conjecture, (k4_xcmplx_0(X1)=np__0|~v1_xboole_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2395, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k2_xcmplx_0(X1,X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_126548, negated_conjecture, (k2_xcmplx_0(X1,np__0)!=np__1|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_131526, negated_conjecture, (k4_xcmplx_0(X1)=np__0|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_146648, negated_conjecture, (esk2_0!=np__1|esk2_0!=np__0)).#trainneg
cnf(i_0_145853, negated_conjecture, (X1=X2|X1!=k3_xcmplx_0(X2,np__1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_146644, plain, (v2_funct_1(k1_fcont_1(X1,X2))|esk4_1(k1_fcont_1(X1,X2))!=np__1|esk4_1(k1_fcont_1(X1,X2))!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_146638, plain, (v2_funct_1(k1_fcont_1(X1,X2))|esk5_1(k1_fcont_1(X1,X2))!=np__1|esk5_1(k1_fcont_1(X1,X2))!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_145809, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=X2|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_146636, plain, (v2_funct_1(esk18_0)|esk5_1(esk18_0)!=np__1|esk5_1(esk18_0)!=np__0)).#trainneg
cnf(i_0_2418, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X2),X2)=k3_xcmplx_0(k2_xcmplx_0(np__1,X1),X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_146632, plain, (esk14_0!=np__1|esk14_0!=np__0)).#trainneg
cnf(i_0_146655, plain, (k4_xcmplx_0(esk22_0)!=np__1|k4_xcmplx_0(esk22_0)!=np__0)).#trainneg
cnf(i_0_146760, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__1,np__1))!=np__0)).#trainneg
cnf(i_0_146673, plain, (k2_xcmplx_0(esk22_0,esk22_0)!=np__1|k2_xcmplx_0(esk22_0,esk22_0)!=np__0)).#trainneg
cnf(i_0_145927, negated_conjecture, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2)=k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116646, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_141515, negated_conjecture, (k3_xcmplx_0(np__1,X1)=k3_xcmplx_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_126212, negated_conjecture, (k1_fcont_1(X1,X2)=k1_fcont_1(X3,X4)|k2_xcmplx_0(X2,k3_xcmplx_0(X1,esk3_3(X3,X4,k1_fcont_1(X1,X2))))!=k2_xcmplx_0(k3_xcmplx_0(X3,esk3_3(X3,X4,k1_fcont_1(X1,X2))),X4)|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_xcmplx_0(X1,esk3_3(X3,X4,k1_fcont_1(X1,X2))))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_149864, negated_conjecture, (X1=X2|np__0!=X2|~v1_xboole_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_146642, plain, (v2_funct_1(esk18_0)|esk4_1(esk18_0)!=np__1|esk4_1(esk18_0)!=np__0)).#trainneg
cnf(i_0_2426, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(np__1)),k4_xcmplx_0(X2))=k3_xcmplx_0(k2_xcmplx_0(X2,X1),k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X1,k4_xcmplx_0(np__1)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_146665, plain, (k4_xcmplx_0(X1)!=np__1|k4_xcmplx_0(X1)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_146669, plain, (k1_funct_1(esk13_0,X1)!=np__1|k1_funct_1(esk13_0,X1)!=np__0)).#trainneg
cnf(i_0_146674, plain, (k2_xcmplx_0(esk22_0,esk9_0)!=np__1|k2_xcmplx_0(esk22_0,esk9_0)!=np__0)).#trainneg
cnf(i_0_146680, plain, (k2_xcmplx_0(esk22_0,np__0)!=np__1|k2_xcmplx_0(esk22_0,np__0)!=np__0)).#trainneg
cnf(i_0_146683, plain, (k2_xcmplx_0(esk9_0,esk22_0)!=np__1|k2_xcmplx_0(esk9_0,esk22_0)!=np__0)).#trainneg
cnf(i_0_146684, plain, (k2_xcmplx_0(esk9_0,esk9_0)!=np__1|k2_xcmplx_0(esk9_0,esk9_0)!=np__0)).#trainneg
cnf(i_0_146720, plain, (k2_xcmplx_0(np__0,esk22_0)!=np__1|k2_xcmplx_0(np__0,esk22_0)!=np__0)).#trainneg
cnf(i_0_146743, plain, (k3_xcmplx_0(np__0,esk22_0)!=np__1|k3_xcmplx_0(np__0,esk22_0)!=np__0|~v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_153161, negated_conjecture, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_152901, negated_conjecture, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_153165, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(X1,np__1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_153126, negated_conjecture, (k4_xcmplx_0(X1)=X2|k3_xcmplx_0(X1,np__1)!=k4_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_129673, plain, (X1=k4_xcmplx_0(X2)|X1!=np__0|~v1_xboole_0(X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_146657, plain, (v2_funct_1(esk24_0)|k4_xcmplx_0(esk5_1(esk24_0))!=np__1|k4_xcmplx_0(esk5_1(esk24_0))!=np__0)).#trainneg
cnf(i_0_2441, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k3_xcmplx_0(k2_xcmplx_0(X2,X1),k4_xcmplx_0(np__1))|~v1_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_146744, plain, (k3_xcmplx_0(np__0,esk22_0)!=np__1|k3_xcmplx_0(np__0,esk22_0)!=np__0|~m1_subset_1(k4_xcmplx_0(esk22_0),k5_numbers))).#trainneg
cnf(i_0_146746, plain, (k3_xcmplx_0(np__0,esk9_0)!=np__1|k3_xcmplx_0(np__0,esk9_0)!=np__0|~m1_subset_1(k4_xcmplx_0(esk9_0),k5_numbers))).#trainneg
cnf(i_0_149777, negated_conjecture, (X1=X2|X1!=k3_xcmplx_0(X2,np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_149972, negated_conjecture, (X1=X2|X1!=k2_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_155172, negated_conjecture, (k4_xcmplx_0(X1)=X2|X1!=k4_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_153493, negated_conjecture, (X1=X2|np__0!=X2|~v1_xboole_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_146658, plain, (v2_funct_1(esk13_0)|k4_xcmplx_0(esk5_1(esk13_0))!=np__1|k4_xcmplx_0(esk5_1(esk13_0))!=np__0)).#trainneg
cnf(i_0_153206, negated_conjecture, (X1=X2|k3_xcmplx_0(X1,np__1)!=k3_xcmplx_0(X2,np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_153224, negated_conjecture, (X1=X2|k3_xcmplx_0(X1,k4_xcmplx_0(np__1))!=k4_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_155169, negated_conjecture, (k4_xcmplx_0(X1)=X2|k2_xcmplx_0(X1,np__0)!=k4_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_146659, plain, (v2_funct_1(esk24_0)|k4_xcmplx_0(esk4_1(esk24_0))!=np__1|k4_xcmplx_0(esk4_1(esk24_0))!=np__0)).#trainneg
cnf(i_0_145606, negated_conjecture, (X1=np__0|k4_xcmplx_0(X1)!=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_146788, plain, (v2_xxreal_0(esk22_0)|k6_xcmplx_0(np__1,esk22_0)!=np__0|~v1_xreal_0(k6_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_146660, plain, (v2_funct_1(esk13_0)|k4_xcmplx_0(esk4_1(esk13_0))!=np__1|k4_xcmplx_0(esk4_1(esk13_0))!=np__0)).#trainneg
cnf(i_0_129630, plain, (v2_funct_1(esk24_0)|k2_xcmplx_0(esk5_1(esk24_0),np__1)!=np__0|~v1_xreal_0(k2_xcmplx_0(esk5_1(esk24_0),np__1)))).#trainneg
cnf(i_0_117933, plain, (k4_xcmplx_0(X1)=X2|X3=np__0|k3_xcmplx_0(X2,X3)!=np__0|k3_xcmplx_0(X1,X3)!=np__0|~v1_xreal_0(k3_xcmplx_0(X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_116655, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_124453, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_153417, negated_conjecture, (X1=X2|X1!=k3_xcmplx_0(np__1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40616, plain, (v3_xxreal_0(X2)|v2_xxreal_0(X1)|v1_xboole_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_140016, plain, (X1=np__0|k2_xcmplx_0(X1,np__0)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_129631, plain, (v2_funct_1(esk13_0)|k2_xcmplx_0(esk5_1(esk13_0),np__1)!=np__0|~v1_xreal_0(k2_xcmplx_0(esk5_1(esk13_0),np__1)))).#trainneg
cnf(i_0_129634, plain, (v2_funct_1(esk24_0)|k2_xcmplx_0(esk4_1(esk24_0),np__1)!=np__0|~v1_xreal_0(k2_xcmplx_0(esk4_1(esk24_0),np__1)))).#trainneg
cnf(i_0_129635, plain, (v2_funct_1(esk13_0)|k2_xcmplx_0(esk4_1(esk13_0),np__1)!=np__0|~v1_xreal_0(k2_xcmplx_0(esk4_1(esk13_0),np__1)))).#trainneg
cnf(i_0_129640, plain, (v2_funct_1(esk24_0)|k2_xcmplx_0(np__1,esk5_1(esk24_0))!=np__0|~v1_xreal_0(k2_xcmplx_0(np__1,esk5_1(esk24_0))))).#trainneg
cnf(i_0_129641, plain, (v2_funct_1(esk13_0)|k2_xcmplx_0(np__1,esk5_1(esk13_0))!=np__0|~v1_xreal_0(k2_xcmplx_0(np__1,esk5_1(esk13_0))))).#trainneg
cnf(i_0_116668, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X1)))=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_156387, negated_conjecture, (X1=X2|k4_xcmplx_0(X1)!=k4_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_157569, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158043, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_157358, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X1)))=np__0|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_157562, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_158130, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_156633, plain, (v2_xxreal_0(esk22_0)|k6_xcmplx_0(np__1,esk22_0)!=np__0)).#trainneg
cnf(i_0_149929, negated_conjecture, (np__1=X1|k2_xcmplx_0(X1,np__0)!=np__1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_156377, negated_conjecture, (X1=X2|k3_xcmplx_0(np__1,k4_xcmplx_0(X1))!=k4_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_157571, negated_conjecture, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_158670, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),k6_xcmplx_0(X1,np__0)))=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_160248, negated_conjecture, (k6_xcmplx_0(X1,X1)=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_160498, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_158906, negated_conjecture, (X1=X2|k4_xcmplx_0(X2)!=np__0|X1!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_159349, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_159205, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_163103, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k6_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_162506, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_162526, negated_conjecture, (esk2_0=X1|k4_xcmplx_0(X1)!=np__0|esk2_0!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_158342, plain, (v2_funct_1(esk24_0)|k2_xcmplx_0(np__1,esk4_1(esk24_0))!=np__0|~v1_xcmplx_0(esk4_1(esk24_0)))).#trainneg
cnf(i_0_159350, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X2),X1))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_159370, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_159917, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_164176, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_165763, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_164405, negated_conjecture, (esk2_0=X1|esk2_0!=np__0|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_164329, negated_conjecture, (esk22_0=esk2_0|v2_xxreal_0(esk22_0)|esk2_0!=np__0)).#trainneg
cnf(i_0_164336, negated_conjecture, (k4_xcmplx_0(esk22_0)=esk2_0|v2_xxreal_0(esk22_0)|esk2_0!=np__0)).#trainneg
cnf(i_0_158348, plain, (v2_funct_1(esk13_0)|k2_xcmplx_0(np__1,esk4_1(esk13_0))!=np__0|~v1_xcmplx_0(esk4_1(esk13_0)))).#trainneg
cnf(i_0_163985, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_163900, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_163793, negated_conjecture, (X1=k6_xcmplx_0(X2,X3)|k4_xcmplx_0(X1)!=k6_xcmplx_0(X3,X2)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_163855, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k4_xcmplx_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_168453, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_168292, negated_conjecture, (X1=k6_xcmplx_0(k4_xcmplx_0(X2),X3)|k4_xcmplx_0(X1)!=k2_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_164738, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_158116, plain, (k3_xcmplx_0(X1,X2)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_165849, negated_conjecture, (k3_xcmplx_0(np__1,k2_xcmplx_0(np__1,np__1))!=np__0)).#trainneg
cnf(i_0_167494, negated_conjecture, (v2_xxreal_0(esk22_0)|esk2_0!=np__0|~v3_xxreal_0(esk2_0))).#trainneg
cnf(i_0_167495, negated_conjecture, (v2_xxreal_0(esk22_0)|esk2_0!=np__0|~v2_xxreal_0(esk2_0))).#trainneg
cnf(i_0_164993, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(X2),X1))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_165014, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_168454, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k2_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_168291, negated_conjecture, (X1=k6_xcmplx_0(k4_xcmplx_0(X2),X3)|k4_xcmplx_0(X1)!=k2_xcmplx_0(X3,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_169811, plain, (k3_xcmplx_0(X1,X2)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_169828, negated_conjecture, (esk22_0!=np__0|esk2_0!=np__0|~v3_xxreal_0(esk2_0))).#trainneg
cnf(i_0_169831, negated_conjecture, (esk22_0!=np__0|esk2_0!=np__0|~v2_xxreal_0(esk2_0))).#trainneg
cnf(i_0_169827, negated_conjecture, (esk2_0!=np__0|~v3_xxreal_0(esk2_0)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_169830, negated_conjecture, (esk2_0!=np__0|~v2_xxreal_0(esk2_0)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_163685, negated_conjecture, (k6_xcmplx_0(X1,X2)=np__0|k6_xcmplx_0(X2,X1)!=np__0|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158770, negated_conjecture, (X1=k4_xcmplx_0(X2)|k4_xcmplx_0(X1)!=X2|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_161662, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(X1,X1))=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_165759, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X2),X1))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_160329, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_167486, negated_conjecture, (k3_xcmplx_0(X1,X2)=esk2_0|v2_xxreal_0(X1)|esk2_0!=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_162507, negated_conjecture, (esk22_0=X1|k4_xcmplx_0(X1)!=np__0|esk22_0!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_167476, negated_conjecture, (k9_xtuple_0(X1)=esk2_0|esk2_0!=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_167498, negated_conjecture, (v2_xxreal_0(esk22_0)|v1_xcmplx_0(esk2_0)|esk2_0!=np__0)).#trainneg
cnf(i_0_163899, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_164680, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|k4_xcmplx_0(X1)!=k6_xcmplx_0(X2,X3)|~v1_xreal_0(k2_xcmplx_0(k4_xcmplx_0(X2),X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_162056, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(k6_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_173101, negated_conjecture, (k4_xcmplx_0(X1)=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k6_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_173064, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X3))|k4_xcmplx_0(X1)!=k2_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_173085, negated_conjecture, (k6_xcmplx_0(X1,X2)=k2_xcmplx_0(k4_xcmplx_0(X3),X4)|k6_xcmplx_0(X2,X1)!=k6_xcmplx_0(X3,X4)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_165636, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_172645, negated_conjecture, (k3_xcmplx_0(X1,X2)=esk2_0|v2_xxreal_0(X1)|esk2_0!=np__0|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_172787, negated_conjecture, (v1_xcmplx_0(esk2_0)|esk22_0!=np__0|esk2_0!=np__0)).#trainneg
cnf(i_0_172786, negated_conjecture, (v1_xcmplx_0(esk2_0)|esk2_0!=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_167552, negated_conjecture, (v2_xxreal_0(esk22_0)|v1_xboole_0(esk2_0)|esk2_0!=np__0)).#trainneg
cnf(i_0_171299, negated_conjecture, (v1_xboole_0(esk2_0)|esk2_0!=np__0|~v1_xboole_0(esk22_0)|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_171296, negated_conjecture, (v1_xboole_0(esk2_0)|esk22_0!=np__0|esk2_0!=np__0|~m1_subset_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_146663, plain, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(k4_xcmplx_0(esk22_0))!=np__1|k4_xcmplx_0(k4_xcmplx_0(esk22_0))!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_170986, negated_conjecture, (X1=k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X3))|k4_xcmplx_0(X1)!=k6_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_159931, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_163854, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k4_xcmplx_0(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_173094, negated_conjecture, (k4_xcmplx_0(X1)=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k6_xcmplx_0(X2,X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158852, negated_conjecture, (k4_xcmplx_0(X1)=X2|k6_xcmplx_0(X1,np__0)!=k4_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_159331, negated_conjecture, (k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_162522, negated_conjecture, (esk4_1(k1_fcont_1(X1,X2))=X3|v2_funct_1(k1_fcont_1(X1,X2))|esk4_1(k1_fcont_1(X1,X2))!=np__0|k4_xcmplx_0(X3)!=np__0|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_163791, negated_conjecture, (k4_xcmplx_0(X1)=k6_xcmplx_0(X2,X3)|X1!=k6_xcmplx_0(X3,X2)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_162501, plain, (X1=k3_xcmplx_0(np__0,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_174331, negated_conjecture, (v1_xboole_0(esk2_0)|esk2_0!=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_174332, negated_conjecture, (v1_xboole_0(esk2_0)|esk22_0!=np__0|esk2_0!=np__0)).#trainneg
cnf(i_0_174336, negated_conjecture, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(esk2_0)!=np__1|esk2_0!=np__0)).#trainneg
cnf(i_0_174341, plain, (v2_xxreal_0(esk22_0)|k6_xcmplx_0(esk22_0,np__0)!=np__1)).#trainneg
cnf(i_0_174343, plain, (v2_xxreal_0(esk22_0)|k3_xcmplx_0(np__1,esk22_0)!=np__1)).#trainneg
cnf(i_0_174344, plain, (v2_xxreal_0(esk22_0)|k3_xcmplx_0(esk22_0,np__1)!=np__1)).#trainneg
cnf(i_0_176165, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_175746, negated_conjecture, (k4_xcmplx_0(X1)=k6_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k2_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_174088, negated_conjecture, (k6_xcmplx_0(np__0,X1)=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k6_xcmplx_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_175744, negated_conjecture, (k4_xcmplx_0(X1)=k6_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k2_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_174451, negated_conjecture, (k4_xcmplx_0(X1)=k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X3))|X1!=k6_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_175745, negated_conjecture, (k4_xcmplx_0(X1)=k6_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k2_xcmplx_0(X3,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_174744, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k4_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_160483, plain, (X1=k3_xcmplx_0(X2,esk22_0)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_172729, negated_conjecture, (esk22_0=X1|esk22_0!=np__0|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_167497, negated_conjecture, (k4_xcmplx_0(esk2_0)=esk22_0|v2_xxreal_0(esk22_0)|esk2_0!=np__0)).#trainneg
cnf(i_0_167833, negated_conjecture, (k4_xcmplx_0(esk2_0)=np__0|v2_xxreal_0(esk22_0)|esk2_0!=np__0)).#trainneg
cnf(i_0_167540, negated_conjecture, (v2_xxreal_0(esk22_0)|m1_subset_1(esk2_0,k1_numbers)|esk2_0!=np__0)).#trainneg
cnf(i_0_167550, negated_conjecture, (v2_xxreal_0(esk22_0)|v1_xreal_0(k4_xcmplx_0(esk2_0))|esk2_0!=np__0)).#trainneg
cnf(i_0_119564, plain, (v2_xxreal_0(esk22_0)|m1_subset_1(k4_xcmplx_0(esk22_0),k1_numbers))).#trainneg
cnf(i_0_120276, plain, (v2_xxreal_0(esk22_0)|v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_173872, negated_conjecture, (k4_xcmplx_0(X1)=k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_173847, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X3))|k4_xcmplx_0(X1)!=k2_xcmplx_0(X3,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_174075, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k2_xcmplx_0(k4_xcmplx_0(X3),X4)|k2_xcmplx_0(X1,X2)!=k6_xcmplx_0(X3,X4)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_171628, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X1),np__0)=np__0|~v1_xreal_0(k6_xcmplx_0(X1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_164112, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X2),X1))=k2_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_175760, negated_conjecture, (k4_xcmplx_0(X1)=k6_xcmplx_0(X2,X2)|X1!=np__0|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_160484, plain, (X1=k3_xcmplx_0(X2,esk9_0)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_180987, negated_conjecture, (k4_xcmplx_0(X1)!=np__1|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_179287, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|esk2_0!=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_179288, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|esk22_0!=np__0|esk2_0!=np__0)).#trainneg
cnf(i_0_179303, plain, (v2_xxreal_0(esk22_0)|v1_funct_1(k4_xcmplx_0(esk22_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_179304, plain, (v2_xxreal_0(esk22_0)|v1_relat_1(k4_xcmplx_0(esk22_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_176521, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_179342, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_176283, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_160490, plain, (X1=k3_xcmplx_0(X2,np__1)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_178599, negated_conjecture, (k9_xtuple_0(X1)=esk22_0|esk22_0!=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_178578, plain, (X1=k3_xcmplx_0(np__0,esk22_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_179306, plain, (v2_xxreal_0(esk22_0)|v1_xxreal_0(k4_xcmplx_0(esk22_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_179308, plain, (v3_ordinal1(k4_xcmplx_0(esk22_0))|v2_xxreal_0(esk22_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_179309, plain, (v1_rat_1(k4_xcmplx_0(esk22_0))|v2_xxreal_0(esk22_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_179310, plain, (v1_int_1(k4_xcmplx_0(esk22_0))|v2_xxreal_0(esk22_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_176282, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_176281, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_163932, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(X1))=k2_xcmplx_0(np__0,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_160493, plain, (X1=k3_xcmplx_0(X2,np__0)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_162510, negated_conjecture, (esk14_0=X1|k4_xcmplx_0(X1)!=np__0|esk14_0!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2467, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_171289, plain, (k3_xcmplx_0(np__0,X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_155098, negated_conjecture, (k4_xcmplx_0(np__1)=X1|k4_xcmplx_0(X1)!=np__1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_181714, negated_conjecture, (v1_funct_1(esk2_0)|esk2_0!=np__0|~v4_funct_1(k1_numbers)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_181715, negated_conjecture, (v1_relat_1(esk2_0)|esk2_0!=np__0|~v4_funct_1(k1_numbers)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_181717, negated_conjecture, (v1_xxreal_0(esk2_0)|esk2_0!=np__0|~v2_membered(k1_numbers)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_181719, negated_conjecture, (v3_ordinal1(esk2_0)|esk2_0!=np__0|~v3_ordinal1(k1_numbers)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_181720, negated_conjecture, (v1_rat_1(esk2_0)|esk2_0!=np__0|~v4_membered(k1_numbers)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_183782, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_186225, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_186118, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_182726, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_178583, plain, (X1=k3_xcmplx_0(X2,esk22_0)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_185132, plain, (k3_xcmplx_0(np__0,esk22_0)=np__0)).#trainneg
cnf(i_0_187109, negated_conjecture, (esk22_0=esk14_0|v2_xxreal_0(esk22_0)|esk14_0!=np__0)).#trainneg
cnf(i_0_181725, negated_conjecture, (v1_funct_1(esk2_0)|esk22_0!=np__0|esk2_0!=np__0|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_181721, negated_conjecture, (v1_int_1(esk2_0)|esk2_0!=np__0|~v5_membered(k1_numbers)|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_181726, negated_conjecture, (v1_relat_1(esk2_0)|esk22_0!=np__0|esk2_0!=np__0|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_181728, negated_conjecture, (v1_xxreal_0(esk2_0)|esk22_0!=np__0|esk2_0!=np__0|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_181731, negated_conjecture, (v1_rat_1(esk2_0)|esk22_0!=np__0|esk2_0!=np__0|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_181732, negated_conjecture, (v1_int_1(esk2_0)|esk22_0!=np__0|esk2_0!=np__0|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_187000, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(np__0,X1),X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_183783, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_182003, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_184053, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_181622, plain, (X1=k3_xcmplx_0(X2,esk9_0)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_187198, negated_conjecture, (esk14_0=X1|esk14_0!=np__0|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2474, plain, (v2_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k4_xcmplx_0(X2))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_181617, plain, (X1=k3_xcmplx_0(np__0,esk9_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_129750, plain, (v2_xxreal_0(esk22_0)|v1_xboole_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_129764, plain, (v2_xxreal_0(esk22_0)|v1_xboole_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_179334, plain, (v2_xxreal_0(esk22_0)|v1_xreal_0(k6_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_186999, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(np__0,X1),X2)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_183652, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_172627, negated_conjecture, (k6_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_194657, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_193133, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_185061, plain, (X1=k3_xcmplx_0(X2,np__1)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2475, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_190540, plain, (esk22_0=np__0|k4_xcmplx_0(esk22_0)!=np__0)).#trainneg
cnf(i_0_188017, negated_conjecture, (v1_xxreal_0(esk2_0)|esk2_0!=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_190625, negated_conjecture, (v1_xxreal_0(esk2_0)|esk22_0!=np__0|esk2_0!=np__0)).#trainneg
cnf(i_0_190502, plain, (esk22_0=np__0|k3_xcmplx_0(np__1,esk22_0)!=np__0)).#trainneg
cnf(i_0_190548, plain, (esk22_0=np__0|k2_xcmplx_0(esk22_0,np__0)!=np__0)).#trainneg
cnf(i_0_189833, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_188037, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_187108, plain, (X1=k3_xcmplx_0(X2,np__0)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2476, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v3_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_190544, plain, (k3_xcmplx_0(esk22_0,np__0)=np__0)).#trainneg
cnf(i_0_190615, plain, (k6_xcmplx_0(esk22_0,esk22_0)=np__0)).#trainneg
cnf(i_0_194117, plain, (k3_xcmplx_0(np__0,esk9_0)=np__0)).#trainneg
cnf(i_0_179336, plain, (v2_xxreal_0(esk22_0)|v1_xreal_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_179337, plain, (v2_xxreal_0(esk22_0)|v1_xreal_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_194850, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,np__0),X2)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_188801, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158118, plain, (k3_xcmplx_0(X1,X2)=k9_xtuple_0(X3)|v2_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2477, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v3_xxreal_0(k4_xcmplx_0(X2))|~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_193814, negated_conjecture, (k9_xtuple_0(X1)=esk14_0|esk14_0!=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_190560, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))=np__0)).#trainneg
cnf(i_0_190480, plain, (k3_xcmplx_0(X1,esk22_0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_188096, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3)))=k6_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_195552, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_201983, plain, (k3_xcmplx_0(X1,X2)=k9_xtuple_0(X3)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2478, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k4_xcmplx_0(X2))|~v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_199360, plain, (~v2_xxreal_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_199418, plain, (k3_xcmplx_0(np__1,esk9_0)!=np__0)).#trainneg
cnf(i_0_199355, plain, (v1_xboole_0(esk22_0)|~v7_ordinal1(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_199389, plain, (v3_xxreal_0(k4_xcmplx_0(esk22_0))|~v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_199388, plain, (v1_xboole_0(k4_xcmplx_0(esk22_0))|~v7_ordinal1(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_197120, plain, (esk22_0=np__0|k2_xcmplx_0(np__0,esk22_0)!=np__0)).#trainneg
cnf(i_0_195816, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_200616, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_198748, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_195551, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3)))).#trainneg
cnf(i_0_200863, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_156552, negated_conjecture, (X1=X2|X1!=np__0|~v1_xboole_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2479, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(X2)|~v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_199262, plain, (esk22_0=np__0|k3_xcmplx_0(esk22_0,np__1)!=np__0)).#trainneg
cnf(i_0_196844, plain, (v2_xxreal_0(esk22_0)|~v2_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_198931, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_141511, negated_conjecture, (k3_xcmplx_0(np__0,X2)=k3_xcmplx_0(X2,X1)|X1!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_124031, plain, (np__0=X1|X2=np__0|k3_xcmplx_0(X2,np__0)!=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_142580, negated_conjecture, (X1=np__0|X2=X3|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X3,X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_154367, negated_conjecture, (X1=X2|k3_xcmplx_0(X1,np__1)!=k4_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158121, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(X3)|v2_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(X3)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_2480, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(np__1))=k6_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_204945, plain, (esk22_0!=np__0|~v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_204939, plain, (k4_xcmplx_0(esk22_0)!=np__0|~v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_204938, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_204937, plain, (~v2_xxreal_0(esk22_0)|~m1_subset_1(k4_xcmplx_0(esk22_0),k5_numbers))).#trainneg
cnf(i_0_204934, plain, (v1_xboole_0(esk22_0)|~m1_subset_1(k4_xcmplx_0(esk22_0),k5_numbers))).#trainneg
cnf(i_0_210242, plain, (v2_xxreal_0(esk22_0)|~v2_xxreal_0(k6_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_210244, plain, (v2_xxreal_0(esk22_0)|~v2_xxreal_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_204087, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_211261, negated_conjecture, (X1=np__0|k3_xcmplx_0(X1,X2)!=X1|X2!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_208798, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_208797, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_211361, negated_conjecture, (k3_xcmplx_0(X1,X2)=np__0|X2!=np__0|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_209892, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,np__0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_160495, plain, (X1=k3_xcmplx_0(X2,k2_xcmplx_0(X3,X4))|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_199460, plain, (k3_xcmplx_0(esk9_0,np__0)=np__0)).#trainneg
cnf(i_0_210245, plain, (v2_xxreal_0(esk22_0)|~v2_xxreal_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_210265, plain, (v2_xxreal_0(esk22_0)|~v3_xxreal_0(k6_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_210236, negated_conjecture, (v2_xxreal_0(esk22_0)|esk2_0!=np__0|~v2_xxreal_0(k4_xcmplx_0(esk2_0)))).#trainneg
cnf(i_0_212009, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=k4_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_160496, plain, (X1=k3_xcmplx_0(X2,k3_xcmplx_0(X3,X4))|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_164347, negated_conjecture, (k4_xcmplx_0(X1)=esk2_0|esk2_0!=np__0|X1!=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2504, plain, (v1_xreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_199531, plain, (k6_xcmplx_0(esk9_0,esk9_0)=np__0)).#trainneg
cnf(i_0_211165, negated_conjecture, (k3_xcmplx_0(np__1,X1)!=np__1|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_204921, plain, (v1_xboole_0(k4_xcmplx_0(esk22_0))|~m1_subset_1(k4_xcmplx_0(esk22_0),k5_numbers))).#trainneg
cnf(i_0_202520, plain, (k4_xcmplx_0(esk22_0)=np__0|k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__0)!=np__0)).#trainneg
cnf(i_0_199268, plain, (k3_xcmplx_0(k4_xcmplx_0(esk22_0),np__0)=np__0)).#trainneg
cnf(i_0_212704, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_214666, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_217349, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_217017, negated_conjecture, (X1=X2|X1!=k4_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_217014, negated_conjecture, (X1=X2|X1!=k4_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_212723, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(X3)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_2520, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v3_xxreal_0(X2)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_218499, plain, (~v7_ordinal1(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_218504, plain, (~v2_xxreal_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_216711, plain, (k3_xcmplx_0(esk9_0,np__1)!=np__0)).#trainneg
cnf(i_0_218535, plain, (v3_xxreal_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_213418, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_216601, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,np__0)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_212340, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_213452, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_214964, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_158119, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(X3)|v3_xxreal_0(X3)|v2_xxreal_0(X3)|v2_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2522, plain, (v3_xxreal_0(X2)|~v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_222980, plain, (~m1_subset_1(k4_xcmplx_0(esk9_0),k5_numbers))).#trainneg
cnf(i_0_218651, plain, (k4_xcmplx_0(esk22_0)=np__0|k6_xcmplx_0(np__0,esk22_0)!=np__0)).#trainneg
cnf(i_0_223316, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_228433, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_222628, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_222646, negated_conjecture, (X1=k6_xcmplx_0(X2,X2)|X1!=np__0|~v1_xreal_0(k6_xcmplx_0(X2,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158122, plain, (k3_xcmplx_0(X1,X2)=k1_relset_1(X3,X4)|v2_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v4_relat_1(X4,X3)|~v1_xboole_0(X3)|~v1_relat_1(X4))).#trainneg
cnf(i_0_2523, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_218814, plain, (k4_xcmplx_0(esk22_0)=np__0|k4_xcmplx_0(k4_xcmplx_0(esk22_0))!=np__0)).#trainneg
cnf(i_0_218642, negated_conjecture, (k4_xcmplx_0(esk22_0)=np__0|k2_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))!=np__0)).#trainneg
cnf(i_0_218655, plain, (k4_xcmplx_0(esk22_0)=np__0|k4_xcmplx_0(k2_xcmplx_0(esk22_0,np__0))!=np__0)).#trainneg
cnf(i_0_221183, negated_conjecture, (k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X1)))=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_216006, negated_conjecture, (k3_xcmplx_0(X1,X2)=np__0|X1!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_213964, negated_conjecture, (X1=np__0|k3_xcmplx_0(X2,X1)!=X1|X2!=np__0|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_216944, negated_conjecture, (X1=X2|k2_xcmplx_0(X1,np__0)!=k6_xcmplx_0(X2,np__0)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_222658, negated_conjecture, (X1=k4_xcmplx_0(X2)|X1!=np__0|X2!=np__0|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_169796, plain, (k3_xcmplx_0(X1,esk22_0)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_181252, negated_conjecture, (v1_xreal_0(k6_xcmplx_0(X2,X2))|X1!=np__0|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2540, plain, (X1=k1_fcont_1(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|k1_seq_1(X1,esk3_3(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)),X1))!=k6_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)),X1)),X3),X4)|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)),X1)),X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)),X1)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_213839, negated_conjecture, (k3_xcmplx_0(X1,np__1)!=np__1|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_199370, plain, (k2_xcmplx_0(k4_xcmplx_0(esk22_0),esk22_0)=np__0)).#trainneg
cnf(i_0_199371, plain, (k2_xcmplx_0(esk22_0,k4_xcmplx_0(esk22_0))=np__0)).#trainneg
cnf(i_0_237139, plain, (~v3_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_234245, plain, (k4_xcmplx_0(esk22_0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_234890, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_234704, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_230412, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_232564, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_230431, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(X3)|v3_xxreal_0(X3)|v2_xxreal_0(X1)|v2_xxreal_0(X3)|~v7_ordinal1(X1)|~v1_xboole_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2559, plain, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_234246, plain, (k4_xcmplx_0(esk22_0)=np__0|k2_xcmplx_0(esk22_0,np__0)!=np__0)).#trainneg
cnf(i_0_234282, negated_conjecture, (k4_xcmplx_0(esk22_0)=np__0|~v1_xboole_0(k2_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_234247, plain, (k4_xcmplx_0(esk22_0)=np__0|k6_xcmplx_0(esk22_0,np__0)!=np__0)).#trainneg
cnf(i_0_234248, plain, (k4_xcmplx_0(esk22_0)=np__0|k2_xcmplx_0(np__0,esk22_0)!=np__0)).#trainneg
cnf(i_0_234249, plain, (k4_xcmplx_0(esk22_0)=np__0|k3_xcmplx_0(np__1,esk22_0)!=np__0)).#trainneg
cnf(i_0_234250, plain, (k4_xcmplx_0(esk22_0)=np__0|k3_xcmplx_0(esk22_0,np__1)!=np__0)).#trainneg
cnf(i_0_233904, negated_conjecture, (X1=k6_xcmplx_0(X2,X2)|X1!=np__0|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_231502, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_233922, plain, (k3_xcmplx_0(X1,X2)=k1_relset_1(X3,X4)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v4_relat_1(X4,X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_relat_1(X4))).#trainneg
cnf(i_0_2561, plain, (k6_xcmplx_0(k2_xcmplx_0(X1,np__1),np__1)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_234266, negated_conjecture, (k4_xcmplx_0(esk22_0)=np__0|k4_xcmplx_0(k2_xcmplx_0(np__0,esk22_0))!=np__0)).#trainneg
cnf(i_0_232258, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_214299, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_169797, plain, (k3_xcmplx_0(X1,esk9_0)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_158331, negated_conjecture, (X1=np__0|k2_xcmplx_0(np__0,X1)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2562, plain, (k6_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(k4_xcmplx_0(X3),X2))|~v1_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_237112, plain, (k2_xcmplx_0(np__0,esk22_0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_237291, negated_conjecture, (k2_xcmplx_0(esk22_0,np__0)=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_237290, plain, (k2_xcmplx_0(esk22_0,np__0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_213440, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_227825, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_244819, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_237731, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_238289, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,k2_xcmplx_0(X1,X2)),X2)=np__0|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_244813, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_169803, plain, (k3_xcmplx_0(X1,np__1)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2568, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X3),X2))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_245872, plain, (k2_xcmplx_0(np__0,esk22_0)!=np__1|esk22_0!=np__0)).#trainneg
cnf(i_0_237305, plain, (~v3_xxreal_0(k2_xcmplx_0(esk22_0,np__0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_237307, plain, (~v3_xxreal_0(k2_xcmplx_0(np__0,esk22_0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_237306, plain, (~v3_xxreal_0(k6_xcmplx_0(esk22_0,np__0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k6_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_237308, plain, (~v3_xxreal_0(k3_xcmplx_0(np__1,esk22_0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_224660, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,np__0)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_216599, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_215000, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_169806, plain, (k3_xcmplx_0(X1,np__0)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_237309, plain, (~v3_xxreal_0(k3_xcmplx_0(esk22_0,np__1))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_240742, negated_conjecture, (k4_xcmplx_0(esk22_0)=np__0|~v1_xboole_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_217489, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_238620, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_235110, negated_conjecture, (k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0)))=np__0|v2_funct_1(k1_fcont_1(X1,np__0))|X1!=np__0|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,np__0))))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,np__0)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_172630, negated_conjecture, (k3_xcmplx_0(X1,esk22_0)=esk2_0|v2_xxreal_0(X1)|esk2_0!=np__0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2573, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),k2_xcmplx_0(k4_xcmplx_0(X3),X4))=k2_xcmplx_0(X1,k2_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(X2),X3),X4))|~v1_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X3),X4))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_234731, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_249565, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_251175, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_249991, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_250371, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_251337, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_172631, negated_conjecture, (k3_xcmplx_0(X1,esk9_0)=esk2_0|v2_xxreal_0(X1)|esk2_0!=np__0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2577, plain, (X1=k1_fcont_1(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4))|k1_seq_1(X1,esk3_3(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4),X1))!=k2_xcmplx_0(k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4),X1)),X3),X4)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4),X1)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k2_xcmplx_0(k4_xcmplx_0(X3),X4))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_258091, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_252420, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X3,X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_256793, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_172637, negated_conjecture, (k3_xcmplx_0(X1,np__1)=esk2_0|v2_xxreal_0(X1)|esk2_0!=np__0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2578, plain, (X1=k1_fcont_1(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4))|k1_funct_1(X1,esk3_3(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4),X1))!=k2_xcmplx_0(k6_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4),X1)),X3),X4)|~v3_valued_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X4),X1)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k2_xcmplx_0(k4_xcmplx_0(X3),X4))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_245439, plain, (k6_xcmplx_0(np__0,esk22_0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_245655, negated_conjecture, (k6_xcmplx_0(np__0,esk22_0)=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_237137, plain, (v2_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0)))|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_224509, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_261282, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X3,X1),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_267142, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_172640, negated_conjecture, (k3_xcmplx_0(X1,np__0)=esk2_0|v2_xxreal_0(X1)|esk2_0!=np__0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2579, plain, (X1=k1_fcont_1(k3_xcmplx_0(X2,X3),k2_xcmplx_0(k4_xcmplx_0(X4),X5))|k1_seq_1(X1,esk3_3(k3_xcmplx_0(X2,X3),k2_xcmplx_0(k4_xcmplx_0(X4),X5),X1))!=k2_xcmplx_0(k6_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,esk3_3(k3_xcmplx_0(X2,X3),k2_xcmplx_0(k4_xcmplx_0(X4),X5),X1))),X4),X5)|~v1_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,esk3_3(k3_xcmplx_0(X2,X3),k2_xcmplx_0(k4_xcmplx_0(X4),X5),X1))))|~v1_xcmplx_0(esk3_3(k3_xcmplx_0(X2,X3),k2_xcmplx_0(k4_xcmplx_0(X4),X5),X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k2_xcmplx_0(k4_xcmplx_0(X4),X5))|~v1_xreal_0(k3_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_269736, plain, (esk22_0!=np__0|~v3_xxreal_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_269816, negated_conjecture, (~v3_xxreal_0(k4_xcmplx_0(esk22_0))|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_269827, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_269828, plain, (k4_xcmplx_0(k4_xcmplx_0(esk22_0))!=np__0|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_269847, plain, (v2_xxreal_0(k2_xcmplx_0(esk22_0,np__0))|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_269849, plain, (v2_xxreal_0(k2_xcmplx_0(np__0,esk22_0))|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_158124, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(esk22_0)|v2_xxreal_0(esk22_0)|v2_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_160544, negated_conjecture, (np__1=X1|k2_xcmplx_0(np__0,X1)!=np__1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2581, plain, (k2_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X2),X3),X1)=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X2),X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_269694, plain, (k6_xcmplx_0(esk22_0,np__0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_269774, negated_conjecture, (k6_xcmplx_0(esk22_0,np__0)=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_267960, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_256799, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_250709, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_158833, negated_conjecture, (X1=X2|k4_xcmplx_0(X1)!=np__0|~v1_xboole_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_213857, negated_conjecture, (k4_xcmplx_0(X1)!=k4_xcmplx_0(np__1)|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_269848, plain, (v2_xxreal_0(k6_xcmplx_0(esk22_0,np__0))|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k6_xcmplx_0(esk22_0,np__0)))).#trainneg
cnf(i_0_269850, plain, (v2_xxreal_0(k3_xcmplx_0(np__1,esk22_0))|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_269851, plain, (v2_xxreal_0(k3_xcmplx_0(esk22_0,np__1))|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0))|~v1_xreal_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_209891, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_255075, negated_conjecture, (k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X1)))=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X1)|~v1_xcmplx_0(k2_xcmplx_0(X3,k4_xcmplx_0(X1)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_253779, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,np__0)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_270180, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_258016, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(k2_xcmplx_0(X2,X3)))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_280043, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X3))|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_276265, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(esk22_0)|v2_xxreal_0(esk22_0)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_279498, negated_conjecture, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(np__1)!=esk22_0|esk2_0!=np__0)).#trainneg
cnf(i_0_2595, plain, (k6_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_276152, plain, (v1_xboole_0(k4_xcmplx_0(esk22_0))|esk22_0!=np__0)).#trainneg
cnf(i_0_277229, negated_conjecture, (k3_xcmplx_0(np__1,esk22_0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_279496, negated_conjecture, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(np__1)!=esk2_0|esk22_0!=np__0|esk2_0!=np__0)).#trainneg
cnf(i_0_276176, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(esk22_0,np__0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_276177, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k6_xcmplx_0(esk22_0,np__0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_276178, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk22_0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_276179, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k3_xcmplx_0(np__1,esk22_0))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_276180, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k3_xcmplx_0(esk22_0,np__1))|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_274279, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_280240, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(np__0,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_275067, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_281182, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_272319, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_280927, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2597, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,k4_xcmplx_0(X2)),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_276162, negated_conjecture, (v1_xboole_0(k4_xcmplx_0(esk22_0))|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_199375, plain, (k2_xcmplx_0(esk22_0,np__0)=k2_xcmplx_0(np__0,esk22_0))).#trainneg
cnf(i_0_276200, plain, (k2_xcmplx_0(esk22_0,np__0)!=np__0|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_272939, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_256999, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,X3),k4_xcmplx_0(X1))=k6_xcmplx_0(k6_xcmplx_0(X2,X1),X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_290991, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,k6_xcmplx_0(X1,X2)),X2)=np__0|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_286891, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),np__0)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_213163, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_211227, negated_conjecture, (X1=np__0|X2=np__0|k3_xcmplx_0(X2,X1)!=k3_xcmplx_0(X1,X3)|X3!=np__0|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2598, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,np__0),X2)=k2_xcmplx_0(X1,k2_xcmplx_0(np__0,X2))|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_284157, plain, (k4_xcmplx_0(esk22_0)!=np__1|esk22_0!=np__0)).#trainneg
cnf(i_0_289763, negated_conjecture, (k4_xcmplx_0(esk22_0)!=np__1|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_289808, plain, (~v3_xxreal_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_284434, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_284444, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_217506, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(np__0,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_249899, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_159425, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(X1),np__0)=k6_xcmplx_0(X2,X1)|X2!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_171841, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X1))))=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_162491, plain, (X1=k3_xcmplx_0(esk22_0,X2)|v2_xxreal_0(esk22_0)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_289997, plain, (esk22_0=np__1|k2_xcmplx_0(np__0,esk22_0)!=np__1)).#trainneg
cnf(i_0_2599, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k6_xcmplx_0(k4_xcmplx_0(X2),X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_284168, negated_conjecture, (k4_xcmplx_0(esk22_0)=esk22_0|esk22_0!=np__0)).#trainneg
cnf(i_0_290029, plain, (k2_xcmplx_0(np__0,esk22_0)!=np__0|~v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_284454, plain, (~v2_xxreal_0(esk22_0)|~v1_xboole_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_116661, negated_conjecture, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_292848, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_291729, negated_conjecture, (X1=k2_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_291730, negated_conjecture, (X1=k2_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(X3,X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_292668, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_297437, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_295316, negated_conjecture, (k6_xcmplx_0(X2,X1)=k4_xcmplx_0(X1)|X2!=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_292661, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_291502, negated_conjecture, (k2_xcmplx_0(np__0,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_201968, plain, (k3_xcmplx_0(X1,esk22_0)=k9_xtuple_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2605, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k4_xcmplx_0(k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_296725, negated_conjecture, (k4_xcmplx_0(np__1)!=esk22_0|esk22_0!=np__0)).#trainneg
cnf(i_0_276210, plain, (k4_xcmplx_0(k4_xcmplx_0(esk22_0))!=np__0|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_201969, plain, (k3_xcmplx_0(X1,esk9_0)=k9_xtuple_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_296692, negated_conjecture, (esk22_0=X1|k4_xcmplx_0(X1)!=esk22_0|esk22_0!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2656, plain, (v1_xboole_0(esk5_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k9_xtuple_0(X1),esk5_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_289979, plain, (k2_xcmplx_0(np__0,esk22_0)=esk22_0)).#trainneg
cnf(i_0_284167, negated_conjecture, (k4_xcmplx_0(esk22_0)=esk2_0|esk2_0!=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_284173, negated_conjecture, (k4_xcmplx_0(esk22_0)=esk14_0|esk14_0!=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_293748, plain, (k3_xcmplx_0(np__1,esk22_0)!=np__0|~v2_xxreal_0(esk22_0)|~v1_xreal_0(k3_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_296998, plain, (k3_xcmplx_0(esk22_0,np__1)!=np__0|~v2_xxreal_0(esk22_0)|~v1_xreal_0(k3_xcmplx_0(esk22_0,np__1)))).#trainneg
cnf(i_0_153152, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(np__1,X1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_174444, negated_conjecture, (k4_xcmplx_0(X1)=k6_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X3))|X1!=k6_xcmplx_0(X2,X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_201975, plain, (k3_xcmplx_0(X1,np__1)=k9_xtuple_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_199375, plain, (k2_xcmplx_0(esk22_0,np__0)=esk22_0)).#trainneg
cnf(i_0_303644, negated_conjecture, (esk22_0!=np__0|esk14_0!=np__0|~v2_xxreal_0(esk14_0))).#trainneg
cnf(i_0_303645, negated_conjecture, (esk22_0!=np__0|esk14_0!=np__0|~v3_xxreal_0(esk14_0))).#trainneg
cnf(i_0_303974, plain, (k3_xcmplx_0(np__1,esk22_0)!=np__0|~v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_303995, plain, (k3_xcmplx_0(esk22_0,np__1)!=np__0|~v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_302707, plain, (k6_xcmplx_0(esk22_0,np__0)!=np__0|~v2_xxreal_0(esk22_0)|~v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_299088, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k6_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_298716, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k6_xcmplx_0(X2,X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_298932, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k6_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_299304, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k6_xcmplx_0(X3,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_299087, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k6_xcmplx_0(X2,X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_298715, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k6_xcmplx_0(X2,X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_201978, plain, (k3_xcmplx_0(X1,np__0)=k9_xtuple_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2657, plain, (v1_xboole_0(esk4_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k9_xtuple_0(X1),esk4_1(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_303042, negated_conjecture, (k4_xcmplx_0(esk2_0)=esk22_0|esk2_0!=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_299305, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X2)|X1!=np__0|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_299059, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_298933, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X2)|X1!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_298687, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_298931, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k6_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_299091, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X2))|X1!=np__0|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_300501, negated_conjecture, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_236224, negated_conjecture, (v1_xreal_0(k6_xcmplx_0(X1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2673, plain, (k1_relset_1(k2_zfmisc_1(X1,X2),X3)=np__0|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_302846, plain, (k3_xcmplx_0(esk22_0,np__1)=esk22_0)).#trainneg
cnf(i_0_303504, negated_conjecture, (v1_xcmplx_0(esk14_0)|esk14_0!=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_298719, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X2))|X1!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_299303, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X3)|X1!=k6_xcmplx_0(X3,X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_305335, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_304860, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_299111, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),k2_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(k6_xcmplx_0(X3,X2),X4)|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_209910, negated_conjecture, (X1=k9_xtuple_0(X2)|X1!=np__0|~v1_xboole_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2677, plain, (v1_xreal_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_306087, negated_conjecture, (k6_xcmplx_0(esk22_0,np__0)=esk22_0)).#trainneg
cnf(i_0_306263, negated_conjecture, (v1_xboole_0(esk14_0)|esk22_0!=np__0|esk14_0!=np__0|~m1_subset_1(esk14_0,k5_numbers))).#trainneg
cnf(i_0_309076, plain, (k3_xcmplx_0(np__1,esk22_0)=esk22_0)).#trainneg
cnf(i_0_298664, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_298739, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),k2_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(k6_xcmplx_0(X3,X2),X4)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_298666, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_299057, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_212708, plain, (k3_xcmplx_0(X1,esk22_0)=k4_xcmplx_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2684, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k6_xcmplx_0(X2,X1),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_171557, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X1),np__0)=np__0|~v1_xreal_0(k6_xcmplx_0(X1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_173086, negated_conjecture, (k6_xcmplx_0(X1,X2)=k2_xcmplx_0(k4_xcmplx_0(X3),X4)|k6_xcmplx_0(X2,X1)!=k6_xcmplx_0(X3,X4)|~v1_xcmplx_0(X2)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_310220, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_315746, negated_conjecture, (X1=k6_xcmplx_0(X2,X2)|X1!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X2,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_311138, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_311093, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_310221, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_212709, plain, (k3_xcmplx_0(X1,esk9_0)=k4_xcmplx_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2687, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(k2_xcmplx_0(X1,X2)),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_303503, negated_conjecture, (k4_xcmplx_0(esk14_0)=esk22_0|esk14_0!=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_306262, negated_conjecture, (v3_xxreal_0(esk14_0)|v1_xboole_0(esk14_0)|esk22_0!=np__0|esk14_0!=np__0)).#trainneg
cnf(i_0_313632, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(np__0,X1))=k6_xcmplx_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_317193, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_321910, negated_conjecture, (np__0=X1|X2=np__0|k6_xcmplx_0(X2,X2)!=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_316552, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_317195, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_317190, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_317064, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_212715, plain, (k3_xcmplx_0(X1,np__1)=k4_xcmplx_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2691, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_309179, plain, (k4_xcmplx_0(k4_xcmplx_0(esk22_0))=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_199476, plain, (k3_xcmplx_0(np__0,k4_xcmplx_0(esk9_0))=np__0)).#trainneg
cnf(i_0_316002, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,k4_xcmplx_0(X3)),X2)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_322595, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_322676, negated_conjecture, (X1=np__0|k6_xcmplx_0(X1,X1)!=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_327945, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),k4_xcmplx_0(X2))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_325422, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_326153, negated_conjecture, (k2_xcmplx_0(X1,k4_xcmplx_0(X1))=k3_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_324050, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_212718, plain, (k3_xcmplx_0(X1,np__0)=k4_xcmplx_0(X2)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_327366, plain, (k2_xcmplx_0(k4_xcmplx_0(esk9_0),np__0)!=np__0)).#trainneg
cnf(i_0_327220, plain, (k4_xcmplx_0(k4_xcmplx_0(esk9_0))!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_303014, plain, (k6_xcmplx_0(np__0,esk22_0)=k4_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_327178, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(esk9_0))!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_204920, plain, (v3_xxreal_0(k4_xcmplx_0(esk22_0))|v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_216717, plain, (k3_xcmplx_0(k4_xcmplx_0(esk9_0),np__0)=np__0)).#trainneg
cnf(i_0_322531, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_326231, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=k3_xcmplx_0(np__0,X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_331215, negated_conjecture, (k6_xcmplx_0(X1,X1)=k3_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_328584, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_209907, negated_conjecture, (X1=k4_xcmplx_0(X2)|X1!=np__0|~v1_xboole_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_327037, negated_conjecture, (X1=k4_xcmplx_0(esk22_0)|esk22_0!=np__0|X1!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2706, plain, (k6_xcmplx_0(X3,k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_332051, plain, (k6_xcmplx_0(np__0,esk9_0)!=np__0)).#trainneg
cnf(i_0_332387, plain, (k4_xcmplx_0(k4_xcmplx_0(esk9_0))!=np__0)).#trainneg
cnf(i_0_332042, negated_conjecture, (k2_xcmplx_0(np__0,k4_xcmplx_0(esk9_0))!=np__0)).#trainneg
cnf(i_0_322383, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),np__0)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_325496, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_325497, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X4)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_327730, negated_conjecture, (k2_xcmplx_0(X3,k6_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k6_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_322534, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_322399, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_160497, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2707, plain, (k6_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_336354, plain, (k6_xcmplx_0(esk9_0,np__0)!=np__0)).#trainneg
cnf(i_0_332055, plain, (k4_xcmplx_0(k2_xcmplx_0(esk9_0,np__0))!=np__0)).#trainneg
cnf(i_0_336373, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(np__0,esk9_0))!=np__0)).#trainneg
cnf(i_0_334296, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_335757, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_162505, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_335842, negated_conjecture, (k9_xtuple_0(X1)=k4_xcmplx_0(esk22_0)|k9_xtuple_0(X1)!=np__0|esk22_0!=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2712, plain, (k6_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_332105, plain, (v1_xcmplx_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_332101, plain, (k4_xcmplx_0(k4_xcmplx_0(esk22_0))=esk22_0)).#trainneg
cnf(i_0_332192, plain, (k4_xcmplx_0(k3_xcmplx_0(esk9_0,np__1))!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_332193, plain, (k4_xcmplx_0(k3_xcmplx_0(np__1,esk9_0))!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_334292, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_347402, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_342684, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_342296, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_346742, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(esk22_0),k4_xcmplx_0(esk22_0)))=np__0)).#trainneg
cnf(i_0_342294, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_347374, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X2)|X1!=k6_xcmplx_0(X2,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_347375, negated_conjecture, (X1=k2_xcmplx_0(k4_xcmplx_0(X2),X2)|X1!=k6_xcmplx_0(X2,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_341764, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_346699, negated_conjecture, (esk22_0=np__0|~v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_346750, negated_conjecture, (k4_xcmplx_0(esk22_0)!=np__0|esk22_0!=np__1)).#trainneg
cnf(i_0_346751, negated_conjecture, (k4_xcmplx_0(np__1)!=esk22_0|k4_xcmplx_0(esk22_0)!=np__0)).#trainneg
cnf(i_0_346653, plain, (esk22_0!=np__1|~v1_xboole_0(k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_342295, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_347406, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_352879, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(esk22_0),k4_xcmplx_0(esk22_0))=np__0)).#trainneg
cnf(i_0_352442, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_353717, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_352202, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_353476, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_351270, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_209916, negated_conjecture, (X1=k2_xcmplx_0(X2,X3)|X1!=np__0|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_353977, negated_conjecture, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(np__1)!=esk22_0)).#trainneg
cnf(i_0_353975, negated_conjecture, (k4_xcmplx_0(np__1)!=esk22_0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_347016, negated_conjecture, (~v1_xboole_0(k3_xcmplx_0(esk9_0,np__1))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_347029, negated_conjecture, (~v1_xboole_0(k3_xcmplx_0(np__1,esk9_0))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_355118, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X1))|~v1_xcmplx_0(k6_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_355131, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_353822, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),X1)=k6_xcmplx_0(X1,X1)|~v1_xreal_0(k6_xcmplx_0(X1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_363707, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X2))|X1!=k6_xcmplx_0(X2,X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_363706, negated_conjecture, (X1=k2_xcmplx_0(X2,k4_xcmplx_0(X2))|X1!=k6_xcmplx_0(X2,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_230416, plain, (k3_xcmplx_0(X1,esk22_0)=k4_xcmplx_0(X2)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_199425, negated_conjecture, (np__0=X1|k3_xcmplx_0(X1,esk9_0)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_360092, negated_conjecture, (~v1_xboole_0(k3_xcmplx_0(esk9_0,np__1)))).#trainneg
cnf(i_0_360098, negated_conjecture, (~v1_xboole_0(k3_xcmplx_0(np__1,esk9_0)))).#trainneg
cnf(i_0_357420, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_356193, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_349682, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X2)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_230417, plain, (k3_xcmplx_0(X1,esk9_0)=k4_xcmplx_0(X2)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_112041, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))!=np__1)).#trainneg
cnf(i_0_332300, plain, (k3_xcmplx_0(k4_xcmplx_0(esk9_0),np__1)!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_342670, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_342285, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_352242, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_326152, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X1)))=k3_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_337675, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_348933, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X2))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X2))|~v1_xreal_0(k6_xcmplx_0(X3,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_353517, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_371081, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_370912, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_230423, plain, (k3_xcmplx_0(X1,np__1)=k4_xcmplx_0(X2)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_125910, plain, (v2_xxreal_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(np__1,k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_369002, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_367806, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_354770, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_322530, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),k4_xcmplx_0(X2))=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_322826, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X3)),X1)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_230426, plain, (k3_xcmplx_0(X1,np__0)=k4_xcmplx_0(X2)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_125909, plain, (v2_xxreal_0(esk22_0)|~v1_xboole_0(k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__1)))).#trainneg
cnf(i_0_326191, negated_conjecture, (k4_xcmplx_0(X1)=np__0|k4_xcmplx_0(X2)=np__0|k3_xcmplx_0(X1,X2)!=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_358047, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_356187, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_158131, plain, (v2_xxreal_0(X2)|v1_relat_1(k9_xtuple_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v4_relat_1(X1,np__0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_216725, negated_conjecture, (np__0=X1|k3_xcmplx_0(esk9_0,X1)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2718, plain, (k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(X2),X1))=k2_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_366177, plain, (k3_xcmplx_0(esk9_0,k4_xcmplx_0(np__1))!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_366184, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),esk9_0)!=np__0|~v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_324951, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_380648, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_373423, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_326156, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X1)))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_326005, negated_conjecture, (k4_xcmplx_0(X1)=np__0|X2=np__0|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_379143, negated_conjecture, (k4_xcmplx_0(X1)=np__0|k4_xcmplx_0(X2)=np__0|k3_xcmplx_0(X2,X1)!=k6_xcmplx_0(X2,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158132, plain, (v2_xxreal_0(X2)|v1_xboole_0(k9_xtuple_0(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v4_relat_1(X1,np__0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2719, plain, (k2_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k4_xcmplx_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_323063, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,X3),k4_xcmplx_0(X1))=k6_xcmplx_0(k6_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_326030, negated_conjecture, (k4_xcmplx_0(X1)=np__0|k2_xcmplx_0(X1,k4_xcmplx_0(X1))!=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_333847, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k2_xcmplx_0(X1,X2)))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__0)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158284, plain, (v2_xxreal_0(X1)|v1_xboole_0(k3_xcmplx_0(X2,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2724, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X2),X1))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_349288, negated_conjecture, (k4_xcmplx_0(esk22_0)=esk22_0|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_284174, plain, (X1=k4_xcmplx_0(esk22_0)|esk22_0!=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_332210, plain, (v1_xboole_0(k4_xcmplx_0(esk22_0))|k4_xcmplx_0(esk22_0)!=np__0)).#trainneg
cnf(i_0_346587, plain, (m1_subset_1(k4_xcmplx_0(esk22_0),k1_numbers))).#trainneg
cnf(i_0_386051, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X3),k4_xcmplx_0(X2))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X3))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_322821, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X3)),X1)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_326190, negated_conjecture, (k4_xcmplx_0(X1)=np__0|k4_xcmplx_0(X2)=np__0|k3_xcmplx_0(X2,X1)!=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_325515, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,k2_xcmplx_0(X1,X2)))=np__0|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_316003, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,k4_xcmplx_0(X3)),X2)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_389095, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),np__0)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_390106, negated_conjecture, (k4_xcmplx_0(X1)=np__0|k4_xcmplx_0(X2)=np__0|k3_xcmplx_0(X1,X2)!=k6_xcmplx_0(X2,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_387115, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v1_xboole_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_388671, plain, (v1_funct_1(k4_xcmplx_0(esk22_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_388672, plain, (v1_relat_1(k4_xcmplx_0(esk22_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_309082, plain, (k4_xcmplx_0(np__1)=np__1|esk22_0=np__0|k4_xcmplx_0(esk22_0)!=esk22_0)).#trainneg
cnf(i_0_388674, plain, (v1_xxreal_0(k4_xcmplx_0(esk22_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_388676, plain, (v3_ordinal1(k4_xcmplx_0(esk22_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_388677, plain, (v1_rat_1(k4_xcmplx_0(esk22_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_388678, plain, (v1_int_1(k4_xcmplx_0(esk22_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_388679, plain, (v1_xboole_0(k4_xcmplx_0(esk22_0))|~m1_subset_1(k1_numbers,k4_xcmplx_0(esk22_0)))).#trainneg
cnf(i_0_388826, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X3),k4_xcmplx_0(X2))=k2_xcmplx_0(X3,k6_xcmplx_0(X1,X2))|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_364184, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X3)),X1)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_325408, negated_conjecture, (k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X3))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_369132, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(k2_xcmplx_0(X2,X3)))=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_385756, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X3),X2)=k6_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_387087, plain, (k3_xcmplx_0(X1,X2)=np__0|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_373937, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(np__1,k4_xcmplx_0(esk22_0))!=np__0|~v1_xreal_0(k2_xcmplx_0(np__1,k4_xcmplx_0(esk22_0))))).#trainneg
cnf(i_0_378974, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__1)!=np__0|~v1_xreal_0(k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__1)))).#trainneg
cnf(i_0_346551, plain, (v3_xxreal_0(k4_xcmplx_0(esk22_0))|v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_328109, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1)),X3)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_393486, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),k4_xcmplx_0(X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_392150, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k2_xcmplx_0(X2,X1),np__0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_392998, negated_conjecture, (k2_xcmplx_0(X2,k6_xcmplx_0(X1,k4_xcmplx_0(X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_398295, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_377514, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_392549, plain, (X1=k3_xcmplx_0(X2,X3)|v2_xxreal_0(X3)|~v7_ordinal1(X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2749, plain, (X1=k1_fcont_1(X2,k2_xcmplx_0(X3,X4))|k1_seq_1(X1,esk3_3(X2,k2_xcmplx_0(X3,X4),X1))!=k2_xcmplx_0(X4,k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,X4),X1)),X3))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,X4),X1)),X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,X4),X1)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_218514, plain, (k2_xcmplx_0(k4_xcmplx_0(esk9_0),esk9_0)=np__0)).#trainneg
cnf(i_0_218515, plain, (k2_xcmplx_0(esk9_0,k4_xcmplx_0(esk9_0))=np__0)).#trainneg
cnf(i_0_218519, plain, (k2_xcmplx_0(esk9_0,np__0)=k2_xcmplx_0(np__0,esk9_0))).#trainneg
cnf(i_0_397632, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(np__1,k4_xcmplx_0(esk22_0))!=np__0)).#trainneg
cnf(i_0_397662, plain, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__1)!=np__0)).#trainneg
cnf(i_0_338391, negated_conjecture, (v1_xreal_0(k2_xcmplx_0(X1,k4_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_398765, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3)))=k6_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_324950, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_397424, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X3))=k6_xcmplx_0(k2_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_384449, negated_conjecture, (k4_xcmplx_0(X1)=np__0|k2_xcmplx_0(k4_xcmplx_0(X1),np__0)!=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_397616, plain, (k3_xcmplx_0(X1,X2)=np__0|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_405119, plain, (esk9_0=np__1|k2_xcmplx_0(np__0,esk9_0)!=np__1)).#trainneg
cnf(i_0_404886, plain, (~v3_xxreal_0(k2_xcmplx_0(np__0,esk9_0)))).#trainneg
cnf(i_0_404628, plain, (~v3_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_404880, plain, (v7_ordinal1(k2_xcmplx_0(np__0,esk9_0)))).#trainneg
cnf(i_0_326001, negated_conjecture, (k4_xcmplx_0(X1)=np__0|X2=np__0|k3_xcmplx_0(X2,k4_xcmplx_0(X1))!=k2_xcmplx_0(X1,k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_348932, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X2))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X3,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_324494, negated_conjecture, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),np__0)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_317063, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(X2)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_373293, negated_conjecture, (k4_xcmplx_0(k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3)))=k6_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_409709, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_392534, plain, (X1=k3_xcmplx_0(esk22_0,X2)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_405101, plain, (k2_xcmplx_0(np__0,esk9_0)=esk9_0)).#trainneg
cnf(i_0_408328, plain, (~v3_xxreal_0(k6_xcmplx_0(esk9_0,np__0))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0))|~v1_xreal_0(k6_xcmplx_0(esk9_0,np__0)))).#trainneg
cnf(i_0_408330, plain, (~v3_xxreal_0(k3_xcmplx_0(np__1,esk9_0))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0))|~v1_xreal_0(k3_xcmplx_0(np__1,esk9_0)))).#trainneg
cnf(i_0_408331, plain, (~v3_xxreal_0(k3_xcmplx_0(esk9_0,np__1))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0))|~v1_xreal_0(k3_xcmplx_0(esk9_0,np__1)))).#trainneg
cnf(i_0_404372, plain, (v2_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_409708, negated_conjecture, (X1=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X3,X4))|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X2)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_408252, negated_conjecture, (k4_xcmplx_0(X1)=np__0|k2_xcmplx_0(k4_xcmplx_0(X1),np__0)!=k3_xcmplx_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_407759, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X3)),X2)=k6_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_407523, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X2),k2_xcmplx_0(X1,X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_407521, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X2),k2_xcmplx_0(X1,X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_407270, negated_conjecture, (k2_xcmplx_0(X2,k2_xcmplx_0(k4_xcmplx_0(X1),X3))=k2_xcmplx_0(k6_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_407278, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k2_xcmplx_0(k6_xcmplx_0(X1,X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_392535, plain, (X1=k3_xcmplx_0(esk9_0,X2)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_218519, plain, (k2_xcmplx_0(esk9_0,np__0)=esk9_0)).#trainneg
cnf(i_0_411890, plain, (k3_xcmplx_0(esk9_0,np__1)=esk9_0)).#trainneg
cnf(i_0_411984, negated_conjecture, (k6_xcmplx_0(esk9_0,np__0)=esk9_0)).#trainneg
cnf(i_0_412092, plain, (k6_xcmplx_0(np__0,esk9_0)=k4_xcmplx_0(esk9_0))).#trainneg
cnf(i_0_412157, plain, (~v1_xboole_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0)))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_412176, plain, (v2_xxreal_0(k3_xcmplx_0(np__1,esk9_0))|~v1_xcmplx_0(k4_xcmplx_0(esk9_0))|~v1_xreal_0(k3_xcmplx_0(np__1,esk9_0)))).#trainneg
cnf(i_0_414442, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(k6_xcmplx_0(X2,X1),X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_416684, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(X1,k4_xcmplx_0(X1)))=k2_xcmplx_0(k6_xcmplx_0(X1,X1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_414455, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,k2_xcmplx_0(X1,X2)),X2)=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_321341, negated_conjecture, (k4_xcmplx_0(k3_xcmplx_0(np__0,k4_xcmplx_0(X1)))=k2_xcmplx_0(k4_xcmplx_0(X1),X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_406632, negated_conjecture, (k2_xcmplx_0(X2,k2_xcmplx_0(X1,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_351167, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_415688, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,k4_xcmplx_0(X2)),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_392541, plain, (X1=k3_xcmplx_0(np__1,X2)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_416974, plain, (k4_xcmplx_0(np__1)=np__1|k4_xcmplx_0(esk9_0)!=esk9_0)).#trainneg
cnf(i_0_2778, plain, (X1=k1_fcont_1(X2,k2_xcmplx_0(X3,X4))|k1_seq_1(X1,esk3_3(X2,k2_xcmplx_0(X3,X4),X1))!=k2_xcmplx_0(X3,k2_xcmplx_0(X4,k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,X4),X1))))|~v1_xcmplx_0(k3_xcmplx_0(X2,esk3_3(X2,k2_xcmplx_0(X3,X4),X1)))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_2(X1,k1_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_417368, plain, (v1_xcmplx_0(k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_416968, plain, (k3_xcmplx_0(np__1,esk9_0)=esk9_0)).#trainneg
cnf(i_0_418070, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),X2)=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_419889, negated_conjecture, (X1=k6_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(X2,k4_xcmplx_0(X3))|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_423071, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_418282, negated_conjecture, (k2_xcmplx_0(X1,k3_xcmplx_0(np__0,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_427234, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X2,X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_426962, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X2,X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_429340, negated_conjecture, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X2))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_427430, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_427158, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_427227, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X2,X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_424892, plain, (X1=esk22_0|v2_xxreal_0(esk22_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_216963, negated_conjecture, (np__0=X1|k4_xcmplx_0(k4_xcmplx_0(X1))!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_412157, plain, (~v1_xboole_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_366184, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),esk9_0)!=np__0)).#trainneg
cnf(i_0_366177, plain, (k3_xcmplx_0(esk9_0,k4_xcmplx_0(np__1))!=np__0)).#trainneg
cnf(i_0_332300, plain, (k3_xcmplx_0(k4_xcmplx_0(esk9_0),np__1)!=np__0)).#trainneg
cnf(i_0_327178, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(esk9_0))!=np__0)).#trainneg
cnf(i_0_404628, plain, (~v3_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0)))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_426955, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X2,X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_427426, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_427154, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_427180, negated_conjecture, (X1=k6_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(k4_xcmplx_0(X3),X2)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_430240, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X3,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_433975, negated_conjecture, (X1=k6_xcmplx_0(X2,X2)|X1!=k2_xcmplx_0(X2,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_429786, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_433957, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_392544, plain, (X1=k3_xcmplx_0(np__0,X2)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_404372, plain, (v2_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0)))|~v1_xreal_0(k4_xcmplx_0(k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_427428, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_427119, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_424907, plain, (X1=k3_xcmplx_0(np__1,X2)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2807, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k2_xcmplx_0(X2,k6_xcmplx_0(X1,X3))|~v1_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X3)))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_432998, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X3,X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_436812, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X3),X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_432229, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(X3))|X1!=k2_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_427452, negated_conjecture, (X1=k6_xcmplx_0(X2,X3)|X1!=k2_xcmplx_0(k4_xcmplx_0(X3),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_427378, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_387089, plain, (k3_xcmplx_0(X1,X2)=k9_xtuple_0(X3)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_235981, negated_conjecture, (np__0=X1|k6_xcmplx_0(X1,np__0)!=np__0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_429896, negated_conjecture, (k6_xcmplx_0(np__0,k4_xcmplx_0(X1))=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_436801, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_427423, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4))).#trainneg
cnf(i_0_387092, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(X3)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2818, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_427409, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_440575, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_427106, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_439836, plain, (k3_xcmplx_0(X1,X2)=k9_xtuple_0(X3)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2830, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k2_xcmplx_0(X3,X1))=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_417364, plain, (k4_xcmplx_0(k4_xcmplx_0(esk9_0))=esk9_0)).#trainneg
cnf(i_0_427151, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_435964, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_439308, negated_conjecture, (X1=k6_xcmplx_0(X2,X2)|X1!=k2_xcmplx_0(X2,k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_427137, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_440955, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(X3)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2833, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_434120, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(k2_xcmplx_0(X3,X4)))|X1!=k2_xcmplx_0(X3,k2_xcmplx_0(X4,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_445764, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X1)|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_446512, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_443594, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_443089, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X2)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_447507, negated_conjecture, (X1=k6_xcmplx_0(X2,k4_xcmplx_0(k2_xcmplx_0(X3,X4)))|X1!=k2_xcmplx_0(X3,k2_xcmplx_0(X2,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_444798, negated_conjecture, (k4_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(esk9_0),k4_xcmplx_0(esk9_0)))=np__0)).#trainneg
cnf(i_0_387091, plain, (k3_xcmplx_0(X1,X2)=k2_xcmplx_0(X3,X4)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xboole_0(X4)|~v1_xboole_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_306247, negated_conjecture, (esk22_0=X1|k2_xcmplx_0(X1,np__0)!=esk22_0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2841, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(k4_xcmplx_0(X3),X1))=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_444810, negated_conjecture, (k4_xcmplx_0(esk9_0)=esk22_0|esk22_0!=esk9_0|esk22_0!=np__0)).#trainneg
cnf(i_0_444643, plain, (m1_subset_1(k4_xcmplx_0(esk9_0),k1_numbers))).#trainneg
cnf(i_0_436343, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),k4_xcmplx_0(X4))|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_448477, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_448840, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k6_xcmplx_0(X2,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_448143, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X2)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_387103, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X3,X4)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xboole_0(X4)|~v1_xboole_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_451777, negated_conjecture, (esk22_0!=np__0|esk22_0!=esk9_0)).#trainneg
cnf(i_0_452133, plain, (~m1_subset_1(k1_numbers,k4_xcmplx_0(esk9_0)))).#trainneg
cnf(i_0_452125, plain, (v1_funct_1(k4_xcmplx_0(esk9_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_449907, negated_conjecture, (~v2_xxreal_0(k6_xcmplx_0(k4_xcmplx_0(esk9_0),k4_xcmplx_0(esk9_0)))|~v1_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(esk9_0),k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_449920, negated_conjecture, (~v3_xxreal_0(k6_xcmplx_0(k4_xcmplx_0(esk9_0),k4_xcmplx_0(esk9_0)))|~v1_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(esk9_0),k4_xcmplx_0(esk9_0))))).#trainneg
cnf(i_0_450390, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(esk9_0),k4_xcmplx_0(esk9_0))=np__0)).#trainneg
cnf(i_0_448839, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(k6_xcmplx_0(X2,X4))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_452480, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),np__0)|X1!=k2_xcmplx_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_448292, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),np__0)|X1!=k2_xcmplx_0(X2,X3)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_448652, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),np__0)|X1!=k2_xcmplx_0(X2,X3)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_387090, plain, (k3_xcmplx_0(X1,X2)=k4_xcmplx_0(X3)|v3_xxreal_0(X3)|v2_xxreal_0(X3)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_309305, negated_conjecture, (esk22_0=X1|k3_xcmplx_0(X1,np__1)!=esk22_0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_452126, plain, (v1_relat_1(k4_xcmplx_0(esk9_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_452128, plain, (v1_xxreal_0(k4_xcmplx_0(esk9_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_452130, plain, (v3_ordinal1(k4_xcmplx_0(esk9_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_452131, plain, (v1_rat_1(k4_xcmplx_0(esk9_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_452132, plain, (v1_int_1(k4_xcmplx_0(esk9_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_346564, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))=esk22_0)).#trainneg
cnf(i_0_453054, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X3,k4_xcmplx_0(X3))|X2!=np__0|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_427384, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(X2,k2_xcmplx_0(X3,k4_xcmplx_0(X4)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(k4_xcmplx_0(X4))|~v1_xreal_0(X4))).#trainneg
cnf(i_0_447963, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),np__0)|X1!=k2_xcmplx_0(X3,X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_438506, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X3,k6_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_453167, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X2),X2)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_387093, plain, (k3_xcmplx_0(X1,X2)=k1_relset_1(X3,X4)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v4_relat_1(X4,X3)|~v1_xboole_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_relat_1(X4))).#trainneg
cnf(i_0_2847, plain, (k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_458093, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X2),X2)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X2,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_457217, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_458078, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X3)|X1!=k2_xcmplx_0(X2,k6_xcmplx_0(X3,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_457222, negated_conjecture, (X1=k6_xcmplx_0(k2_xcmplx_0(X2,X3),X4)|X1!=k2_xcmplx_0(k6_xcmplx_0(X2,X4),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_457825, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X1)=k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(X1,k6_xcmplx_0(X2,X1)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_454112, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X3,X4)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2849, plain, (k6_xcmplx_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_418995, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(k6_xcmplx_0(X1,X3),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_459050, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_459523, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,X1),X2)=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_459521, negated_conjecture, (k6_xcmplx_0(np__0,X2)=X1|~v1_xboole_0(X2)|~v1_xboole_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_313373, negated_conjecture, (esk22_0=X1|k3_xcmplx_0(np__1,X1)!=esk22_0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2852, plain, (v1_xreal_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_444620, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(esk9_0))=esk9_0)).#trainneg
cnf(i_0_463448, negated_conjecture, (k6_xcmplx_0(X1,k6_xcmplx_0(X1,X2))=X2|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_459054, negated_conjecture, (k2_xcmplx_0(X2,k6_xcmplx_0(X1,X2))=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_459052, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,X2),X1)=X1|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_463240, negated_conjecture, (k6_xcmplx_0(X1,k2_xcmplx_0(X1,X2))=k4_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2858, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k6_xcmplx_0(X2,k4_xcmplx_0(X1)),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_193781, plain, (k3_xcmplx_0(X1,esk9_0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_196808, plain, (k3_xcmplx_0(X1,np__1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_466263, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X2),k2_xcmplx_0(X2,X1))=X1|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_458967, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=X1|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_458327, plain, (k3_xcmplx_0(X1,X2)=k1_relset_1(X3,X4)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v4_relat_1(X4,X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_relat_1(X4))).#trainneg
cnf(i_0_198901, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_463892, negated_conjecture, (k6_xcmplx_0(np__0,X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_216839, negated_conjecture, (v2_xxreal_0(esk22_0)|v1_xboole_0(k4_xcmplx_0(esk2_0))|esk2_0!=np__0|~m1_subset_1(k4_xcmplx_0(esk2_0),k5_numbers))).#trainneg
cnf(i_0_458910, negated_conjecture, (k6_xcmplx_0(X2,k2_xcmplx_0(X1,X2))=k4_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_470024, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),k4_xcmplx_0(X2))=X1|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_397601, plain, (k3_xcmplx_0(esk22_0,X1)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_277266, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__0)=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_277267, negated_conjecture, (k2_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_459155, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X1),X1)=k4_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_465096, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X2,X1),k6_xcmplx_0(X1,X3))=k6_xcmplx_0(X2,X3)|~v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_473837, negated_conjecture, (k6_xcmplx_0(X1,k6_xcmplx_0(X2,X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xreal_0(k6_xcmplx_0(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_397602, plain, (k3_xcmplx_0(esk9_0,X1)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_472572, negated_conjecture, (k2_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))!=np__1|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_472434, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(esk22_0),np__0)=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_473871, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2)=k6_xcmplx_0(X1,np__0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_465240, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(X2))=X1|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_464778, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2))=X2|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_476629, negated_conjecture, (k6_xcmplx_0(X1,k2_xcmplx_0(k6_xcmplx_0(k6_xcmplx_0(X1,X2),X2),X2))=X2|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_475704, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(k6_xcmplx_0(X1,X2),X2),np__0)=X1|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_476585, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k6_xcmplx_0(X1,np__0))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_397608, plain, (k3_xcmplx_0(np__1,X1)=np__0|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_475328, negated_conjecture, (k6_xcmplx_0(X1,np__0)=k2_xcmplx_0(X1,np__0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_474541, negated_conjecture, (v2_xxreal_0(esk22_0)|k2_xcmplx_0(np__0,esk2_0)!=np__1|esk2_0!=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_459485, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),k2_xcmplx_0(X2,X3))=X1|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_459217, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X3)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_471171, negated_conjecture, (k6_xcmplx_0(k6_xcmplx_0(X1,X2),X1)=k4_xcmplx_0(X2)|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_471599, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X1,X3)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_468946, negated_conjecture, (k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(X2,X1))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_483721, negated_conjecture, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X1,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_432263, plain, (k3_xcmplx_0(X1,X2)=esk22_0|v2_xxreal_0(esk22_0)|v2_xxreal_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_474719, negated_conjecture, (k3_xcmplx_0(np__1,k4_xcmplx_0(esk22_0))=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_476579, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X1,np__0))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_485078, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=np__0|X3=X1|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X1,X2)))!=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_485143, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=np__1|X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))!=X1|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_485230, negated_conjecture, (X2=esk5_1(k1_fcont_1(X1,X3))|X1=np__0|v2_funct_1(k1_fcont_1(X1,X3))|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X3)))|~v1_xcmplx_0(X2)|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X3)))|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_485076, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=np__0|X3=X1|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X3,esk5_1(k1_fcont_1(X1,X2)))!=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_485055, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=np__0|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))!=esk5_1(k1_fcont_1(X1,X2))|X1!=np__0|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_486448, plain, (k3_xcmplx_0(X1,X2)=esk22_0|v2_xxreal_0(esk22_0)|v2_xxreal_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_486466, negated_conjecture, (k3_xcmplx_0(k4_xcmplx_0(esk22_0),np__1)=np__0|~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_187124, negated_conjecture, (k4_xcmplx_0(esk22_0)=esk14_0|v2_xxreal_0(esk22_0)|esk14_0!=np__0)).#trainneg
cnf(i_0_485146, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=np__0|X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))!=k3_xcmplx_0(X1,np__0)|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_469400, negated_conjecture, (k2_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),k2_xcmplx_0(X2,k4_xcmplx_0(X2)))=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_485116, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=np__0|X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))!=k3_xcmplx_0(esk5_1(k1_fcont_1(X1,X2)),np__0)|~v1_xcmplx_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_432262, plain, (k3_xcmplx_0(X1,X2)=esk22_0|v2_xxreal_0(esk22_0)|v2_xxreal_0(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2911, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)=k6_xcmplx_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_277180, plain, (k2_xcmplx_0(k4_xcmplx_0(esk22_0),np__0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_485154, negated_conjecture, (X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))!=X1|esk5_1(k1_fcont_1(X1,X2))!=np__0|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_485157, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=np__0|X1=np__0|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))!=k6_xcmplx_0(X1,X1)|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_476870, negated_conjecture, (k6_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2))),X2),k2_xcmplx_0(k6_xcmplx_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))),X2),X2))=X2|v2_funct_1(k1_fcont_1(X1,X2))|~v1_xreal_0(k3_xcmplx_0(X1,esk5_1(k1_fcont_1(X1,X2))))|~v1_xreal_0(esk4_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_479786, negated_conjecture, (k6_xcmplx_0(X1,np__0)=k2_xcmplx_0(X1,np__0)|~v1_xboole_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_171463, negated_conjecture, (X1=k4_xcmplx_0(esk22_0)|k4_xcmplx_0(X1)!=esk22_0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_171464, negated_conjecture, (X1=k4_xcmplx_0(esk9_0)|k4_xcmplx_0(X1)!=esk9_0|~v1_xreal_0(X1))).#trainneg
cnf(i_0_405191, negated_conjecture, (v2_xxreal_0(esk22_0)|v1_xreal_0(k2_xcmplx_0(esk22_0,esk2_0))|esk2_0!=np__0)).#trainneg
cnf(i_0_405198, negated_conjecture, (v2_xxreal_0(esk22_0)|v1_xreal_0(k2_xcmplx_0(esk2_0,esk22_0))|esk2_0!=np__0)).#trainneg
cnf(i_0_456360, negated_conjecture, (k6_xcmplx_0(k4_xcmplx_0(esk22_0),np__0)=np__0|esk22_0!=np__0)).#trainneg
cnf(i_0_481046, negated_conjecture, (k6_xcmplx_0(X2,k2_xcmplx_0(k4_xcmplx_0(X1),X2))=X1|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_487709, negated_conjecture, (esk5_1(k1_fcont_1(X1,X2))=esk5_1(k1_fcont_1(X1,X3))|X1=np__0|v2_funct_1(k1_fcont_1(X1,X3))|v2_funct_1(k1_fcont_1(X1,X2))|k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X2)))!=k3_xcmplx_0(X1,esk4_1(k1_fcont_1(X1,X3)))|~v1_xreal_0(esk5_1(k1_fcont_1(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_492374, negated_conjecture, (k6_xcmplx_0(X1,k2_xcmplx_0(X1,k4_xcmplx_0(X1)))=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_492330, negated_conjecture, (k6_xcmplx_0(X1,k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=X2|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_487753, negated_conjecture, (X1=esk5_1(k1_fcont_1(X2,X3))|X2=np__0|v2_funct_1(k1_fcont_1(X2,X3))|k3_xcmplx_0(X1,X2)!=k3_xcmplx_0(X2,esk4_1(k1_fcont_1(X2,X3)))|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_489772, plain, (k3_xcmplx_0(X1,X2)=esk22_0|v2_xxreal_0(esk22_0)|v2_xxreal_0(X2)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_488704, negated_conjecture, (v2_xxreal_0(esk22_0)|k4_xcmplx_0(np__1)!=esk14_0|esk22_0!=np__0|esk14_0!=np__0)).#trainneg
cnf(i_0_2913, plain, (k6_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_490520, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(esk22_0))!=np__1|esk22_0!=np__0)).#trainneg
cnf(i_0_488486, negated_conjecture, (v2_xxreal_0(esk22_0)|esk14_0!=np__0|~v3_xxreal_0(esk14_0))).#trainneg
cnf(i_0_488487, negated_conjecture, (v2_xxreal_0(esk22_0)|esk14_0!=np__0|~v2_xxreal_0(esk14_0))).#trainneg
cnf(i_0_466103, negated_conjecture, (k6_xcmplx_0(X2,X2)=k6_xcmplx_0(X1,X1)|~v1_xreal_0(k6_xcmplx_0(X1,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 233
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 405
# Removed in clause preprocessing      : 59
# Initial clauses in saturation        : 346
# Processed clauses                    : 5241
# ...of these trivial                  : 240
# ...subsumed                          : 1919
# ...remaining for further processing  : 3082
# Other redundant clauses eliminated   : 521
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 211
# Backward-rewritten                   : 295
# Generated clauses                    : 494063
# ...of the previous two non-redundant : 104962
# ...frozen by parental guidance       : 353587
# ...of these subsequently unfrozen    : 0
# ...aggressively subsumed             : 23772
# Contextual simplify-reflections      : 899
# Paramodulations                      : 493349
# Factorizations                       : 0
# NegExts                              : 0
# Equation resolutions                 : 714
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2576
#    Positive orientable unit clauses  : 247
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 68
#    Non-unit-clauses                  : 2261
# Current number of unprocessed clauses: 99498
# ...number of literals in the above   : 597250
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3588
# Proof object given clauses           : 279
# Proof search given clauses           : 3082
# Clause-clause subsumption calls (NU) : 12062477
# Rec. Clause-clause subsumption calls : 696943
# Non-unit clause-clause subsumptions  : 22267
# Unit Clause-clause subsumption calls : 39304
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 88
# BW rewrite match successes           : 64
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12895535

# -------------------------------------------------
# User time                : 141.861 s
# System time              : 1.601 s
# Total time               : 143.461 s
# Maximum resident set size: 516844 pages
141.86user 1.65system 4:50.69elapsed 49%CPU (0avgtext+0avgdata 517104maxresident)k
56inputs+1224outputs (0major+329421minor)pagefaults 0swaps
