2020-05-30 19:19:23.350266: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_53
2020-05-30 19:19:23.434792: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-30 19:19:23.940763: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-30 19:19:24.854345: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1504113 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_53' loaded (query=128; context=256; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t120_sin_cos9, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))&r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1)))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k7_real_1(np__1,k5_square_1(X2))))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t120_sin_cos9)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t4_subset)).
fof(d7_fdiff_1, axiom, ![X1]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X2]:(r2_fdiff_1(X1,X2)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k2_fdiff_1(X1,X2)<=>(k1_relset_1(k1_numbers,X3)=X2&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(X3,X4)=k1_fdiff_1(X1,X4)))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', d7_fdiff_1)).
fof(dt_k2_fdiff_1, axiom, ![X1, X2]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v1_funct_1(k2_fdiff_1(X1,X2))&m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k2_fdiff_1)).
fof(t50_integr13, conjecture, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))&r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1)))=>(r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),X1),X2)=k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k7_real_1(np__1,k5_square_1(X2)))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t50_integr13)).
fof(involutiveness_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(k1_real_1(X1))=X1), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', involutiveness_k1_real_1)).
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k1_real_1)).
fof(t9_fdiff_1, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r2_fdiff_1(X2,X1)<=>(r1_tarski(X1,k1_relset_1(k1_numbers,X2))&![X3]:(m1_subset_1(X3,k1_numbers)=>(r2_hidden(X3,X1)=>r1_fdiff_1(X2,X3))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t9_fdiff_1)).
fof(redefinition_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k7_real_1)).
fof(l3_integr13, axiom, ((m1_subset_1(k1_real_1(np__1),k1_numbers)&m1_subset_1(k1_real_1(k10_real_1(np__1,np__2)),k1_numbers))&m1_subset_1(k10_real_1(np__1,np__2),k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', l3_integr13)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', spc1_numerals)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k1_partfun1)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k1_partfun1)).
fof(redefinition_k32_valued_1, axiom, ![X1, X2, X3]:(((v3_membered(X2)&v1_funct_1(X3))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k32_valued_1(X1,X2,X3)=k30_valued_1(X3)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k32_valued_1)).
fof(involutiveness_k32_valued_1, axiom, ![X1, X2, X3]:(((v3_membered(X2)&v1_funct_1(X3))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k32_valued_1(X1,X2,k32_valued_1(X1,X2,X3))=X3), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', involutiveness_k32_valued_1)).
fof(dt_k32_valued_1, axiom, ![X1, X2, X3]:(((v3_membered(X2)&v1_funct_1(X3))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(v1_funct_1(k32_valued_1(X1,X2,X3))&m1_subset_1(k32_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k32_valued_1)).
fof(dt_k24_sin_cos, axiom, ((v1_funct_1(k24_sin_cos)&v1_funct_2(k24_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k24_sin_cos)).
fof(dt_k2_sin_cos9, axiom, (v1_funct_1(k2_sin_cos9)&m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k2_sin_cos9)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', fc3_membered)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k5_numbers)).
fof(commutativity_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k7_real_1(X1,X2)=k7_real_1(X2,X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', commutativity_k7_real_1)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc14_membered)).
fof(dt_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k1_real_1(X1),k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k1_real_1)).
fof(cc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>v1_xreal_0(X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc3_int_1)).
fof(fc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_int_1(k4_xcmplx_0(X1)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', fc3_int_1)).
fof(cc2_int_1, axiom, ![X1]:(v7_ordinal1(X1)=>v1_int_1(X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc2_int_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc2_relset_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', fc6_relat_1)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k1_relset_1)).
fof(d6_valued_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_valued_0(X1))=>k30_valued_1(X1)=k24_valued_1(X1,k4_xcmplx_0(np__1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', d6_valued_1)).
fof(redefinition_k26_valued_1, axiom, ![X1, X2, X3, X4]:((((v3_membered(X2)&v1_funct_1(X3))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_xreal_0(X4))=>k26_valued_1(X1,X2,X3,X4)=k24_valued_1(X3,X4)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k26_valued_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc1_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k1_relset_1)).
fof(t20_fdiff_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>![X2]:((v3_rcomp_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>((r1_tarski(X2,k1_relset_1(k1_numbers,k26_valued_1(k1_numbers,k1_numbers,X3,X1)))&r2_fdiff_1(X3,X2))=>(r2_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X3,X1),X2)&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(k2_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X3,X1),X2),X4)=k8_real_1(X1,k1_fdiff_1(X3,X4))))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t20_fdiff_1)).
fof(t8_valued_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_valued_0(X1))=>(k9_xtuple_0(k30_valued_1(X1))=k9_xtuple_0(X1)&![X2]:k1_funct_1(k30_valued_1(X1),X2)=k4_xcmplx_0(k1_funct_1(X1,X2)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t8_valued_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', reflexivity_r1_tarski)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc18_valued_0)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k1_seq_1)).
fof(commutativity_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k8_real_1(X2,X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', commutativity_k8_real_1)).
fof(t15_fdiff_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>![X2]:(m1_subset_1(X2,k1_numbers)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r1_fdiff_1(X3,X1)=>(r1_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X3,X2),X1)&k1_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X3,X2),X1)=k8_real_1(X2,k1_fdiff_1(X3,X1))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', t15_fdiff_1)).
fof(dt_k30_valued_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_valued_0(X1))=>((v1_relat_1(k30_valued_1(X1))&v1_funct_1(k30_valued_1(X1)))&v1_valued_0(k30_valued_1(X1)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k30_valued_1)).
fof(cc30_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v5_relat_1(X1,k1_numbers))=>(v1_relat_1(X1)&v3_valued_0(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc30_valued_0)).
fof(dt_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k10_real_1(X1,X2),k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k10_real_1)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', redefinition_k8_real_1)).
fof(dt_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k4_xcmplx_0(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k4_xcmplx_0)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc12_membered)).
fof(cc10_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v1_valued_0(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_valued_0(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc10_valued_0)).
fof(cc6_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v3_valued_0(X1))=>(v1_relat_1(X1)&v1_valued_0(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc6_valued_0)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc5_membered)).
fof(spc2_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', spc2_arithm)).
fof(commutativity_k3_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', commutativity_k3_xcmplx_0)).
fof(dt_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k7_real_1(X1,X2),k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k7_real_1)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc13_membered)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc21_membered)).
fof(dt_k2_rcomp_1, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k2_rcomp_1)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', cc4_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k5_numbers)).
fof(dt_k1_fdiff_1, axiom, ![X1, X2]:(((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))&v1_xreal_0(X2))=>m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k1_fdiff_1)).
fof(dt_k5_square_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k5_square_1(X1),k1_numbers)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t50_integr13', dt_k5_square_1)).
fof(c_0_64, plain, ![X9, X10]:((r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X9)|(~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X9,k2_rcomp_1(k1_real_1(np__1),np__1)))|(~v3_rcomp_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X10,k1_numbers)|(~r2_hidden(X10,X9)|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X9),X10)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X10)),k7_real_1(np__1,k5_square_1(X10)))))|(~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X9,k2_rcomp_1(k1_real_1(np__1),np__1)))|(~v3_rcomp_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t120_sin_cos9])])])])).
fof(c_0_65, plain, ![X37, X38, X39]:(~r2_hidden(X37,X38)|~m1_subset_1(X38,k1_zfmisc_1(X39))|m1_subset_1(X37,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_66, plain, ![X15, X16, X17, X18]:(((k1_relset_1(k1_numbers,X17)=X16|X17!=k2_fdiff_1(X15,X16)|(~v1_funct_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X15,X16)|(~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(~m1_subset_1(X18,k1_numbers)|(~r2_hidden(X18,X16)|k1_seq_1(X17,X18)=k1_fdiff_1(X15,X18))|X17!=k2_fdiff_1(X15,X16)|(~v1_funct_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X15,X16)|(~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&((m1_subset_1(esk3_3(X15,X16,X17),k1_numbers)|k1_relset_1(k1_numbers,X17)!=X16|X17=k2_fdiff_1(X15,X16)|(~v1_funct_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X15,X16)|(~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk3_3(X15,X16,X17),X16)|k1_relset_1(k1_numbers,X17)!=X16|X17=k2_fdiff_1(X15,X16)|(~v1_funct_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X15,X16)|(~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(k1_seq_1(X17,esk3_3(X15,X16,X17))!=k1_fdiff_1(X15,esk3_3(X15,X16,X17))|k1_relset_1(k1_numbers,X17)!=X16|X17=k2_fdiff_1(X15,X16)|(~v1_funct_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X15,X16)|(~v1_funct_1(X15)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_fdiff_1])])])])])).
fof(c_0_67, plain, ![X34, X35]:((v1_funct_1(k2_fdiff_1(X34,X35))|(~v1_funct_1(X34)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(m1_subset_1(k2_fdiff_1(X34,X35),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_funct_1(X34)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_fdiff_1])])])).
cnf(c_0_68, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X2),X1)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k7_real_1(np__1,k5_square_1(X1))))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X2,k2_rcomp_1(k1_real_1(np__1),np__1))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_69, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_70, negated_conjecture, ~(![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))&r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1)))=>(r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),X1),X2)=k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k7_real_1(np__1,k5_square_1(X2))))))))), inference(assume_negation,[status(cth)],[t50_integr13])).
fof(c_0_71, plain, ![X64]:(~m1_subset_1(X64,k1_numbers)|k1_real_1(k1_real_1(X64))=X64), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k1_real_1])])).
fof(c_0_72, plain, ![X117]:(~m1_subset_1(X117,k1_numbers)|k1_real_1(X117)=k4_xcmplx_0(X117)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
cnf(c_0_73, plain, (k1_seq_1(X3,X1)=k1_fdiff_1(X4,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k7_real_1(np__1,k5_square_1(X2))))|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(csr,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_77, plain, ![X25, X26, X27]:(((r1_tarski(X25,k1_relset_1(k1_numbers,X26))|~r2_fdiff_1(X26,X25)|(~v1_funct_1(X26)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X27,k1_numbers)|(~r2_hidden(X27,X25)|r1_fdiff_1(X26,X27))|~r2_fdiff_1(X26,X25)|(~v1_funct_1(X26)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k1_numbers)))))&((m1_subset_1(esk4_2(X25,X26),k1_numbers)|~r1_tarski(X25,k1_relset_1(k1_numbers,X26))|r2_fdiff_1(X26,X25)|(~v1_funct_1(X26)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k1_numbers))))&((r2_hidden(esk4_2(X25,X26),X25)|~r1_tarski(X25,k1_relset_1(k1_numbers,X26))|r2_fdiff_1(X26,X25)|(~v1_funct_1(X26)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k1_numbers))))&(~r1_fdiff_1(X26,esk4_2(X25,X26))|~r1_tarski(X25,k1_relset_1(k1_numbers,X26))|r2_fdiff_1(X26,X25)|(~v1_funct_1(X26)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_fdiff_1])])])])])).
fof(c_0_78, negated_conjecture, ((v3_rcomp_1(esk1_0)&m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))&((r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))&r1_tarski(esk1_0,k2_rcomp_1(k1_real_1(np__1),np__1)))&((m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0))&((r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0))&(k1_seq_1(k2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0),esk2_0)!=k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_70])])])])).
fof(c_0_79, plain, ![X125, X126]:(~m1_subset_1(X125,k1_numbers)|~v1_xreal_0(X126)|k7_real_1(X125,X126)=k2_xcmplx_0(X125,X126)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_real_1])])).
cnf(c_0_80, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_81, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_82, plain, (m1_subset_1(k1_real_1(np__1),k1_numbers)), inference(split_conjunct,[status(thm)],[l3_integr13])).
cnf(c_0_83, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_84, plain, ![X58, X59, X60, X61, X62, X63]:(~v1_funct_1(X62)|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X58,X59)))|~v1_funct_1(X63)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X60,X61)))|k1_partfun1(X58,X59,X60,X61,X62,X63)=k3_relat_1(X62,X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_85, plain, (k1_seq_1(k2_fdiff_1(X1,X2),X3)=k1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_73]), c_0_74]), c_0_75])).
cnf(c_0_86, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X3),X2)|~r2_hidden(X2,X1)|~r2_hidden(X2,X3)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~r1_tarski(X3,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v3_rcomp_1(X3)), inference(spm,[status(thm)],[c_0_76, c_0_76])).
cnf(c_0_87, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_88, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_89, negated_conjecture, (r1_tarski(esk1_0,k2_rcomp_1(k1_real_1(np__1),np__1))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_90, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_91, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_92, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_93, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_94, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_95, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
cnf(c_0_96, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_81]), c_0_83])])).
fof(c_0_97, plain, ![X44, X45, X46, X47, X48, X49]:((v1_funct_1(k1_partfun1(X44,X45,X46,X47,X48,X49))|(~v1_funct_1(X48)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X44,X45)))|~v1_funct_1(X49)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))))&(m1_subset_1(k1_partfun1(X44,X45,X46,X47,X48,X49),k1_zfmisc_1(k2_zfmisc_1(X44,X47)))|(~v1_funct_1(X48)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X44,X45)))|~v1_funct_1(X49)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_98, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_99, plain, ![X102, X103, X104]:(~v3_membered(X103)|~v1_funct_1(X104)|~m1_subset_1(X104,k1_zfmisc_1(k2_zfmisc_1(X102,X103)))|k32_valued_1(X102,X103,X104)=k30_valued_1(X104)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k32_valued_1])])).
fof(c_0_100, plain, ![X70, X71, X72]:(~v3_membered(X71)|~v1_funct_1(X72)|~m1_subset_1(X72,k1_zfmisc_1(k2_zfmisc_1(X70,X71)))|k32_valued_1(X70,X71,k32_valued_1(X70,X71,X72))=X72), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k32_valued_1])])).
fof(c_0_101, plain, ![X50, X51, X52]:((v1_funct_1(k32_valued_1(X50,X51,X52))|(~v3_membered(X51)|~v1_funct_1(X52)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))))&(m1_subset_1(k32_valued_1(X50,X51,X52),k1_zfmisc_1(k2_zfmisc_1(X50,k1_numbers)))|(~v3_membered(X51)|~v1_funct_1(X52)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k32_valued_1])])])).
cnf(c_0_102, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X2,X3)|~r2_hidden(X2,X1)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X3)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X3,k2_rcomp_1(k1_real_1(np__1),np__1))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v3_rcomp_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_69]), c_0_87])).
cnf(c_0_103, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_91])]), c_0_92])])).
cnf(c_0_104, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k2_xcmplx_0(np__1,k5_square_1(X2))))|~v1_xreal_0(k5_square_1(X2))|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_93]), c_0_83])])).
cnf(c_0_105, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])])).
cnf(c_0_106, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_107, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))), inference(spm,[status(thm)],[c_0_98, c_0_98])).
cnf(c_0_108, plain, (k32_valued_1(X3,X1,X2)=k30_valued_1(X2)|~v3_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_109, plain, (k32_valued_1(X3,X1,k32_valued_1(X3,X1,X2))=X2|~v3_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_110, plain, (v1_funct_1(k32_valued_1(X1,X2,X3))|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_111, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X2,esk1_0)|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_89]), c_0_103]), c_0_90]), c_0_91])])).
cnf(c_0_112, plain, (v1_funct_1(k24_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k24_sin_cos])).
cnf(c_0_113, plain, (v1_funct_1(k2_sin_cos9)), inference(split_conjunct,[status(thm)],[dt_k2_sin_cos9])).
cnf(c_0_114, plain, (m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k24_sin_cos])).
cnf(c_0_115, plain, (m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k2_sin_cos9])).
cnf(c_0_116, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k2_xcmplx_0(np__1,k5_square_1(X2))))|~v1_xreal_0(k5_square_1(X2))|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(rw,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_117, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X8)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X10)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_118, plain, (k30_valued_1(k32_valued_1(X1,X2,X3))=X3|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(k32_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])).
cnf(c_0_119, plain, (m1_subset_1(k32_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_120, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_121, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),X5),X6)=k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),X6)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X6,esk1_0)|~r2_hidden(X6,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X5,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_107]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_122, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),X5),X6)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X6)),k2_xcmplx_0(np__1,k5_square_1(X6))))|~v1_xreal_0(k5_square_1(X6))|~r2_hidden(X6,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X5,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_107]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_123, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_114]), c_0_112])])).
fof(c_0_124, plain, ![X79, X80]:(~v6_membered(X79)|(~m1_subset_1(X80,X79)|v7_ordinal1(X80))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_125, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_126, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_127, plain, (k30_valued_1(k32_valued_1(X1,k1_numbers,X2))=X2|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120])])).
fof(c_0_128, plain, ![X133, X134]:(~m1_subset_1(X133,k1_numbers)|~v1_xreal_0(X134)|k7_real_1(X133,X134)=k7_real_1(X134,X133)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k7_real_1])])).
fof(c_0_129, plain, ![X53, X54]:(~v3_membered(X53)|(~m1_subset_1(X54,X53)|v1_xreal_0(X54))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_130, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),X5),X6)=k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),X6)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X6,esk1_0)|~r2_hidden(X6,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X5,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)), inference(rw,[status(thm)],[c_0_121, c_0_105])).
cnf(c_0_131, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_seq_1(k2_fdiff_1(k1_partfun1(X3,X4,X5,X6,k2_sin_cos9,k24_sin_cos),X7),X2)|~v1_xreal_0(k5_square_1(X2))|~r2_hidden(X2,X1)|~r2_hidden(X2,X7)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r1_tarski(X7,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v3_rcomp_1(X7)), inference(spm,[status(thm)],[c_0_116, c_0_122])).
cnf(c_0_132, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_115]), c_0_113])])).
fof(c_0_133, plain, ![X153]:(~m1_subset_1(X153,k1_numbers)|m1_subset_1(k1_real_1(X153),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_real_1])])).
fof(c_0_134, plain, ![X143]:(~v1_int_1(X143)|v1_xreal_0(X143)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_int_1])])).
fof(c_0_135, plain, ![X142]:((v1_xcmplx_0(k4_xcmplx_0(X142))|~v1_int_1(X142))&(v1_int_1(k4_xcmplx_0(X142))|~v1_int_1(X142))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_int_1])])])).
fof(c_0_136, plain, ![X95]:(~v7_ordinal1(X95)|v1_int_1(X95)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_int_1])])).
cnf(c_0_137, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_138, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_139, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_125, c_0_126])).
fof(c_0_140, plain, ![X55, X56, X57]:((v4_relat_1(X57,X55)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))&(v5_relat_1(X57,X56)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_141, plain, ![X138, X139]:(~v1_relat_1(X138)|(~m1_subset_1(X139,k1_zfmisc_1(X138))|v1_relat_1(X139))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_142, plain, ![X111, X112]:v1_relat_1(k2_zfmisc_1(X111,X112)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_143, plain, (k32_valued_1(X1,k1_numbers,X2)=k30_valued_1(X2)|~v1_funct_1(k32_valued_1(X1,k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(k32_valued_1(X1,k1_numbers,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_109]), c_0_120])])).
cnf(c_0_144, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_145, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_146, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_fdiff_1(k1_partfun1(X3,X4,X5,X6,k2_sin_cos9,k24_sin_cos),X2)|~v1_xreal_0(k5_square_1(X2))|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X2,esk1_0)|~r2_hidden(X2,X7)|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X7,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v3_rcomp_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_132])).
cnf(c_0_147, negated_conjecture, (r1_tarski(esk1_0,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_81]), c_0_83])])).
cnf(c_0_148, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_149, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_150, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_151, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_152, plain, (v7_ordinal1(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_139])])).
fof(c_0_153, plain, ![X118, X119]:(~v1_relat_1(X119)|~v4_relat_1(X119,X118)|m1_subset_1(k1_relset_1(X118,X119),k1_zfmisc_1(X118))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
cnf(c_0_154, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_155, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_156, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_157, plain, (k32_valued_1(X1,k1_numbers,X2)=k30_valued_1(X2)|~v1_funct_1(k32_valued_1(X1,k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_119]), c_0_120])])).
cnf(c_0_158, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k7_real_1(np__1,k5_square_1(X2))))|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(rw,[status(thm)],[c_0_76, c_0_105])).
cnf(c_0_159, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(spm,[status(thm)],[c_0_144, c_0_93])).
cnf(c_0_160, plain, (v1_xreal_0(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_83]), c_0_120])])).
cnf(c_0_161, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X2)|~v1_xreal_0(k5_square_1(X2))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X2,esk1_0)|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_92]), c_0_147]), c_0_114]), c_0_115]), c_0_90]), c_0_91])])).
cnf(c_0_162, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_148]), c_0_120])])).
fof(c_0_163, plain, ![X206]:(~v1_relat_1(X206)|~v1_funct_1(X206)|~v1_valued_0(X206)|k30_valued_1(X206)=k24_valued_1(X206,k4_xcmplx_0(np__1))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_valued_1])])).
fof(c_0_164, plain, ![X90, X91, X92, X93]:(~v3_membered(X91)|~v1_funct_1(X92)|~m1_subset_1(X92,k1_zfmisc_1(k2_zfmisc_1(X90,X91)))|~v1_xreal_0(X93)|k26_valued_1(X90,X91,X92,X93)=k24_valued_1(X92,X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k26_valued_1])])).
cnf(c_0_165, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(spm,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_166, plain, (v1_int_1(np__1)), inference(spm,[status(thm)],[c_0_151, c_0_152])).
fof(c_0_167, plain, ![X67, X68, X69]:(~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68)))|v1_relat_1(X69)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_168, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_169, plain, (v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(spm,[status(thm)],[c_0_154, c_0_74])).
cnf(c_0_170, plain, (v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_74]), c_0_156])])).
cnf(c_0_171, plain, (k1_relset_1(k1_numbers,X1)=X2|X1!=k2_fdiff_1(X3,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_172, plain, ![X168, X169]:(~v1_relat_1(X169)|~v4_relat_1(X169,X168)|k1_relset_1(X168,X169)=k9_xtuple_0(X169)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_173, plain, (v4_relat_1(k32_valued_1(X1,X2,X3),X1)|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_154, c_0_119])).
cnf(c_0_174, plain, (k32_valued_1(X1,k1_numbers,X2)=k30_valued_1(X2)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_110]), c_0_120])])).
cnf(c_0_175, plain, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X2)),k2_xcmplx_0(k5_square_1(X2),np__1)))|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k5_square_1(X2),k1_numbers)|~v3_rcomp_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_160])])).
cnf(c_0_176, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),X5),X6)=k1_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos),X6)|~v1_xreal_0(k5_square_1(X6))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X6,esk1_0)|~r2_hidden(X6,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X5,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_107]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_177, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_80]), c_0_148])).
fof(c_0_178, plain, ![X11, X12, X13, X14]:((r2_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X13,X11),X12)|(~r1_tarski(X12,k1_relset_1(k1_numbers,k26_valued_1(k1_numbers,k1_numbers,X13,X11)))|~r2_fdiff_1(X13,X12))|(~v1_funct_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(k1_numbers)))|~m1_subset_1(X11,k1_numbers))&(~m1_subset_1(X14,k1_numbers)|(~r2_hidden(X14,X12)|k1_seq_1(k2_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X13,X11),X12),X14)=k8_real_1(X11,k1_fdiff_1(X13,X14)))|(~r1_tarski(X12,k1_relset_1(k1_numbers,k26_valued_1(k1_numbers,k1_numbers,X13,X11)))|~r2_fdiff_1(X13,X12))|(~v1_funct_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(k1_numbers)))|~m1_subset_1(X11,k1_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_fdiff_1])])])])).
cnf(c_0_179, plain, (k30_valued_1(X1)=k24_valued_1(X1,k4_xcmplx_0(np__1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_180, plain, (k26_valued_1(X3,X1,X2,X4)=k24_valued_1(X2,X4)|~v3_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_xreal_0(X4)), inference(split_conjunct,[status(thm)],[c_0_164])).
cnf(c_0_181, plain, (v1_xreal_0(k4_xcmplx_0(np__1))), inference(spm,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_182, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_183, plain, (m1_subset_1(k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_170])).
cnf(c_0_184, plain, (k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))=X2|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_171]), c_0_74]), c_0_75])).
fof(c_0_185, plain, ![X32, X33]:((k9_xtuple_0(k30_valued_1(X32))=k9_xtuple_0(X32)|(~v1_relat_1(X32)|~v1_funct_1(X32)|~v1_valued_0(X32)))&(k1_funct_1(k30_valued_1(X32),X33)=k4_xcmplx_0(k1_funct_1(X32,X33))|(~v1_relat_1(X32)|~v1_funct_1(X32)|~v1_valued_0(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_valued_1])])])])).
cnf(c_0_186, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_187, plain, (v4_relat_1(k30_valued_1(X1),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_120])])).
cnf(c_0_188, negated_conjecture, (k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k2_xcmplx_0(k5_square_1(X1),np__1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X1,esk1_0)|~r2_hidden(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X2,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k5_square_1(X1),k1_numbers)|~v3_rcomp_1(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176]), c_0_114]), c_0_115])]), c_0_177])).
cnf(c_0_189, plain, (r2_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X1,X2),X3)|~r1_tarski(X3,k1_relset_1(k1_numbers,k26_valued_1(k1_numbers,k1_numbers,X1,X2)))|~r2_fdiff_1(X1,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_178])).
cnf(c_0_190, plain, (k26_valued_1(X1,X2,X3,k4_xcmplx_0(np__1))=k30_valued_1(X3)|~v3_membered(X2)|~v1_valued_0(X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_181])]), c_0_182])).
cnf(c_0_191, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(spm,[status(thm)],[c_0_183, c_0_184])).
cnf(c_0_192, plain, (k9_xtuple_0(k30_valued_1(X1))=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_193, plain, (k1_relset_1(X1,k30_valued_1(X2))=k9_xtuple_0(k30_valued_1(X2))|~v1_relat_1(k30_valued_1(X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))), inference(spm,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_194, negated_conjecture, (k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k2_xcmplx_0(k5_square_1(X1),np__1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_92]), c_0_147]), c_0_90]), c_0_91])])).
cnf(c_0_195, plain, (r2_fdiff_1(k30_valued_1(X1),X2)|~v1_valued_0(X1)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k30_valued_1(X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_96]), c_0_120])]), c_0_191])).
cnf(c_0_196, plain, (k1_relset_1(X1,k30_valued_1(X2))=k9_xtuple_0(X2)|~v1_valued_0(X2)|~v1_relat_1(k30_valued_1(X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193]), c_0_182])).
fof(c_0_197, plain, ![X86, X87]:((~m1_subset_1(X86,k1_zfmisc_1(X87))|r1_tarski(X86,X87))&(~r1_tarski(X86,X87)|m1_subset_1(X86,k1_zfmisc_1(X87)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_198, plain, ![X156]:r1_tarski(X156,X156), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_199, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1),X2)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X2,esk1_0)|~r2_hidden(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k5_square_1(X2),k1_numbers)|~v3_rcomp_1(X1)), inference(spm,[status(thm)],[c_0_175, c_0_194])).
fof(c_0_200, plain, ![X99, X100, X101]:(~v3_membered(X100)|(~m1_subset_1(X101,k1_zfmisc_1(k2_zfmisc_1(X99,X100)))|v3_valued_0(X101))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
cnf(c_0_201, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_202, plain, (r2_fdiff_1(k30_valued_1(X1),X2)|~v1_valued_0(X1)|~v1_relat_1(k30_valued_1(X1))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)), inference(spm,[status(thm)],[c_0_195, c_0_196])).
cnf(c_0_203, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_204, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_205, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_198])).
cnf(c_0_206, negated_conjecture, (k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k7_real_1(np__1,k5_square_1(X1))))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X1,esk1_0)|~r2_hidden(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~r1_tarski(X2,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k5_square_1(X1),k1_numbers)|~v3_rcomp_1(X2)), inference(spm,[status(thm)],[c_0_158, c_0_199])).
fof(c_0_207, plain, ![X122, X123]:(~v1_relat_1(X122)|~v1_funct_1(X122)|~v3_valued_0(X122)|m1_subset_1(k1_seq_1(X122,X123),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
cnf(c_0_208, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_200])).
fof(c_0_209, plain, ![X77, X78]:(~m1_subset_1(X77,k1_numbers)|~v1_xreal_0(X78)|k8_real_1(X77,X78)=k8_real_1(X78,X77)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k8_real_1])])).
fof(c_0_210, plain, ![X29, X30, X31]:((r1_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X31,X30),X29)|~r1_fdiff_1(X31,X29)|(~v1_funct_1(X31)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X30,k1_numbers)|~m1_subset_1(X29,k1_numbers))&(k1_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X31,X30),X29)=k8_real_1(X30,k1_fdiff_1(X31,X29))|~r1_fdiff_1(X31,X29)|(~v1_funct_1(X31)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X30,k1_numbers)|~m1_subset_1(X29,k1_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_fdiff_1])])])])).
cnf(c_0_211, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_fdiff_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_108]), c_0_120])])).
cnf(c_0_212, plain, (r2_fdiff_1(k30_valued_1(X1),X2)|~v1_valued_0(X1)|~v1_relat_1(k30_valued_1(X1))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k9_xtuple_0(X1)))|~v3_rcomp_1(X2)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
fof(c_0_213, plain, ![X155]:(((v1_relat_1(k30_valued_1(X155))|(~v1_relat_1(X155)|~v1_funct_1(X155)|~v1_valued_0(X155)))&(v1_funct_1(k30_valued_1(X155))|(~v1_relat_1(X155)|~v1_funct_1(X155)|~v1_valued_0(X155))))&(v1_valued_0(k30_valued_1(X155))|(~v1_relat_1(X155)|~v1_funct_1(X155)|~v1_valued_0(X155)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k30_valued_1])])])).
fof(c_0_214, plain, ![X94]:((v1_relat_1(X94)|(~v1_relat_1(X94)|~v5_relat_1(X94,k1_numbers)))&(v3_valued_0(X94)|(~v1_relat_1(X94)|~v5_relat_1(X94,k1_numbers)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc30_valued_0])])])).
cnf(c_0_215, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_216, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_204, c_0_205])).
cnf(c_0_217, negated_conjecture, (k1_real_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k7_real_1(np__1,k5_square_1(X1))))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_92]), c_0_147]), c_0_90]), c_0_91])])).
fof(c_0_218, plain, ![X73, X74]:(~m1_subset_1(X73,k1_numbers)|~v1_xreal_0(X74)|m1_subset_1(k10_real_1(X73,X74),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_real_1])])).
cnf(c_0_219, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_220, plain, (v3_valued_0(k24_sin_cos)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_114]), c_0_120])])).
cnf(c_0_221, plain, (v1_relat_1(k24_sin_cos)), inference(spm,[status(thm)],[c_0_182, c_0_114])).
fof(c_0_222, plain, ![X65, X66]:(~m1_subset_1(X65,k1_numbers)|~v1_xreal_0(X66)|k8_real_1(X65,X66)=k3_xcmplx_0(X65,X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
cnf(c_0_223, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_209])).
cnf(c_0_224, plain, (k1_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X1,X2),X3)=k8_real_1(X2,k1_fdiff_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_210])).
fof(c_0_225, plain, ![X149]:(~v1_xcmplx_0(X149)|v1_xcmplx_0(k4_xcmplx_0(X149))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_xcmplx_0])])).
fof(c_0_226, plain, ![X131, X132]:(~v1_membered(X131)|(~m1_subset_1(X132,X131)|v1_xcmplx_0(X132))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
cnf(c_0_227, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~v1_relat_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212]), c_0_103]), c_0_91])])).
cnf(c_0_228, plain, (v1_relat_1(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_213])).
cnf(c_0_229, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_154, c_0_106])).
fof(c_0_230, plain, ![X221, X222]:(~v1_relat_1(X221)|~v1_valued_0(X221)|(~m1_subset_1(X222,k1_zfmisc_1(X221))|v1_valued_0(X222))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc10_valued_0])])])).
fof(c_0_231, plain, ![X121]:((v1_relat_1(X121)|(~v1_relat_1(X121)|~v3_valued_0(X121)))&(v1_valued_0(X121)|(~v1_relat_1(X121)|~v3_valued_0(X121)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_valued_0])])])).
cnf(c_0_232, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_233, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2)), inference(spm,[status(thm)],[c_0_215, c_0_216])).
cnf(c_0_234, negated_conjecture, (k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k7_real_1(np__1,k5_square_1(X1)))=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k7_real_1(np__1,k5_square_1(X1))),k1_numbers)|~m1_subset_1(k5_square_1(X1),k1_numbers)), inference(spm,[status(thm)],[c_0_80, c_0_217])).
cnf(c_0_235, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_218])).
cnf(c_0_236, plain, (m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220]), c_0_221]), c_0_112])])).
cnf(c_0_237, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_222])).
cnf(c_0_238, plain, (k8_real_1(k1_fdiff_1(X1,X2),X3)=k1_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X1,X3),X2)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(spm,[status(thm)],[c_0_223, c_0_224])).
cnf(c_0_239, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_240, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_80]), c_0_148])).
cnf(c_0_241, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_226])).
cnf(c_0_242, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_148, c_0_81])).
fof(c_0_243, plain, ![X128]:(~v3_membered(X128)|v1_membered(X128)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_244, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228]), c_0_182])).
cnf(c_0_245, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_186, c_0_229])).
cnf(c_0_246, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))))), inference(spm,[status(thm)],[c_0_204, c_0_92])).
cnf(c_0_247, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_230])).
cnf(c_0_248, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_231])).
cnf(c_0_249, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233]), c_0_156])])).
cnf(c_0_250, negated_conjecture, (k1_seq_1(k2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0),esk2_0)!=k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_251, negated_conjecture, (k10_real_1(k1_seq_1(k24_sin_cos,k1_seq_1(k2_sin_cos9,X1)),k7_real_1(np__1,k5_square_1(X1)))=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),X1))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235]), c_0_236])])).
cnf(c_0_252, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_78])).
fof(c_0_253, plain, ![X224]:(~v1_xcmplx_0(X224)|k3_xcmplx_0(X224,k4_xcmplx_0(np__1))=k4_xcmplx_0(X224)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc2_arithm])])).
fof(c_0_254, plain, ![X115, X116]:(~v1_xcmplx_0(X115)|~v1_xcmplx_0(X116)|k3_xcmplx_0(X115,X116)=k3_xcmplx_0(X116,X115)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_xcmplx_0])])).
cnf(c_0_255, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_256, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_237]), c_0_177]), c_0_177])).
cnf(c_0_257, plain, (k1_fdiff_1(k26_valued_1(k1_numbers,k1_numbers,X1,X2),X3)=k3_xcmplx_0(k1_fdiff_1(X1,X3),X2)|~r1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X3),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_177]), c_0_177])).
cnf(c_0_258, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_239, c_0_240])).
cnf(c_0_259, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_241, c_0_242])).
cnf(c_0_260, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_243])).
cnf(c_0_261, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245]), c_0_246]), c_0_112]), c_0_113]), c_0_114]), c_0_115])]), c_0_182])).
cnf(c_0_262, plain, (v1_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_valued_0(k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_106]), c_0_156])])).
cnf(c_0_263, plain, (v1_valued_0(k2_zfmisc_1(X1,k1_numbers))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_249]), c_0_156])])).
cnf(c_0_264, negated_conjecture, (k1_seq_1(k2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0),esk2_0)!=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251]), c_0_252])).
cnf(c_0_265, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_69, c_0_90])).
cnf(c_0_266, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_253])).
cnf(c_0_267, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_268, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(spm,[status(thm)],[c_0_255, c_0_166])).
cnf(c_0_269, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_256]), c_0_177])).
cnf(c_0_270, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k4_xcmplx_0(np__1))=k1_fdiff_1(k30_valued_1(X1),X2)|~v1_valued_0(X1)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_190]), c_0_96]), c_0_120])])).
cnf(c_0_271, plain, (v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_258, c_0_81])).
cnf(c_0_272, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260]), c_0_120])])).
cnf(c_0_273, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262]), c_0_263]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
fof(c_0_274, plain, ![X140, X141]:(~m1_subset_1(X140,k1_numbers)|~v1_xreal_0(X141)|m1_subset_1(k7_real_1(X140,X141),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_real_1])])).
cnf(c_0_275, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1)))), inference(spm,[status(thm)],[c_0_204, c_0_89])).
fof(c_0_276, plain, ![X201, X202]:(~v2_membered(X201)|(~m1_subset_1(X202,X201)|v1_xxreal_0(X202))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
fof(c_0_277, plain, ![X83, X84]:(~v3_membered(X83)|(~m1_subset_1(X84,k1_zfmisc_1(X83))|v3_membered(X84))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_278, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_fdiff_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_108]), c_0_120])])).
cnf(c_0_279, negated_conjecture, (k1_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk2_0)!=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_funct_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_fdiff_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)|~m1_subset_1(k32_valued_1(k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_85]), c_0_265]), c_0_252])).
cnf(c_0_280, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267]), c_0_268])])).
cnf(c_0_281, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k1_fdiff_1(X1,X2))=k1_fdiff_1(k30_valued_1(X1),X2)|~v1_valued_0(X1)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270]), c_0_96])])).
cnf(c_0_282, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_271, c_0_272])).
cnf(c_0_283, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_106]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_284, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_285, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_286, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_81]), c_0_83])])).
fof(c_0_287, plain, ![X161, X162]:(~v1_xxreal_0(X161)|~v1_xxreal_0(X162)|m1_subset_1(k2_rcomp_1(X161,X162),k1_zfmisc_1(k1_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_rcomp_1])])).
cnf(c_0_288, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_276])).
fof(c_0_289, plain, ![X113]:(~v3_membered(X113)|v2_membered(X113)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_290, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_277])).
cnf(c_0_291, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_292, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~v1_relat_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_212]), c_0_103]), c_0_91])])).
cnf(c_0_293, negated_conjecture, (k1_fdiff_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk2_0)!=k1_real_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_funct_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_fdiff_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)|~m1_subset_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_108]), c_0_120])])).
cnf(c_0_294, plain, (k1_fdiff_1(k30_valued_1(X1),X2)=k4_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_valued_0(X1)|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281]), c_0_282])).
cnf(c_0_295, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_284]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_296, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_285]), c_0_120])])).
cnf(c_0_297, negated_conjecture, (m1_subset_1(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_69, c_0_286])).
cnf(c_0_298, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_287])).
cnf(c_0_299, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers)), inference(spm,[status(thm)],[c_0_288, c_0_138])).
cnf(c_0_300, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_289])).
cnf(c_0_301, plain, (v3_membered(k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_291]), c_0_120])])).
cnf(c_0_302, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_228]), c_0_182])).
cnf(c_0_303, negated_conjecture, (~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0)|~v1_funct_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_fdiff_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)|~m1_subset_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0),k1_numbers)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_294]), c_0_295])]), c_0_81])).
cnf(c_0_304, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(spm,[status(thm)],[c_0_296, c_0_144])).
cnf(c_0_305, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_145, c_0_297])).
cnf(c_0_306, plain, (v3_membered(k2_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_298]), c_0_120])])).
cnf(c_0_307, plain, (v1_xxreal_0(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_300]), c_0_301])])).
cnf(c_0_308, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_245]), c_0_246]), c_0_112]), c_0_113]), c_0_114]), c_0_115])]), c_0_182])).
cnf(c_0_309, negated_conjecture, (~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0)|~v1_funct_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r2_fdiff_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),esk1_0)|~m1_subset_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0),k1_numbers)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304]), c_0_160])])).
cnf(c_0_310, plain, (m1_subset_1(k30_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v3_membered(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(spm,[status(thm)],[c_0_119, c_0_108])).
cnf(c_0_311, negated_conjecture, (v1_xreal_0(X1)|~v1_xxreal_0(k4_xcmplx_0(np__1))|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_307])])).
cnf(c_0_312, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers)), inference(spm,[status(thm)],[c_0_288, c_0_96])).
cnf(c_0_313, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_262]), c_0_263]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_314, negated_conjecture, (~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0)|~v1_funct_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_212]), c_0_103]), c_0_91])]), c_0_182])).
cnf(c_0_315, plain, (m1_subset_1(k30_valued_1(k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v3_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_106]), c_0_284])).
cnf(c_0_316, negated_conjecture, (v1_xreal_0(X1)|~v2_membered(k1_numbers)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_311, c_0_312])).
cnf(c_0_317, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_106]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_318, negated_conjecture, (~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0)|~v1_funct_1(k30_valued_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0),k1_numbers)|~m1_subset_1(esk1_0,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315]), c_0_120]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_319, plain, (v1_funct_1(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_213])).
fof(c_0_320, plain, ![X88, X89]:(~v1_funct_1(X88)|~m1_subset_1(X88,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X89)|m1_subset_1(k1_fdiff_1(X88,X89),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_fdiff_1])])).
cnf(c_0_321, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_300]), c_0_120])])).
cnf(c_0_322, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_284]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_323, plain, (r1_fdiff_1(X3,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~r2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_324, negated_conjecture, (~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0),k1_numbers)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_245]), c_0_246]), c_0_112]), c_0_113]), c_0_114]), c_0_115])]), c_0_319]), c_0_182])).
cnf(c_0_325, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_320])).
cnf(c_0_326, negated_conjecture, (v1_xreal_0(esk2_0)), inference(spm,[status(thm)],[c_0_321, c_0_322])).
cnf(c_0_327, plain, (r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)), inference(csr,[status(thm)],[c_0_323, c_0_69])).
cnf(c_0_328, negated_conjecture, (~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_325]), c_0_326])])).
cnf(c_0_329, negated_conjecture, (r1_fdiff_1(X1,esk2_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_322]), c_0_90]), c_0_91])])).
cnf(c_0_330, negated_conjecture, (~v1_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_329]), c_0_103])])).
cnf(c_0_331, negated_conjecture, (~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_262]), c_0_263]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
fof(c_0_332, plain, ![X152]:(~m1_subset_1(X152,k1_numbers)|m1_subset_1(k5_square_1(X152),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_square_1])])).
cnf(c_0_333, negated_conjecture, (~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_106]), c_0_112]), c_0_113]), c_0_114]), c_0_115])])).
cnf(c_0_334, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_332])).
cnf(c_0_335, negated_conjecture, (~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_sin_cos9,k24_sin_cos))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_334]), c_0_295])])).
cnf(c_0_336, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_284]), c_0_112]), c_0_113]), c_0_114]), c_0_115])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 337
# Proof object clause steps            : 217
# Proof object formula steps           : 120
# Proof object conjectures             : 64
# Proof object clause conjectures      : 61
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 83
# Proof object initial formulas used   : 64
# Proof object generating inferences   : 128
# Proof object simplifying inferences  : 253
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 159
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 229
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 222
# Processed clauses                    : 3217
# ...of these trivial                  : 118
# ...subsumed                          : 1173
# ...remaining for further processing  : 1926
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 144
# Backward-rewritten                   : 132
# Generated clauses                    : 20619
# ...of the previous two non-trivial   : 19748
# Contextual simplify-reflections      : 429
# Paramodulations                      : 20613
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1647
#    Positive orientable unit clauses  : 272
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 1371
# Current number of unprocessed clauses: 16598
# ...number of literals in the above   : 174286
# Current number of archived formulas  : 0
# Current number of archived clauses   : 279
# Clause-clause subsumption calls (NU) : 1396520
# Rec. Clause-clause subsumption calls : 83170
# Non-unit clause-clause subsumptions  : 1716
# Unit Clause-clause subsumption calls : 12054
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 80
# BW rewrite match successes           : 33
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1142914

# -------------------------------------------------
# User time                : 48.307 s
# System time              : 0.968 s
# Total time               : 49.275 s
# Maximum resident set size: 470424 pages
48.30user 1.01system 0:49.32elapsed 99%CPU (0avgtext+0avgdata 470424maxresident)k
0inputs+224outputs (0major+403438minor)pagefaults 0swaps
