2020-05-31 15:52:25.559486: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55
2020-05-31 15:52:25.649601: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 15:52:25.872811: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 15:52:26.744016: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 15:52:30.541227: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 4981711 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 6.381 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d5_orders_2)).
fof(d5_tarski, axiom, ![X1, X2]:k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d5_tarski)).
fof(t9_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X4,X5),X3))))<=>k1_relset_1(X2,X3)=X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t9_relset_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', commutativity_k2_tarski)).
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', free_g1_orders_2)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', dt_u1_orders_2)).
fof(abstractness_v1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>(v1_orders_2(X1)=>X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', abstractness_v1_orders_2)).
fof(t10_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X5,X4),X3))))<=>k2_relset_1(X2,X3)=X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t10_relset_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((((~(v2_struct_0(X1))&v3_orders_2(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', redefinition_r3_orders_2)).
fof(reflexivity_r3_orders_2, axiom, ![X1, X2, X3]:(((((~(v2_struct_0(X1))&v3_orders_2(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>r3_orders_2(X1,X2,X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', reflexivity_r3_orders_2)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', existence_m1_subset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t1_subset)).
fof(d1_lattice3, axiom, ![X1, X2]:((v3_lattices(X2)&l3_lattices(X2))=>(X2=k1_lattice3(X1)<=>(u1_struct_0(X2)=k9_setfam_1(X1)&![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(k1_binop_1(u2_lattices(X2),X3,X4)=k4_subset_1(X1,X3,X4)&k1_binop_1(u1_lattices(X2),X3,X4)=k9_subset_1(X1,X3,X4))))))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d1_lattice3)).
fof(redefinition_k9_setfam_1, axiom, ![X1]:k9_setfam_1(X1)=k1_zfmisc_1(X1), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', redefinition_k9_setfam_1)).
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k1_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d1_wellord2)).
fof(dt_k1_lattice3, axiom, ![X1]:(v3_lattices(k1_lattice3(X1))&l3_lattices(k1_lattice3(X1))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', dt_k1_lattice3)).
fof(d2_lattice3, axiom, ![X1]:(((~(v2_struct_0(X1))&v10_lattices(X1))&l3_lattices(X1))=>k3_lattice3(X1)=g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d2_lattice3)).
fof(fc2_lattice3, axiom, ![X1]:(v3_lattices(k1_lattice3(X1))&v10_lattices(k1_lattice3(X1))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', fc2_lattice3)).
fof(fc1_lattice3, axiom, ![X1]:(~(v2_struct_0(k1_lattice3(X1)))&v3_lattices(k1_lattice3(X1))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', fc1_lattice3)).
fof(dt_k3_yellow_1, axiom, ![X1]:(v1_orders_2(k3_yellow_1(X1))&l1_orders_2(k3_yellow_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', dt_k3_yellow_1)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d2_yellow_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', cc1_relset_1)).
fof(fc7_yellow_1, axiom, ![X1]:((((~(v2_struct_0(k3_yellow_1(X1)))&v1_orders_2(k3_yellow_1(X1)))&v3_orders_2(k3_yellow_1(X1)))&v4_orders_2(k3_yellow_1(X1)))&v5_orders_2(k3_yellow_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', fc7_yellow_1)).
fof(t2_yellow_1, axiom, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(r3_orders_2(k3_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3)))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t2_yellow_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', redefinition_k1_relset_1)).
fof(dt_k3_lattice3, axiom, ![X1]:(((~(v2_struct_0(X1))&v10_lattices(X1))&l3_lattices(X1))=>((((v1_orders_2(k3_lattice3(X1))&v3_orders_2(k3_lattice3(X1)))&v4_orders_2(k3_lattice3(X1)))&v5_orders_2(k3_lattice3(X1)))&l1_orders_2(k3_lattice3(X1)))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', dt_k3_lattice3)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', redefinition_k2_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', cc2_relset_1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k1_relat_1(X1)=k2_xboole_0(k9_xtuple_0(X1),k10_xtuple_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d6_relat_1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', idempotence_k2_xboole_0)).
fof(t56_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_yellow_0(g1_orders_2(X2,k1_toler_1(u1_orders_2(X1),X2)),X1)&m1_yellow_0(g1_orders_2(X2,k1_toler_1(u1_orders_2(X1),X2)),X1)))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t56_yellow_0)).
fof(redefinition_k1_toler_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>k1_toler_1(X1,X2)=k2_wellord1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', redefinition_k1_toler_1)).
fof(t7_wellord2, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t7_wellord2)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', dt_k1_wellord2)).
fof(t5_yellow_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(v4_yellow_0(k2_yellow_1(X2),k3_yellow_1(X1))&m1_yellow_0(k2_yellow_1(X2),k3_yellow_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t5_yellow_1)).
fof(d1_yellow_1, axiom, ![X1]:k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', d1_yellow_1)).
fof(redefinition_k1_yellow_1, axiom, ![X1]:k1_yellow_1(X1)=k1_wellord2(X1), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', redefinition_k1_yellow_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t5_yellow_1', t3_subset)).
fof(c_0_38, plain, ![X27, X28, X29]:((~r1_orders_2(X27,X28,X29)|r2_hidden(k4_tarski(X28,X29),u1_orders_2(X27))|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))&(~r2_hidden(k4_tarski(X28,X29),u1_orders_2(X27))|r1_orders_2(X27,X28,X29)|~m1_subset_1(X29,u1_struct_0(X27))|~m1_subset_1(X28,u1_struct_0(X27))|~l1_orders_2(X27))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_orders_2])])])])).
fof(c_0_39, plain, ![X71, X72]:k4_tarski(X71,X72)=k2_tarski(k2_tarski(X71,X72),k1_tarski(X71)), inference(variable_rename,[status(thm)],[d5_tarski])).
fof(c_0_40, plain, ![X38, X39, X40, X42, X43]:(((r2_hidden(esk9_3(X38,X39,X40),X39)|k1_relset_1(X39,X40)=X39|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X39,X38))))&(~r2_hidden(k4_tarski(esk9_3(X38,X39,X40),X42),X40)|k1_relset_1(X39,X40)=X39|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X39,X38)))))&(k1_relset_1(X39,X40)!=X39|(~r2_hidden(X43,X39)|r2_hidden(k4_tarski(X43,esk10_4(X38,X39,X40,X43)),X40))|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X39,X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_relset_1])])])])])])).
cnf(c_0_41, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42, plain, (k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_43, plain, ![X86, X87]:k2_tarski(X86,X87)=k2_tarski(X87,X86), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_44, plain, ![X67, X68, X69, X70]:((X67=X69|g1_orders_2(X67,X68)!=g1_orders_2(X69,X70)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X67,X67))))&(X68=X70|g1_orders_2(X67,X68)!=g1_orders_2(X69,X70)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X67,X67))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
fof(c_0_45, plain, ![X76]:(~l1_orders_2(X76)|m1_subset_1(u1_orders_2(X76),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X76),u1_struct_0(X76))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_46, plain, (k1_relset_1(X2,X3)=X2|~r2_hidden(k4_tarski(esk9_3(X1,X2,X3),X4),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),u1_orders_2(X1))|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_50, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_51, plain, ![X85]:(~l1_orders_2(X85)|(~v1_orders_2(X85)|X85=g1_orders_2(u1_struct_0(X85),u1_orders_2(X85)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_orders_2])])).
cnf(c_0_52, plain, (k1_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,X2,X3),X4),k1_tarski(esk9_3(X1,X2,X3))),X3)), inference(rw,[status(thm)],[c_0_46, c_0_42])).
cnf(c_0_53, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),u1_orders_2(X1))|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_56, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_57, plain, (X1=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~l1_orders_2(X1)|~v1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_58, plain, ![X31, X32, X33, X35, X36]:(((r2_hidden(esk7_3(X31,X32,X33),X32)|k2_relset_1(X32,X33)=X32|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32))))&(~r2_hidden(k4_tarski(X35,esk7_3(X31,X32,X33)),X33)|k2_relset_1(X32,X33)=X32|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))))&(k2_relset_1(X32,X33)!=X32|(~r2_hidden(X36,X32)|r2_hidden(k4_tarski(esk8_4(X31,X32,X33,X36),X36),X33))|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_relset_1])])])])])])).
cnf(c_0_59, plain, (k1_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,X2,X3),X4),k1_tarski(esk9_3(X1,X2,X3))),X3)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_60, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54])])).
fof(c_0_61, plain, ![X62, X63, X64]:((~r3_orders_2(X62,X63,X64)|r1_orders_2(X62,X63,X64)|(v2_struct_0(X62)|~v3_orders_2(X62)|~l1_orders_2(X62)|~m1_subset_1(X63,u1_struct_0(X62))|~m1_subset_1(X64,u1_struct_0(X62))))&(~r1_orders_2(X62,X63,X64)|r3_orders_2(X62,X63,X64)|(v2_struct_0(X62)|~v3_orders_2(X62)|~l1_orders_2(X62)|~m1_subset_1(X63,u1_struct_0(X62))|~m1_subset_1(X64,u1_struct_0(X62))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_62, plain, ![X188, X189, X190]:(v2_struct_0(X188)|~v3_orders_2(X188)|~l1_orders_2(X188)|~m1_subset_1(X189,u1_struct_0(X188))|~m1_subset_1(X190,u1_struct_0(X188))|r3_orders_2(X188,X189,X189)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r3_orders_2])])])).
fof(c_0_63, plain, ![X139]:m1_subset_1(esk12_1(X139),X139), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_64, plain, (r1_orders_2(X3,X1,X2)|~r2_hidden(k4_tarski(X1,X2),u1_orders_2(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_65, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56])])).
cnf(c_0_66, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=X1|~v1_orders_2(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_67, plain, (k2_relset_1(X3,X4)=X3|~r2_hidden(k4_tarski(X1,esk7_3(X2,X3,X4)),X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (k1_relset_1(X1,X2)=X1|~r2_hidden(k2_tarski(k1_tarski(esk9_3(X3,X1,X2)),k2_tarski(esk9_3(X3,X1,X2),X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_54])])).
cnf(c_0_69, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_54])])).
cnf(c_0_70, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_72, plain, (m1_subset_1(esk12_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_73, plain, ![X124, X125]:(~r2_hidden(X124,X125)|m1_subset_1(X124,X125)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_74, plain, (r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),u1_orders_2(X3))), inference(rw,[status(thm)],[c_0_64, c_0_42])).
cnf(c_0_75, plain, (u1_orders_2(X1)=X2|X1!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66])])).
fof(c_0_76, plain, ![X9, X10, X11, X12]:(((u1_struct_0(X10)=k9_setfam_1(X9)|X10!=k1_lattice3(X9)|(~v3_lattices(X10)|~l3_lattices(X10)))&((k1_binop_1(u2_lattices(X10),X11,X12)=k4_subset_1(X9,X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(X9))|~m1_subset_1(X11,k1_zfmisc_1(X9))|X10!=k1_lattice3(X9)|(~v3_lattices(X10)|~l3_lattices(X10)))&(k1_binop_1(u1_lattices(X10),X11,X12)=k9_subset_1(X9,X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(X9))|~m1_subset_1(X11,k1_zfmisc_1(X9))|X10!=k1_lattice3(X9)|(~v3_lattices(X10)|~l3_lattices(X10)))))&((m1_subset_1(esk3_2(X9,X10),k1_zfmisc_1(X9))|u1_struct_0(X10)!=k9_setfam_1(X9)|X10=k1_lattice3(X9)|(~v3_lattices(X10)|~l3_lattices(X10)))&((m1_subset_1(esk4_2(X9,X10),k1_zfmisc_1(X9))|u1_struct_0(X10)!=k9_setfam_1(X9)|X10=k1_lattice3(X9)|(~v3_lattices(X10)|~l3_lattices(X10)))&(k1_binop_1(u2_lattices(X10),esk3_2(X9,X10),esk4_2(X9,X10))!=k4_subset_1(X9,esk3_2(X9,X10),esk4_2(X9,X10))|k1_binop_1(u1_lattices(X10),esk3_2(X9,X10),esk4_2(X9,X10))!=k9_subset_1(X9,esk3_2(X9,X10),esk4_2(X9,X10))|u1_struct_0(X10)!=k9_setfam_1(X9)|X10=k1_lattice3(X9)|(~v3_lattices(X10)|~l3_lattices(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_lattice3])])])])])).
fof(c_0_77, plain, ![X24]:k9_setfam_1(X24)=k1_zfmisc_1(X24), inference(variable_rename,[status(thm)],[redefinition_k9_setfam_1])).
fof(c_0_78, plain, ![X18, X19, X20, X21]:(((k1_relat_1(X19)=X18|X19!=k1_wellord2(X18)|~v1_relat_1(X19))&((~r2_hidden(k4_tarski(X20,X21),X19)|r1_tarski(X20,X21)|(~r2_hidden(X20,X18)|~r2_hidden(X21,X18))|X19!=k1_wellord2(X18)|~v1_relat_1(X19))&(~r1_tarski(X20,X21)|r2_hidden(k4_tarski(X20,X21),X19)|(~r2_hidden(X20,X18)|~r2_hidden(X21,X18))|X19!=k1_wellord2(X18)|~v1_relat_1(X19))))&(((r2_hidden(esk5_2(X18,X19),X18)|k1_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19))&(r2_hidden(esk6_2(X18,X19),X18)|k1_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19)))&((~r2_hidden(k4_tarski(esk5_2(X18,X19),esk6_2(X18,X19)),X19)|~r1_tarski(esk5_2(X18,X19),esk6_2(X18,X19))|k1_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19))&(r2_hidden(k4_tarski(esk5_2(X18,X19),esk6_2(X18,X19)),X19)|r1_tarski(esk5_2(X18,X19),esk6_2(X18,X19))|k1_relat_1(X19)!=X18|X19=k1_wellord2(X18)|~v1_relat_1(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
cnf(c_0_79, plain, (k2_relset_1(X3,X4)=X3|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(k2_tarski(k2_tarski(X1,esk7_3(X2,X3,X4)),k1_tarski(X1)),X4)), inference(rw,[status(thm)],[c_0_67, c_0_42])).
cnf(c_0_80, plain, (k1_relset_1(X1,u1_orders_2(X2))=X1|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r1_orders_2(X2,esk9_3(X3,X1,u1_orders_2(X2)),esk9_3(X3,X1,u1_orders_2(X2)))|~l1_orders_2(X2)|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69])])).
cnf(c_0_81, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r3_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_82, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_83, plain, (m1_subset_1(esk12_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_84, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_85, plain, (r2_hidden(esk9_3(X1,X2,X3),X2)|k1_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_86, plain, (r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),u1_orders_2(X3))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_87, plain, (u1_orders_2(g1_orders_2(X1,X2))=X2|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_75])])).
cnf(c_0_88, plain, (u1_struct_0(X1)=k9_setfam_1(X2)|X1!=k1_lattice3(X2)|~v3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_89, plain, (k9_setfam_1(X1)=k1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_90, plain, ![X50]:(v3_lattices(k1_lattice3(X50))&l3_lattices(k1_lattice3(X50))), inference(variable_rename,[status(thm)],[dt_k1_lattice3])).
cnf(c_0_91, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_92, plain, (X2=k1_wellord2(X1)|~r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X2)|~r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_93, plain, (k2_relset_1(X3,X4)=X3|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(k2_tarski(k2_tarski(X1,esk7_3(X2,X3,X4)),k1_tarski(X1)),X4)), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_94, plain, (k1_relset_1(X1,u1_orders_2(X2))=X1|v2_struct_0(X2)|~l1_orders_2(X2)|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r3_orders_2(X2,esk9_3(X3,X1,u1_orders_2(X2)),esk9_3(X3,X1,u1_orders_2(X2)))|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81])])).
cnf(c_0_95, plain, (r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_96, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_97, plain, (k1_relset_1(X2,X3)=X2|r2_hidden(esk9_3(X1,X2,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_98, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,X2))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),X2)|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~v1_orders_2(g1_orders_2(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87])])).
fof(c_0_99, plain, ![X45]:(v2_struct_0(X45)|~v10_lattices(X45)|~l3_lattices(X45)|k3_lattice3(X45)=g1_orders_2(u1_struct_0(X45),k2_lattice3(X45))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_lattice3])])])).
cnf(c_0_100, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|X1!=k1_lattice3(X2)|~v3_lattices(X1)|~l3_lattices(X1)), inference(rw,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_101, plain, (l3_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_102, plain, (v3_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_90])).
fof(c_0_103, plain, ![X59]:(v3_lattices(k1_lattice3(X59))&v10_lattices(k1_lattice3(X59))), inference(variable_rename,[status(thm)],[fc2_lattice3])).
fof(c_0_104, plain, ![X84]:(~v2_struct_0(k1_lattice3(X84))&v3_lattices(k1_lattice3(X84))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_lattice3])])).
fof(c_0_105, plain, ![X46]:(v1_orders_2(k3_yellow_1(X46))&l1_orders_2(k3_yellow_1(X46))), inference(variable_rename,[status(thm)],[dt_k3_yellow_1])).
fof(c_0_106, plain, ![X30]:k3_yellow_1(X30)=k3_lattice3(k1_lattice3(X30)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
cnf(c_0_107, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_108, plain, (X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)|~r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|~r2_hidden(k2_tarski(k2_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),k1_tarski(esk5_2(X1,X2))),X2)), inference(rw,[status(thm)],[c_0_92, c_0_42])).
fof(c_0_109, plain, ![X143, X144, X145]:(~m1_subset_1(X145,k1_zfmisc_1(k2_zfmisc_1(X143,X144)))|v1_relat_1(X145)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_110, plain, (k2_relset_1(X1,X2)=X1|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,esk7_3(X4,X1,X2))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))), inference(rw,[status(thm)],[c_0_93, c_0_54])).
cnf(c_0_111, plain, (k1_relset_1(X1,u1_orders_2(X2))=X1|v2_struct_0(X2)|~l1_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(csr,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_112, plain, (m1_subset_1(esk9_3(X1,X2,X3),X2)|k1_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97])])).
cnf(c_0_113, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~r1_orders_2(X2,X3,X4)|~l1_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_53])])).
cnf(c_0_114, plain, (v2_struct_0(X1)|k3_lattice3(X1)=g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_115, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|X1!=k1_lattice3(X2)|~v3_lattices(X1)|~l3_lattices(X1)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_116, plain, (l3_lattices(k1_lattice3(X1))), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_117, plain, (v3_lattices(k1_lattice3(X1))), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_118, plain, (v10_lattices(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_119, plain, (~v2_struct_0(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_120, plain, (l1_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_121, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_122, plain, (v1_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_105])).
fof(c_0_123, plain, ![X136]:((((~v2_struct_0(k3_yellow_1(X136))&v1_orders_2(k3_yellow_1(X136)))&v3_orders_2(k3_yellow_1(X136)))&v4_orders_2(k3_yellow_1(X136)))&v5_orders_2(k3_yellow_1(X136))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc7_yellow_1])])).
cnf(c_0_124, plain, (u1_struct_0(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_56])])).
cnf(c_0_125, plain, (X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)|~r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|~r2_hidden(k2_tarski(k2_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),k1_tarski(esk5_2(X1,X2))),X2)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_126, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_109])).
fof(c_0_127, plain, ![X47, X48, X49]:((~r3_orders_2(k3_yellow_1(X47),X48,X49)|r1_tarski(X48,X49)|~m1_subset_1(X49,u1_struct_0(k3_yellow_1(X47)))|~m1_subset_1(X48,u1_struct_0(k3_yellow_1(X47))))&(~r1_tarski(X48,X49)|r3_orders_2(k3_yellow_1(X47),X48,X49)|~m1_subset_1(X49,u1_struct_0(k3_yellow_1(X47)))|~m1_subset_1(X48,u1_struct_0(k3_yellow_1(X47))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_yellow_1])])])])).
cnf(c_0_128, plain, (k2_relset_1(X1,u1_orders_2(X2))=X1|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r1_orders_2(X2,esk7_3(X3,X1,u1_orders_2(X2)),esk7_3(X3,X1,u1_orders_2(X2)))|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_69])])).
cnf(c_0_129, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|k2_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_130, plain, ![X57, X58]:(~v1_relat_1(X58)|~v4_relat_1(X58,X57)|k1_relset_1(X57,X58)=k9_xtuple_0(X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_131, plain, (k1_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_132, plain, (r1_orders_2(g1_orders_2(X1,X3),X4,X5)|~r1_orders_2(g1_orders_2(X2,X3),X4,X5)|~l1_orders_2(g1_orders_2(X1,X3))|~l1_orders_2(g1_orders_2(X2,X3))|~v1_orders_2(g1_orders_2(X1,X3))|~m1_subset_1(X5,u1_struct_0(g1_orders_2(X1,X3)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X3)))|~m1_subset_1(X5,u1_struct_0(g1_orders_2(X2,X3)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X2,X3)))|~v1_orders_2(g1_orders_2(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_87])])).
cnf(c_0_133, plain, (g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))=k3_lattice3(X1)|v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_134, plain, (u1_struct_0(k1_lattice3(X1))=k1_zfmisc_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_115]), c_0_116]), c_0_117])])])).
cnf(c_0_135, plain, (v10_lattices(k1_lattice3(X1))), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_136, plain, (~v2_struct_0(k1_lattice3(X1))), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_137, plain, (l1_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(rw,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_138, plain, (v1_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(rw,[status(thm)],[c_0_122, c_0_121])).
cnf(c_0_139, plain, (v3_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_140, plain, (u1_struct_0(X1)=X2|X1!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_66])])).
cnf(c_0_141, plain, (~v2_struct_0(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_123])).
fof(c_0_142, plain, ![X25]:(((((v1_orders_2(k3_lattice3(X25))|(v2_struct_0(X25)|~v10_lattices(X25)|~l3_lattices(X25)))&(v3_orders_2(k3_lattice3(X25))|(v2_struct_0(X25)|~v10_lattices(X25)|~l3_lattices(X25))))&(v4_orders_2(k3_lattice3(X25))|(v2_struct_0(X25)|~v10_lattices(X25)|~l3_lattices(X25))))&(v5_orders_2(k3_lattice3(X25))|(v2_struct_0(X25)|~v10_lattices(X25)|~l3_lattices(X25))))&(l1_orders_2(k3_lattice3(X25))|(v2_struct_0(X25)|~v10_lattices(X25)|~l3_lattices(X25)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_lattice3])])])])).
cnf(c_0_143, plain, (r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X2)|r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_144, plain, (X1=k1_wellord2(X2)|k1_relat_1(X1)!=X2|~r2_hidden(k2_tarski(k1_tarski(esk5_2(X2,X1)),k2_tarski(esk5_2(X2,X1),esk6_2(X2,X1))),X1)|~v1_relat_1(X1)|~r1_tarski(esk5_2(X2,X1),esk6_2(X2,X1))), inference(rw,[status(thm)],[c_0_125, c_0_54])).
cnf(c_0_145, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_146, plain, (r3_orders_2(k3_yellow_1(X3),X1,X2)|~r1_tarski(X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(X3)))|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_147, plain, (k2_relset_1(X1,u1_orders_2(X2))=X1|v2_struct_0(X2)|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r3_orders_2(X2,esk7_3(X3,X1,u1_orders_2(X2)),esk7_3(X3,X1,u1_orders_2(X2)))|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_81])])).
cnf(c_0_148, plain, (k2_relset_1(X2,X3)=X2|r2_hidden(esk7_3(X1,X2,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_149, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_150, plain, (k1_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_56])])).
cnf(c_0_151, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X4)|v2_struct_0(g1_orders_2(X5,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X5,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X5,X2))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X5,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X5,X2)))|~r3_orders_2(g1_orders_2(X5,X2),X3,X4)|~v3_orders_2(g1_orders_2(X5,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_81])])).
cnf(c_0_152, plain, (g1_orders_2(k1_zfmisc_1(X1),k2_lattice3(k1_lattice3(X1)))=k3_lattice3(k1_lattice3(X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_135]), c_0_116])]), c_0_136])])).
cnf(c_0_153, plain, (l1_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_154, plain, (v1_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_155, plain, (v3_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(rw,[status(thm)],[c_0_139, c_0_121])).
cnf(c_0_156, plain, (u1_struct_0(g1_orders_2(X1,X2))=X1|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_140])])).
cnf(c_0_157, plain, (~v2_struct_0(k3_lattice3(k1_lattice3(X1)))), inference(rw,[status(thm)],[c_0_141, c_0_121])).
cnf(c_0_158, plain, (r1_tarski(X2,X3)|~r3_orders_2(k3_yellow_1(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_159, plain, (v1_orders_2(k3_lattice3(X1))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_160, plain, (l1_orders_2(k3_lattice3(X1))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_161, plain, (X2=k1_wellord2(X1)|r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|r2_hidden(k2_tarski(k2_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),k1_tarski(esk5_2(X1,X2))),X2)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(rw,[status(thm)],[c_0_143, c_0_42])).
cnf(c_0_162, plain, (u1_orders_2(X1)=k1_wellord2(X2)|k1_relat_1(u1_orders_2(X1))!=X2|~v1_relat_1(u1_orders_2(X1))|~r1_tarski(esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~r1_orders_2(X1,esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~l1_orders_2(X1)|~m1_subset_1(esk6_2(X2,u1_orders_2(X1)),u1_struct_0(X1))|~m1_subset_1(esk5_2(X2,u1_orders_2(X1)),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_60])])).
cnf(c_0_163, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_56])])).
cnf(c_0_164, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~r1_tarski(X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X3))))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(k1_lattice3(X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_121]), c_0_121]), c_0_121])).
fof(c_0_165, plain, ![X51, X52]:(~v1_relat_1(X52)|~v5_relat_1(X52,X51)|k2_relset_1(X51,X52)=k10_xtuple_0(X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_166, plain, (k2_relset_1(X1,u1_orders_2(X2))=X1|v2_struct_0(X2)|~l1_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(csr,[status(thm)],[c_0_147, c_0_95])).
cnf(c_0_167, plain, (m1_subset_1(esk7_3(X1,X2,X3),X2)|k2_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_148])])).
fof(c_0_168, plain, ![X155, X156, X157]:((v4_relat_1(X157,X155)|~m1_subset_1(X157,k1_zfmisc_1(k2_zfmisc_1(X155,X156))))&(v5_relat_1(X157,X156)|~m1_subset_1(X157,k1_zfmisc_1(k2_zfmisc_1(X155,X156))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_169, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_170, plain, (k1_relset_1(u1_struct_0(g1_orders_2(X1,X2)),X2)=u1_struct_0(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_87])])).
cnf(c_0_171, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(X2)),X3,X4)|v2_struct_0(X2)|~r3_orders_2(X2,X3,X4)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~l1_orders_2(X2)|~v3_orders_2(X2)|~v1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v1_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_66])])).
cnf(c_0_172, plain, (u1_orders_2(k3_lattice3(k1_lattice3(X1)))=k2_lattice3(k1_lattice3(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_152]), c_0_153]), c_0_154])])])).
cnf(c_0_173, plain, (v3_orders_2(k3_lattice3(k1_lattice3(X1)))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_174, plain, (u1_struct_0(k3_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_152]), c_0_153]), c_0_154])])])).
cnf(c_0_175, plain, (~v2_struct_0(k3_lattice3(k1_lattice3(X1)))), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_176, plain, (r1_tarski(X2,X3)|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X1))))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_158, c_0_121]), c_0_121]), c_0_121])).
cnf(c_0_177, plain, (u1_orders_2(k3_lattice3(X1))=k2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_133])])).
cnf(c_0_178, plain, (v2_struct_0(X1)|v1_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_179, plain, (v2_struct_0(X1)|l1_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(evalgc,[status(thm)],[c_0_160])).
cnf(c_0_180, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),u1_orders_2(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_54])])).
cnf(c_0_181, plain, (X2=k1_wellord2(X1)|r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|r2_hidden(k2_tarski(k2_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),k1_tarski(esk5_2(X1,X2))),X2)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[c_0_161])).
cnf(c_0_182, plain, (u1_orders_2(X1)=k1_wellord2(X2)|k1_relat_1(u1_orders_2(X1))!=X2|~r1_orders_2(X1,esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~l1_orders_2(X1)|~r1_tarski(esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~m1_subset_1(esk6_2(X2,u1_orders_2(X1)),u1_struct_0(X1))|~m1_subset_1(esk5_2(X2,u1_orders_2(X1)),u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_183, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~r1_tarski(X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X3))))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(k1_lattice3(X3))))), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_184, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_185, plain, (k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167])])).
cnf(c_0_186, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_187, plain, (u1_struct_0(g1_orders_2(X1,X2))=k9_xtuple_0(X2)|v2_struct_0(g1_orders_2(X1,X2))|~v4_relat_1(X2,u1_struct_0(g1_orders_2(X1,X2)))|~v1_relat_1(X2)|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170])])).
cnf(c_0_188, plain, (v1_relat_1(X2)|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_87])])).
cnf(c_0_189, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_190, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))),X3,X4)|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_153]), c_0_173]), c_0_154]), c_0_174]), c_0_174])]), c_0_175])])).
cnf(c_0_191, plain, (r1_tarski(X2,X3)|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X1))))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X1))))), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_192, plain, (u1_orders_2(k3_lattice3(X1))=k2_lattice3(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_177, c_0_178]), c_0_179])).
cnf(c_0_193, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_172]), c_0_153]), c_0_174]), c_0_174])])])).
cnf(c_0_194, plain, (X1=k1_wellord2(X2)|r2_hidden(k2_tarski(k1_tarski(esk5_2(X2,X1)),k2_tarski(esk5_2(X2,X1),esk6_2(X2,X1))),X1)|r1_tarski(esk5_2(X2,X1),esk6_2(X2,X1))|k1_relat_1(X1)!=X2|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_181, c_0_54])).
cnf(c_0_195, plain, (u1_orders_2(X1)=k1_wellord2(X2)|v2_struct_0(X1)|k1_relat_1(u1_orders_2(X1))!=X2|~l1_orders_2(X1)|~r1_tarski(esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~m1_subset_1(esk6_2(X2,u1_orders_2(X1)),u1_struct_0(X1))|~m1_subset_1(esk5_2(X2,u1_orders_2(X1)),u1_struct_0(X1))|~r3_orders_2(X1,esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_81])])).
cnf(c_0_196, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_183, c_0_174]), c_0_174])])).
cnf(c_0_197, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_198, plain, (k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_56])])).
cnf(c_0_199, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_200, plain, (u1_struct_0(g1_orders_2(X1,X2))=k9_xtuple_0(X2)|v2_struct_0(g1_orders_2(X1,X2))|~v4_relat_1(X2,u1_struct_0(g1_orders_2(X1,X2)))|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))), inference(csr,[status(thm)],[c_0_187, c_0_188])).
cnf(c_0_201, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_202, plain, (u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))=k1_wellord2(X3)|k1_relat_1(u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))!=X3|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~r1_tarski(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))))|~m1_subset_1(esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),k1_zfmisc_1(X2))|~m1_subset_1(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_190])])).
cnf(c_0_203, plain, (r1_tarski(X1,X2)|~r3_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_191, c_0_174]), c_0_174])])).
cnf(c_0_204, plain, (g1_orders_2(u1_struct_0(k3_lattice3(X1)),k2_lattice3(X1))=k3_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_192])])).
cnf(c_0_205, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_206, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|r1_tarski(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~v1_relat_1(k2_lattice3(k1_lattice3(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_207, plain, (v1_relat_1(k2_lattice3(k1_lattice3(X1)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_172]), c_0_153])])])).
cnf(c_0_208, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~r1_tarski(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196]), c_0_172]), c_0_172]), c_0_153]), c_0_173]), c_0_172]), c_0_172]), c_0_172]), c_0_174]), c_0_172]), c_0_174]), c_0_172]), c_0_172])]), c_0_175])])).
fof(c_0_209, plain, ![X83]:(~v1_relat_1(X83)|k1_relat_1(X83)=k2_xboole_0(k9_xtuple_0(X83),k10_xtuple_0(X83))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
cnf(c_0_210, plain, (k10_xtuple_0(u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_211, plain, (v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_56])])).
cnf(c_0_212, plain, (k9_xtuple_0(u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v4_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_66])])).
cnf(c_0_213, plain, (v4_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_56])])).
fof(c_0_214, plain, ![X121]:k2_xboole_0(X121,X121)=X121, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_215, plain, (u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))=k1_wellord2(X3)|k1_relat_1(u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))!=X3|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))))|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),k1_zfmisc_1(X2))|~m1_subset_1(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),k1_zfmisc_1(X2))), inference(csr,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_216, plain, (g1_orders_2(u1_struct_0(k3_lattice3(X1)),k2_lattice3(X1))=k3_lattice3(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_204, c_0_178]), c_0_179])).
cnf(c_0_217, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_218, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|r1_orders_2(k3_lattice3(k1_lattice3(X1)),esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_207])]), c_0_208])).
cnf(c_0_219, plain, (k1_relat_1(X1)=k2_xboole_0(k9_xtuple_0(X1),k10_xtuple_0(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_209])).
cnf(c_0_220, plain, (k10_xtuple_0(u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_210, c_0_163]), c_0_211])).
cnf(c_0_221, plain, (k9_xtuple_0(u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)), inference(csr,[status(thm)],[c_0_212, c_0_213])).
cnf(c_0_222, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_223, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_172]), c_0_172]), c_0_172]), c_0_172]), c_0_153]), c_0_154]), c_0_172]), c_0_174]), c_0_172]), c_0_174]), c_0_172]), c_0_172]), c_0_135]), c_0_116])]), c_0_136])])).
cnf(c_0_224, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X1)),esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218]), c_0_153]), c_0_173]), c_0_174]), c_0_174])]), c_0_175])])).
cnf(c_0_225, plain, (k2_xboole_0(k9_xtuple_0(X1),k10_xtuple_0(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_226, plain, (k10_xtuple_0(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_172]), c_0_174]), c_0_153]), c_0_173])]), c_0_175])])).
cnf(c_0_227, plain, (k9_xtuple_0(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_172]), c_0_174]), c_0_153]), c_0_173]), c_0_154])]), c_0_175])])).
cnf(c_0_228, plain, (k2_xboole_0(X1,X1)=X1), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_229, plain, (r2_hidden(esk6_2(X1,X2),X1)|X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
fof(c_0_230, plain, ![X167, X168]:((v4_yellow_0(g1_orders_2(X168,k1_toler_1(u1_orders_2(X167),X168)),X167)|~m1_subset_1(X168,k1_zfmisc_1(u1_struct_0(X167)))|~l1_orders_2(X167))&(m1_yellow_0(g1_orders_2(X168,k1_toler_1(u1_orders_2(X167),X168)),X167)|~m1_subset_1(X168,k1_zfmisc_1(u1_struct_0(X167)))|~l1_orders_2(X167))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_yellow_0])])])])).
fof(c_0_231, plain, ![X78, X79]:(~v1_relat_1(X78)|k1_toler_1(X78,X79)=k2_wellord1(X78,X79)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_toler_1])])).
fof(c_0_232, plain, ![X15, X16]:(~r1_tarski(X15,X16)|k2_wellord1(k1_wellord2(X16),X15)=k1_wellord2(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_wellord2])])).
fof(c_0_233, plain, ![X130]:v1_relat_1(k1_wellord2(X130)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
cnf(c_0_234, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_223, c_0_224])).
cnf(c_0_235, plain, (k1_relat_1(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226]), c_0_227]), c_0_228]), c_0_207])])])).
cnf(c_0_236, plain, (X2=k1_wellord2(X1)|r2_hidden(esk6_2(X1,X2),X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_237, plain, (r2_hidden(esk5_2(X1,X2),X1)|X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_238, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_230])).
cnf(c_0_239, plain, (k1_toler_1(X1,X2)=k2_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_231])).
cnf(c_0_240, plain, (k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_241, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_233])).
cnf(c_0_242, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_zfmisc_1(X1)!=X2|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_234, c_0_235])])).
cnf(c_0_243, plain, (m1_subset_1(esk6_2(X1,X2),X1)|X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_236])])).
cnf(c_0_244, plain, (X2=k1_wellord2(X1)|r2_hidden(esk5_2(X1,X2),X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[c_0_237])).
fof(c_0_245, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(v4_yellow_0(k2_yellow_1(X2),k3_yellow_1(X1))&m1_yellow_0(k2_yellow_1(X2),k3_yellow_1(X1))))), inference(assume_negation,[status(cth)],[t5_yellow_1])).
fof(c_0_246, plain, ![X17]:k2_yellow_1(X17)=g1_orders_2(X17,k1_yellow_1(X17)), inference(variable_rename,[status(thm)],[d1_yellow_1])).
fof(c_0_247, plain, ![X77]:k1_yellow_1(X77)=k1_wellord2(X77), inference(variable_rename,[status(thm)],[redefinition_k1_yellow_1])).
cnf(c_0_248, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~l1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_249, plain, (k2_wellord1(X1,X2)=k1_toler_1(X1,X2)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_250, plain, (k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_251, plain, (v1_relat_1(k1_wellord2(X1))), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_252, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(k1_zfmisc_1(X1),k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243]), c_0_235]), c_0_207])])])).
cnf(c_0_253, plain, (m1_subset_1(esk5_2(X1,X2),X1)|X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_244])])).
cnf(c_0_254, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_230])).
fof(c_0_255, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))&(~v4_yellow_0(k2_yellow_1(esk2_0),k3_yellow_1(esk1_0))|~m1_yellow_0(k2_yellow_1(esk2_0),k3_yellow_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_245])])])).
cnf(c_0_256, plain, (k2_yellow_1(X1)=g1_orders_2(X1,k1_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_246])).
cnf(c_0_257, plain, (k1_yellow_1(X1)=k1_wellord2(X1)), inference(split_conjunct,[status(thm)],[c_0_247])).
cnf(c_0_258, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(X3,X1)),g1_orders_2(X2,X3))|~l1_orders_2(g1_orders_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,X3))))|~v1_orders_2(g1_orders_2(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_87])])).
cnf(c_0_259, plain, (k1_toler_1(k1_wellord2(X1),X2)=k1_wellord2(X2)|~r1_tarski(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_250]), c_0_251])])])).
cnf(c_0_260, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_235]), c_0_207])])])).
fof(c_0_261, plain, ![X103, X104]:((~m1_subset_1(X103,k1_zfmisc_1(X104))|r1_tarski(X103,X104))&(~r1_tarski(X103,X104)|m1_subset_1(X103,k1_zfmisc_1(X104)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_262, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~l1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_254])).
cnf(c_0_263, negated_conjecture, (~v4_yellow_0(k2_yellow_1(esk2_0),k3_yellow_1(esk1_0))|~m1_yellow_0(k2_yellow_1(esk2_0),k3_yellow_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_255])).
cnf(c_0_264, plain, (k2_yellow_1(X1)=g1_orders_2(X1,k1_wellord2(X1))), inference(rw,[status(thm)],[c_0_256, c_0_257])).
cnf(c_0_265, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X3,k1_wellord2(X2)))|~l1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~v1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X3,k1_wellord2(X2)))))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_259])])).
cnf(c_0_266, plain, (g1_orders_2(k1_zfmisc_1(X1),k1_wellord2(k1_zfmisc_1(X1)))=k3_lattice3(k1_lattice3(X1))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_260])])).
cnf(c_0_267, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_268, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(X3,X1)),g1_orders_2(X2,X3))|~l1_orders_2(g1_orders_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,X3))))|~v1_orders_2(g1_orders_2(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_87])])).
cnf(c_0_269, negated_conjecture, (~v4_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0)))|~m1_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_263, c_0_121]), c_0_121]), c_0_264]), c_0_264])).
cnf(c_0_270, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),k3_lattice3(k1_lattice3(X2)))|~r1_tarski(X1,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_266]), c_0_153]), c_0_154]), c_0_174])])])).
cnf(c_0_271, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_272, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X3,k1_wellord2(X2)))|~l1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~v1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X3,k1_wellord2(X2)))))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_259])])).
cnf(c_0_273, negated_conjecture, (~v4_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0)))|~m1_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0)))), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_274, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),k3_lattice3(k1_lattice3(X2)))|~r1_tarski(X1,k1_zfmisc_1(X2))), inference(csr,[status(thm)],[c_0_270, c_0_271])).
cnf(c_0_275, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),k3_lattice3(k1_lattice3(X2)))|~r1_tarski(X1,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_266]), c_0_153]), c_0_154]), c_0_174])])])).
cnf(c_0_276, negated_conjecture, (~m1_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0)))|~r1_tarski(esk2_0,k1_zfmisc_1(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_274])])).
cnf(c_0_277, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),k3_lattice3(k1_lattice3(X2)))|~r1_tarski(X1,k1_zfmisc_1(X2))), inference(csr,[status(thm)],[c_0_275, c_0_271])).
cnf(c_0_278, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_279, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_255])).
cnf(c_0_280, negated_conjecture, (~r1_tarski(esk2_0,k1_zfmisc_1(esk1_0))), inference(csr,[status(thm)],[c_0_276, c_0_277])).
cnf(c_0_281, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_282, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_283, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281]), c_0_282])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 284
# Proof object clause steps            : 207
# Proof object formula steps           : 77
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 55
# Proof object initial formulas used   : 38
# Proof object generating inferences   : 66
# Proof object simplifying inferences  : 145
# Training examples: 119 positive, 1344 negative
# Training: Positive examples begin
cnf(c_0_117, plain, (v3_lattices(k1_lattice3(X1)))).# trainpos
cnf(c_0_273, negated_conjecture, (~v4_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0)))|~m1_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0))))).# trainpos
cnf(c_0_282, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).# trainpos
cnf(c_0_116, plain, (l3_lattices(k1_lattice3(X1)))).# trainpos
cnf(c_0_251, plain, (v1_relat_1(k1_wellord2(X1)))).# trainpos
cnf(c_0_135, plain, (v10_lattices(k1_lattice3(X1)))).# trainpos
cnf(c_0_136, plain, (~v2_struct_0(k1_lattice3(X1)))).# trainpos
cnf(c_0_175, plain, (~v2_struct_0(k3_lattice3(k1_lattice3(X1))))).# trainpos
cnf(c_0_228, plain, (k2_xboole_0(X1,X1)=X1)).# trainpos
cnf(c_0_83, plain, (m1_subset_1(esk12_1(X1),X1))).# trainpos
cnf(c_0_154, plain, (v1_orders_2(k3_lattice3(k1_lattice3(X1))))).# trainpos
cnf(c_0_96, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_173, plain, (v3_orders_2(k3_lattice3(k1_lattice3(X1))))).# trainpos
cnf(c_0_153, plain, (l1_orders_2(k3_lattice3(k1_lattice3(X1))))).# trainpos
cnf(c_0_271, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_55, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).# trainpos
cnf(c_0_54, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).# trainpos
cnf(c_0_107, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).# trainpos
cnf(c_0_145, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_249, plain, (k2_wellord1(X1,X2)=k1_toler_1(X1,X2)|~v1_relat_1(X1))).# trainpos
cnf(c_0_281, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_225, plain, (k2_xboole_0(k9_xtuple_0(X1),k10_xtuple_0(X1))=k1_relat_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_250, plain, (k2_wellord1(k1_wellord2(X2),X1)=k1_wellord2(X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_199, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_201, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_66, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=X1|~v1_orders_2(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_197, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_169, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_178, plain, (v2_struct_0(X1)|v1_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).# trainpos
cnf(c_0_179, plain, (v2_struct_0(X1)|l1_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).# trainpos
cnf(c_0_93, plain, (k2_relset_1(X3,X4)=X3|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(k2_tarski(k2_tarski(X1,esk7_3(X2,X3,X4)),k1_tarski(X1)),X4))).# trainpos
cnf(c_0_115, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|X1!=k1_lattice3(X2)|~v3_lattices(X1)|~l3_lattices(X1))).# trainpos
cnf(c_0_56, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1))).# trainpos
cnf(c_0_244, plain, (X2=k1_wellord2(X1)|r2_hidden(esk5_2(X1,X2),X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2))).# trainpos
cnf(c_0_236, plain, (X2=k1_wellord2(X1)|r2_hidden(esk6_2(X1,X2),X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2))).# trainpos
cnf(c_0_125, plain, (X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2)|~r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|~r2_hidden(k2_tarski(k2_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),k1_tarski(esk5_2(X1,X2))),X2))).# trainpos
cnf(c_0_133, plain, (g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))=k3_lattice3(X1)|v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1))).# trainpos
cnf(c_0_248, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~l1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_59, plain, (k1_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,X2,X3),X4),k1_tarski(esk9_3(X1,X2,X3))),X3))).# trainpos
cnf(c_0_262, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~l1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_148, plain, (k2_relset_1(X2,X3)=X2|r2_hidden(esk7_3(X1,X2,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_97, plain, (k1_relset_1(X2,X3)=X2|r2_hidden(esk9_3(X1,X2,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).# trainpos
cnf(c_0_191, plain, (r1_tarski(X2,X3)|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X1))))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X1)))))).# trainpos
cnf(c_0_183, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~r1_tarski(X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(X3))))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(k1_lattice3(X3)))))).# trainpos
cnf(c_0_81, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r3_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_217, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_82, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_53, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),u1_orders_2(X1))|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_86, plain, (r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),u1_orders_2(X3)))).# trainpos
cnf(c_0_181, plain, (X2=k1_wellord2(X1)|r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|r2_hidden(k2_tarski(k2_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),k1_tarski(esk5_2(X1,X2))),X2)|k1_relat_1(X2)!=X1|~v1_relat_1(X2))).# trainpos
cnf(c_0_259, plain, (k1_toler_1(k1_wellord2(X1),X2)=k1_wellord2(X2)|~r1_tarski(X2,X1))).# trainpos
cnf(c_0_124, plain, (u1_struct_0(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1))).# trainpos
cnf(c_0_134, plain, (u1_struct_0(k1_lattice3(X1))=k1_zfmisc_1(X1))).# trainpos
cnf(c_0_60, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).# trainpos
cnf(c_0_95, plain, (r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_65, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_163, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_68, plain, (k1_relset_1(X1,X2)=X1|~r2_hidden(k2_tarski(k1_tarski(esk9_3(X3,X1,X2)),k2_tarski(esk9_3(X3,X1,X2),X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_112, plain, (m1_subset_1(esk9_3(X1,X2,X3),X2)|k1_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).# trainpos
cnf(c_0_152, plain, (g1_orders_2(k1_zfmisc_1(X1),k2_lattice3(k1_lattice3(X1)))=k3_lattice3(k1_lattice3(X1)))).# trainpos
cnf(c_0_140, plain, (u1_struct_0(X1)=X2|X1!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v1_orders_2(X1))).# trainpos
cnf(c_0_180, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),u1_orders_2(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_75, plain, (u1_orders_2(X1)=X2|X1!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v1_orders_2(X1))).# trainpos
cnf(c_0_69, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).# trainpos
cnf(c_0_162, plain, (u1_orders_2(X1)=k1_wellord2(X2)|k1_relat_1(u1_orders_2(X1))!=X2|~v1_relat_1(u1_orders_2(X1))|~r1_tarski(esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~r1_orders_2(X1,esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~l1_orders_2(X1)|~m1_subset_1(esk6_2(X2,u1_orders_2(X1)),u1_struct_0(X1))|~m1_subset_1(esk5_2(X2,u1_orders_2(X1)),u1_struct_0(X1)))).# trainpos
cnf(c_0_156, plain, (u1_struct_0(g1_orders_2(X1,X2))=X1|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).# trainpos
cnf(c_0_87, plain, (u1_orders_2(g1_orders_2(X1,X2))=X2|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).# trainpos
cnf(c_0_128, plain, (k2_relset_1(X1,u1_orders_2(X2))=X1|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r1_orders_2(X2,esk7_3(X3,X1,u1_orders_2(X2)),esk7_3(X3,X1,u1_orders_2(X2)))|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2)))).# trainpos
cnf(c_0_80, plain, (k1_relset_1(X1,u1_orders_2(X2))=X1|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r1_orders_2(X2,esk9_3(X3,X1,u1_orders_2(X2)),esk9_3(X3,X1,u1_orders_2(X2)))|~l1_orders_2(X2)|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2)))).# trainpos
cnf(c_0_195, plain, (u1_orders_2(X1)=k1_wellord2(X2)|v2_struct_0(X1)|k1_relat_1(u1_orders_2(X1))!=X2|~l1_orders_2(X1)|~r1_tarski(esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~m1_subset_1(esk6_2(X2,u1_orders_2(X1)),u1_struct_0(X1))|~m1_subset_1(esk5_2(X2,u1_orders_2(X1)),u1_struct_0(X1))|~r3_orders_2(X1,esk5_2(X2,u1_orders_2(X1)),esk6_2(X2,u1_orders_2(X1)))|~v3_orders_2(X1))).# trainpos
cnf(c_0_174, plain, (u1_struct_0(k3_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1))).# trainpos
cnf(c_0_147, plain, (k2_relset_1(X1,u1_orders_2(X2))=X1|v2_struct_0(X2)|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r3_orders_2(X2,esk7_3(X3,X1,u1_orders_2(X2)),esk7_3(X3,X1,u1_orders_2(X2)))|~v3_orders_2(X2))).# trainpos
cnf(c_0_98, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,X2))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),X2)|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~v1_orders_2(g1_orders_2(X1,X2)))).# trainpos
cnf(c_0_172, plain, (u1_orders_2(k3_lattice3(k1_lattice3(X1)))=k2_lattice3(k1_lattice3(X1)))).# trainpos
cnf(c_0_94, plain, (k1_relset_1(X1,u1_orders_2(X2))=X1|v2_struct_0(X2)|~l1_orders_2(X2)|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r3_orders_2(X2,esk9_3(X3,X1,u1_orders_2(X2)),esk9_3(X3,X1,u1_orders_2(X2)))|~v3_orders_2(X2))).# trainpos
cnf(c_0_177, plain, (u1_orders_2(k3_lattice3(X1))=k2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).# trainpos
cnf(c_0_258, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(X3,X1)),g1_orders_2(X2,X3))|~l1_orders_2(g1_orders_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,X3))))|~v1_orders_2(g1_orders_2(X2,X3)))).# trainpos
cnf(c_0_268, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(X3,X1)),g1_orders_2(X2,X3))|~l1_orders_2(g1_orders_2(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,X3))))|~v1_orders_2(g1_orders_2(X2,X3)))).# trainpos
cnf(c_0_203, plain, (r1_tarski(X1,X2)|~r3_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_196, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_265, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X3,k1_wellord2(X2)))|~l1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~v1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X3,k1_wellord2(X2)))))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_272, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X3,k1_wellord2(X2)))|~l1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~v1_orders_2(g1_orders_2(X3,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X3,k1_wellord2(X2)))))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_193, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_207, plain, (v1_relat_1(k2_lattice3(k1_lattice3(X1))))).# trainpos
cnf(c_0_208, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~r1_tarski(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_113, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~r1_orders_2(X2,X3,X4)|~l1_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).# trainpos
cnf(c_0_204, plain, (g1_orders_2(u1_struct_0(k3_lattice3(X1)),k2_lattice3(X1))=k3_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).# trainpos
cnf(c_0_132, plain, (r1_orders_2(g1_orders_2(X1,X3),X4,X5)|~r1_orders_2(g1_orders_2(X2,X3),X4,X5)|~l1_orders_2(g1_orders_2(X1,X3))|~l1_orders_2(g1_orders_2(X2,X3))|~v1_orders_2(g1_orders_2(X1,X3))|~m1_subset_1(X5,u1_struct_0(g1_orders_2(X1,X3)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X3)))|~m1_subset_1(X5,u1_struct_0(g1_orders_2(X2,X3)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X2,X3)))|~v1_orders_2(g1_orders_2(X2,X3)))).# trainpos
cnf(c_0_213, plain, (v4_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_211, plain, (v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_151, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X4)|v2_struct_0(g1_orders_2(X5,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X5,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X5,X2))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X5,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X5,X2)))|~r3_orders_2(g1_orders_2(X5,X2),X3,X4)|~v3_orders_2(g1_orders_2(X5,X2)))).# trainpos
cnf(c_0_131, plain, (k1_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X2))))).# trainpos
cnf(c_0_150, plain, (k1_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_253, plain, (m1_subset_1(esk5_2(X1,X2),X1)|X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2))).# trainpos
cnf(c_0_243, plain, (m1_subset_1(esk6_2(X1,X2),X1)|X2=k1_wellord2(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2))).# trainpos
cnf(c_0_171, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(X2)),X3,X4)|v2_struct_0(X2)|~r3_orders_2(X2,X3,X4)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~l1_orders_2(X2)|~v3_orders_2(X2)|~v1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v1_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).# trainpos
cnf(c_0_188, plain, (v1_relat_1(X2)|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).# trainpos
cnf(c_0_206, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|r1_tarski(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~v1_relat_1(k2_lattice3(k1_lattice3(X1))))).# trainpos
cnf(c_0_224, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X1)),esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2)).# trainpos
cnf(c_0_170, plain, (k1_relset_1(u1_struct_0(g1_orders_2(X1,X2)),X2)=u1_struct_0(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).# trainpos
cnf(c_0_187, plain, (u1_struct_0(g1_orders_2(X1,X2))=k9_xtuple_0(X2)|v2_struct_0(g1_orders_2(X1,X2))|~v4_relat_1(X2,u1_struct_0(g1_orders_2(X1,X2)))|~v1_relat_1(X2)|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).# trainpos
cnf(c_0_167, plain, (m1_subset_1(esk7_3(X1,X2,X3),X2)|k2_relset_1(X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_212, plain, (k9_xtuple_0(u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v4_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1))).# trainpos
cnf(c_0_185, plain, (k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1)))))).# trainpos
cnf(c_0_227, plain, (k9_xtuple_0(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1))).# trainpos
cnf(c_0_198, plain, (k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_190, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))),X3,X4)|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_210, plain, (k10_xtuple_0(u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_202, plain, (u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))=k1_wellord2(X3)|k1_relat_1(u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))!=X3|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~r1_tarski(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))))|~m1_subset_1(esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(esk6_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),k1_zfmisc_1(X2))|~m1_subset_1(esk5_2(X3,u1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))),k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_223, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),esk5_2(X2,k2_lattice3(k1_lattice3(X1))),esk6_2(X2,k2_lattice3(k1_lattice3(X1))))|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_226, plain, (k10_xtuple_0(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1))).# trainpos
cnf(c_0_235, plain, (k1_relat_1(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1))).# trainpos
cnf(c_0_242, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_zfmisc_1(X1)!=X2|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_252, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(k1_zfmisc_1(X1),k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_260, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_266, plain, (g1_orders_2(k1_zfmisc_1(X1),k1_wellord2(k1_zfmisc_1(X1)))=k3_lattice3(k1_lattice3(X1)))).# trainpos
cnf(c_0_275, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),k3_lattice3(k1_lattice3(X2)))|~r1_tarski(X1,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).# trainpos
cnf(c_0_270, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),k3_lattice3(k1_lattice3(X2)))|~r1_tarski(X1,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).# trainpos
cnf(c_0_276, negated_conjecture, (~m1_yellow_0(g1_orders_2(esk2_0,k1_wellord2(esk2_0)),k3_lattice3(k1_lattice3(esk1_0)))|~r1_tarski(esk2_0,k1_zfmisc_1(esk1_0)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_102, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_202, plain, (k5_numbers=k4_ordinal1)).#trainneg
cnf(i_0_151, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_195, plain, (v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_212, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_225, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_226, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_283, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_208, plain, (v1_relat_1(k1_wellord2(k1_xboole_0)))).#trainneg
cnf(i_0_207, plain, (v1_xboole_0(k1_wellord2(k1_xboole_0)))).#trainneg
cnf(i_0_201, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_236, plain, (~v2_struct_0(X1)|~l1_orders_2(X1)|~v1_lattice3(X1))).#trainneg
cnf(i_0_91, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_211, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_210, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_209, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_104, plain, (l1_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_127, plain, (l2_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_68, plain, (v1_relat_2(k1_wellord2(X1)))).#trainneg
cnf(i_0_67, plain, (v4_relat_2(k1_wellord2(X1)))).#trainneg
cnf(i_0_106, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_205, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_231, plain, (~v2_struct_0(X1)|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_66, plain, (v8_relat_2(k1_wellord2(X1)))).#trainneg
cnf(i_0_287, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_154, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_128, plain, (l1_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_230, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_96, plain, (v1_xboole_0(esk11_1(X1)))).#trainneg
cnf(i_0_125, plain, (v4_lattice3(k1_lattice3(X1)))).#trainneg
cnf(i_0_215, plain, (v17_lattices(k1_lattice3(X1)))).#trainneg
cnf(i_0_114, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_82, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).#trainneg
cnf(i_0_92, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_98, plain, (k3_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_273, plain, (v3_card_1(k1_tarski(X1),np__1))).#trainneg
cnf(i_0_65, plain, (v1_partfun1(k1_wellord2(X1),X1))).#trainneg
cnf(i_0_218, plain, (~v1_subset_1(esk14_1(X1),X1))).#trainneg
cnf(i_0_133, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_198, plain, (l1_struct_0(X1)|~l2_lattices(X1))).#trainneg
cnf(i_0_158, plain, (l1_struct_0(X1)|~l1_lattices(X1))).#trainneg
cnf(i_0_229, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_307, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_306, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_295, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_238, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_316, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_121, plain, (v4_orders_2(k3_lattice3(k1_lattice3(X1))))).#trainneg
cnf(i_0_120, plain, (v5_orders_2(k3_lattice3(k1_lattice3(X1))))).#trainneg
cnf(i_0_94, plain, (v3_lattice3(k3_lattice3(k1_lattice3(X1))))).#trainneg
cnf(i_0_222, plain, (v1_relat_1(esk15_2(X1,X2)))).#trainneg
cnf(i_0_337, plain, (v1_relat_1(esk21_2(X1,X2)))).#trainneg
cnf(i_0_223, plain, (v1_xboole_0(esk15_2(X1,X2)))).#trainneg
cnf(i_0_334, plain, (v1_funct_1(esk21_2(X1,X2)))).#trainneg
cnf(i_0_97, plain, (m1_subset_1(esk11_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_219, plain, (m1_subset_1(esk14_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_52, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).#trainneg
cnf(i_0_206, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_303, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk18_1(X1)))).#trainneg
cnf(i_0_203, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_268, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_wellord2(X1)))).#trainneg
cnf(i_0_270, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk17_1(X1)))).#trainneg
cnf(i_0_156, plain, (v1_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_155, plain, (v2_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_142, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_55, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1))).#trainneg
cnf(i_0_220, plain, (v5_relat_1(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_335, plain, (v5_relat_1(esk21_2(X1,X2),X2))).#trainneg
cnf(i_0_221, plain, (v4_relat_1(esk15_2(X1,X2),X1))).#trainneg
cnf(i_0_336, plain, (v4_relat_1(esk21_2(X1,X2),X1))).#trainneg
cnf(i_0_25, plain, (v1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_79, plain, (v3_orders_2(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_78, plain, (v4_orders_2(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_77, plain, (v5_orders_2(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_24, plain, (l1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_240, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_83, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_310, plain, (v1_xboole_0(X1)|v3_card_1(esk19_1(X1),np__1))).#trainneg
cnf(i_0_286, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_304, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk18_1(X1)))).#trainneg
cnf(i_0_119, plain, (v1_funct_1(u2_lattices(X1))|~l2_lattices(X1))).#trainneg
cnf(i_0_252, plain, (v6_relat_2(X1)|~v1_relat_1(X1)|~v2_wellord1(X1))).#trainneg
cnf(i_0_250, plain, (v1_wellord1(X1)|~v1_relat_1(X1)|~v2_wellord1(X1))).#trainneg
cnf(i_0_176, plain, (v1_funct_1(u1_lattices(X1))|~l1_lattices(X1))).#trainneg
cnf(i_0_278, plain, (v6_relat_2(k1_wellord2(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_308, plain, (v1_finset_1(k1_wellord2(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_152, plain, (m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_315, plain, (m1_yellow_0(esk20_1(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_105, plain, (v1_relat_1(k2_wellord1(X1,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_199, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_200, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_64, plain, (m1_subset_1(k1_wellord2(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_254, plain, (v1_relat_2(X1)|~v1_relat_1(X1)|~v2_wellord1(X1))).#trainneg
cnf(i_0_253, plain, (v4_relat_2(X1)|~v1_relat_1(X1)|~v2_wellord1(X1))).#trainneg
cnf(i_0_251, plain, (v8_relat_2(X1)|~v1_relat_1(X1)|~v2_wellord1(X1))).#trainneg
cnf(i_0_274, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_276, plain, (v2_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_265, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_237, plain, (v1_relat_2(u1_orders_2(X1))|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_63, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1))).#trainneg
cnf(i_0_17, plain, (k1_relat_1(X1)=X2|X1!=k1_wellord2(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_275, plain, (v3_yellow_0(X1)|~l1_orders_2(X1)|~v1_yellow_0(X1)|~v2_yellow_0(X1))).#trainneg
cnf(i_0_197, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_89, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_193, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_192, plain, (v2_struct_0(X1)|v2_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_139, plain, (v2_struct_0(X1)|v4_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_138, plain, (v2_struct_0(X1)|v5_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_271, plain, (v1_xboole_0(X1)|m1_subset_1(esk17_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_305, plain, (v1_xboole_0(X1)|m1_subset_1(esk18_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_311, plain, (v1_xboole_0(X1)|m1_subset_1(esk19_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_224, plain, (m1_subset_1(esk15_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_338, plain, (m1_subset_1(esk21_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_217, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_228, plain, (k9_subset_1(X2,X3,X3)=X3|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_131, plain, (v3_orders_2(X1)|~v3_orders_2(X2)|~l1_orders_2(X2)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_143, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_137, plain, (v2_struct_0(X1)|v6_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_116, plain, (v4_orders_2(X1)|~v4_orders_2(X2)|~l1_orders_2(X2)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_136, plain, (v2_struct_0(X1)|v7_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_190, plain, (v2_struct_0(X1)|v13_lattices(X1)|~l3_lattices(X1)|~v15_lattices(X1))).#trainneg
cnf(i_0_183, plain, (v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v2_struct_0(k3_lattice3(X1)))).#trainneg
cnf(i_0_234, plain, (v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v1_xboole_0(esk16_1(X1)))).#trainneg
cnf(i_0_135, plain, (v2_struct_0(X1)|v8_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_134, plain, (v2_struct_0(X1)|v9_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_266, plain, (v2_struct_0(X1)|v3_yellow_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_189, plain, (v2_struct_0(X1)|v14_lattices(X1)|~l3_lattices(X1)|~v15_lattices(X1))).#trainneg
cnf(i_0_141, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_302, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_257, plain, (v2_struct_0(X1)|v15_lattices(X1)|~l3_lattices(X1)|~v17_lattices(X1))).#trainneg
cnf(i_0_258, plain, (v2_struct_0(X1)|v11_lattices(X1)|~l3_lattices(X1)|~v17_lattices(X1))).#trainneg
cnf(i_0_256, plain, (v2_struct_0(X1)|v16_lattices(X1)|~l3_lattices(X1)|~v17_lattices(X1))).#trainneg
cnf(i_0_264, plain, (v1_partfun1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_orders_2(X1))).#trainneg
cnf(i_0_59, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_101, plain, (v5_orders_2(X1)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_178, plain, (v1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_227, plain, (v2_struct_0(X1)|v1_relat_1(k8_filter_1(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_213, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_153, plain, (k8_filter_1(X1)=k2_lattice3(X1)|v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_300, plain, (v2_struct_0(X1)|v15_lattices(X1)|~l3_lattices(X1)|~v13_lattices(X1)|~v14_lattices(X1))).#trainneg
cnf(i_0_243, plain, (v8_relat_2(u1_orders_2(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)|~v2_orders_2(X1))).#trainneg
cnf(i_0_260, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_290, plain, (v2_struct_0(X1)|v12_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v11_lattices(X1))).#trainneg
cnf(i_0_93, plain, (m1_subset_1(k1_toler_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_177, plain, (l1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_22, plain, (v2_struct_0(X1)|v3_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_21, plain, (v2_struct_0(X1)|v4_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_20, plain, (v2_struct_0(X1)|v5_orders_2(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_214, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(k9_subset_1(X2,X3,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_47, plain, (v2_struct_0(X1)|v1_relat_2(k2_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_46, plain, (v2_struct_0(X1)|v4_relat_2(k2_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_165, plain, (v2_wellord1(X1)|~v1_relat_1(X1)|~v1_relat_2(X1)|~v4_relat_2(X1)|~v8_relat_2(X1)|~v6_relat_2(X1)|~v1_wellord1(X1))).#trainneg
cnf(i_0_249, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_272, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_109, plain, (v2_struct_0(X1)|v1_lattice3(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_108, plain, (v2_struct_0(X1)|v2_lattice3(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_45, plain, (v2_struct_0(X1)|v8_relat_2(k2_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_86, plain, (k9_subset_1(X2,X1,X3)=k9_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_233, plain, (v2_struct_0(X1)|v18_lattices(esk16_1(X1),X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_232, plain, (v2_struct_0(X1)|v19_lattices(esk16_1(X1),X1)|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_74, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_81, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_90, plain, (g3_lattices(u1_struct_0(X1),u2_lattices(X1),u1_lattices(X1))=X1|~v3_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_118, plain, (v1_funct_2(u2_lattices(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l2_lattices(X1))).#trainneg
cnf(i_0_175, plain, (v1_funct_2(u1_lattices(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l1_lattices(X1))).#trainneg
cnf(i_0_48, plain, (v2_struct_0(X1)|v1_partfun1(k2_lattice3(X1),u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_280, plain, (v6_orders_2(X2,X1)|~l1_orders_2(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_157, plain, (k4_subset_1(X2,X1,X1)=X1|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_184, plain, (v2_struct_0(X1)|v3_lattice3(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1))).#trainneg
cnf(i_0_146, plain, (v2_struct_0(X1)|v1_yellow_0(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)|~v13_lattices(X1))).#trainneg
cnf(i_0_169, plain, (v2_struct_0(X1)|v2_yellow_0(k3_lattice3(X1))|~l3_lattices(X1)|~v10_lattices(X1)|~v14_lattices(X1))).#trainneg
cnf(i_0_235, plain, (v2_struct_0(X1)|m1_subset_1(esk16_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_312, plain, (v4_relat_1(k3_xboole_0(X1,X3),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_289, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),X2),np__1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_87, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_239, plain, (v3_card_1(k9_xtuple_0(X2),X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_card_1(X1)|~v3_card_1(X2,X1))).#trainneg
cnf(i_0_281, plain, (v2_struct_0(X1)|v10_lattices(X1)|~l3_lattices(X1)|~v4_lattices(X1)|~v5_lattices(X1)|~v6_lattices(X1)|~v7_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1))).#trainneg
cnf(i_0_168, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_167, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_204, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5, plain, (X2=k1_lattice3(X1)|m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(X1))|u1_struct_0(X2)!=k1_zfmisc_1(X1)|~v3_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_4, plain, (X2=k1_lattice3(X1)|m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(X1))|u1_struct_0(X2)!=k1_zfmisc_1(X1)|~v3_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_117, plain, (m1_subset_1(u2_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_174, plain, (m1_subset_1(u1_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_43, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_75, plain, (k4_subset_1(X2,X1,X3)=k4_subset_1(X2,X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_44, plain, (v2_struct_0(X1)|m1_subset_1(k2_lattice3(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l3_lattices(X1)|~v10_lattices(X1))).#trainneg
cnf(i_0_293, plain, (v3_orders_2(g1_orders_2(X1,X2))|~v1_relat_2(X2)|~v1_partfun1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_298, plain, (v4_orders_2(g1_orders_2(X1,X2))|~v8_relat_2(X2)|~v1_partfun1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_296, plain, (v5_orders_2(g1_orders_2(X1,X2))|~v4_relat_2(X2)|~v1_partfun1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_15, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),X3)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)|~r1_tarski(X1,X2)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_314, plain, (v2_struct_0(X2)|v18_lattices(X1,X2)|~l3_lattices(X2)|~v10_lattices(X2)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_3, plain, (X1=k1_lattice3(X2)|u1_struct_0(X1)!=k1_zfmisc_1(X2)|k1_binop_1(u2_lattices(X1),esk3_2(X2,X1),esk4_2(X2,X1))!=k4_subset_1(X2,esk3_2(X2,X1),esk4_2(X2,X1))|k1_binop_1(u1_lattices(X1),esk3_2(X2,X1),esk4_2(X2,X1))!=k9_subset_1(X2,esk3_2(X2,X1),esk4_2(X2,X1))|~v3_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_313, plain, (v2_struct_0(X2)|v19_lattices(X1,X2)|~l3_lattices(X2)|~v10_lattices(X2)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_261, plain, (v2_struct_0(X1)|v21_lattices(X2,X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v19_lattices(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_277, plain, (v2_struct_0(X1)|v20_lattices(X2,X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v18_lattices(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7, plain, (k1_binop_1(u2_lattices(X1),X2,X3)=k4_subset_1(X4,X2,X3)|X1!=k1_lattice3(X4)|~v3_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X3,k1_zfmisc_1(X4))|~m1_subset_1(X2,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_16, plain, (r1_tarski(X1,X2)|X3!=k1_wellord2(X4)|~v1_relat_1(X3)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),X3))).#trainneg
cnf(i_0_6, plain, (k1_binop_1(u1_lattices(X1),X2,X3)=k9_subset_1(X4,X2,X3)|X1!=k1_lattice3(X4)|~v3_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X3,k1_zfmisc_1(X4))|~m1_subset_1(X2,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_32, plain, (r2_hidden(k2_tarski(k2_tarski(X3,esk10_4(X4,X1,X2,X3)),k1_tarski(X3)),X2)|k1_relset_1(X1,X2)!=X1|~r2_hidden(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X4))))).#trainneg
cnf(i_0_29, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_4(X4,X1,X2,X3),X3),k1_tarski(esk8_4(X4,X1,X2,X3))),X2)|k2_relset_1(X1,X2)!=X1|~r2_hidden(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X1))))).#trainneg
cnf(i_0_263, plain, (v3_lattices(g3_lattices(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_262, plain, (l3_lattices(g3_lattices(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_285, plain, (v1_xboole_0(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v2_struct_0(g3_lattices(X1,X2,X3))|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_162, plain, (X1=X2|g3_lattices(X3,X4,X1)!=g3_lattices(X5,X6,X2)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_funct_2(X4,k2_zfmisc_1(X3,X3),X3)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3))))).#trainneg
cnf(i_0_163, plain, (X1=X2|g3_lattices(X3,X1,X4)!=g3_lattices(X5,X2,X6)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_funct_2(X4,k2_zfmisc_1(X3,X3),X3)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3))))).#trainneg
cnf(i_0_164, plain, (X1=X2|g3_lattices(X1,X3,X4)!=g3_lattices(X2,X5,X6)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,k2_zfmisc_1(X1,X1),X1)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_339, plain, (esk13_0=k1_xboole_0)).#trainneg
cnf(i_0_666, plain, (X2=k1_wellord2(X1)|~v1_xboole_0(X1)|k1_relat_1(X2)!=X1|~v1_relat_1(X2))).#trainneg
cnf(i_0_340, plain, (k1_wellord2(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_408, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_477, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_374, plain, (~m1_subset_1(k4_ordinal1,k4_ordinal1))).#trainneg
cnf(i_0_343, plain, (l1_struct_0(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_341, plain, (esk11_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_893, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),X3)|v1_xboole_0(X4)|X3!=k1_wellord2(X4)|~r2_hidden(X1,X4)|~v1_relat_1(X3)|~r1_tarski(X1,X2)|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_397, plain, (X1=X2|g1_orders_2(X1,k1_wellord2(X1))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_704, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk15_2(k1_tarski(X1),k1_tarski(X1))),np__1))).#trainneg
cnf(i_0_708, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k1_toler_1(X2,k1_tarski(X1))),np__1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_700, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk21_2(k1_tarski(X1),k1_tarski(X1))),np__1))).#trainneg
cnf(i_0_352, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_375, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk14_1(X2)))).#trainneg
cnf(i_0_376, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_410, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_346, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_406, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_529, plain, (v1_xboole_0(k1_wellord2(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_347, plain, (esk15_2(X1,X2)=k1_xboole_0)).#trainneg
cnf(i_0_348, plain, (k2_xboole_0(k1_xboole_0,X1)=X1)).#trainneg
cnf(i_0_908, plain, (r2_hidden(k2_tarski(k2_tarski(X1,k2_tarski(k2_tarski(X2,X3),k1_tarski(X2))),k1_tarski(X1)),X4)|X4!=k1_wellord2(u1_orders_2(X5))|~r2_hidden(X1,u1_orders_2(X5))|~v1_relat_1(X4)|~r1_tarski(X1,k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)))|~r1_orders_2(X5,X2,X3)|~l1_orders_2(X5)|~m1_subset_1(X3,u1_struct_0(X5))|~m1_subset_1(X2,u1_struct_0(X5)))).#trainneg
cnf(i_0_911, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_396, plain, (k1_wellord2(X1)=X2|g1_orders_2(X1,k1_wellord2(X1))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_414, plain, (k1_relat_1(k1_wellord2(X1))=X1)).#trainneg
cnf(i_0_711, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k1_wellord2(k1_tarski(X1))),np__1))).#trainneg
cnf(i_0_416, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_712, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k9_subset_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X2,X3)),np__1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_721, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k4_subset_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X2,X3)),np__1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_901, plain, (r1_tarski(X1,X2)|X3!=k1_wellord2(X4)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),X3)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~v1_relat_1(X3))).#trainneg
cnf(i_0_379, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_368, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0)).#trainneg
cnf(i_0_449, plain, (X1=X2|g1_orders_2(X1,esk21_2(X1,X1))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_491, plain, (v5_relat_1(k1_wellord2(X1),X1))).#trainneg
cnf(i_0_500, plain, (v4_relat_1(k1_wellord2(X1),X1))).#trainneg
cnf(i_0_358, plain, (esk14_1(k2_zfmisc_1(X1,X1))=X2|g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_359, plain, (esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=X2|g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_541, plain, (X1=X2|g1_orders_2(X1,k1_toler_1(X3,X1))!=g1_orders_2(X2,X4)|~v1_relat_1(X3))).#trainneg
cnf(i_0_437, plain, (v1_relat_1(k1_toler_1(X1,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_544, plain, (v1_xboole_0(k1_toler_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_378, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_461, plain, (k9_subset_1(X1,X2,X2)=X2)).#trainneg
cnf(i_0_849, plain, (v3_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v1_relat_2(u1_orders_2(X1))|~v1_partfun1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_736, plain, (k1_relset_1(X1,X2)=X1|~r2_hidden(k2_tarski(k2_tarski(X4,esk9_3(X3,X1,X2)),k1_tarski(esk9_3(X3,X1,X2))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_579, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_399, plain, (v1_relat_1(esk14_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_381, plain, (X1=X2|g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_382, plain, (X1=X2|g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_695, plain, (v4_relat_1(k1_xboole_0,X2)|~v4_relat_1(X1,X2)|~v1_relat_1(k1_xboole_0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_501, plain, (v13_lattices(X1)|v2_struct_0(X1)|~l3_lattices(X1)|~v17_lattices(X1))).#trainneg
cnf(i_0_502, plain, (v14_lattices(X1)|v2_struct_0(X1)|~l3_lattices(X1)|~v17_lattices(X1))).#trainneg
cnf(i_0_906, plain, (~r2_hidden(u1_orders_2(X3),k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_898, plain, (k1_binop_1(u2_lattices(k1_lattice3(X1)),X2,X3)=k4_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_510, plain, (k3_xboole_0(X1,X2)=k9_subset_1(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_745, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_451, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_580, plain, (m1_subset_1(k3_xboole_0(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_616, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_889, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk9_3(X2,X3,X4)),k1_tarski(X1)),X5)|k1_relset_1(X3,X4)=X3|X5!=k1_wellord2(X3)|~r2_hidden(X1,X3)|~v1_relat_1(X5)|~r1_tarski(X1,esk9_3(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_540, plain, (k1_toler_1(X1,X2)=X3|g1_orders_2(X2,k1_toler_1(X1,X2))!=g1_orders_2(X4,X3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_645, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_404, plain, (v3_ordinal1(esk12_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_413, plain, (l1_orders_2(esk20_1(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_448, plain, (esk21_2(X1,X1)=X2|g1_orders_2(X1,esk21_2(X1,X1))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_353, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk14_1(X1)))).#trainneg
cnf(i_0_354, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk12_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_688, plain, (v3_yellow_0(k3_lattice3(X1))|v2_struct_0(X1)|~v1_yellow_0(k3_lattice3(X1))|~l1_orders_2(k3_lattice3(X1))|~v14_lattices(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_395, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,k1_wellord2(X2)))).#trainneg
cnf(i_0_400, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_626, plain, (k9_subset_1(X1,X3,X2)=k3_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_907, plain, (~v1_xboole_0(u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_904, plain, (k1_binop_1(u1_lattices(k1_lattice3(X1)),X2,X3)=k9_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_829, plain, (u1_struct_0(X1)=X2|v2_struct_0(X1)|g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))!=g1_orders_2(X2,X3)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_910, plain, (k1_relset_1(X1,u1_orders_2(X2))=X1|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r1_orders_2(X2,esk9_3(X3,X1,u1_orders_2(X2)),X4)|~l1_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(X2)),u1_struct_0(X2)))).#trainneg
cnf(i_0_679, negated_conjecture, (k4_subset_1(k1_zfmisc_1(esk1_0),X1,X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_903, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|X3!=k1_wellord2(X4)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~v1_relat_1(X3)|~m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),X3))).#trainneg
cnf(i_0_867, plain, (v4_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v8_relat_2(u1_orders_2(X1))|~v1_partfun1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_662, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_473, plain, (v1_xboole_0(esk14_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_789, plain, (v1_relat_1(u2_lattices(X1))|~l2_lattices(X1))).#trainneg
cnf(i_0_636, plain, (X1=X2|g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X3))!=g1_orders_2(X2,X4)|~v5_relat_1(X3,k2_zfmisc_1(X1,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_403, plain, (v3_ordinal1(esk14_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_801, plain, (v1_relat_1(u1_lattices(X1))|~l1_lattices(X1))).#trainneg
cnf(i_0_394, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X1))|~r2_hidden(X2,k1_wellord2(X1)))).#trainneg
cnf(i_0_734, plain, (~v1_xboole_0(k9_xtuple_0(X1))|~v3_card_1(X1,np__1)|~v1_card_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_682, plain, (k4_subset_1(X1,X2,X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v4_relat_1(X3,X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_683, plain, (k4_subset_1(X1,X2,X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v5_relat_1(X3,X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_411, plain, (v3_ordinal1(k1_wellord2(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_419, plain, (v3_ordinal1(esk17_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_705, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk14_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))),np__1))).#trainneg
cnf(i_0_732, plain, (v1_xboole_0(k4_subset_1(X1,X2,X3))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_905, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_807, plain, (k4_subset_1(X1,X3,X2)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_554, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_627, plain, (m1_subset_1(k9_subset_1(X1,X3,X2),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1002, plain, (X1=u1_struct_0(X2)|g1_orders_2(X1,k1_wellord2(X1))!=X2|~l1_orders_2(X2)|~v1_orders_2(X2))).#trainneg
cnf(i_0_949, plain, (v8_relat_2(k1_xboole_0))).#trainneg
cnf(i_0_950, plain, (v4_relat_2(k1_xboole_0))).#trainneg
cnf(i_0_97, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_420, plain, (esk17_1(k2_zfmisc_1(X1,X1))=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_955, plain, (v1_partfun1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_951, plain, (v1_relat_2(k1_xboole_0))).#trainneg
cnf(i_0_1070, plain, (X1=k1_wellord2(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_421, plain, (X1=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_422, plain, (v1_relat_1(esk17_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1038, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))),np__1))).#trainneg
cnf(i_0_1005, plain, (g1_orders_2(k1_zfmisc_1(X1),u1_orders_2(k1_lattice3(X1)))=k1_lattice3(X1)|~l1_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_952, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_1110, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_424, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk17_1(X2)))).#trainneg
cnf(i_0_1142, plain, (v5_relat_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_953, plain, (v6_relat_2(k1_xboole_0)|~v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_1143, plain, (v4_relat_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_1029, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1165, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk1_0))|v1_xboole_0(esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_425, plain, (v3_ordinal1(esk18_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_426, plain, (esk18_1(k2_zfmisc_1(X1,X1))=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_1001, plain, (X1=u1_struct_0(X2)|v2_struct_0(X2)|g1_orders_2(X1,k1_wellord2(X1))!=k3_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_1151, plain, (X1=u1_struct_0(X2)|v2_struct_0(X2)|g1_orders_2(X1,k1_toler_1(X3,X1))!=k3_lattice3(X2)|~v1_relat_1(X3)|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_1041, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1181, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_221, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_957, plain, (v1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_220, plain, (v5_relat_1(k1_xboole_0,X2))).#trainneg
cnf(i_0_427, plain, (X1=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_958, plain, (v3_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_959, plain, (v4_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_960, plain, (v5_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_428, plain, (v1_relat_1(esk18_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_430, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk18_1(X2)))).#trainneg
cnf(i_0_1394, plain, (u1_struct_0(X1)=X2|v2_struct_0(X1)|k3_lattice3(X1)!=g1_orders_2(X2,X3)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_1012, plain, (m1_subset_1(u1_orders_2(k1_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1))))|~l1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_915, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~r2_hidden(k2_tarski(k2_tarski(X3,X2),k1_tarski(X2)),u1_orders_2(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1215, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_xboole_0(X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_961, plain, (l1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_704, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k1_xboole_0),np__1))).#trainneg
cnf(i_0_431, plain, (v3_ordinal1(esk19_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1339, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1201, plain, (v13_lattices(k1_lattice3(X1)))).#trainneg
cnf(i_0_1338, plain, (v3_yellow_0(k3_lattice3(X1))|v2_struct_0(X1)|~v14_lattices(X1)|~v10_lattices(X1)|~l3_lattices(X1)|~v13_lattices(X1))).#trainneg
cnf(i_0_432, plain, (esk19_1(k2_zfmisc_1(X1,X1))=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_433, plain, (X1=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_947, plain, (u1_struct_0(X1)=X2|g3_lattices(u1_struct_0(X1),X3,u2_lattices(X1))!=g3_lattices(X2,X4,X5)|~v1_funct_2(X3,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_lattices(X1))|~v1_funct_1(X3)|~m1_subset_1(u2_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_934, plain, (r2_hidden(k2_tarski(k1_tarski(esk8_4(X2,X3,X4,X1)),k2_tarski(X1,esk8_4(X2,X3,X4,X1))),X4)|k2_relset_1(X3,X4)!=X3|~r2_hidden(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_967, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l3_lattices(X1))).#trainneg
cnf(i_0_946, plain, (u1_struct_0(X1)=X2|g3_lattices(u1_struct_0(X1),X3,u1_lattices(X1))!=g3_lattices(X2,X4,X5)|~v1_funct_2(X3,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u1_lattices(X1))|~v1_funct_1(X3)|~m1_subset_1(u1_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_1115, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k9_subset_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X3,X2)),np__1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_1009, plain, (g3_lattices(k1_zfmisc_1(X1),u2_lattices(k1_lattice3(X1)),u1_lattices(k1_lattice3(X1)))=k1_lattice3(X1))).#trainneg
cnf(i_0_1043, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_1202, plain, (v14_lattices(k1_lattice3(X1)))).#trainneg
cnf(i_0_1003, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,k1_xboole_0)!=g1_orders_2(X1,X2))).#trainneg
cnf(i_0_434, plain, (v1_relat_1(esk19_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1247, plain, (u1_struct_0(X1)=k1_xboole_0|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_446, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,esk21_2(X1,X2)))).#trainneg
cnf(i_0_380, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk19_1(X1)))).#trainneg
cnf(i_0_1071, plain, (k1_wellord2(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1417, plain, (X1=esk14_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_436, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk19_1(X2)))).#trainneg
cnf(i_0_445, plain, (v3_ordinal1(esk21_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1004, plain, (v1_partfun1(u1_orders_2(k1_lattice3(X1)),k1_zfmisc_1(X1))|~v2_orders_2(k1_lattice3(X1))|~l1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_1023, plain, (k2_relset_1(X1,u1_orders_2(X2))=X1|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~r1_orders_2(X2,X3,esk7_3(X4,X1,u1_orders_2(X2)))|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X4,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_1158, negated_conjecture, (m1_subset_1(esk9_3(X1,esk2_0,X2),k1_zfmisc_1(esk1_0))|k1_relset_1(esk2_0,X2)=esk2_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk2_0,X1))))).#trainneg
cnf(i_0_1221, plain, (k9_subset_1(X4,X2,X1)=k9_subset_1(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_980, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,esk2_0),k1_tarski(X1)),X2)|X2!=k1_wellord2(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~r2_hidden(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v1_relat_1(X2)|~r1_tarski(X1,esk2_0))).#trainneg
cnf(i_0_1183, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_subset_1(X2,X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1045, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_1418, plain, (esk14_1(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1108, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,k1_xboole_0)!=g1_orders_2(X2,X1))).#trainneg
cnf(i_0_447, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk21_2(X2,X3)))).#trainneg
cnf(i_0_962, plain, (k2_wellord1(k1_xboole_0,X1)=k1_wellord2(X1)|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_611, plain, (~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~r2_hidden(X2,u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1134, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1246, plain, (X1=u1_struct_0(X2)|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1155, plain, (X1=k1_toler_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_452, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(k1_relat_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_453, plain, (v1_xboole_0(k10_xtuple_0(X1))|~v1_xboole_0(k1_relat_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1214, plain, (k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1103, plain, (u1_orders_2(k1_lattice3(X1))=X2|g1_orders_2(k1_zfmisc_1(X1),u1_orders_2(k1_lattice3(X1)))!=g1_orders_2(X3,X2)|~l1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_1448, plain, (m1_subset_1(k2_xboole_0(X3,X2),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1048, plain, (k4_subset_1(X1,X2,X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1175, plain, (k1_relset_1(X1,X2)=X1|~r2_hidden(k2_tarski(k1_tarski(esk9_3(X4,X1,X2)),k2_tarski(X3,esk9_3(X4,X1,X2))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X4))))).#trainneg
cnf(i_0_488, plain, (v5_relat_1(esk14_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_1321, plain, (X2=k1_wellord2(esk14_1(X1))|~v1_xboole_0(X1)|k1_relat_1(X2)!=esk14_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_468, plain, (v1_xboole_0(esk21_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_969, plain, (k1_toler_1(k1_xboole_0,X1)=k1_wellord2(X1)|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_474, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1525, plain, (k1_wellord2(X1)=k1_wellord2(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_478, plain, (v1_xboole_0(k1_wellord2(X1))|~v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_484, plain, (v5_relat_1(esk19_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1516, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1567, negated_conjecture, (v1_xboole_0(esk2_0)|m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_485, plain, (v5_relat_1(esk18_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1769, negated_conjecture, (~m1_yellow_0(g1_orders_2(esk2_0,k1_xboole_0),k3_lattice3(k1_lattice3(esk1_0)))|~v4_yellow_0(g1_orders_2(esk2_0,k1_xboole_0),k3_lattice3(k1_lattice3(esk1_0)))|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1755, plain, (u1_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1548, plain, (X1=k1_zfmisc_1(X2)|g1_orders_2(X1,k1_wellord2(X1))!=k1_lattice3(X2)|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2)))).#trainneg
cnf(i_0_486, plain, (v5_relat_1(esk17_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_489, plain, (v5_relat_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_1625, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v2_struct_0(X1)|k3_lattice3(X1)!=X2|~v10_lattices(X1)|~l3_lattices(X1)|~l1_orders_2(X2)|~v1_orders_2(X2))).#trainneg
cnf(i_0_1739, plain, (u1_struct_0(X1)=X2|g3_lattices(u1_struct_0(X1),u2_lattices(X1),u1_lattices(X1))!=g3_lattices(X2,X3,X4)|~l1_lattices(X1)|~v1_funct_1(u2_lattices(X1))|~m1_subset_1(u2_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_1709, plain, (u1_struct_0(X1)=X2|g3_lattices(u1_struct_0(X1),u1_lattices(X1),u2_lattices(X1))!=g3_lattices(X2,X3,X4)|~v1_funct_1(u1_lattices(X1))|~l2_lattices(X1)|~m1_subset_1(u1_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_1926, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1866, negated_conjecture, (~v1_xboole_0(X1)|X1!=k1_wellord2(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~r2_hidden(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v1_relat_1(X1)|~r1_tarski(X2,esk2_0))).#trainneg
cnf(i_0_493, plain, (v4_relat_1(esk19_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1513, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0))).#trainneg
cnf(i_0_1900, plain, (~v1_subset_1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1776, plain, (v1_finset_1(k1_xboole_0)|~v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_494, plain, (v4_relat_1(esk18_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_495, plain, (v4_relat_1(esk17_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1743, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k3_xboole_0(X3,X2)),np__1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_2029, plain, (k1_relset_1(X1,X2)=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|X2!=k1_wellord2(X1)|~r2_hidden(esk9_3(X3,X1,X2),X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1574, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X3,X2)),u1_orders_2(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1498, plain, (v1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1499, plain, (l1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1846, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|esk21_2(X1,X2)!=k1_wellord2(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~r2_hidden(X3,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~r1_tarski(X3,esk2_0))).#trainneg
cnf(i_0_497, plain, (v4_relat_1(esk14_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_1777, plain, (v1_partfun1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1637, plain, (~v2_struct_0(g1_orders_2(k1_zfmisc_1(X1),u1_orders_2(k1_lattice3(X1))))|~l1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_2046, plain, (X1=esk12_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1527, plain, (u1_struct_0(X1)=k1_wellord2(X2)|~v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_498, plain, (v4_relat_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_503, plain, (k2_relset_1(X1,X2)=k2_relset_1(X3,X2)|~v5_relat_1(X2,X1)|~v1_relat_1(X2)|~v5_relat_1(X2,X3))).#trainneg
cnf(i_0_1624, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X2)|k3_lattice3(X1)!=k3_lattice3(X2)|~v10_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_1552, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|X1!=k1_lattice3(X2)|~l1_orders_2(X1)|~v1_orders_2(X1)|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2)))).#trainneg
cnf(i_0_1812, plain, (k2_relset_1(X1,u1_orders_2(X2))=X1|v2_struct_0(X2)|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X4,X1,u1_orders_2(X2)),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(X2),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(X3,u1_struct_0(X2))|~r3_orders_2(X2,X3,esk7_3(X4,X1,u1_orders_2(X2)))|~v3_orders_2(X2))).#trainneg
cnf(i_0_1626, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|v2_struct_0(X1)|k3_lattice3(X1)!=k3_lattice3(k1_lattice3(X2))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_1480, plain, (X1=k1_zfmisc_1(X2)|g1_orders_2(X1,k1_toler_1(X3,X1))!=k3_lattice3(k1_lattice3(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1535, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|X1!=k3_lattice3(k1_lattice3(X2))|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_1821, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k9_subset_1(X4,X2,X3)),np__1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))|~m1_subset_1(X3,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_1750, plain, (k1_zfmisc_1(X1)=k1_xboole_0|k3_lattice3(k1_lattice3(X1))!=g1_orders_2(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_504, plain, (k2_xboole_0(k9_xtuple_0(X1),k2_relset_1(X2,X1))=k1_relat_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_1504, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k1_xboole_0,X1))).#trainneg
cnf(i_0_1946, plain, (~v3_card_1(X1,np__1)|~v1_card_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(k1_relat_1(X1)))).#trainneg
cnf(i_0_1500, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_505, plain, (k2_xboole_0(k1_relset_1(X2,X1),k10_xtuple_0(X1))=k1_relat_1(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_506, plain, (k1_relset_1(X1,X2)=k1_relset_1(X3,X2)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X3))).#trainneg
cnf(i_0_1969, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k3_xboole_0(X2,X3)),np__1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_1814, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,esk9_3(X2,esk2_0,X3)),k1_tarski(X1)),X4)|k1_relset_1(esk2_0,X3)=esk2_0|X4!=k1_wellord2(k1_zfmisc_1(esk1_0))|~r2_hidden(X1,k1_zfmisc_1(esk1_0))|~v1_relat_1(X4)|~r1_tarski(X1,esk9_3(X2,esk2_0,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk2_0,X2))))).#trainneg
cnf(i_0_1735, plain, (k1_zfmisc_1(X1)=X2|g3_lattices(k1_zfmisc_1(X1),X3,u1_lattices(k1_lattice3(X1)))!=g3_lattices(X2,X4,X5)|~l1_lattices(k1_lattice3(X1))|~v1_funct_2(X3,k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_1506, plain, (k1_xboole_0=X1|g1_orders_2(X2,k1_xboole_0)!=g1_orders_2(X3,X1))).#trainneg
cnf(i_0_507, plain, (k9_subset_1(X2,X1,k1_xboole_0)=k1_xboole_0|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1630, plain, (~v1_xboole_0(k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)))|~r2_hidden(X2,u1_orders_2(k1_lattice3(X1)))|~l1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_1269, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1240, plain, (v1_xboole_0(k1_wellord2(X1))|m1_subset_1(X2,k2_zfmisc_1(X1,X1))|~m1_subset_1(X2,k1_wellord2(X1)))).#trainneg
cnf(i_0_1580, plain, (r1_orders_2(X1,X2,esk10_4(X3,X4,u1_orders_2(X1),X2))|~l1_orders_2(X1)|~m1_subset_1(esk10_4(X3,X4,u1_orders_2(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|k1_relset_1(X4,u1_orders_2(X1))!=X4|~r2_hidden(X2,X4)|~m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(X4,X3))))).#trainneg
cnf(i_0_1975, plain, (m1_subset_1(k3_xboole_0(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1492, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=k3_lattice3(k1_lattice3(X2))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2047, plain, (esk12_1(k1_zfmisc_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_512, plain, (v1_orders_2(g1_orders_2(X1,esk21_2(X1,X1))))).#trainneg
cnf(i_0_1250, plain, (v3_ordinal1(k3_xboole_0(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_513, plain, (v1_orders_2(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_514, plain, (v1_orders_2(g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_1706, plain, (k1_zfmisc_1(X1)=X2|g3_lattices(k1_zfmisc_1(X1),X3,u2_lattices(k1_lattice3(X1)))!=g3_lattices(X2,X4,X5)|~v1_funct_2(X3,k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1))|~v1_funct_1(X3)|~l2_lattices(k1_lattice3(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_1820, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k9_subset_1(X4,X2,X3)),np__1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))|~m1_subset_1(X3,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_1159, negated_conjecture, (m1_subset_1(esk7_3(X1,esk2_0,X2),k1_zfmisc_1(esk1_0))|k2_relset_1(esk2_0,X2)=esk2_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0))))).#trainneg
cnf(i_0_1358, plain, (k9_subset_1(X1,X2,X3)=k9_subset_1(X4,X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_2118, plain, (u1_struct_0(g1_orders_2(X1,k1_wellord2(X1)))=X1)).#trainneg
cnf(i_0_2223, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_2161, negated_conjecture, (v1_xboole_0(esk12_1(esk2_0))|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_2214, plain, (~v1_xboole_0(X1)|k1_relset_1(X4,k1_xboole_0)!=X4|~r2_hidden(X2,X4))).#trainneg
cnf(i_0_515, plain, (v1_orders_2(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_2363, plain, (~v3_card_1(k1_xboole_0,np__1)|~v1_card_1(np__1)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2158, negated_conjecture, (v3_ordinal1(esk12_1(esk2_0))|v1_xboole_0(esk2_0)|~v3_ordinal1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_517, plain, (v1_orders_2(g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_518, plain, (v1_orders_2(g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_2114, plain, (u1_struct_0(k3_lattice3(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2139, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),X4)|~r1_orders_2(g1_orders_2(X3,X4),X1,X2)|~l1_orders_2(g1_orders_2(X3,X4))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X3,X4)))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X3,X4)))|~v1_orders_2(g1_orders_2(X3,X4)))).#trainneg
cnf(i_0_2274, plain, (k1_relset_1(X1,k1_wellord2(X1))=X1)).#trainneg
cnf(i_0_2142, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),X4)|~r1_orders_2(g1_orders_2(X3,X4),X1,X2)|~l1_orders_2(g1_orders_2(X3,X4))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X3,X4)))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X3,X4)))|~v1_orders_2(g1_orders_2(X3,X4)))).#trainneg
cnf(i_0_2165, plain, (g1_orders_2(k1_xboole_0,u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))=g1_orders_2(k1_xboole_0,k1_xboole_0)|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2215, plain, (~v1_xboole_0(X1)|k2_relset_1(X3,k1_xboole_0)!=X3|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_521, plain, (v1_xboole_0(esk21_2(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2164, plain, (v1_partfun1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~v2_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2182, plain, (l1_orders_2(g1_orders_2(k1_xboole_0,u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2061, plain, (v1_partfun1(k1_wellord2(X2),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_522, plain, (v1_xboole_0(esk19_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2123, plain, (m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_orders_2(X1,X2)),u1_struct_0(g1_orders_2(X1,X2)))))|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_2433, negated_conjecture, (k1_relset_1(esk2_0,X1)=esk2_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,X2)))|X1!=k1_wellord2(k1_zfmisc_1(esk1_0))|~r2_hidden(esk9_3(X2,esk2_0,X1),k1_zfmisc_1(esk1_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_2141, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),X4)|~r1_orders_2(g1_orders_2(X3,X4),X2,X1)|~l1_orders_2(g1_orders_2(X3,X4))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X3,X4)))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X3,X4)))|~v1_orders_2(g1_orders_2(X3,X4)))).#trainneg
cnf(i_0_2273, plain, (k1_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|u1_orders_2(X1)!=k1_wellord2(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2198, plain, (u1_struct_0(X1)=X2|X1!=g3_lattices(X2,X3,X4)|~l1_lattices(X1)|~l2_lattices(X1)|~l3_lattices(X1)|~v3_lattices(X1))).#trainneg
cnf(i_0_2448, plain, (k9_subset_1(X1,k1_xboole_0,X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2172, plain, (m1_subset_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2177, plain, (v3_orders_2(g1_orders_2(k1_xboole_0,u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v1_relat_2(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_partfun1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2178, plain, (v4_orders_2(g1_orders_2(k1_xboole_0,u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v8_relat_2(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_partfun1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2075, plain, (v5_relat_1(k1_wellord2(X2),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2076, plain, (v4_relat_1(k1_wellord2(X2),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_526, plain, (v1_xboole_0(esk14_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_527, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2134, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,X2))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),X2)|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~v1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_2454, plain, (k9_subset_1(X1,k1_xboole_0,X2)=k1_xboole_0)).#trainneg
cnf(i_0_2128, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,X2))|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),X2)|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~v1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_2143, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),X4)|~r1_orders_2(g1_orders_2(X3,X4),X1,X2)|~l1_orders_2(g1_orders_2(X3,X4))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X3,X4)))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X3,X4)))|~v1_orders_2(g1_orders_2(X3,X4)))).#trainneg
cnf(i_0_2170, plain, (v1_funct_2(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2534, plain, (k1_zfmisc_1(X1)=k1_xboole_0|g1_orders_2(k1_xboole_0,u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))!=k3_lattice3(k1_lattice3(X1))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_530, plain, (v2_struct_0(X1)|v1_relat_1(k2_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2162, negated_conjecture, (m1_subset_1(X1,esk1_0)|v1_xboole_0(esk2_0)|~r2_hidden(X1,esk12_1(esk2_0)))).#trainneg
cnf(i_0_2171, plain, (v1_funct_2(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2174, plain, (m1_subset_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_533, plain, (v3_ordinal1(k1_relset_1(X2,X1))|~v5_ordinal1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_534, plain, (v3_ordinal1(k1_toler_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_2603, negated_conjecture, (r2_hidden(k2_tarski(k2_tarski(X1,esk7_3(X2,esk2_0,X3)),k1_tarski(X1)),X4)|k2_relset_1(esk2_0,X3)=esk2_0|X4!=k1_wellord2(k1_zfmisc_1(esk1_0))|~r2_hidden(X1,k1_zfmisc_1(esk1_0))|~v1_relat_1(X4)|~r1_tarski(X1,esk7_3(X2,esk2_0,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,esk2_0))))).#trainneg
cnf(i_0_2278, plain, (k1_relset_1(X1,k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4))=X1|k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4)!=k1_wellord2(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2591, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k9_subset_1(X2,X3,k1_toler_1(X4,k1_tarski(X1)))),np__1)|~m1_subset_1(k1_toler_1(X4,k1_tarski(X1)),k1_zfmisc_1(X2))|~v1_relat_1(X4))).#trainneg
cnf(i_0_2600, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k9_subset_1(X2,X3,k4_subset_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X4,X5))),np__1)|~m1_subset_1(k4_subset_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X4,X5),k1_zfmisc_1(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_2649, plain, (g1_orders_2(X1,k2_lattice3(g1_orders_2(X1,k1_wellord2(X1))))=k3_lattice3(g1_orders_2(X1,k1_wellord2(X1)))|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_2766, plain, (k1_relset_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_537, plain, (v1_xboole_0(k1_toler_1(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X2,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_2715, plain, (~v3_card_1(k1_xboole_0,np__1)|~v1_funct_1(k1_xboole_0)|~v1_xboole_0(np__1))).#trainneg
cnf(i_0_2673, plain, (u1_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_2661, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_538, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,k1_toler_1(X3,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_539, plain, (v1_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_2789, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2788, plain, (m1_subset_1(k2_lattice3(k1_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_2723, plain, (g1_orders_2(u1_struct_0(X1),k2_lattice3(k3_lattice3(X1)))=k3_lattice3(k3_lattice3(X1))|v2_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~v10_lattices(k3_lattice3(X1))|~l3_lattices(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2664, plain, (l1_orders_2(g1_orders_2(X1,u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))))).#trainneg
cnf(i_0_2654, plain, (m1_subset_1(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_2647, plain, (g1_orders_2(X1,u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))=g1_orders_2(X1,k1_wellord2(X1)))).#trainneg
cnf(i_0_2705, negated_conjecture, (esk12_1(esk2_0)=k1_xboole_0|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_2711, plain, (k9_xtuple_0(k1_xboole_0)!=X1|~v1_xboole_0(X2)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_543, plain, (v4_relat_1(k1_toler_1(X1,X2),X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3030, plain, (v1_relat_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3031, plain, (v1_xboole_0(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2059, plain, (v6_relat_2(k1_wellord2(X2))|~v3_ordinal1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2060, plain, (v1_finset_1(k1_wellord2(X2))|~v1_finset_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_545, plain, (v5_relat_1(k1_toler_1(X1,X2),X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_546, plain, (l1_orders_2(g1_orders_2(X1,esk21_2(X1,X1))))).#trainneg
cnf(i_0_2758, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),X3)|v2_struct_0(g1_orders_2(X4,X3))|~l1_orders_2(g1_orders_2(X4,X3))|~v1_orders_2(g1_orders_2(X4,X3))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X4,X3)))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X4,X3)))|~r3_orders_2(g1_orders_2(X4,X3),X1,X2)|~v3_orders_2(g1_orders_2(X4,X3)))).#trainneg
cnf(i_0_2756, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_lattice3(k1_lattice3(X3)))|~r1_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1031, plain, (v1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_2850, plain, (k10_xtuple_0(k1_xboole_0)!=X1|~v1_xboole_0(X2)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_547, plain, (l1_orders_2(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_2704, negated_conjecture, (X1=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(X1)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_548, plain, (l1_orders_2(g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_549, plain, (l1_orders_2(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_2721, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(k3_lattice3(X1)))=k3_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2761, plain, (k9_xtuple_0(k1_wellord2(X1))=X1)).#trainneg
cnf(i_0_2969, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),X3)|v2_struct_0(g1_orders_2(X4,X3))|~l1_orders_2(g1_orders_2(X4,X3))|~v1_orders_2(g1_orders_2(X4,X3))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X4,X3)))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X4,X3)))|~r3_orders_2(g1_orders_2(X4,X3),X2,X1)|~v3_orders_2(g1_orders_2(X4,X3)))).#trainneg
cnf(i_0_2991, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,k1_xboole_0))))|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_2804, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(k2_lattice3(k1_lattice3(X2)),X1)),k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1306, plain, (m1_subset_1(X1,X2)|~v5_relat_1(X3,X2)|~r2_hidden(X1,k10_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1032, plain, (l1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_3032, plain, (v4_relat_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2826, plain, (u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|g1_orders_2(X1,k1_xboole_0)!=g1_orders_2(k1_xboole_0,k1_xboole_0)|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_551, plain, (l1_orders_2(g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_3034, plain, (v5_relat_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_552, plain, (l1_orders_2(g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_556, plain, (v1_xboole_0(esk21_2(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2886, plain, (r1_orders_2(k3_lattice3(X1),X2,X3)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),k2_lattice3(X1))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2782, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),X3)|v2_struct_0(g1_orders_2(X4,X3))|~l1_orders_2(g1_orders_2(X4,X3))|~v1_orders_2(g1_orders_2(X4,X3))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X4,X3)))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X4,X3)))|~r3_orders_2(g1_orders_2(X4,X3),X1,X2)|~v3_orders_2(g1_orders_2(X4,X3)))).#trainneg
cnf(i_0_2805, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(k2_lattice3(k1_lattice3(X2)),X1)),k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_2757, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),u1_orders_2(k1_lattice3(X3)))|~r1_orders_2(k1_lattice3(X3),X1,X2)|~l1_orders_2(k1_lattice3(X3))|~v1_orders_2(k1_lattice3(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_2897, plain, (r1_orders_2(k3_lattice3(X1),X2,X3)|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),k2_lattice3(X1))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2680, plain, (g1_orders_2(k1_zfmisc_1(X1),k2_lattice3(k3_lattice3(k1_lattice3(X1))))=k3_lattice3(k3_lattice3(k1_lattice3(X1)))|~v10_lattices(k3_lattice3(k1_lattice3(X1)))|~l3_lattices(k3_lattice3(k1_lattice3(X1))))).#trainneg
cnf(i_0_3298, plain, (v2_struct_0(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3668, plain, (m1_yellow_0(g1_orders_2(k1_xboole_0,k1_xboole_0),g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_3273, plain, (k10_xtuple_0(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_2712, plain, (k1_relset_1(X4,k1_xboole_0)!=X1|~v1_xboole_0(X2)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_557, plain, (v1_xboole_0(esk19_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3482, negated_conjecture, (m1_subset_1(k1_xboole_0,esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_3653, plain, (k9_xtuple_0(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_3508, plain, (v6_relat_2(k1_wellord2(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3524, plain, (v1_finset_1(k1_wellord2(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3552, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),k2_lattice3(k1_lattice3(X1)))|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3283, plain, (g1_orders_2(k1_xboole_0,u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))=g1_orders_2(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_3551, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),u1_orders_2(X1))|~r3_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3731, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),u1_orders_2(X1))|~r3_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1407, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3431, plain, (v1_relat_1(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1)))))).#trainneg
cnf(i_0_2851, plain, (k2_relset_1(X4,k1_xboole_0)!=X1|~v1_xboole_0(X2)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_561, plain, (v1_xboole_0(esk14_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3282, plain, (v1_partfun1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~v2_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_562, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_566, plain, (v3_ordinal1(k9_subset_1(X1,X2,X3))|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3550, plain, (v2_struct_0(k3_lattice3(X1))|r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),k2_lattice3(X1))|v2_struct_0(X1)|~r3_orders_2(k3_lattice3(X1),X2,X3)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_3540, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~r1_tarski(X1,k1_xboole_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,k1_xboole_0)))))).#trainneg
cnf(i_0_3450, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),X1)),g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3732, plain, (r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),k2_lattice3(k1_lattice3(X1)))|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3451, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),X1)),g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3578, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_lattice3(k1_lattice3(X3)))|~r1_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_3297, plain, (k1_relset_1(k1_xboole_0,u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k1_xboole_0|u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_xboole_0)).#trainneg
cnf(i_0_3767, plain, (k1_zfmisc_1(X1)=k1_xboole_0|k3_lattice3(k3_lattice3(k1_lattice3(X1)))!=g1_orders_2(k1_xboole_0,k1_xboole_0)|~v10_lattices(k3_lattice3(k1_lattice3(X1)))|~l3_lattices(k3_lattice3(k1_lattice3(X1))))).#trainneg
cnf(i_0_567, plain, (v1_orders_2(g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_3290, plain, (m1_subset_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))))).#trainneg
cnf(i_0_568, plain, (l1_orders_2(g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_569, plain, (k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3)=X4|g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3))!=g1_orders_2(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_3274, plain, (k1_relset_1(X1,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_3729, plain, (v2_struct_0(k3_lattice3(X1))|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),k2_lattice3(X1))|v2_struct_0(X1)|~r3_orders_2(k3_lattice3(X1),X2,X3)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_3574, plain, (k1_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r1_orders_2(k3_lattice3(k1_lattice3(X2)),esk9_3(X3,X1,k2_lattice3(k1_lattice3(X2))),esk9_3(X3,X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(esk9_3(X3,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3659, plain, (v2_struct_0(k3_lattice3(X1))|r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X3)),k2_lattice3(X1))|v2_struct_0(X1)|~r3_orders_2(k3_lattice3(X1),X3,X2)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_3662, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X3)),k2_lattice3(k1_lattice3(X1)))|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3364, plain, (k1_relset_1(k1_zfmisc_1(X1),k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1)|k2_lattice3(k1_lattice3(X1))!=k1_wellord2(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_570, plain, (X1=X2|g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X3,X4))!=g1_orders_2(X2,X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_571, plain, (v1_relat_1(k9_subset_1(k2_zfmisc_1(X1,X2),X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_572, plain, (v1_xboole_0(k9_subset_1(k2_zfmisc_1(X1,X2),X3,X4))|~v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3644, plain, (g1_orders_2(X1,u1_orders_2(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1)))))=k3_lattice3(g1_orders_2(X1,k1_wellord2(X1)))|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_3652, plain, (k1_relset_1(X1,k1_wellord2(X2))=X2|~v4_relat_1(k1_wellord2(X2),X1))).#trainneg
cnf(i_0_3661, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X3)),u1_orders_2(X1))|~r3_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3549, plain, (v2_struct_0(k3_lattice3(k3_lattice3(X1)))|r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),k2_lattice3(k3_lattice3(X1)))|v2_struct_0(X1)|~r3_orders_2(k3_lattice3(k3_lattice3(X1)),X2,X3)|~l1_orders_2(k3_lattice3(k3_lattice3(X1)))|~v3_orders_2(k3_lattice3(k3_lattice3(X1)))|~v1_orders_2(k3_lattice3(k3_lattice3(X1)))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(k3_lattice3(X1))))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k3_lattice3(X1))))|~v10_lattices(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(k3_lattice3(X1))|~l3_lattices(X1))).#trainneg
cnf(i_0_2826, plain, (u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|g1_orders_2(X1,k1_xboole_0)!=g1_orders_2(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_573, plain, (v4_relat_1(k9_subset_1(k2_zfmisc_1(X1,X2),X3,X4),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3031, plain, (v1_xboole_0(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))))).#trainneg
cnf(i_0_574, plain, (v1_xboole_0(k9_subset_1(k2_zfmisc_1(X1,X2),X3,X4))|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_575, plain, (v5_relat_1(k9_subset_1(k2_zfmisc_1(X1,X2),X3,X4),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3854, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3))).#trainneg
cnf(i_0_3934, plain, (m1_subset_1(X1,X2)|~v4_relat_1(k1_wellord2(X3),X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_3030, plain, (v1_relat_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))))).#trainneg
cnf(i_0_578, plain, (v1_xboole_0(k9_subset_1(X1,X2,X3))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3804, plain, (k2_relset_1(X1,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_581, plain, (k2_relset_1(X1,X2)=X1|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(esk7_3(X4,X1,X2),X3)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X1))))).#trainneg
cnf(i_0_3912, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_relat_1(k2_lattice3(k1_lattice3(X1)))!=X2|~m1_subset_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,k2_lattice3(k1_lattice3(X1))),k1_zfmisc_1(esk6_2(X2,k2_lattice3(k1_lattice3(X1))))))).#trainneg
cnf(i_0_2850, plain, (k1_xboole_0!=X1|~v1_xboole_0(X2)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_583, plain, (k2_relset_1(X1,X2)=X1|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X3,esk7_3(X4,X1,X2))),X2))).#trainneg
cnf(i_0_3034, plain, (v5_relat_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0))).#trainneg
cnf(i_0_3032, plain, (v4_relat_1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0))).#trainneg
cnf(i_0_585, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,esk21_2(X1,X1))))).#trainneg
cnf(i_0_586, plain, (v1_xboole_0(X1)|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v2_struct_0(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_4273, plain, (u1_struct_0(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~l1_orders_2(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~v1_orders_2(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_4008, plain, (v4_yellow_0(g1_orders_2(k1_xboole_0,k1_xboole_0),g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_587, plain, (v1_xboole_0(X1)|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v2_struct_0(g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_2063, plain, (v1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2064, plain, (v3_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_588, plain, (v1_xboole_0(X1)|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v2_struct_0(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_590, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_4049, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(k1_wellord2(X2),X1)),g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_4076, plain, (v1_orders_2(g1_orders_2(X1,k9_subset_1(X4,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_3996, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(X2)),X3,X4)|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~r3_orders_2(k3_lattice3(X2),X3,X4)|~v10_lattices(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(k3_lattice3(X2)))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X2))))).#trainneg
cnf(i_0_1507, plain, (X1=X2|g1_orders_2(X1,k1_xboole_0)!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_1780, plain, (v3_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4087, plain, (v3_ordinal1(u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))))).#trainneg
cnf(i_0_591, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_2065, plain, (v4_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2066, plain, (v5_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_593, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_595, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_4115, plain, (l1_orders_2(g1_orders_2(X1,k9_subset_1(X4,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_3737, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(k2_lattice3(k1_lattice3(X2)),X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_2990, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(k1_lattice3(X2)),X1)),k1_lattice3(X2))|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3853, plain, (k1_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),esk9_3(X3,X1,k2_lattice3(k1_lattice3(X2))),esk9_3(X3,X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(esk9_3(X3,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1781, plain, (v4_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_596, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(esk21_2(X1,X2),X1))).#trainneg
cnf(i_0_1784, plain, (k2_wellord1(k1_xboole_0,X2)=k1_wellord2(X2)|~r1_tarski(X2,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2067, plain, (l1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2793, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r2_hidden(k2_tarski(k2_tarski(X3,X2),k1_tarski(X2)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4023, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(k1_wellord2(X2),X1)),g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1782, plain, (v5_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3288, plain, (v1_funct_2(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2174, plain, (m1_subset_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1790, plain, (k1_toler_1(k1_xboole_0,X2)=k1_wellord2(X2)|~r1_tarski(X2,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3289, plain, (v1_funct_2(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_601, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(esk14_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_602, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_2808, plain, (k2_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~r1_orders_2(k3_lattice3(k1_lattice3(X2)),X3,esk7_3(X4,X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(esk7_3(X4,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2))|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2993, plain, (u1_struct_0(g3_lattices(X1,X2,X3))=X1|~l3_lattices(g3_lattices(X1,X2,X3))|~v3_lattices(g3_lattices(X1,X2,X3)))).#trainneg
cnf(i_0_2810, plain, (k2_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~r1_orders_2(k3_lattice3(k1_lattice3(X2)),esk7_3(X3,X1,k2_lattice3(k1_lattice3(X2))),esk7_3(X3,X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(esk7_3(X3,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2))|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_3394, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_lattice3(k3_lattice3(X3)))|v2_struct_0(X3)|~r1_orders_2(k3_lattice3(k3_lattice3(X3)),X1,X2)|~l1_orders_2(k3_lattice3(k3_lattice3(X3)))|~v1_orders_2(k3_lattice3(k3_lattice3(X3)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(k3_lattice3(X3))))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(k3_lattice3(X3))))|~v10_lattices(k3_lattice3(X3))|~v10_lattices(X3)|~l3_lattices(k3_lattice3(X3))|~l3_lattices(X3))).#trainneg
cnf(i_0_2670, plain, (X1=k1_zfmisc_1(X2)|g1_orders_2(X1,u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))!=k3_lattice3(k1_lattice3(X2)))).#trainneg
cnf(i_0_3472, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),u1_orders_2(g1_orders_2(X3,k1_wellord2(X3))))|~r1_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_2817, plain, (k1_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~m1_subset_1(esk9_3(X3,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2))|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_3556, plain, (v2_struct_0(g1_orders_2(X1,X2))|r2_hidden(k2_tarski(k2_tarski(X3,X3),k1_tarski(X3)),X2)|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2))))).#trainneg
cnf(i_0_1799, plain, (k1_wellord2(X1)=esk14_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4357, plain, (u1_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_606, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k9_subset_1(k2_zfmisc_1(X1,X2),X3,X4),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4747, plain, (v1_xboole_0(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_610, plain, (v3_ordinal1(u1_orders_2(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_613, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4562, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_4728, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_4778, plain, (u1_orders_2(X1)=k2_lattice3(X2)|v2_struct_0(X2)|X1!=k3_lattice3(X2)|~l1_orders_2(X1)|~v1_orders_2(X1)|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_4677, plain, (k2_wellord1(k1_xboole_0,X1)=k1_wellord2(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4748, plain, (v5_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_614, plain, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~r2_hidden(X1,u1_orders_2(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_4743, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_partfun1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4879, plain, (m1_subset_1(k1_wellord2(X2),k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_615, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_617, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4858, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4603, plain, (r1_orders_2(k3_lattice3(X1),X2,X3)|v2_struct_0(X1)|~r3_orders_2(k3_lattice3(X1),X2,X3)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1))))).#trainneg
cnf(i_0_4859, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),u1_orders_2(k1_lattice3(X1)))|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4709, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(X2)),X3,X4)|v2_struct_0(X2)|~r1_orders_2(k3_lattice3(X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~l1_orders_2(k3_lattice3(X2))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(k3_lattice3(X2)))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X2)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_4692, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X2))).#trainneg
cnf(i_0_4856, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),u1_orders_2(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4844, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k1_wellord2(X3))|~r1_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_4750, plain, (k1_toler_1(k1_xboole_0,X1)=k1_wellord2(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_622, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4735, plain, (v3_ordinal1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4746, plain, (v4_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4850, negated_conjecture, (k1_relset_1(esk2_0,k2_lattice3(k1_lattice3(esk1_0)))=esk2_0|~m1_subset_1(k2_lattice3(k1_lattice3(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk2_0,X1))))).#trainneg
cnf(i_0_4852, plain, (v2_struct_0(k3_lattice3(X1))|r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),k2_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_4687, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X3,X2)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4769, plain, (m1_yellow_0(g1_orders_2(X1,k1_toler_1(k2_lattice3(X2),X1)),k3_lattice3(X2))|v2_struct_0(X2)|~l1_orders_2(k3_lattice3(X2))|~v1_orders_2(k3_lattice3(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_lattice3(X2))))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_4798, plain, (m1_subset_1(k2_lattice3(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_lattice3(X1)),u1_struct_0(k3_lattice3(X1)))))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_4847, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),u1_orders_2(g1_orders_2(X3,k1_wellord2(X3))))|~r1_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_4791, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_lattice3(X3))|v2_struct_0(X3)|~r1_orders_2(k3_lattice3(X3),X1,X2)|~l1_orders_2(k3_lattice3(X3))|~v1_orders_2(k3_lattice3(X3))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X3)))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(X3)))|~v10_lattices(X3)|~l3_lattices(X3))).#trainneg
cnf(i_0_4885, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_wellord2(X3))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4496, negated_conjecture, (k1_xboole_0!=X1|~v1_xboole_0(X2)|X1!=k1_wellord2(k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~r2_hidden(X3,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~v1_relat_1(X1)|~r1_tarski(X3,esk2_0))).#trainneg
cnf(i_0_3293, plain, (m1_subset_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_630, plain, (v3_ordinal1(k2_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4768, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(k2_lattice3(X2),X1)),k3_lattice3(X2))|v2_struct_0(X2)|~l1_orders_2(k3_lattice3(X2))|~v1_orders_2(k3_lattice3(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_lattice3(X2))))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_4845, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),u1_orders_2(g1_orders_2(X3,k1_wellord2(X3))))|~r1_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_4416, plain, (k2_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~m1_subset_1(esk7_3(X3,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2))|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),esk7_3(X3,X1,k2_lattice3(k1_lattice3(X2))),esk7_3(X3,X1,k2_lattice3(k1_lattice3(X2)))))).#trainneg
cnf(i_0_4508, plain, (g1_orders_2(X1,k2_lattice3(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1)))))=k3_lattice3(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|v2_struct_0(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v10_lattices(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~l3_lattices(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_631, plain, (v1_orders_2(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_4749, plain, (k1_relset_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_632, plain, (l1_orders_2(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_633, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5161, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5204, plain, (r1_orders_2(k1_lattice3(X1),X2,X2)|~l1_orders_2(k1_lattice3(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_634, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5452, plain, (v1_xboole_0(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_5007, plain, (u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_635, plain, (k2_relset_1(k2_zfmisc_1(X1,X1),X2)=X3|g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2))!=g1_orders_2(X4,X3)|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_637, plain, (v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5248, plain, (r1_orders_2(X1,X2,X2)|v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_5191, plain, (k2_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~m1_subset_1(esk7_3(X4,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2))|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),X3,esk7_3(X4,X1,k2_lattice3(k1_lattice3(X2)))))).#trainneg
cnf(i_0_5453, plain, (v5_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_638, plain, (v1_xboole_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5448, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_partfun1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_639, plain, (v4_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_640, plain, (v1_xboole_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5315, plain, (r1_orders_2(k3_lattice3(X1),X2,X2)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1))))).#trainneg
cnf(i_0_5405, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),k2_lattice3(X1))|~r1_orders_2(k3_lattice3(X1),X2,X3)|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1))))).#trainneg
cnf(i_0_5385, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k1_wellord2(X3))|~r1_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_2487, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5008, negated_conjecture, (u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2442, plain, (k2_lattice3(k1_lattice3(X1))=k1_xboole_0|g1_orders_2(X2,k1_xboole_0)!=k3_lattice3(k1_lattice3(X1)))).#trainneg
cnf(i_0_641, plain, (v5_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5440, plain, (v3_ordinal1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_644, plain, (v1_xboole_0(k2_relset_1(X1,X2))|~v1_xboole_0(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_646, plain, (m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5055, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(k1_xboole_0,k1_xboole_0))|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_5486, plain, (k2_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=X1|~m1_subset_1(esk7_3(X3,X1,k2_lattice3(k1_lattice3(X2))),k1_zfmisc_1(X2))|~m1_subset_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_4701, plain, (r3_orders_2(g1_orders_2(X1,u1_orders_2(X2)),X3,X4)|v2_struct_0(g1_orders_2(X1,u1_orders_2(X2)))|~l1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v3_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~r1_orders_2(X2,X3,X4)|~l1_orders_2(X2)|~v1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_5256, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X2)),u1_orders_2(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3314, plain, (v2_struct_0(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5451, plain, (v4_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3770, plain, (k2_lattice3(k3_lattice3(k1_lattice3(X1)))=k1_xboole_0|g1_orders_2(X2,k1_xboole_0)!=k3_lattice3(k3_lattice3(k1_lattice3(X1)))|~v10_lattices(k3_lattice3(k1_lattice3(X1)))|~l3_lattices(k3_lattice3(k1_lattice3(X1))))).#trainneg
cnf(i_0_647, plain, (v3_ordinal1(k1_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5547, plain, (v3_ordinal1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v5_ordinal1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_funct_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v4_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_5454, plain, (k1_relset_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_648, plain, (v1_orders_2(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_649, plain, (l1_orders_2(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5128, plain, (u1_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))=u1_struct_0(X1)|~l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_5042, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(k1_xboole_0,k1_xboole_0))|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_2922, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(k1_lattice3(X2)),X1)),k1_lattice3(X2))|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3331, plain, (u1_struct_0(g1_orders_2(X1,k1_toler_1(X2,X1)))=X1|~l1_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_650, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_4889, plain, (v8_relat_2(esk14_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5543, plain, (k9_xtuple_0(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|~v4_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3287, plain, (g3_lattices(k1_xboole_0,u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=g1_orders_2(k1_xboole_0,k1_xboole_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_651, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_652, plain, (k1_relset_1(k2_zfmisc_1(X1,X1),X2)=X3|g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2))!=g1_orders_2(X4,X3)|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5208, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X1)),u1_orders_2(k1_lattice3(X2)))|~l1_orders_2(k1_lattice3(X2))|~v3_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_653, plain, (X1=X2|g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X3))!=g1_orders_2(X2,X4)|~v4_relat_1(X3,k2_zfmisc_1(X1,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_4890, plain, (v4_relat_2(esk14_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5578, plain, (u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_5585, plain, (u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3313, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_654, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_655, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5562, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5915, plain, (g1_orders_2(X1,k2_lattice3(g1_orders_2(X1,k1_toler_1(X2,X1))))=k3_lattice3(g1_orders_2(X1,k1_toler_1(X2,X1)))|v2_struct_0(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v10_lattices(g1_orders_2(X1,k1_toler_1(X2,X1)))|~l3_lattices(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5986, plain, (v8_relat_2(esk14_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5716, negated_conjecture, (u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4470, plain, (X3=k1_wellord2(X1)|k1_xboole_0!=X1|~v1_xboole_0(X2)|k1_relat_1(X3)!=X1|~v1_relat_1(X3))).#trainneg
cnf(i_0_656, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_6008, plain, (~v3_card_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),np__1)|~v1_card_1(np__1)|~v1_funct_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6009, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(k1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3150, plain, (v3_ordinal1(X1)|~v5_ordinal1(k1_wellord2(X1))|~v1_funct_1(k1_wellord2(X1)))).#trainneg
cnf(i_0_657, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_658, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5781, plain, (r3_orders_2(g1_orders_2(X1,X3),X4,X5)|v2_struct_0(g1_orders_2(X1,X3))|~r1_orders_2(g1_orders_2(X2,X3),X4,X5)|~l1_orders_2(g1_orders_2(X1,X3))|~l1_orders_2(g1_orders_2(X2,X3))|~v3_orders_2(g1_orders_2(X1,X3))|~v1_orders_2(g1_orders_2(X1,X3))|~m1_subset_1(X5,u1_struct_0(g1_orders_2(X1,X3)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X3)))|~m1_subset_1(X5,u1_struct_0(g1_orders_2(X2,X3)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X2,X3)))|~v1_orders_2(g1_orders_2(X2,X3)))).#trainneg
cnf(i_0_5635, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X3)|v2_struct_0(g1_orders_2(X4,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X4,X2))|~v3_orders_2(g1_orders_2(X4,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X4,X2))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X4,X2))))).#trainneg
cnf(i_0_5570, plain, (r3_orders_2(k1_lattice3(X1),X2,X2)|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_5786, plain, (r3_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X3)|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~r1_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6099, plain, (v4_relat_2(esk14_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5579, negated_conjecture, (u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_661, plain, (v1_xboole_0(k1_relset_1(X1,X2))|~v1_xboole_0(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4882, plain, (v3_ordinal1(k1_wellord2(X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_663, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6058, plain, (k2_relset_1(X1,u1_orders_2(k1_lattice3(X2)))=X1|~m1_subset_1(u1_orders_2(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~l1_orders_2(k1_lattice3(X2))|~v3_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(k1_lattice3(X2))),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_5796, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(X2)),X3,X3)|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~l1_orders_2(X2)|~v3_orders_2(X2)|~v1_orders_2(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_5730, plain, (v1_xboole_0(k10_xtuple_0(X2))|~v1_xboole_0(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6006, plain, (k2_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k10_xtuple_0(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))=k1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2080, plain, (k1_relat_1(X1)=X2|X1!=k1_wellord2(X3)|~v1_relat_1(X1)|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5826, plain, (v3_ordinal1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v5_ordinal1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_funct_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v4_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4315, plain, (m1_subset_1(X2,k1_zfmisc_1(X1))|~v4_relat_1(k1_wellord2(X2),X1))).#trainneg
cnf(i_0_670, plain, (X1=k1_wellord2(X2)|v1_xboole_0(X1)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~r1_tarski(esk5_2(X2,X1),esk6_2(X2,X1))|~m1_subset_1(k2_tarski(k1_tarski(esk5_2(X2,X1)),k2_tarski(esk5_2(X2,X1),esk6_2(X2,X1))),X1))).#trainneg
cnf(i_0_689, plain, (v3_ordinal1(esk16_1(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_5784, plain, (r3_orders_2(k1_lattice3(X1),X2,X3)|~r1_orders_2(k1_lattice3(X1),X2,X3)|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5779, plain, (r3_orders_2(g1_orders_2(X1,k2_lattice3(X2)),X3,X4)|v2_struct_0(g1_orders_2(X1,k2_lattice3(X2)))|v2_struct_0(X2)|~r1_orders_2(k3_lattice3(X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~l1_orders_2(k3_lattice3(X2))|~v3_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(k3_lattice3(X2)))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X2)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_5481, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k1_wellord2(X3))|~r1_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X2,X1)|~m1_subset_1(X1,X3)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_5710, plain, (v1_xboole_0(k3_xboole_0(X2,X1))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5822, plain, (k9_xtuple_0(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|~v4_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_5816, plain, (v3_ordinal1(X2)|~v3_ordinal1(k1_zfmisc_1(X1))|~v4_relat_1(k1_wellord2(X2),X1))).#trainneg
cnf(i_0_694, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk16_1(X2))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_697, plain, (v4_relat_1(k3_xboole_0(X2,X1),X3)|~v4_relat_1(X1,X3)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5342, plain, (l1_orders_2(g1_orders_2(u1_struct_0(k3_lattice3(X1)),k2_lattice3(X1)))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_5802, plain, (r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X2)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2728, plain, (m1_subset_1(u1_orders_2(k3_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_6649, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_6538, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6119, plain, (v1_funct_1(k1_xboole_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6279, plain, (k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~v3_card_1(k1_xboole_0,np__1)|~v1_card_1(np__1)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_699, plain, (v4_relat_1(k9_subset_1(X4,X1,X2),X3)|~v4_relat_1(X1,X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_701, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk19_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))),np__1)|v1_xboole_0(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))).#trainneg
cnf(i_0_6564, plain, (r1_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X3)|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~r3_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6578, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6242, plain, (k1_relset_1(X1,u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_6562, plain, (r1_orders_2(k1_lattice3(X1),X2,X3)|~r3_orders_2(k1_lattice3(X1),X2,X3)|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6118, plain, (u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=k1_xboole_0|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_702, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk18_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))),np__1)|v1_xboole_0(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))).#trainneg
cnf(i_0_6130, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_partfun1(k1_xboole_0,k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6246, negated_conjecture, (v1_funct_1(esk12_1(esk2_0))|v1_xboole_0(esk2_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_703, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk17_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))),np__1)|v1_xboole_0(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))).#trainneg
cnf(i_0_706, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk12_1(k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1))))),np__1))).#trainneg
cnf(i_0_6328, plain, (r3_orders_2(g1_orders_2(X1,X2),X3,X4)|v2_struct_0(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X5,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X5,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X5,X2))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X5,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X5,X2)))|~r3_orders_2(g1_orders_2(X5,X2),X3,X4)|~v3_orders_2(g1_orders_2(X5,X2)))).#trainneg
cnf(i_0_6677, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(k3_lattice3(X1))))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_6251, negated_conjecture, (v1_relat_1(esk12_1(esk2_0))|v1_xboole_0(esk2_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_2801, plain, (k1_zfmisc_1(X1)=k1_zfmisc_1(X2)|k3_lattice3(k1_lattice3(X1))!=k3_lattice3(k1_lattice3(X2)))).#trainneg
cnf(i_0_709, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k1_relset_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X2)),np__1)|~v4_relat_1(X2,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_6281, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_710, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),k2_relset_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X2)),np__1)|~v5_relat_1(X2,k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_713, plain, (v3_orders_2(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)))|~l1_orders_2(X2)|~v3_orders_2(X2)|~m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_6472, plain, (r1_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X2)|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6334, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(X2)),X3,X3)|v2_struct_0(k3_lattice3(X2))|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~l1_orders_2(k3_lattice3(X2))|~v3_orders_2(k3_lattice3(X2))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~v1_orders_2(k3_lattice3(X2))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X2)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_3542, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~r1_tarski(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,k1_xboole_0))))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_6363, negated_conjecture, (u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_714, plain, (v4_orders_2(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)))|~l1_orders_2(X2)|~v4_orders_2(X2)|~m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_6121, plain, (v1_funct_2(k1_xboole_0,k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6135, plain, (v3_ordinal1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~v5_ordinal1(k1_xboole_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6379, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~v4_relat_1(k1_wellord2(X2),X1))).#trainneg
cnf(i_0_715, plain, (v5_orders_2(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~m1_yellow_0(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_716, plain, (v3_ordinal1(k4_subset_1(X1,X2,X3))|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6547, plain, (r3_orders_2(g1_orders_2(X1,u1_orders_2(X2)),X3,X4)|v2_struct_0(g1_orders_2(X1,u1_orders_2(X2)))|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v3_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(X2))))|~r3_orders_2(X2,X3,X4)|~l1_orders_2(X2)|~v3_orders_2(X2)|~v1_orders_2(g1_orders_2(X1,u1_orders_2(X2)))|~v1_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_5277, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,k1_xboole_0))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5855, plain, (g1_orders_2(u1_struct_0(X1),k2_lattice3(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))=k3_lattice3(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|v2_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v10_lattices(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l3_lattices(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3576, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_lattice3(k1_lattice3(X3)))|~r1_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_6232, plain, (k1_relset_1(k1_zfmisc_1(X1),k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1))).#trainneg
cnf(i_0_5396, plain, (k1_relset_1(X1,k2_lattice3(X2))=X1|v2_struct_0(X2)|~m1_subset_1(k2_lattice3(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r1_orders_2(k3_lattice3(X2),esk9_3(X3,X1,k2_lattice3(X2)),X4)|~v10_lattices(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(k3_lattice3(X2)))|~m1_subset_1(esk9_3(X3,X1,k2_lattice3(X2)),u1_struct_0(k3_lattice3(X2))))).#trainneg
cnf(i_0_2646, plain, (v1_partfun1(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))),X1)|~v2_orders_2(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_5950, plain, (v2_struct_0(g1_orders_2(X1,k1_toler_1(X2,X1)))|r2_hidden(k2_tarski(k2_tarski(X3,X3),k1_tarski(X3)),k1_toler_1(X2,X1))|~l1_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v3_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v1_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~m1_subset_1(X3,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6260, negated_conjecture, (v5_relat_1(esk12_1(esk2_0),k1_xboole_0)|v1_xboole_0(esk2_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4937, plain, (k1_relat_1(X1)=X2|X1!=esk14_1(X3)|~v1_relat_1(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_717, plain, (v1_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_4891, plain, (v1_relat_2(esk14_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6488, plain, (v1_xboole_0(k10_xtuple_0(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6489, plain, (v1_xboole_0(k10_xtuple_0(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_718, plain, (l1_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_719, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_5297, plain, (v4_relat_1(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))),X1))).#trainneg
cnf(i_0_5344, plain, (v3_orders_2(g1_orders_2(u1_struct_0(k3_lattice3(X1)),k2_lattice3(X1)))|v2_struct_0(X1)|~v1_relat_2(k2_lattice3(X1))|~v1_partfun1(k2_lattice3(X1),u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_6360, plain, (r3_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X2)|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5276, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,k1_xboole_0))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6256, negated_conjecture, (k9_xtuple_0(esk12_1(esk2_0))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|v1_xboole_0(esk2_0)|esk12_1(esk2_0)!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_720, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4892, plain, (v1_relat_1(esk14_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6129, plain, (k2_zfmisc_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_722, plain, (k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)=X4|g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3))!=g1_orders_2(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_723, plain, (X1=X2|g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X3,X4))!=g1_orders_2(X2,X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_7147, plain, (v1_relat_2(esk14_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_724, plain, (v1_relat_1(k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6756, plain, (~v3_card_1(k1_xboole_0,np__1)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(np__1))).#trainneg
cnf(i_0_6748, plain, (m1_subset_1(k1_wellord2(X1),k1_zfmisc_1(X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_725, plain, (v1_xboole_0(k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4))|~v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_726, plain, (v4_relat_1(k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6872, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_7216, plain, (v3_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(X1)|~v1_partfun1(k2_lattice3(X1),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_6777, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k1_wellord2(X3))|v2_struct_0(g1_orders_2(X3,k1_wellord2(X3)))|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)|~r3_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X1,X2))).#trainneg
cnf(i_0_7261, plain, (v1_relat_1(esk14_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7172, plain, (k10_xtuple_0(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k1_xboole_0|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4614, plain, (X1=k1_zfmisc_1(X2)|g1_orders_2(X1,k1_xboole_0)!=k3_lattice3(k1_lattice3(X2)))).#trainneg
cnf(i_0_727, plain, (v1_xboole_0(k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4))|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_728, plain, (v5_relat_1(k4_subset_1(k2_zfmisc_1(X1,X2),X3,X4),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_733, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_6776, plain, (k2_relset_1(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))=X1|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(esk7_3(X4,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X2)|~m1_subset_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(X3,X2)|~r3_orders_2(g1_orders_2(X2,k1_wellord2(X2)),X3,esk7_3(X4,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))))).#trainneg
cnf(i_0_6854, plain, (r3_orders_2(g1_orders_2(X1,k2_lattice3(X2)),X3,X4)|v2_struct_0(k3_lattice3(X2))|v2_struct_0(g1_orders_2(X1,k2_lattice3(X2)))|v2_struct_0(X2)|~r3_orders_2(k3_lattice3(X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~l1_orders_2(k3_lattice3(X2))|~v3_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~v3_orders_2(k3_lattice3(X2))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~v1_orders_2(k3_lattice3(X2))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(k3_lattice3(X2)))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X2)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_5376, plain, (k2_relset_1(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))=X1|~m1_subset_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),esk7_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),esk7_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))))|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X2))).#trainneg
cnf(i_0_6486, plain, (v1_xboole_0(k10_xtuple_0(k1_wellord2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7185, plain, (k10_xtuple_0(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k1_xboole_0|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_735, plain, (v3_card_1(k1_relset_1(X3,X1),X2)|~v3_card_1(X1,X2)|~v1_card_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X3))).#trainneg
cnf(i_0_740, plain, (k1_relset_1(X1,X2)=X1|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(k2_tarski(k2_tarski(esk9_3(X3,X1,X2),X4),k1_tarski(esk9_3(X3,X1,X2))),X2))).#trainneg
cnf(i_0_746, plain, (v3_ordinal1(esk3_2(X1,X2))|X2=k1_lattice3(X1)|~v3_ordinal1(k1_zfmisc_1(X1))|u1_struct_0(X2)!=k1_zfmisc_1(X1)|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_6778, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k1_wellord2(X3))|v2_struct_0(g1_orders_2(X3,k1_wellord2(X3)))|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)|~r3_orders_2(g1_orders_2(X3,k1_wellord2(X3)),X1,X2))).#trainneg
cnf(i_0_2985, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~l1_orders_2(g1_orders_2(X3,X2))|~v1_orders_2(g1_orders_2(X3,X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X3,X2)))))).#trainneg
cnf(i_0_5434, plain, (v5_relat_1(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))),X1))).#trainneg
cnf(i_0_7126, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk12_1(esk2_0)))|v1_xboole_0(esk2_0)|~v1_relat_1(esk12_1(esk2_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3794, plain, (u1_struct_0(X1)=k1_zfmisc_1(X2)|X1!=k3_lattice3(k3_lattice3(k1_lattice3(X2)))|~l1_orders_2(X1)|~v1_orders_2(X1)|~v10_lattices(k3_lattice3(k1_lattice3(X2)))|~l3_lattices(k3_lattice3(k1_lattice3(X2))))).#trainneg
cnf(i_0_747, plain, (v1_orders_2(g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_3648, plain, (v3_card_1(X1,X2)|~v3_card_1(k1_wellord2(X1),X2)|~v1_card_1(X2)|~v1_funct_1(k1_wellord2(X1)))).#trainneg
cnf(i_0_748, plain, (l1_orders_2(g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_749, plain, (v1_xboole_0(X1)|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v2_struct_0(g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X2)))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_4815, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(k3_lattice3(X2))),X3,X4)|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k3_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k3_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(k3_lattice3(X2)))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(k3_lattice3(X2)))))|~r1_orders_2(k3_lattice3(k3_lattice3(X2)),X3,X4)|~l1_orders_2(k3_lattice3(k3_lattice3(X2)))|~v1_orders_2(k3_lattice3(k3_lattice3(X2)))|~v10_lattices(k3_lattice3(X2))|~v10_lattices(X2)|~l3_lattices(k3_lattice3(X2))|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(k3_lattice3(k3_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(k3_lattice3(X2)))))).#trainneg
cnf(i_0_3555, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))|~r3_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3642, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(k3_lattice3(k3_lattice3(X1))))=k3_lattice3(k3_lattice3(X1))|v2_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~v10_lattices(k3_lattice3(X1))|~l3_lattices(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_5803, plain, (v2_struct_0(k3_lattice3(X1))|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X2)),k2_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_3560, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~r1_tarski(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(g1_orders_2(X2,k1_xboole_0))))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2904, plain, (k2_relset_1(X1,k2_lattice3(X2))=X1|v2_struct_0(k3_lattice3(X2))|v2_struct_0(X2)|~l1_orders_2(k3_lattice3(X2))|~v3_orders_2(k3_lattice3(X2))|~m1_subset_1(esk7_3(X3,X1,k2_lattice3(X2)),u1_struct_0(k3_lattice3(X2)))|~m1_subset_1(k2_lattice3(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_750, plain, (v1_xboole_0(X1)|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|~v1_partfun1(esk3_2(k2_zfmisc_1(X1,X2),X3),X1)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_4930, plain, (v6_relat_2(esk14_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_751, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk3_2(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X2)),np__1)|X2=k1_lattice3(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_752, plain, (esk3_2(k2_zfmisc_1(X1,X1),X2)=X3|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X2))!=g1_orders_2(X4,X3)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_5403, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k2_tarski(X3,X2),k1_tarski(X2)),k2_lattice3(X1))|~r1_orders_2(k3_lattice3(X1),X2,X3)|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1))))).#trainneg
cnf(i_0_6828, plain, (u1_orders_2(k1_lattice3(X1))=k1_wellord2(X2)|k1_relat_1(u1_orders_2(k1_lattice3(X1)))!=X2|~l1_orders_2(k1_lattice3(X1))|~r1_tarski(esk5_2(X2,u1_orders_2(k1_lattice3(X1))),esk6_2(X2,u1_orders_2(k1_lattice3(X1))))|~m1_subset_1(esk6_2(X2,u1_orders_2(k1_lattice3(X1))),k1_zfmisc_1(X1))|~m1_subset_1(esk5_2(X2,u1_orders_2(k1_lattice3(X1))),k1_zfmisc_1(X1))|~r3_orders_2(k1_lattice3(X1),esk5_2(X2,u1_orders_2(k1_lattice3(X1))),esk6_2(X2,u1_orders_2(k1_lattice3(X1))))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_3065, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X4,X3)),X2)|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2))))).#trainneg
cnf(i_0_2007, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_xboole_0(X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_7360, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_wellord2(X3))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_7645, negated_conjecture, (k10_xtuple_0(esk12_1(esk2_0))=k1_xboole_0|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4613, plain, (X1=k1_zfmisc_1(X2)|g1_orders_2(X1,k1_xboole_0)!=k3_lattice3(k3_lattice3(k1_lattice3(X2)))|~v10_lattices(k3_lattice3(k1_lattice3(X2)))|~l3_lattices(k3_lattice3(k1_lattice3(X2))))).#trainneg
cnf(i_0_753, plain, (X1=X2|X3=k1_lattice3(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X3))!=g1_orders_2(X2,X4)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_7661, plain, (v3_card_1(X1,X2)|~v3_card_1(k1_xboole_0,X2)|~v1_card_1(X2)|~v1_funct_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7348, plain, (v4_relat_1(k1_wellord2(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7350, plain, (v5_relat_1(k1_wellord2(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_754, plain, (v1_relat_1(esk3_2(k2_zfmisc_1(X1,X2),X3))|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_755, plain, (v1_xboole_0(esk3_2(k2_zfmisc_1(X1,X2),X3))|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_7344, plain, (k1_wellord2(X1)=X2|g1_orders_2(X3,k1_wellord2(X1))!=g1_orders_2(X4,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7809, plain, (v6_relat_2(esk14_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7646, negated_conjecture, (k10_xtuple_0(esk12_1(esk2_0))=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_756, plain, (v4_relat_1(esk3_2(k2_zfmisc_1(X1,X2),X3),X1)|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_4931, plain, (v1_finset_1(esk14_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6280, negated_conjecture, (v1_xboole_0(esk2_0)|esk12_1(esk2_0)!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v3_card_1(esk12_1(esk2_0),np__1)|~v1_card_1(np__1)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_757, plain, (v1_xboole_0(esk3_2(k2_zfmisc_1(X1,X2),X3))|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_758, plain, (v5_relat_1(esk3_2(k2_zfmisc_1(X1,X2),X3),X2)|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_7543, plain, (v3_card_1(X1,X2)|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v3_card_1(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))),X2)|~v1_card_1(X2)|~v1_funct_1(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1)))))).#trainneg
cnf(i_0_7495, plain, (k2_relset_1(X1,k1_wellord2(X2))=X1|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),esk7_3(X3,X1,k1_wellord2(X2)),esk7_3(X3,X1,k1_wellord2(X2)))|~m1_subset_1(k1_wellord2(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~m1_subset_1(esk7_3(X3,X1,k1_wellord2(X2)),X2))).#trainneg
cnf(i_0_7467, plain, (k2_relset_1(X1,k1_wellord2(X2))=X1|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~r3_orders_2(g1_orders_2(X2,k1_wellord2(X2)),X3,esk7_3(X4,X1,k1_wellord2(X2)))|~m1_subset_1(k1_wellord2(X2),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(esk7_3(X4,X1,k1_wellord2(X2)),X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_6728, plain, (m1_subset_1(esk14_1(X1),k1_zfmisc_1(X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1586, plain, (u1_orders_2(X1)=k2_lattice3(k1_lattice3(X2))|X1!=k3_lattice3(k1_lattice3(X2))|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_762, plain, (v1_xboole_0(esk3_2(X1,X2))|X2=k1_lattice3(X1)|~v1_xboole_0(X1)|u1_struct_0(X2)!=k1_zfmisc_1(X1)|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_6610, plain, (~v3_card_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),np__1)|~v1_card_1(np__1)|~v1_funct_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6254, negated_conjecture, (v3_ordinal1(esk12_1(esk2_0))|v1_xboole_0(esk2_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_763, plain, (m1_subset_1(X1,X2)|X3=k1_lattice3(X2)|~r2_hidden(X1,esk3_2(X2,X3))|u1_struct_0(X3)!=k1_zfmisc_1(X2)|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_764, plain, (v3_ordinal1(esk4_2(X1,X2))|X2=k1_lattice3(X1)|~v3_ordinal1(k1_zfmisc_1(X1))|u1_struct_0(X2)!=k1_zfmisc_1(X1)|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_7878, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))),X3,X3)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_7883, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))),X3,X3)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2)))))|~l1_orders_2(k1_lattice3(X2))|~v3_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_7600, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)))|~l1_orders_2(X2)|~v1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_7608, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),X1)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_7007, plain, (m1_yellow_0(g1_orders_2(X1,k1_xboole_0),g1_orders_2(X2,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_765, plain, (v1_orders_2(g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_1794, plain, (k1_relat_1(X1)=X2|X1!=k1_xboole_0|~v1_relat_1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7335, plain, (v1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7336, plain, (l1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_766, plain, (l1_orders_2(g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_767, plain, (v1_xboole_0(X1)|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v2_struct_0(g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X2)))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_7387, plain, (u1_struct_0(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_7705, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)),u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))|~r3_orders_2(g1_orders_2(X1,k1_wellord2(X1)),X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_7234, plain, (v4_yellow_0(g1_orders_2(X1,k1_xboole_0),g1_orders_2(X2,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7173, negated_conjecture, (k10_xtuple_0(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3784, plain, (u1_orders_2(X1)=k2_lattice3(k3_lattice3(k1_lattice3(X2)))|X1!=k3_lattice3(k3_lattice3(k1_lattice3(X2)))|~l1_orders_2(X1)|~v1_orders_2(X1)|~v10_lattices(k3_lattice3(k1_lattice3(X2)))|~l3_lattices(k3_lattice3(k1_lattice3(X2))))).#trainneg
cnf(i_0_768, plain, (v1_xboole_0(X1)|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|~v1_partfun1(esk4_2(k2_zfmisc_1(X1,X2),X3),X1)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_6257, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|v1_xboole_0(esk2_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_partfun1(esk12_1(esk2_0),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6611, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(k1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_7333, plain, (v3_ordinal1(k1_wellord2(X1))|~v3_ordinal1(k1_zfmisc_1(X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7338, plain, (v3_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_partfun1(k1_wellord2(X2),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_769, plain, (v13_struct_0(g1_orders_2(k1_tarski(X1),esk4_2(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)),X2)),np__1)|X2=k1_lattice3(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(k1_tarski(X1),k1_tarski(X1)))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_770, plain, (esk4_2(k2_zfmisc_1(X1,X1),X2)=X3|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X2))!=g1_orders_2(X4,X3)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_2738, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(k3_lattice3(X1))))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_6236, plain, (k1_relset_1(X1,u1_orders_2(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1)))))=X1|v2_struct_0(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~l1_orders_2(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~v3_orders_2(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_7186, negated_conjecture, (k10_xtuple_0(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=esk12_1(esk2_0)|v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1330, plain, (X2=k1_wellord2(esk12_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1)|k1_relat_1(X2)!=esk12_1(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_771, plain, (X1=X2|X3=k1_lattice3(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X3))!=g1_orders_2(X2,X4)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_7339, plain, (v4_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_partfun1(k1_wellord2(X2),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_772, plain, (v1_relat_1(esk4_2(k2_zfmisc_1(X1,X2),X3))|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_773, plain, (v1_xboole_0(esk4_2(k2_zfmisc_1(X1,X2),X3))|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_8273, plain, (g1_orders_2(u1_struct_0(X1),k2_lattice3(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))))=k3_lattice3(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(X1)|~v10_lattices(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~l3_lattices(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_8222, plain, (k1_relat_1(X1)=k1_xboole_0|X1!=k1_xboole_0|~v1_relat_1(X1))).#trainneg
cnf(i_0_774, plain, (v4_relat_1(esk4_2(k2_zfmisc_1(X1,X2),X3),X1)|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_775, plain, (v1_xboole_0(esk4_2(k2_zfmisc_1(X1,X2),X3))|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_776, plain, (v5_relat_1(esk4_2(k2_zfmisc_1(X1,X2),X3),X2)|X3=k1_lattice3(k2_zfmisc_1(X1,X2))|u1_struct_0(X3)!=k1_zfmisc_1(k2_zfmisc_1(X1,X2))|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_8271, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))))=g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_8088, plain, (k2_relset_1(X1,k1_wellord2(X2))=X1|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(k1_wellord2(X2),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~m1_subset_1(esk7_3(X3,X1,k1_wellord2(X2)),X2)|~r3_orders_2(g1_orders_2(X2,k1_wellord2(X2)),esk7_3(X3,X1,k1_wellord2(X2)),esk7_3(X3,X1,k1_wellord2(X2))))).#trainneg
cnf(i_0_8268, plain, (k2_relset_1(X1,u1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2))))=X1|v2_struct_0(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|~v3_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2))),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_8326, plain, (u1_struct_0(g1_orders_2(X1,k2_lattice3(g1_orders_2(X1,k1_wellord2(X1)))))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_8291, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~l1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v2_struct_0(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_8251, plain, (u1_struct_0(g1_orders_2(X1,k1_wellord2(X2)))=X1|~l1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_780, plain, (v1_xboole_0(esk4_2(X1,X2))|X2=k1_lattice3(X1)|~v1_xboole_0(X1)|u1_struct_0(X2)!=k1_zfmisc_1(X1)|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_7957, plain, (v3_card_1(X1,X2)|~v3_card_1(k1_xboole_0,X2)|~v1_funct_1(k1_xboole_0)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_781, plain, (m1_subset_1(X1,X2)|X3=k1_lattice3(X2)|~r2_hidden(X1,esk4_2(X2,X3))|u1_struct_0(X3)!=k1_zfmisc_1(X2)|~l3_lattices(X3)|~v3_lattices(X3))).#trainneg
cnf(i_0_782, plain, (v3_ordinal1(u2_lattices(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_8342, plain, (k2_relset_1(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))=X1|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r3_orders_2(g1_orders_2(X2,k1_wellord2(X2)),esk7_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),esk7_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))))|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X2))).#trainneg
cnf(i_0_8320, plain, (u1_struct_0(g1_orders_2(X1,k2_lattice3(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))))=X1|v2_struct_0(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v10_lattices(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~l3_lattices(k3_lattice3(g1_orders_2(X1,k1_wellord2(X1))))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_8318, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(X2)),X3,X3)|v2_struct_0(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~l1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|~v3_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(X2)))|~v1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_8269, plain, (k1_relset_1(X1,u1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2))))=X1|v2_struct_0(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|v2_struct_0(X2)|~l1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|~v3_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2)))),u1_struct_0(X2))|~m1_subset_1(u1_orders_2(g1_orders_2(u1_struct_0(X2),k2_lattice3(X2))),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_8026, plain, (v1_finset_1(esk14_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7935, negated_conjecture, (k9_xtuple_0(esk12_1(esk2_0))=k1_relat_1(esk12_1(esk2_0))|v1_xboole_0(esk2_0)|~v1_relat_1(esk12_1(esk2_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_8564, plain, (k1_xboole_0=X1|k1_wellord2(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_786, plain, (v1_xboole_0(u2_lattices(X1))|~v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l2_lattices(X1))).#trainneg
cnf(i_0_8377, plain, (v3_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X2)|~v1_partfun1(k1_xboole_0,X1))).#trainneg
cnf(i_0_787, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)),u1_struct_0(X2)))|~r2_hidden(X1,u2_lattices(X2))|~l2_lattices(X2))).#trainneg
cnf(i_0_8822, plain, (u1_struct_0(g1_orders_2(X1,k1_xboole_0))=X1)).#trainneg
cnf(i_0_790, plain, (v1_xboole_0(u2_lattices(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l2_lattices(X1))).#trainneg
cnf(i_0_8469, plain, (v4_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X2)|~v1_partfun1(k1_xboole_0,X1))).#trainneg
cnf(i_0_791, plain, (v4_relat_1(u2_lattices(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l2_lattices(X1))).#trainneg
cnf(i_0_792, plain, (v1_xboole_0(u2_lattices(X1))|~v1_xboole_0(u1_struct_0(X1))|~l2_lattices(X1))).#trainneg
cnf(i_0_2802, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,esk10_4(X3,X4,k2_lattice3(k1_lattice3(X1)),X2))|k1_relset_1(X4,k2_lattice3(k1_lattice3(X1)))!=X4|~r2_hidden(X2,X4)|~m1_subset_1(esk10_4(X3,X4,k2_lattice3(k1_lattice3(X1)),X2),k1_zfmisc_1(X1))|~m1_subset_1(k2_lattice3(k1_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5756, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k10_xtuple_0(X3))|~v5_relat_1(X3,X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2781, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),u1_orders_2(k1_lattice3(X3)))|~r1_orders_2(k1_lattice3(X3),X1,X2)|~l1_orders_2(k1_lattice3(X3))|~v1_orders_2(k1_lattice3(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_6238, plain, (k1_relset_1(u1_struct_0(X1),u1_orders_2(k3_lattice3(X1)))=u1_struct_0(X1)|v2_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_8119, plain, (v4_relat_1(esk14_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_793, plain, (v5_relat_1(u2_lattices(X1),u1_struct_0(X1))|~l2_lattices(X1))).#trainneg
cnf(i_0_794, plain, (v3_ordinal1(u1_lattices(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_798, plain, (v1_xboole_0(u1_lattices(X1))|~v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))|~l1_lattices(X1))).#trainneg
cnf(i_0_8121, plain, (v5_relat_1(esk14_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6015, plain, (u1_struct_0(X1)=k1_xboole_0|g3_lattices(u1_struct_0(X1),u1_lattices(X1),u2_lattices(X1))!=g1_orders_2(k1_xboole_0,k1_xboole_0)|~l1_lattices(X1)|~l2_lattices(X1)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_799, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)),u1_struct_0(X2)))|~r2_hidden(X1,u1_lattices(X2))|~l1_lattices(X2))).#trainneg
cnf(i_0_802, plain, (v1_xboole_0(u1_lattices(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_lattices(X1))).#trainneg
cnf(i_0_9456, plain, (u1_struct_0(g1_orders_2(X1,k1_wellord2(X2)))=X2|v2_struct_0(g1_orders_2(X1,k1_wellord2(X2)))|~v4_relat_1(k1_wellord2(X2),u1_struct_0(g1_orders_2(X1,k1_wellord2(X2))))|~l1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v3_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_orders_2(g1_orders_2(X1,k1_wellord2(X2))))).#trainneg
cnf(i_0_9330, plain, (k2_relset_1(X1,u1_orders_2(k1_lattice3(X2)))=X1|~m1_subset_1(u1_orders_2(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~r1_orders_2(k1_lattice3(X2),X3,esk7_3(X4,X1,u1_orders_2(k1_lattice3(X2))))|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(esk7_3(X4,X1,u1_orders_2(k1_lattice3(X2))),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9180, plain, (v2_struct_0(g1_orders_2(X1,k1_xboole_0))|r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),k1_xboole_0)|~v3_orders_2(g1_orders_2(X1,k1_xboole_0))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9350, plain, (k1_relset_1(u1_struct_0(X1),k2_lattice3(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_3540, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~r1_tarski(X1,k1_xboole_0)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9467, plain, (k1_relset_1(u1_struct_0(k3_lattice3(X1)),k2_lattice3(X1))=u1_struct_0(k3_lattice3(X1))|v2_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2991, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~r1_tarski(X1,k1_xboole_0)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9158, plain, (v4_relat_1(u1_orders_2(g1_orders_2(X1,k1_xboole_0)),X1))).#trainneg
cnf(i_0_9017, negated_conjecture, (k1_relat_1(esk12_1(esk2_0))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|v1_xboole_0(esk2_0)|esk12_1(esk2_0)!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_3542, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~v1_xboole_0(X3)|~r1_tarski(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_803, plain, (v4_relat_1(u1_lattices(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_lattices(X1))).#trainneg
cnf(i_0_804, plain, (v1_xboole_0(u1_lattices(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_lattices(X1))).#trainneg
cnf(i_0_805, plain, (v5_relat_1(u1_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1))).#trainneg
cnf(i_0_9147, plain, (u1_struct_0(g1_orders_2(X1,u1_orders_2(g1_orders_2(X1,k1_xboole_0))))=X1)).#trainneg
cnf(i_0_9331, plain, (k2_relset_1(X1,u1_orders_2(k1_lattice3(X2)))=X1|~m1_subset_1(u1_orders_2(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r1_orders_2(k1_lattice3(X2),esk7_3(X3,X1,u1_orders_2(k1_lattice3(X2))),esk7_3(X3,X1,u1_orders_2(k1_lattice3(X2))))|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(k1_lattice3(X2))),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3560, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~v1_xboole_0(X3)|~r1_tarski(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9466, plain, (k1_relset_1(X1,X2)=u1_struct_0(g1_orders_2(X3,X2))|v2_struct_0(g1_orders_2(X3,X2))|~v4_relat_1(X2,X1)|~v4_relat_1(X2,u1_struct_0(g1_orders_2(X3,X2)))|~v1_relat_1(X2)|~l1_orders_2(g1_orders_2(X3,X2))|~v3_orders_2(g1_orders_2(X3,X2))|~v1_orders_2(g1_orders_2(X3,X2)))).#trainneg
cnf(i_0_9340, plain, (k9_xtuple_0(u1_orders_2(k3_lattice3(X1)))=u1_struct_0(X1)|v2_struct_0(X1)|~v4_relat_1(u1_orders_2(k3_lattice3(X1)),u1_struct_0(X1))|~v1_relat_1(u1_orders_2(k3_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_9293, plain, (r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,esk10_4(X3,X4,k2_lattice3(k1_lattice3(X1)),X2))|~m1_subset_1(esk10_4(X3,X4,k2_lattice3(k1_lattice3(X1)),X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|k1_relset_1(X4,k2_lattice3(k1_lattice3(X1)))!=X4|~r2_hidden(X2,X4)|~m1_subset_1(k2_lattice3(k1_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(X4,X3))))).#trainneg
cnf(i_0_5200, plain, (k1_relset_1(X1,u1_orders_2(k1_lattice3(X2)))=X1|~m1_subset_1(u1_orders_2(k1_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~l1_orders_2(k1_lattice3(X2))|~v3_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(k1_lattice3(X2))),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_5238, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_orders_2(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_9159, plain, (v5_relat_1(u1_orders_2(g1_orders_2(X1,k1_xboole_0)),X1))).#trainneg
cnf(i_0_9510, plain, (k1_zfmisc_1(X1)=k1_xboole_0|g3_lattices(k1_zfmisc_1(X1),u1_lattices(k1_lattice3(X1)),u2_lattices(k1_lattice3(X1)))!=g1_orders_2(k1_xboole_0,k1_xboole_0)|~l1_lattices(k1_lattice3(X1))|~l2_lattices(k1_lattice3(X1))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_9008, negated_conjecture, (v3_ordinal1(k1_relat_1(esk12_1(esk2_0)))|v1_xboole_0(esk2_0)|~v5_ordinal1(esk12_1(esk2_0))|~v1_funct_1(esk12_1(esk2_0))|~v1_relat_1(esk12_1(esk2_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_811, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(u1_orders_2(X2),X1)))|~l1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_3553, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),u1_orders_2(k1_lattice3(X1)))|~r3_orders_2(k1_lattice3(X1),X2,X3)|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5175, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))),X3,X4)|~r1_orders_2(k1_lattice3(X2),X3,X4)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))))|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2)))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2)))))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_8788, plain, (u1_struct_0(g1_orders_2(X1,u1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))))=X1|~l1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7532, plain, (k1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_9500, plain, (k1_zfmisc_1(X1)=k1_xboole_0|g3_lattices(k1_zfmisc_1(X1),u1_lattices(k3_lattice3(k1_lattice3(X1))),u2_lattices(k3_lattice3(k1_lattice3(X1))))!=g1_orders_2(k1_xboole_0,k1_xboole_0)|~l1_lattices(k3_lattice3(k1_lattice3(X1)))|~l2_lattices(k3_lattice3(k1_lattice3(X1)))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_818, plain, (v3_ordinal1(k2_lattice3(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_8565, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))).#trainneg
cnf(i_0_8566, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))).#trainneg
cnf(i_0_822, plain, (v1_xboole_0(k2_lattice3(X1))|v2_struct_0(X1)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_823, plain, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X1,k2_lattice3(X2))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_9860, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_orders_2(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_9601, plain, (k9_xtuple_0(k2_lattice3(X1))=u1_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~v4_relat_1(k2_lattice3(X1),u1_struct_0(k3_lattice3(X1)))|~v1_relat_1(k2_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_9805, plain, (k1_relset_1(X1,u1_orders_2(X2))=u1_struct_0(X2)|v2_struct_0(X2)|~v4_relat_1(u1_orders_2(X2),u1_struct_0(X2))|~v4_relat_1(u1_orders_2(X2),X1)|~l1_orders_2(X2)|~v3_orders_2(X2)|~v1_orders_2(X2))).#trainneg
cnf(i_0_9635, plain, (k9_xtuple_0(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_9571, plain, (k9_xtuple_0(k2_lattice3(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v4_relat_1(k2_lattice3(X1),u1_struct_0(X1))|~v1_relat_1(k2_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_1040, plain, (X1=X2|g1_orders_2(X1,k2_zfmisc_1(X1,X1))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_6016, plain, (u1_struct_0(X1)=k1_xboole_0|g3_lattices(u1_struct_0(X1),u2_lattices(X1),u1_lattices(X1))!=g1_orders_2(k1_xboole_0,k1_xboole_0)|~l1_lattices(X1)|~l2_lattices(X1)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_10072, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_10073, plain, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_7340, plain, (v5_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_partfun1(k1_wellord2(X2),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9609, plain, (k1_relset_1(X1,k2_lattice3(X2))=u1_struct_0(k3_lattice3(X2))|v2_struct_0(X2)|~v4_relat_1(k2_lattice3(X2),X1)|~v4_relat_1(k2_lattice3(X2),u1_struct_0(k3_lattice3(X2)))|~v1_relat_1(k2_lattice3(X2))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_9579, plain, (k1_relset_1(X1,k2_lattice3(X2))=u1_struct_0(X2)|v2_struct_0(X2)|~v4_relat_1(k2_lattice3(X2),X1)|~v4_relat_1(k2_lattice3(X2),u1_struct_0(X2))|~v1_relat_1(k2_lattice3(X2))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_9629, plain, (k9_xtuple_0(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1)|~v4_relat_1(k2_lattice3(k1_lattice3(X1)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9153, plain, (v1_orders_2(g1_orders_2(X1,u1_orders_2(g1_orders_2(X1,k1_xboole_0)))))).#trainneg
cnf(i_0_9869, negated_conjecture, (v3_ordinal1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|v1_xboole_0(esk2_0)|~v5_ordinal1(esk12_1(esk2_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|esk12_1(esk2_0)!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_9504, plain, (X1=k1_xboole_0|g3_lattices(X1,u1_lattices(g1_orders_2(X1,k1_xboole_0)),u2_lattices(g1_orders_2(X1,k1_xboole_0)))!=g1_orders_2(k1_xboole_0,k1_xboole_0)|~l1_lattices(g1_orders_2(X1,k1_xboole_0))|~l2_lattices(g1_orders_2(X1,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_828, plain, (k2_lattice3(X1)=X2|v2_struct_0(X1)|g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))!=g1_orders_2(X3,X2)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_9934, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),u1_orders_2(k1_lattice3(X3)))|~r3_orders_2(k1_lattice3(X3),X1,X2)|~l1_orders_2(k1_lattice3(X3))|~v3_orders_2(k1_lattice3(X3))|~v1_orders_2(k1_lattice3(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_9807, plain, (k1_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=k1_zfmisc_1(X2)|~v4_relat_1(k2_lattice3(k1_lattice3(X2)),k1_zfmisc_1(X2))|~v4_relat_1(k2_lattice3(k1_lattice3(X2)),X1))).#trainneg
cnf(i_0_9813, plain, (k1_relset_1(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))=X2|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~v4_relat_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),X1))).#trainneg
cnf(i_0_9917, plain, (k2_relset_1(X1,k1_wellord2(X1))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~m1_subset_1(k1_wellord2(X1),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).#trainneg
cnf(i_0_9154, plain, (l1_orders_2(g1_orders_2(X1,u1_orders_2(g1_orders_2(X1,k1_xboole_0)))))).#trainneg
cnf(i_0_6125, plain, (k1_xboole_0=X1|g3_lattices(k1_xboole_0,u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)!=g3_lattices(X1,X2,X3)|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_832, plain, (v4_relat_1(k2_lattice3(X1),u1_struct_0(X1))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_7341, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k1_wellord2(X3),X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_7337, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k1_wellord2(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_834, plain, (v5_relat_1(k2_lattice3(X1),u1_struct_0(X1))|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_4553, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(k1_wellord2(X2),X1)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_5282, plain, (u1_orders_2(X1)=k1_xboole_0|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3998, plain, (k1_relset_1(X1,k2_lattice3(X2))=X1|v2_struct_0(X2)|~m1_subset_1(k2_lattice3(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r3_orders_2(k3_lattice3(X2),esk9_3(X3,X1,k2_lattice3(X2)),esk9_3(X3,X1,k2_lattice3(X2)))|~v10_lattices(X2)|~l3_lattices(X2)|~m1_subset_1(esk9_3(X3,X1,k2_lattice3(X2)),u1_struct_0(k3_lattice3(X2))))).#trainneg
cnf(i_0_8287, plain, (g1_orders_2(u1_struct_0(X1),k2_lattice3(k3_lattice3(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))))=k3_lattice3(k3_lattice3(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))))|v2_struct_0(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(X1)|~v10_lattices(k3_lattice3(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))))|~v10_lattices(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~l3_lattices(k3_lattice3(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1))))|~l3_lattices(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_835, plain, (v3_orders_2(g1_orders_2(X1,esk21_2(X1,X1)))|~v1_relat_2(esk21_2(X1,X1))|~v1_partfun1(esk21_2(X1,X1),X1))).#trainneg
cnf(i_0_836, plain, (v3_orders_2(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v1_relat_2(esk19_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk19_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_837, plain, (v3_orders_2(g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v1_relat_2(esk18_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk18_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_10405, plain, (k9_xtuple_0(k2_lattice3(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_10420, plain, (v4_relat_1(k2_lattice3(k1_lattice3(X1)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10152, plain, (k9_xtuple_0(X2)=u1_struct_0(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_10406, plain, (k1_relset_1(X1,k2_lattice3(X2))=u1_struct_0(X2)|v2_struct_0(X2)|~v4_relat_1(k2_lattice3(X2),X1)|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_7345, plain, (X1=X2|g1_orders_2(X1,k1_wellord2(X3))!=g1_orders_2(X2,X4)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_10272, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|v1_xboole_0(esk2_0)|esk12_1(esk2_0)!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_10402, plain, (k1_zfmisc_1(X1)=k1_xboole_0|g3_lattices(k1_xboole_0,u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)),k1_xboole_0)!=k1_lattice3(X1)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_838, plain, (v3_orders_2(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v1_relat_2(esk17_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk17_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_840, plain, (v3_orders_2(g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1))))|~v1_relat_2(esk14_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk14_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_841, plain, (v3_orders_2(g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~v1_relat_2(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))|~v1_partfun1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X1))).#trainneg
cnf(i_0_10150, plain, (k9_xtuple_0(k2_lattice3(X1))=u1_struct_0(k3_lattice3(X1))|v2_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v1_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_10135, plain, (~v1_xboole_0(k2_lattice3(k1_lattice3(X1))))).#trainneg
cnf(i_0_10185, plain, (k1_relset_1(X1,X3)=u1_struct_0(g1_orders_2(X2,X3))|v2_struct_0(g1_orders_2(X2,X3))|~v4_relat_1(X3,X1)|~l1_orders_2(g1_orders_2(X2,X3))|~v3_orders_2(g1_orders_2(X2,X3))|~v1_orders_2(g1_orders_2(X2,X3)))).#trainneg
cnf(i_0_10395, plain, (k2_relset_1(X1,k1_wellord2(X1))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_10359, plain, (l1_orders_2(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~v4_relat_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_6898, plain, (v3_orders_2(g1_orders_2(X1,k1_toler_1(k1_xboole_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_7417, plain, (k9_xtuple_0(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_843, plain, (v3_orders_2(g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v1_relat_2(esk4_2(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(esk4_2(k2_zfmisc_1(X1,X1),X2),X1)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_844, plain, (v3_orders_2(g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v1_relat_2(esk3_2(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(esk3_2(k2_zfmisc_1(X1,X1),X2),X1)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_845, plain, (v3_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v1_relat_2(k1_toler_1(X2,X1))|~v1_partfun1(k1_toler_1(X2,X1),X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10358, plain, (v1_orders_2(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~v4_relat_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_10182, plain, (k1_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=k1_zfmisc_1(X2)|~v4_relat_1(k2_lattice3(k1_lattice3(X2)),X1))).#trainneg
cnf(i_0_10477, plain, (g1_orders_2(X1,k1_wellord2(X1))=g1_orders_2(X1,k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10419, plain, (v4_relat_1(k2_lattice3(g1_orders_2(X1,k1_wellord2(X1))),X1)|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v10_lattices(g1_orders_2(X1,k1_wellord2(X1)))|~l3_lattices(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_3569, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)))),X4)|X4!=k1_wellord2(k2_lattice3(k1_lattice3(X5)))|~r2_hidden(X1,k2_lattice3(k1_lattice3(X5)))|~v1_relat_1(X4)|~r1_tarski(X1,k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)))|~r1_orders_2(k3_lattice3(k1_lattice3(X5)),X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X5))|~m1_subset_1(X2,k1_zfmisc_1(X5)))).#trainneg
cnf(i_0_6964, plain, (v4_orders_2(g1_orders_2(X1,k1_toler_1(k1_xboole_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_7528, plain, (k9_xtuple_0(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6258, negated_conjecture, (v4_relat_1(esk12_1(esk2_0),k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|v1_xboole_0(esk2_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_847, plain, (v3_orders_2(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v1_relat_2(k1_relset_1(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X1),X2),X1)|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_848, plain, (v3_orders_2(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v1_relat_2(k2_relset_1(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X1),X2),X1)|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5058, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~r1_tarski(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4169, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X3,X2)),k2_lattice3(X1))|~r3_orders_2(k3_lattice3(X1),X2,X3)|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1))))).#trainneg
cnf(i_0_1595, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_subset_1(X5,X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X5)))).#trainneg
cnf(i_0_6981, plain, (v5_orders_2(g1_orders_2(X1,k1_toler_1(k1_xboole_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_9654, negated_conjecture, (k2_zfmisc_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|v1_xboole_0(esk2_0)|esk12_1(esk2_0)!=k1_xboole_0|~v1_relat_1(esk12_1(esk2_0))|esk12_1(esk2_0)!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_851, plain, (v3_orders_2(g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v1_relat_2(k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3))|~v1_partfun1(k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_852, plain, (v3_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v1_relat_2(k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3))|~v1_partfun1(k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_853, plain, (v4_orders_2(g1_orders_2(X1,esk21_2(X1,X1)))|~v8_relat_2(esk21_2(X1,X1))|~v1_partfun1(esk21_2(X1,X1),X1))).#trainneg
cnf(i_0_6860, plain, (r3_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))),X3,X4)|v2_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~v3_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_854, plain, (v4_orders_2(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v8_relat_2(esk19_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk19_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_855, plain, (v4_orders_2(g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v8_relat_2(esk18_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk18_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_856, plain, (v4_orders_2(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v8_relat_2(esk17_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk17_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_10798, plain, (k2_relset_1(k1_zfmisc_1(X1),k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1))).#trainneg
cnf(i_0_11142, plain, (v3_orders_2(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~v4_relat_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k2_zfmisc_1(X1,X1))|~v1_relat_2(X2)|~v1_partfun1(X2,X1))).#trainneg
cnf(i_0_11146, plain, (v3_orders_2(g1_orders_2(X1,X2))|~v4_relat_1(k1_wellord2(X2),k2_zfmisc_1(X1,X1))|~v1_relat_2(X2)|~v1_partfun1(X2,X1))).#trainneg
cnf(i_0_10816, plain, (k2_relset_1(k1_zfmisc_1(X1),u1_orders_2(k1_lattice3(X1)))=k1_zfmisc_1(X1)|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_10696, plain, (~v1_xboole_0(k1_relat_1(k2_lattice3(k1_lattice3(X1)))))).#trainneg
cnf(i_0_10892, plain, (l1_orders_2(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~v4_relat_1(k1_wellord2(X2),k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_10980, plain, (k1_wellord2(X1)=k1_wellord2(X2)|g1_orders_2(X1,k1_wellord2(X1))!=g1_orders_2(X2,k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_858, plain, (v4_orders_2(g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1))))|~v8_relat_2(esk14_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk14_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_859, plain, (v4_orders_2(g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~v8_relat_2(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))|~v1_partfun1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X1))).#trainneg
cnf(i_0_861, plain, (v4_orders_2(g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v8_relat_2(esk4_2(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(esk4_2(k2_zfmisc_1(X1,X1),X2),X1)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_10927, plain, (v1_orders_2(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~v4_relat_1(k1_wellord2(X2),k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_10792, plain, (k2_relset_1(X2,u1_orders_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~v5_relat_1(u1_orders_2(X1),X2)|~v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11043, plain, (k1_wellord2(X1)=X2|g1_orders_2(X1,k1_xboole_0)!=g1_orders_2(X3,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10908, plain, (k1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k2_zfmisc_1(k1_xboole_0,k1_xboole_0)|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_10783, plain, (k1_zfmisc_1(X1)=k1_xboole_0|g3_lattices(k1_xboole_0,k1_xboole_0,k1_xboole_0)!=k1_lattice3(X1)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_862, plain, (v4_orders_2(g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v8_relat_2(esk3_2(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(esk3_2(k2_zfmisc_1(X1,X1),X2),X1)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_5043, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_wellord2(X3)))|~r1_tarski(X1,X2)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_863, plain, (v4_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v8_relat_2(k1_toler_1(X2,X1))|~v1_partfun1(k1_toler_1(X2,X1),X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_864, plain, (v4_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(X1)|~v8_relat_2(k2_lattice3(X1))|~v1_partfun1(k2_lattice3(X1),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_10840, plain, (k1_relset_1(X1,u1_orders_2(k1_lattice3(X2)))=k1_zfmisc_1(X2)|~v4_relat_1(u1_orders_2(k1_lattice3(X2)),X1)|~l1_orders_2(k1_lattice3(X2))|~v3_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2)))).#trainneg
cnf(i_0_11095, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_xboole_0))|~r1_tarski(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10801, plain, (k2_relset_1(u1_struct_0(g1_orders_2(X1,X2)),X2)=u1_struct_0(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_10815, plain, (k2_relset_1(X1,u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_10805, plain, (k2_relset_1(u1_struct_0(X1),u1_orders_2(k3_lattice3(X1)))=u1_struct_0(X1)|v2_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_10714, plain, (k9_xtuple_0(u1_orders_2(k1_lattice3(X1)))=k1_zfmisc_1(X1)|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_10931, plain, (l1_orders_2(g1_orders_2(X1,k1_zfmisc_1(X2)))|~v4_relat_1(k2_lattice3(k1_lattice3(X2)),k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_9844, plain, (r1_tarski(X1,esk10_4(X2,X3,k2_lattice3(k1_lattice3(X4)),X1))|~m1_subset_1(esk10_4(X2,X3,k2_lattice3(k1_lattice3(X4)),X1),k1_zfmisc_1(X4))|~m1_subset_1(X1,k1_zfmisc_1(X4))|k1_relset_1(X3,k2_lattice3(k1_lattice3(X4)))!=X3|~r2_hidden(X1,X3)|~m1_subset_1(k2_lattice3(k1_lattice3(X4)),k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_865, plain, (v4_orders_2(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v8_relat_2(k1_relset_1(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X1),X2),X1)|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5048, plain, (m1_yellow_0(g1_orders_2(X1,k1_wellord2(X3)),g1_orders_2(X2,k1_wellord2(X2)))|~r1_tarski(X1,X2)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_866, plain, (v4_orders_2(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v8_relat_2(k2_relset_1(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X1),X2),X1)|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_869, plain, (v4_orders_2(g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v8_relat_2(k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3))|~v1_partfun1(k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_7869, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X2,X3)),k2_lattice3(X1))|~r1_orders_2(k3_lattice3(X1),X3,X2)|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X1)))|~m1_subset_1(X3,u1_struct_0(k3_lattice3(X1))))).#trainneg
cnf(i_0_7411, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),k1_wellord2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_870, plain, (v4_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v8_relat_2(k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3))|~v1_partfun1(k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_10899, plain, (v3_ordinal1(k1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v5_ordinal1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_funct_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_11123, plain, (v3_ordinal1(k1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))))|~v5_ordinal1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_funct_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_5056, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X1)),g1_orders_2(X2,k1_wellord2(X3)))|~r1_tarski(X1,X2)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_871, plain, (v5_orders_2(g1_orders_2(X1,esk21_2(X1,X1)))|~v4_relat_2(esk21_2(X1,X1))|~v1_partfun1(esk21_2(X1,X1),X1))).#trainneg
cnf(i_0_872, plain, (v5_orders_2(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v4_relat_2(esk19_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk19_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_4708, plain, (r1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))),X3,X4)|~r1_orders_2(k3_lattice3(k1_lattice3(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2))))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k2_lattice3(k1_lattice3(X2)))))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3106, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),X3)|v2_struct_0(g1_orders_2(X4,X3))|~l1_orders_2(g1_orders_2(X4,X3))|~v1_orders_2(g1_orders_2(X4,X3))|~m1_subset_1(X2,u1_struct_0(g1_orders_2(X4,X3)))|~m1_subset_1(X1,u1_struct_0(g1_orders_2(X4,X3)))|~r3_orders_2(g1_orders_2(X4,X3),X1,X2)|~v3_orders_2(g1_orders_2(X4,X3)))).#trainneg
cnf(i_0_5666, plain, (k2_relset_1(X1,u1_orders_2(k3_lattice3(X2)))=X1|v2_struct_0(X2)|~l1_orders_2(k3_lattice3(X2))|~m1_subset_1(esk7_3(X3,X1,u1_orders_2(k3_lattice3(X2))),u1_struct_0(k3_lattice3(X2)))|~m1_subset_1(u1_orders_2(k3_lattice3(X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_5552, plain, (l1_orders_2(g1_orders_2(X1,k10_xtuple_0(X2)))|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_4866, plain, (v2_struct_0(g1_orders_2(X1,X2))|r2_hidden(k2_tarski(k2_tarski(esk9_3(X3,u1_struct_0(g1_orders_2(X1,X2)),X4),esk9_3(X3,u1_struct_0(g1_orders_2(X1,X2)),X4)),k1_tarski(esk9_3(X3,u1_struct_0(g1_orders_2(X1,X2)),X4))),X2)|k1_relset_1(u1_struct_0(g1_orders_2(X1,X2)),X4)=u1_struct_0(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_orders_2(X1,X2)),X3))))).#trainneg
cnf(i_0_11298, plain, (k2_lattice3(k1_lattice3(X1))!=k1_xboole_0)).#trainneg
cnf(i_0_873, plain, (v5_orders_2(g1_orders_2(X1,esk18_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v4_relat_2(esk18_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk18_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_874, plain, (v5_orders_2(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1))|~v4_relat_2(esk17_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk17_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_876, plain, (v5_orders_2(g1_orders_2(X1,esk14_1(k2_zfmisc_1(X1,X1))))|~v4_relat_2(esk14_1(k2_zfmisc_1(X1,X1)))|~v1_partfun1(esk14_1(k2_zfmisc_1(X1,X1)),X1))).#trainneg
cnf(i_0_877, plain, (v5_orders_2(g1_orders_2(X1,esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~v4_relat_2(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))|~v1_partfun1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X1))).#trainneg
cnf(i_0_879, plain, (v5_orders_2(g1_orders_2(X1,esk4_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v4_relat_2(esk4_2(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(esk4_2(k2_zfmisc_1(X1,X1),X2),X1)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_880, plain, (v5_orders_2(g1_orders_2(X1,esk3_2(k2_zfmisc_1(X1,X1),X2)))|X2=k1_lattice3(k2_zfmisc_1(X1,X1))|~v4_relat_2(esk3_2(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(esk3_2(k2_zfmisc_1(X1,X1),X2),X1)|u1_struct_0(X2)!=k1_zfmisc_1(k2_zfmisc_1(X1,X1))|~l3_lattices(X2)|~v3_lattices(X2))).#trainneg
cnf(i_0_11285, plain, (v3_orders_2(g1_orders_2(X1,X2))|~v1_relat_2(X2)|~v1_partfun1(X2,X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_11535, plain, (k2_relset_1(u1_struct_0(X1),k2_lattice3(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_11357, plain, (k10_xtuple_0(k2_lattice3(X1))=u1_struct_0(k3_lattice3(X1))|v2_struct_0(k3_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(k3_lattice3(X1))|~v3_orders_2(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_11653, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k2_tarski(esk7_3(X2,X1,X3),esk7_3(X2,X1,X3)),k1_tarski(esk7_3(X2,X1,X3))),k1_wellord2(X1))|k2_relset_1(X1,X3)=X1|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1))))).#trainneg
cnf(i_0_11271, plain, (k1_relat_1(k2_lattice3(k1_lattice3(X1)))=k1_zfmisc_1(X1)|~v5_relat_1(k2_lattice3(k1_lattice3(X1)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11359, plain, (k10_xtuple_0(X2)=u1_struct_0(g1_orders_2(X1,X2))|v2_struct_0(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2))|~v3_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_10987, plain, (u1_orders_2(X1)=k1_wellord2(X2)|X1!=g1_orders_2(X2,k1_xboole_0)|~l1_orders_2(X1)|~v1_orders_2(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_881, plain, (v5_orders_2(g1_orders_2(X1,k1_toler_1(X2,X1)))|~v4_relat_2(k1_toler_1(X2,X1))|~v1_partfun1(k1_toler_1(X2,X1),X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_882, plain, (v5_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|v2_struct_0(X1)|~v4_relat_2(k2_lattice3(X1))|~v1_partfun1(k2_lattice3(X1),u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_883, plain, (v5_orders_2(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v4_relat_2(k1_relset_1(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X1),X2),X1)|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_11507, plain, (k10_xtuple_0(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_11344, plain, (k2_relset_1(X1,k2_lattice3(k1_lattice3(X2)))=k1_zfmisc_1(X2)|~v5_relat_1(k2_lattice3(k1_lattice3(X2)),X1))).#trainneg
cnf(i_0_11649, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k2_tarski(esk12_1(X1),esk12_1(X1)),k1_tarski(esk12_1(X1))),k1_wellord2(X1)))).#trainneg
cnf(i_0_11347, plain, (k2_relset_1(X1,X3)=u1_struct_0(g1_orders_2(X2,X3))|v2_struct_0(g1_orders_2(X2,X3))|~v5_relat_1(X3,X1)|~l1_orders_2(g1_orders_2(X2,X3))|~v3_orders_2(g1_orders_2(X2,X3))|~v1_orders_2(g1_orders_2(X2,X3)))).#trainneg
cnf(i_0_11659, plain, (v2_struct_0(g1_orders_2(k1_numbers,k1_wellord2(k1_numbers)))|r2_hidden(k2_tarski(k2_tarski(np__1,np__1),k1_tarski(np__1)),k1_wellord2(k1_numbers)))).#trainneg
cnf(i_0_884, plain, (v5_orders_2(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v4_relat_2(k2_relset_1(k2_zfmisc_1(X1,X1),X2))|~v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X1),X2),X1)|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_10036, plain, (k2_zfmisc_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~v1_relat_1(u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_885, plain, (v5_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v4_relat_2(u1_orders_2(X1))|~v1_partfun1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_887, plain, (v5_orders_2(g1_orders_2(X1,k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v4_relat_2(k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3))|~v1_partfun1(k9_subset_1(k2_zfmisc_1(X1,X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_11745, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X3,X2)),u1_orders_2(X1))|~r3_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_11761, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X2)),k1_wellord2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_11747, plain, (r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X3,X2)),k2_lattice3(k1_lattice3(X1)))|~r3_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11282, plain, (v3_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_partfun1(k1_xboole_0,X1))).#trainneg
cnf(i_0_4862, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5375, plain, (k2_relset_1(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))=X1|~m1_subset_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),X3,esk7_3(X4,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))))|~m1_subset_1(esk7_3(X4,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_11660, plain, (v2_struct_0(g1_orders_2(k4_ordinal1,k1_wellord2(k4_ordinal1)))|r2_hidden(k2_tarski(k2_tarski(np__1,np__1),k1_tarski(np__1)),k1_wellord2(k4_ordinal1)))).#trainneg
cnf(i_0_888, plain, (v5_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v4_relat_2(k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3))|~v1_partfun1(k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_11383, plain, (k2_zfmisc_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~v1_relat_1(u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6248, negated_conjecture, (v1_funct_2(esk12_1(esk2_0),k2_zfmisc_1(k1_xboole_0,k1_xboole_0),k1_xboole_0)|v1_xboole_0(esk2_0)|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_890, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk7_3(X2,X3,X4))),X5)|k2_relset_1(X3,X4)=X3|X5!=k1_wellord2(X3)|~r2_hidden(X1,X3)|~v1_relat_1(X5)|~r1_tarski(X1,esk7_3(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_891, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk6_2(X2,X3)),k1_tarski(X1)),X4)|X3=k1_wellord2(X2)|X4!=k1_wellord2(X2)|~r2_hidden(X1,X2)|~v1_relat_1(X4)|~r1_tarski(X1,esk6_2(X2,X3))|k1_relat_1(X3)!=X2|~v1_relat_1(X3))).#trainneg
cnf(i_0_12199, plain, (v2_struct_0(g1_orders_2(k1_numbers,k1_wellord2(k1_numbers)))|~v1_xboole_0(k1_wellord2(k1_numbers)))).#trainneg
cnf(i_0_12220, plain, (k2_zfmisc_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_892, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk5_2(X2,X3)),k1_tarski(X1)),X4)|X3=k1_wellord2(X2)|X4!=k1_wellord2(X2)|~r2_hidden(X1,X2)|~v1_relat_1(X4)|~r1_tarski(X1,esk5_2(X2,X3))|k1_relat_1(X3)!=X2|~v1_relat_1(X3))).#trainneg
cnf(i_0_12350, plain, (k1_xboole_0=X2|v2_struct_0(g1_orders_2(X2,k1_xboole_0))|~v1_partfun1(k1_xboole_0,X2))).#trainneg
cnf(i_0_12401, plain, (v2_struct_0(g1_orders_2(k4_ordinal1,k1_wellord2(k4_ordinal1)))|~v1_xboole_0(k1_wellord2(k4_ordinal1)))).#trainneg
cnf(i_0_12198, plain, (v2_struct_0(g1_orders_2(k1_numbers,k1_wellord2(k1_numbers)))|~r2_hidden(k1_wellord2(k1_numbers),k2_tarski(k2_tarski(np__1,np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_894, plain, (X1=k1_lattice3(X2)|k1_binop_1(u2_lattices(X1),esk3_2(X2,X1),esk4_2(X2,X1))!=k2_xboole_0(esk3_2(X2,X1),esk4_2(X2,X1))|k1_binop_1(u1_lattices(X1),esk3_2(X2,X1),esk4_2(X2,X1))!=k9_subset_1(X2,esk3_2(X2,X1),esk4_2(X2,X1))|u1_struct_0(X1)!=k1_zfmisc_1(X2)|~l3_lattices(X1)|~v3_lattices(X1)|~m1_subset_1(esk4_2(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(esk3_2(X2,X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_899, plain, (r1_tarski(X1,X2)|X3!=k1_wellord2(X4)|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),X3)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~v1_relat_1(X3))).#trainneg
cnf(i_0_11964, plain, (k10_xtuple_0(k2_lattice3(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~v5_relat_1(k2_lattice3(X1),u1_struct_0(X1))|~v1_relat_1(k2_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_12419, plain, (k2_zfmisc_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_wellord2(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_12400, plain, (v2_struct_0(g1_orders_2(k4_ordinal1,k1_wellord2(k4_ordinal1)))|~r2_hidden(k1_wellord2(k4_ordinal1),k2_tarski(k2_tarski(np__1,np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_909, plain, (r1_tarski(X1,X2)|u1_orders_2(X3)!=k1_wellord2(X4)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~v1_relat_1(u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_919, plain, (r1_orders_2(X1,X2,X3)|v1_xboole_0(u1_orders_2(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),u1_orders_2(X1)))).#trainneg
cnf(i_0_12050, plain, (k10_xtuple_0(k1_wellord2(X1))=X1|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1))))).#trainneg
cnf(i_0_1597, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_subset_1(X5,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(X5)))).#trainneg
cnf(i_0_11373, plain, (k1_wellord2(X1)=k2_lattice3(k1_lattice3(X2))|g1_orders_2(X1,k1_xboole_0)!=k3_lattice3(k1_lattice3(X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_921, plain, (m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk10_4(X2,X3,X4,X1))),X4)|k1_relset_1(X3,X4)!=X3|~r2_hidden(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_5061, plain, (v4_yellow_0(g1_orders_2(X1,k1_wellord2(X3)),g1_orders_2(X2,k1_wellord2(X2)))|~r1_tarski(X1,X2)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_924, plain, (r2_hidden(k2_tarski(k2_tarski(X1,k2_tarski(k1_tarski(X2),k2_tarski(X2,esk10_4(X3,X4,X5,X2)))),k1_tarski(X1)),X6)|X6!=k1_wellord2(X5)|~r2_hidden(X1,X5)|~v1_relat_1(X6)|~r1_tarski(X1,k2_tarski(k1_tarski(X2),k2_tarski(X2,esk10_4(X3,X4,X5,X2))))|k1_relset_1(X4,X5)!=X4|~r2_hidden(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))).#trainneg
cnf(i_0_925, plain, (m1_subset_1(k2_tarski(k1_tarski(esk5_2(X1,X2)),k2_tarski(esk5_2(X1,X2),esk6_2(X1,X2))),X2)|X2=k1_wellord2(X1)|r1_tarski(esk5_2(X1,X2),esk6_2(X1,X2))|k1_relat_1(X2)!=X1|~v1_relat_1(X2))).#trainneg
cnf(i_0_12479, plain, (k2_relset_1(X1,X2)=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|X2!=k1_wellord2(X1)|~r2_hidden(esk7_3(X3,X1,X2),X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_12363, plain, (k1_relset_1(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))=X1|v2_struct_0(g1_orders_2(X2,k1_wellord2(X2)))|~m1_subset_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X2))).#trainneg
cnf(i_0_12025, plain, (k10_xtuple_0(u1_orders_2(k1_lattice3(X1)))=k1_zfmisc_1(X1)|~l1_orders_2(k1_lattice3(X1))|~v3_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_12476, plain, (r1_orders_2(g1_orders_2(X1,X2),X3,esk7_3(X4,X5,X6))|k2_relset_1(X5,X6)=X5|~l1_orders_2(g1_orders_2(X1,X2))|~v1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(esk7_3(X4,X5,X6),u1_struct_0(g1_orders_2(X1,X2)))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,X2)))|X2!=k1_wellord2(X5)|~r2_hidden(X3,X5)|~v1_relat_1(X2)|~r1_tarski(X3,esk7_3(X4,X5,X6))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_12379, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X2)),u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12055, plain, (k10_xtuple_0(X1)=u1_struct_0(g1_orders_2(X2,X1))|v2_struct_0(g1_orders_2(X2,X1))|~l1_orders_2(g1_orders_2(X2,X1))|~v1_orders_2(g1_orders_2(X2,X1))|~v1_xboole_0(X1)|~v1_relat_2(X1)|~v1_partfun1(X1,X2))).#trainneg
cnf(i_0_12542, plain, (~v1_xboole_0(k1_wellord2(k1_numbers)))).#trainneg
cnf(i_0_930, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,esk8_4(X2,X3,X4,X1)),k1_tarski(esk8_4(X2,X3,X4,X1))),X4)|k2_relset_1(X3,X4)!=X3|~r2_hidden(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12604, plain, (~v1_xboole_0(k1_wellord2(k4_ordinal1)))).#trainneg
cnf(i_0_928, plain, (r2_hidden(k2_tarski(k2_tarski(X1,k2_tarski(k1_tarski(esk5_2(X2,X3)),k2_tarski(esk5_2(X2,X3),esk6_2(X2,X3)))),k1_tarski(X1)),X4)|X3=k1_wellord2(X2)|r1_tarski(esk5_2(X2,X3),esk6_2(X2,X3))|X4!=k1_wellord2(X3)|~r2_hidden(X1,X3)|~v1_relat_1(X4)|~r1_tarski(X1,k2_tarski(k1_tarski(esk5_2(X2,X3)),k2_tarski(esk5_2(X2,X3),esk6_2(X2,X3))))|k1_relat_1(X3)!=X2|~v1_relat_1(X3))).#trainneg
cnf(i_0_933, plain, (r2_hidden(k2_tarski(k2_tarski(X1,k2_tarski(k2_tarski(X2,esk8_4(X3,X4,X5,X2)),k1_tarski(esk8_4(X3,X4,X5,X2)))),k1_tarski(X1)),X6)|X6!=k1_wellord2(X5)|~r2_hidden(X1,X5)|~v1_relat_1(X6)|~r1_tarski(X1,k2_tarski(k2_tarski(X2,esk8_4(X3,X4,X5,X2)),k1_tarski(esk8_4(X3,X4,X5,X2))))|k2_relset_1(X4,X5)!=X4|~r2_hidden(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_12702, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~m1_subset_1(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),k2_lattice3(k1_lattice3(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12712, plain, (v1_xboole_0(X1)|k10_xtuple_0(k1_wellord2(X1))=X1)).#trainneg
cnf(i_0_12937, plain, (k2_relset_1(X1,k1_wellord2(X1))=X1)).#trainneg
cnf(i_0_12660, plain, (k2_relset_1(X1,k2_lattice3(X2))=u1_struct_0(X2)|v2_struct_0(X2)|~v5_relat_1(k2_lattice3(X2),X1)|~v1_relat_1(k2_lattice3(X2))|~v10_lattices(X2)|~l3_lattices(X2))).#trainneg
cnf(i_0_10963, plain, (u1_orders_2(g1_orders_2(X1,k1_xboole_0))=k1_wellord2(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_936, plain, (v3_lattices(g3_lattices(u1_struct_0(X1),X2,u1_lattices(X1)))|~v1_funct_2(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u1_lattices(X1))|~v1_funct_1(X2)|~m1_subset_1(u1_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_937, plain, (v3_lattices(g3_lattices(u1_struct_0(X1),X2,u2_lattices(X1)))|~v1_funct_2(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_lattices(X1))|~v1_funct_1(X2)|~m1_subset_1(u2_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_12643, plain, (r1_tarski(X1,X2)|v2_struct_0(X3)|k2_lattice3(X3)!=k1_wellord2(X4)|~r2_hidden(X2,X4)|~r2_hidden(X1,X4)|~v1_relat_1(k2_lattice3(X3))|~r1_orders_2(k3_lattice3(X3),X1,X2)|~v10_lattices(X3)|~l3_lattices(X3)|~m1_subset_1(X2,u1_struct_0(k3_lattice3(X3)))|~m1_subset_1(X1,u1_struct_0(k3_lattice3(X3))))).#trainneg
cnf(i_0_12357, plain, (v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v1_xboole_0(u1_orders_2(g1_orders_2(X1,k1_wellord2(X1))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12537, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_zfmisc_1(X1)!=X2|~v1_xboole_0(X2))).#trainneg
cnf(i_0_966, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_12719, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|v2_struct_0(g1_orders_2(X1,k1_wellord2(X1)))|~v5_relat_1(k1_wellord2(X1),X2))).#trainneg
cnf(i_0_12720, plain, (m1_subset_1(X1,X2)|v2_struct_0(g1_orders_2(X3,k1_wellord2(X3)))|~v5_relat_1(k1_wellord2(X3),X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_938, plain, (l3_lattices(g3_lattices(u1_struct_0(X1),X2,u1_lattices(X1)))|~v1_funct_2(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u1_lattices(X1))|~v1_funct_1(X2)|~m1_subset_1(u1_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_939, plain, (l3_lattices(g3_lattices(u1_struct_0(X1),X2,u2_lattices(X1)))|~v1_funct_2(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_lattices(X1))|~v1_funct_1(X2)|~m1_subset_1(u2_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_8319, plain, (v2_struct_0(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|r2_hidden(k2_tarski(k2_tarski(X2,X2),k1_tarski(X2)),k2_lattice3(X1))|v2_struct_0(X1)|~l1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v3_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~v1_orders_2(g1_orders_2(u1_struct_0(X1),k2_lattice3(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_12388, plain, (k2_relset_1(X1,k1_wellord2(X2))=X1|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),X3,esk7_3(X4,X1,k1_wellord2(X2)))|~m1_subset_1(k1_wellord2(X2),k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(esk7_3(X4,X1,k1_wellord2(X2)),X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_13263, plain, (k10_xtuple_0(k1_wellord2(X1))=X1)).#trainneg
cnf(i_0_12536, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(X2)|k1_zfmisc_1(X1)!=X2|k1_xboole_0!=X2|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1010, plain, (v1_funct_2(u2_lattices(k1_lattice3(X1)),k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1))|~l2_lattices(k1_lattice3(X1)))).#trainneg
cnf(i_0_940, plain, (v1_xboole_0(u1_struct_0(X1))|~v1_funct_2(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u1_lattices(X1))|~v1_funct_1(X2)|~v2_struct_0(g3_lattices(u1_struct_0(X1),X2,u1_lattices(X1)))|~m1_subset_1(u1_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_941, plain, (v1_xboole_0(u1_struct_0(X1))|~v1_funct_2(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_lattices(X1))|~v1_funct_1(X2)|~v2_struct_0(g3_lattices(u1_struct_0(X1),X2,u2_lattices(X1)))|~m1_subset_1(u2_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_13631, plain, (v2_struct_0(X1)|r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X2)),k2_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_12539, plain, (k2_lattice3(k1_lattice3(X1))=k1_wellord2(esk14_1(X2))|k1_zfmisc_1(X1)!=esk14_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1011, plain, (v1_funct_2(u1_lattices(k1_lattice3(X1)),k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1))|~l1_lattices(k1_lattice3(X1)))).#trainneg
cnf(i_0_942, plain, (X1=X2|g3_lattices(u1_struct_0(X3),u1_lattices(X3),X1)!=g3_lattices(X4,X5,X2)|~v1_funct_2(X1,k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))|~v1_funct_1(u1_lattices(X3))|~v1_funct_1(X1)|~m1_subset_1(u1_lattices(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~l1_lattices(X3))).#trainneg
cnf(i_0_943, plain, (X1=X2|g3_lattices(u1_struct_0(X3),u2_lattices(X3),X1)!=g3_lattices(X4,X5,X2)|~v1_funct_2(X1,k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))|~v1_funct_1(u2_lattices(X3))|~v1_funct_1(X1)|~m1_subset_1(u2_lattices(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~l2_lattices(X3))).#trainneg
cnf(i_0_2729, plain, (v2_struct_0(k3_lattice3(X1))|m1_subset_1(k2_lattice3(k3_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|v2_struct_0(X1)|~v10_lattices(k3_lattice3(X1))|~l3_lattices(k3_lattice3(X1))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_2963, negated_conjecture, (k1_relset_1(esk2_0,X1)=esk2_0|X1!=k1_wellord2(k1_zfmisc_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,X2)))|~m1_subset_1(esk9_3(X2,esk2_0,X1),k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_10335, plain, (k2_lattice3(X1)=X2|v2_struct_0(X1)|k3_lattice3(X1)!=g1_orders_2(X3,X2)|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_3746, plain, (r1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))),X3,X3)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))))|~v1_orders_2(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(k1_lattice3(X2)))))|~r1_orders_2(k1_lattice3(X2),X3,X3)|~l1_orders_2(k1_lattice3(X2))|~v1_orders_2(k1_lattice3(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_11372, plain, (k1_wellord2(X1)=k2_lattice3(k3_lattice3(k1_lattice3(X2)))|g1_orders_2(X1,k1_xboole_0)!=k3_lattice3(k3_lattice3(k1_lattice3(X2)))|~v1_xboole_0(X1)|~v10_lattices(k3_lattice3(k1_lattice3(X2)))|~l3_lattices(k3_lattice3(k1_lattice3(X2))))).#trainneg
cnf(i_0_1015, plain, (m1_subset_1(u2_lattices(k1_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1))))|~l2_lattices(k1_lattice3(X1)))).#trainneg
cnf(i_0_12615, plain, (v2_struct_0(g1_orders_2(k1_numbers,k1_wellord2(k1_numbers)))|~m1_subset_1(k1_wellord2(k1_numbers),k2_tarski(k2_tarski(np__1,np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_944, plain, (X1=X2|g3_lattices(u1_struct_0(X3),X1,u1_lattices(X3))!=g3_lattices(X4,X2,X5)|~v1_funct_2(X1,k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))|~v1_funct_1(u1_lattices(X3))|~v1_funct_1(X1)|~m1_subset_1(u1_lattices(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~l1_lattices(X3))).#trainneg
cnf(i_0_945, plain, (X1=X2|g3_lattices(u1_struct_0(X3),X1,u2_lattices(X3))!=g3_lattices(X4,X2,X5)|~v1_funct_2(X1,k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))|~v1_funct_1(u2_lattices(X3))|~v1_funct_1(X1)|~m1_subset_1(u2_lattices(X3),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X3)),u1_struct_0(X3))))|~l2_lattices(X3))).#trainneg
cnf(i_0_6327, plain, (r3_orders_2(g1_orders_2(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X3,X4)|v2_struct_0(g1_orders_2(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))))|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))))|~v3_orders_2(g1_orders_2(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))))|~v1_orders_2(g1_orders_2(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_5372, plain, (k1_relset_1(X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))))=X1|~m1_subset_1(u1_orders_2(g1_orders_2(X2,k1_wellord2(X2))),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),esk9_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X4)|~m1_subset_1(X4,X2)|~m1_subset_1(esk9_3(X3,X1,u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))),X2))).#trainneg
cnf(i_0_12097, plain, (m1_subset_1(X1,X2)|v2_struct_0(g1_orders_2(X3,k1_wellord2(X3)))|~v5_relat_1(u1_orders_2(g1_orders_2(X3,k1_wellord2(X3))),X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_7131, plain, (u1_struct_0(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))=X1|~l1_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_13654, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v5_relat_1(k1_wellord2(X1),X2))).#trainneg
cnf(i_0_13655, plain, (m1_subset_1(X1,X2)|~v5_relat_1(k1_wellord2(X3),X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_13946, negated_conjecture, (esk2_0=k1_xboole_0|k1_wellord2(k1_zfmisc_1(esk1_0))!=k1_xboole_0)).#trainneg
cnf(i_0_1016, plain, (m1_subset_1(u1_lattices(k1_lattice3(X1)),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k1_zfmisc_1(X1),k1_zfmisc_1(X1)),k1_zfmisc_1(X1))))|~l1_lattices(k1_lattice3(X1)))).#trainneg
cnf(i_0_1017, plain, (k1_zfmisc_1(X1)=X2|g1_orders_2(k1_zfmisc_1(X1),u1_orders_2(k1_lattice3(X1)))!=g1_orders_2(X2,X3)|~l1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_1018, plain, (m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),u1_orders_2(X3))|~r1_orders_2(X3,X1,X2)|~l1_orders_2(X3)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_14066, plain, (u2_lattices(X1)=X2|g3_lattices(u1_struct_0(X1),u2_lattices(X1),u1_lattices(X1))!=g3_lattices(X3,X2,X4)|~l1_lattices(X1)|~v1_funct_1(u2_lattices(X1))|~m1_subset_1(u2_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l2_lattices(X1))).#trainneg
cnf(i_0_14109, plain, (u1_lattices(X1)=X2|g3_lattices(u1_struct_0(X1),u1_lattices(X1),u2_lattices(X1))!=g3_lattices(X3,X2,X4)|~v1_funct_1(u1_lattices(X1))|~l2_lattices(X1)|~m1_subset_1(u1_lattices(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_lattices(X1))).#trainneg
cnf(i_0_14258, plain, (u1_struct_0(g1_orders_2(X1,X2))=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_14135, plain, (k1_relset_1(X1,k1_wellord2(X2))=X1|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),esk9_3(X3,X1,k1_wellord2(X2)),X4)|~m1_subset_1(k1_wellord2(X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(esk9_3(X3,X1,k1_wellord2(X2)),X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_14125, plain, (r3_orders_2(g1_orders_2(X1,k1_wellord2(X2)),X3,X4)|v2_struct_0(g1_orders_2(X1,k1_wellord2(X2)))|~r1_orders_2(g1_orders_2(X2,k1_wellord2(X2)),X3,X4)|~l1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v3_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~v1_orders_2(g1_orders_2(X1,k1_wellord2(X2)))|~m1_subset_1(X4,u1_struct_0(g1_orders_2(X1,k1_wellord2(X2))))|~m1_subset_1(X3,u1_struct_0(g1_orders_2(X1,k1_wellord2(X2))))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_8200, plain, (l1_orders_2(g1_orders_2(X1,k1_toler_1(k1_xboole_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_10404, plain, (k1_xboole_0=X1|g3_lattices(k1_xboole_0,k1_xboole_0,k1_xboole_0)!=g3_lattices(X1,X2,X3)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1033, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_12676, plain, (v2_struct_0(g1_orders_2(k4_ordinal1,k1_wellord2(k4_ordinal1)))|~m1_subset_1(k1_wellord2(k4_ordinal1),k2_tarski(k2_tarski(np__1,np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_1021, plain, (r2_hidden(k2_tarski(k2_tarski(X1,k2_tarski(k1_tarski(X2),k2_tarski(X2,X3))),k1_tarski(X1)),X4)|X4!=k1_wellord2(u1_orders_2(X5))|~r2_hidden(X1,u1_orders_2(X5))|~v1_relat_1(X4)|~r1_tarski(X1,k2_tarski(k1_tarski(X2),k2_tarski(X2,X3)))|~r1_orders_2(X5,X2,X3)|~l1_orders_2(X5)|~m1_subset_1(X3,u1_struct_0(X5))|~m1_subset_1(X2,u1_struct_0(X5)))).#trainneg
cnf(i_0_1034, plain, (v3_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1)))|~v1_relat_2(k2_zfmisc_1(X1,X1))|~v1_partfun1(k2_zfmisc_1(X1,X1),X1))).#trainneg
cnf(i_0_13982, plain, (k2_lattice3(X1)=u1_orders_2(g1_orders_2(X2,k1_wellord2(X2)))|v2_struct_0(X1)|k3_lattice3(X1)!=g1_orders_2(X2,k1_wellord2(X2))|~v10_lattices(X1)|~l3_lattices(X1))).#trainneg
cnf(i_0_14165, plain, (g1_orders_2(X1,k2_lattice3(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3))))=k3_lattice3(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|v2_struct_0(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v10_lattices(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~l3_lattices(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_13649, plain, (k2_relset_1(X1,k1_wellord2(X2))=X2|~v5_relat_1(k1_wellord2(X2),X1))).#trainneg
cnf(i_0_14257, plain, (v2_struct_0(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|r2_hidden(k2_tarski(k2_tarski(X4,X4),k1_tarski(X4)),k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3))|~l1_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v3_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~v1_orders_2(g1_orders_2(X1,k4_subset_1(k2_zfmisc_1(X1,X1),X2,X3)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_6747, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(X2))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1598, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_subset_1(X5,X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(X5)))).#trainneg
cnf(i_0_11298, plain, (k1_wellord2(k1_zfmisc_1(X1))!=k1_xboole_0)).#trainneg
cnf(i_0_1035, plain, (v4_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1)))|~v8_relat_2(k2_zfmisc_1(X1,X1))|~v1_partfun1(k2_zfmisc_1(X1,X1),X1))).#trainneg
cnf(i_0_1036, plain, (v5_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1)))|~v4_relat_2(k2_zfmisc_1(X1,X1))|~v1_partfun1(k2_zfmisc_1(X1,X1),X1))).#trainneg
cnf(i_0_1037, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2146, plain, (u1_orders_2(k3_lattice3(k1_lattice3(X1)))=k1_wellord2(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3552, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k1_wellord2(k1_zfmisc_1(X3)))|~r3_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_4850, negated_conjecture, (k1_relset_1(esk2_0,k1_wellord2(k1_zfmisc_1(esk1_0)))=esk2_0|~m1_subset_1(k1_wellord2(k1_zfmisc_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(esk2_0,X1))))).#trainneg
cnf(i_0_2756, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k1_wellord2(k1_zfmisc_1(X3)))|~r1_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_2442, plain, (k1_wellord2(k1_zfmisc_1(X1))=k1_xboole_0|g1_orders_2(X2,k1_xboole_0)!=k3_lattice3(k1_lattice3(X1)))).#trainneg
cnf(i_0_1039, plain, (k2_zfmisc_1(X1,X1)=X2|g1_orders_2(X1,k2_zfmisc_1(X1,X1))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_1052, plain, (m1_subset_1(esk9_3(X1,esk14_1(X2),X3),X2)|k1_relset_1(esk14_1(X2),X3)=esk14_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk14_1(X2),X1))))).#trainneg
cnf(i_0_1053, plain, (m1_subset_1(esk7_3(X1,esk14_1(X2),X3),X2)|k2_relset_1(esk14_1(X2),X3)=esk14_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,esk14_1(X2)))))).#trainneg
cnf(i_0_5486, plain, (k2_relset_1(X1,k1_wellord2(k1_zfmisc_1(X2)))=X1|~m1_subset_1(esk7_3(X3,X1,k1_wellord2(k1_zfmisc_1(X2))),k1_zfmisc_1(X2))|~m1_subset_1(k1_wellord2(k1_zfmisc_1(X2)),k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_2804, plain, (v4_yellow_0(g1_orders_2(X1,k1_toler_1(k1_wellord2(k1_zfmisc_1(X2)),X1)),k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3732, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k1_wellord2(k1_zfmisc_1(X3)))|~r3_orders_2(k3_lattice3(k1_lattice3(X3)),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_4858, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X1),k1_tarski(X1)),k1_wellord2(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2789, plain, (r1_orders_2(k3_lattice3(k1_lattice3(X1)),X2,X3)|~r2_hidden(k2_tarski(k2_tarski(X2,X3),k1_tarski(X2)),k1_wellord2(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12940, plain, (k2_relset_1(k1_xboole_0,u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k1_xboole_0|u1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1054, plain, (m1_subset_1(esk6_2(esk14_1(X1),X2),X1)|X2=k1_wellord2(esk14_1(X1))|k1_relat_1(X2)!=esk14_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_12943, plain, (k2_relset_1(k1_xboole_0,u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))=k1_xboole_0|u2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))!=k1_xboole_0|~l2_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_13958, negated_conjecture, (k1_relset_1(esk2_0,u1_orders_2(g1_orders_2(esk2_0,k1_wellord2(esk2_0))))=esk2_0|u1_orders_2(g1_orders_2(esk2_0,k1_wellord2(esk2_0)))!=k1_wellord2(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1055, plain, (m1_subset_1(esk5_2(esk14_1(X1),X2),X1)|X2=k1_wellord2(esk14_1(X1))|k1_relat_1(X2)!=esk14_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1056, plain, (m1_subset_1(k2_tarski(k1_tarski(esk5_2(X1,esk14_1(X2))),k2_tarski(esk5_2(X1,esk14_1(X2)),esk6_2(X1,esk14_1(X2)))),X2)|esk14_1(X2)=k1_wellord2(X1)|r1_tarski(esk5_2(X1,esk14_1(X2)),esk6_2(X1,esk14_1(X2)))|k1_relat_1(esk14_1(X2))!=X1|~v1_relat_1(esk14_1(X2)))).#trainneg
cnf(i_0_1057, plain, (m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk10_4(X2,X3,esk14_1(X4),X1))),X4)|k1_relset_1(X3,esk14_1(X4))!=X3|~r2_hidden(X1,X3)|~m1_subset_1(esk14_1(X4),k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_14688, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~v5_relat_1(k1_wellord2(X2),X1))).#trainneg
cnf(i_0_14687, plain, (v3_ordinal1(X2)|~v3_ordinal1(k1_zfmisc_1(X1))|~v5_relat_1(k1_wellord2(X2),X1))).#trainneg
cnf(i_0_1058, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,esk8_4(X2,X3,esk14_1(X4),X1)),k1_tarski(esk8_4(X2,X3,esk14_1(X4),X1))),X4)|k2_relset_1(X3,esk14_1(X4))!=X3|~r2_hidden(X1,X3)|~m1_subset_1(esk14_1(X4),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1059, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk14_1(X2))|~m1_subset_1(X1,esk14_1(X2)))).#trainneg
cnf(i_0_14819, plain, (u1_orders_2(X1)=k1_wellord2(k1_zfmisc_1(X2))|X1!=k3_lattice3(k1_lattice3(X2))|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_14737, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,u1_orders_2(X3))|~v1_xboole_0(u1_struct_0(X3))|~l1_orders_2(X3))).#trainneg
cnf(i_0_1060, plain, (m1_subset_1(esk9_3(X1,esk12_1(k1_zfmisc_1(X2)),X3),X2)|k1_relset_1(esk12_1(k1_zfmisc_1(X2)),X3)=esk12_1(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk12_1(k1_zfmisc_1(X2)),X1))))).#trainneg
cnf(i_0_15345, negated_conjecture, (v3_ordinal1(esk2_0)|~v5_ordinal1(u1_orders_2(g1_orders_2(esk2_0,k1_wellord2(esk2_0))))|~v1_funct_1(u1_orders_2(g1_orders_2(esk2_0,k1_wellord2(esk2_0))))|u1_orders_2(g1_orders_2(esk2_0,k1_wellord2(esk2_0)))!=k1_wellord2(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_15288, negated_conjecture, (k2_relset_1(k1_xboole_0,esk12_1(esk2_0))=k1_xboole_0|v1_xboole_0(esk2_0)|~l1_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk1_0)|~l3_lattices(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_14381, plain, (k1_zfmisc_1(X1)=X2|k1_lattice3(X1)!=g1_orders_2(X2,X3)|~l1_orders_2(k1_lattice3(X1))|~v1_orders_2(k1_lattice3(X1)))).#trainneg
cnf(i_0_1061, plain, (m1_subset_1(esk7_3(X1,esk12_1(k1_zfmisc_1(X2)),X3),X2)|k2_relset_1(esk12_1(k1_zfmisc_1(X2)),X3)=esk12_1(k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,esk12_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_1062, plain, (m1_subset_1(esk6_2(esk12_1(k1_zfmisc_1(X1)),X2),X1)|X2=k1_wellord2(esk12_1(k1_zfmisc_1(X1)))|k1_relat_1(X2)!=esk12_1(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_15061, plain, (k1_relset_1(X1,k1_wellord2(k1_zfmisc_1(X2)))=X1|~m1_subset_1(k1_wellord2(k1_zfmisc_1(X2)),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~r3_orders_2(k3_lattice3(k1_lattice3(X2)),esk9_3(X3,X1,k1_wellord2(k1_zfmisc_1(X2))),esk9_3(X3,X1,k1_wellord2(k1_zfmisc_1(X2))))|~m1_subset_1(esk9_3(X3,X1,k1_wellord2(k1_zfmisc_1(X2))),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1063, plain, (m1_subset_1(esk5_2(esk12_1(k1_zfmisc_1(X1)),X2),X1)|X2=k1_wellord2(esk12_1(k1_zfmisc_1(X1)))|k1_relat_1(X2)!=esk12_1(k1_zfmisc_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_10279, plain, (v5_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X2)|~v1_partfun1(k1_xboole_0,X1))).#trainneg
cnf(i_0_15351, negated_conjecture, (k1_relset_1(esk2_0,u1_orders_2(g1_orders_2(esk2_0,k1_xboole_0)))=esk2_0|u1_orders_2(g1_orders_2(esk2_0,k1_xboole_0))!=k1_wellord2(k1_zfmisc_1(esk1_0))|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_1064, plain, (m1_subset_1(k2_tarski(k1_tarski(esk5_2(X1,esk12_1(k1_zfmisc_1(X2)))),k2_tarski(esk5_2(X1,esk12_1(k1_zfmisc_1(X2))),esk6_2(X1,esk12_1(k1_zfmisc_1(X2))))),X2)|esk12_1(k1_zfmisc_1(X2))=k1_wellord2(X1)|r1_tarski(esk5_2(X1,esk12_1(k1_zfmisc_1(X2))),esk6_2(X1,esk12_1(k1_zfmisc_1(X2))))|k1_relat_1(esk12_1(k1_zfmisc_1(X2)))!=X1|~v1_relat_1(esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1065, plain, (m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk10_4(X2,X3,esk12_1(k1_zfmisc_1(X4)),X1))),X4)|k1_relset_1(X3,esk12_1(k1_zfmisc_1(X4)))!=X3|~r2_hidden(X1,X3)|~m1_subset_1(esk12_1(k1_zfmisc_1(X4)),k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 192
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 338
# Removed in clause preprocessing      : 36
# Initial clauses in saturation        : 302
# Processed clauses                    : 3025
# ...of these trivial                  : 93
# ...subsumed                          : 1469
# ...remaining for further processing  : 1463
# Other redundant clauses eliminated   : 14
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 70
# Backward-rewritten                   : 115
# Generated clauses                    : 15635
# ...of the previous two non-trivial   : 14276
# Contextual simplify-reflections      : 270
# Paramodulations                      : 15539
# Factorizations                       : 0
# Equation resolutions                 : 96
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1278
#    Positive orientable unit clauses  : 132
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 1128
# Current number of unprocessed clauses: 11243
# ...number of literals in the above   : 65553
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1653
# Proof object given clauses           : 119
# Proof search given clauses           : 1463
# Clause-clause subsumption calls (NU) : 538820
# Rec. Clause-clause subsumption calls : 119661
# Non-unit clause-clause subsumptions  : 1749
# Unit Clause-clause subsumption calls : 11844
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 131
# BW rewrite match successes           : 50
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 376761

# -------------------------------------------------
# User time                : 49.800 s
# System time              : 1.140 s
# Total time               : 50.940 s
# Maximum resident set size: 408620 pages
49.79user 1.22system 0:51.46elapsed 99%CPU (0avgtext+0avgdata 408620maxresident)k
0inputs+664outputs (0major+315526minor)pagefaults 0swaps
