# local_data = 0x1a002d0
#TFS# Evaluating axioms..
# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
#TF#SERVER: Lgb model filter skipped 0 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=0, query=164, total=165).
#TF#SERVER: Sent 30657 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (660 bytes expected).
#TF#SERVER: Received 660 bytes from the TF server.
# Preprocessing time       : 0.840 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 3 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=164, query=136, total=301).
#TF#SERVER: Sent 60826 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1204 bytes expected).
#TF#SERVER: Received 1204 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 34 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=263, query=165, total=429).
#TF#SERVER: Sent 109959 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1716 bytes expected).
#TF#SERVER: Received 1716 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 248 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=384, query=395, total=780).
#TF#SERVER: Sent 261047 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3120 bytes expected).
#TF#SERVER: Received 3120 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 435 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=521, query=590, total=1112).
#TF#SERVER: Sent 444645 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4448 bytes expected).
#TF#SERVER: Received 4448 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 463 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=607, query=566, total=1174).
#TF#SERVER: Sent 415888 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4696 bytes expected).
#TF#SERVER: Received 4696 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 524 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=669, query=555, total=1225).
#TF#SERVER: Sent 454215 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4900 bytes expected).
#TF#SERVER: Received 4900 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 488 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=750, query=547, total=1298).
#TF#SERVER: Sent 420244 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5192 bytes expected).
#TF#SERVER: Received 5192 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 523 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=792, query=518, total=1311).
#TF#SERVER: Sent 408964 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5244 bytes expected).
#TF#SERVER: Received 5244 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 506 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=851, query=551, total=1403).
#TF#SERVER: Sent 402723 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5612 bytes expected).
#TF#SERVER: Received 5612 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 332 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=908, query=705, total=1614).
#TF#SERVER: Sent 435893 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6456 bytes expected).
#TF#SERVER: Received 6456 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 545 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=959, query=488, total=1448).
#TF#SERVER: Sent 443454 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5792 bytes expected).
#TF#SERVER: Received 5792 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 605 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1003, query=511, total=1515).
#TF#SERVER: Sent 420106 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6060 bytes expected).
#TF#SERVER: Received 6060 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 587 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1056, query=465, total=1522).
#TF#SERVER: Sent 460588 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6088 bytes expected).
#TF#SERVER: Received 6088 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 823 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1085, query=382, total=1468).
#TF#SERVER: Sent 510411 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5872 bytes expected).
#TF#SERVER: Received 5872 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 803 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1097, query=326, total=1424).
#TF#SERVER: Sent 457232 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5696 bytes expected).
#TF#SERVER: Received 5696 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 676 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1121, query=377, total=1499).
#TF#SERVER: Sent 489491 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5996 bytes expected).
#TF#SERVER: Received 5996 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 622 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1157, query=562, total=1720).
#TF#SERVER: Sent 546663 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6880 bytes expected).
#TF#SERVER: Received 6880 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 832 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1174, query=295, total=1470).
#TF#SERVER: Sent 481719 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5880 bytes expected).
#TF#SERVER: Received 5880 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc6_ordinal1)).
fof(fc8_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k7_xcmplx_0(X1,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc8_xreal_0)).
fof(cc5_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(m1_subset_1(X2,X1)=>v3_ordinal1(X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc5_ordinal1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_k5_numbers)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/minsubpreds65new/t61_pepin', t6_boole)).
fof(fc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_xreal_0(k4_xcmplx_0(X1)))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc3_xreal_0)).
fof(rqRealDiv__k7_xcmplx_0__r2_r2_r1, axiom, k7_xcmplx_0(np__2,np__2)=np__1, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealDiv__k7_xcmplx_0__r2_r2_r1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc1_xreal_0)).
fof(redefinition_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_k2_nat_1)).
fof(commutativity_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_nat_1(X2,X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', commutativity_k2_nat_1)).
fof(redefinition_k4_power, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&m1_subset_1(X2,k1_numbers))=>k4_power(X1,X2)=k3_power(X1,X2)), file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_k4_power)).
fof(t8_newton, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:(v1_xcmplx_0(X3)=>k1_newton(X3,k2_xcmplx_0(X1,X2))=k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2))))), file('/scratch/mptp/minsubpreds65new/t61_pepin', t8_newton)).
fof(t5_newton, axiom, ![X1]:(v1_xcmplx_0(X1)=>k1_newton(X1,np__1)=X1), file('/scratch/mptp/minsubpreds65new/t61_pepin', t5_newton)).
fof(cc6_card_1, axiom, ![X1]:((v3_ordinal1(X1)&v1_finset_1(X1))=>v7_ordinal1(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc6_card_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc1_numerals)).
fof(cc4_card_1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v1_finset_1(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc4_card_1)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_k6_numbers)).
fof(spc1_arithm, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc1_arithm)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(rqRealDiff__k6_xcmplx_0__r2_r3_rm1, axiom, k6_xcmplx_0(np__2,np__3)=k4_xcmplx_0(np__1), file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealDiff__k6_xcmplx_0__r2_r3_rm1)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc2_numerals)).
fof(fc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_int_1(k4_xcmplx_0(X1)))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc3_int_1)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc3_numerals)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc8_ordinal1)).
fof(t62_power, axiom, k4_power(np__2,np__4)=np__16, file('/scratch/mptp/minsubpreds65new/t61_pepin', t62_power)).
fof(spc4_numerals, axiom, (((v2_xxreal_0(np__4)&m2_subset_1(np__4,k1_numbers,k5_numbers))&m1_subset_1(np__4,k5_numbers))&m1_subset_1(np__4,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc4_numerals)).
fof(fc7_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k6_xcmplx_0(X1,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc7_xreal_0)).
fof(t58_pepin, axiom, ![X1]:(v7_ordinal1(X1)=>(r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k4_xcmplx_0(np__1),k4_pepin(X1))=>v1_int_2(k4_pepin(X1)))), file('/scratch/mptp/minsubpreds65new/t61_pepin', t58_pepin)).
fof(t61_pepin, conjecture, v1_int_2(np__257), file('/scratch/mptp/minsubpreds65new/t61_pepin', t61_pepin)).
fof(d2_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>((r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2))=>k1_xreal_0(X1,X2)=k6_xcmplx_0(X1,X2))&(~(r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2)))=>k1_xreal_0(X1,X2)=k6_numbers)))), file('/scratch/mptp/minsubpreds65new/t61_pepin', d2_xreal_0)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc0_boole)).
fof(spc257_numerals, axiom, (((v2_xxreal_0(np__257)&m2_subset_1(np__257,k1_numbers,k5_numbers))&m1_subset_1(np__257,k5_numbers))&m1_subset_1(np__257,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc257_numerals)).
fof(fc4_int_1, axiom, ![X1, X2]:((v1_int_1(X1)&v1_int_1(X2))=>v1_int_1(k6_xcmplx_0(X1,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc4_int_1)).
fof(ie1_power, axiom, ![X1, X2]:((v1_xreal_0(X1)&v7_ordinal1(X2))=>k3_power(X1,X2)=k1_newton(X1,X2)), file('/scratch/mptp/minsubpreds65new/t61_pepin', ie1_power)).
fof(rqRealDiff__k6_xcmplx_0__r3_r1_r2, axiom, k6_xcmplx_0(np__3,np__1)=np__2, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealDiff__k6_xcmplx_0__r3_r1_r2)).
fof(t16_nat_2, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>(r1_xxreal_0(X1,X2)=>k3_nat_d(k3_power(np__2,X2),k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(X2,X1))))), file('/scratch/mptp/minsubpreds65new/t61_pepin', t16_nat_2)).
fof(t61_power, axiom, k4_power(np__2,np__3)=np__8, file('/scratch/mptp/minsubpreds65new/t61_pepin', t61_power)).
fof(rqRealAdd__k2_xcmplx_0__r1_r1_r2, axiom, k2_xcmplx_0(np__1,np__1)=np__2, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealAdd__k2_xcmplx_0__r1_r1_r2)).
fof(rqRealDiff__k6_xcmplx_0__r257_r1_r256, axiom, k6_xcmplx_0(np__257,np__1)=np__256, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealDiff__k6_xcmplx_0__r257_r1_r256)).
fof(cc2_int_1, axiom, ![X1]:(v7_ordinal1(X1)=>v1_int_1(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc2_int_1)).
fof(fc1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k2_xcmplx_0(X1,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc1_nat_1)).
fof(t53_pepin, axiom, k4_pepin(np__3)=np__257, file('/scratch/mptp/minsubpreds65new/t61_pepin', t53_pepin)).
fof(redefinition_k7_nat_d, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>k7_nat_d(X1,X2)=k1_xreal_0(X1,X2)), file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_k7_nat_d)).
fof(rqLessOrEqual__r1_xxreal_0__r0_r256, axiom, r1_xxreal_0(np__0,np__256), file('/scratch/mptp/minsubpreds65new/t61_pepin', rqLessOrEqual__r1_xxreal_0__r0_r256)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc3_xreal_0)).
fof(fc3_newton, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v7_ordinal1(X2))=>v1_xcmplx_0(k1_newton(X1,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc3_newton)).
fof(rqRealAdd__k2_xcmplx_0__r4_r4_r8, axiom, k2_xcmplx_0(np__4,np__4)=np__8, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealAdd__k2_xcmplx_0__r4_r4_r8)).
fof(t15_int_2, axiom, ![X1]:(v1_int_1(X1)=>![X2]:(v1_int_1(X2)=>![X3]:(v1_int_1(X3)=>(r2_int_1(X1,X2,X3)<=>r1_int_1(X3,k6_xcmplx_0(X1,X2)))))), file('/scratch/mptp/minsubpreds65new/t61_pepin', t15_int_2)).
fof(fc2_newton, axiom, ![X1, X2]:((v1_xreal_0(X1)&v7_ordinal1(X2))=>v1_xreal_0(k1_newton(X1,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc2_newton)).
fof(rqRealDiff__k6_xcmplx_0__r8_r1_r7, axiom, k6_xcmplx_0(np__8,np__1)=np__7, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealDiff__k6_xcmplx_0__r8_r1_r7)).
fof(rqLessOrEqual__r1_xxreal_0__r0_r7, axiom, r1_xxreal_0(np__0,np__7), file('/scratch/mptp/minsubpreds65new/t61_pepin', rqLessOrEqual__r1_xxreal_0__r0_r7)).
fof(rqRealAdd__k2_xcmplx_0__r256_r1_r257, axiom, k2_xcmplx_0(np__256,np__1)=np__257, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealAdd__k2_xcmplx_0__r256_r1_r257)).
fof(rqRealMult__k3_xcmplx_0__r16_r16_r256, axiom, k3_xcmplx_0(np__16,np__16)=np__256, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealMult__k3_xcmplx_0__r16_r16_r256)).
fof(l109_pepin, axiom, r1_nat_d(k4_pepin(np__3),k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,k6_numbers),np__128)),np__1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', l109_pepin)).
fof(spc256_numerals, axiom, (((v2_xxreal_0(np__256)&m2_subset_1(np__256,k1_numbers,k5_numbers))&m1_subset_1(np__256,k5_numbers))&m1_subset_1(np__256,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc256_numerals)).
fof(commutativity_k3_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', commutativity_k3_xcmplx_0)).
fof(redefinition_r1_nat_d, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>(r1_nat_d(X1,X2)<=>r1_int_1(X1,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_r1_nat_d)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r8, axiom, r1_xxreal_0(np__1,np__8), file('/scratch/mptp/minsubpreds65new/t61_pepin', rqLessOrEqual__r1_xxreal_0__r1_r8)).
fof(spc7_numerals, axiom, (((v2_xxreal_0(np__7)&m2_subset_1(np__7,k1_numbers,k5_numbers))&m1_subset_1(np__7,k5_numbers))&m1_subset_1(np__7,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc7_numerals)).
fof(cc7_ordinal1, axiom, ![X1]:(v1_xboole_0(X1)=>v7_ordinal1(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc7_ordinal1)).
fof(spc128_numerals, axiom, (((v2_xxreal_0(np__128)&m2_subset_1(np__128,k1_numbers,k5_numbers))&m1_subset_1(np__128,k5_numbers))&m1_subset_1(np__128,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc128_numerals)).
fof(redefinition_k4_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_k4_nat_1)).
fof(rqRealAdd__k2_xcmplx_0__r4_r3_r7, axiom, k2_xcmplx_0(np__4,np__3)=np__7, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealAdd__k2_xcmplx_0__r4_r3_r7)).
fof(rqRealMult__k3_xcmplx_0__r32_r0_r0, axiom, k3_xcmplx_0(np__32,np__0)=np__0, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealMult__k3_xcmplx_0__r32_r0_r0)).
fof(spc32_numerals, axiom, (((v2_xxreal_0(np__32)&m2_subset_1(np__32,k1_numbers,k5_numbers))&m1_subset_1(np__32,k5_numbers))&m1_subset_1(np__32,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc32_numerals)).
fof(rqRealMult__k3_xcmplx_0__r16_r8_r128, axiom, k3_xcmplx_0(np__16,np__8)=np__128, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealMult__k3_xcmplx_0__r16_r8_r128)).
fof(rqRealAdd__k2_xcmplx_0__r0_r128_r128, axiom, k2_xcmplx_0(np__0,np__128)=np__128, file('/scratch/mptp/minsubpreds65new/t61_pepin', rqRealAdd__k2_xcmplx_0__r0_r128_r128)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', spc0_numerals)).
fof(fc7_card_1, axiom, ~(v1_finset_1(k4_ordinal1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', fc7_card_1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', redefinition_m2_subset_1)).
fof(cc2_finset_1, axiom, ![X1]:(v1_finset_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_finset_1(X2))), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc2_finset_1)).
fof(cc2_xreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc2_xreal_0)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', dt_k5_numbers)).
fof(dt_k13_newton, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&m1_subset_1(X2,k5_numbers))=>m2_subset_1(k13_newton(X1,X2),k1_numbers,k5_numbers)), file('/scratch/mptp/minsubpreds65new/t61_pepin', dt_k13_newton)).
fof(cc1_finset_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_finset_1(X1)), file('/scratch/mptp/minsubpreds65new/t61_pepin', cc1_finset_1)).
fof(c_0_75, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_76, plain, ![X52, X53]:(~v1_xreal_0(X52)|~v1_xreal_0(X53)|v1_xreal_0(k7_xcmplx_0(X52,X53))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_xreal_0])])).
fof(c_0_77, plain, ![X13, X14]:(~v3_ordinal1(X13)|(~m1_subset_1(X14,X13)|v3_ordinal1(X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_ordinal1])])])).
cnf(c_0_78, plain, (v3_ordinal1(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_79, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_80, plain, ![X86]:(~v1_xboole_0(X86)|X86=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_81, plain, ![X45]:((v1_xcmplx_0(k4_xcmplx_0(X45))|~v1_xreal_0(X45))&(v1_xreal_0(k4_xcmplx_0(X45))|~v1_xreal_0(X45))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xreal_0])])])).
cnf(c_0_82, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_83, plain, (k7_xcmplx_0(np__2,np__2)=np__1), inference(split_conjunct,[status(thm)],[rqRealDiv__k7_xcmplx_0__r2_r2_r1])).
fof(c_0_84, plain, ![X5]:(~m1_subset_1(X5,k1_numbers)|v1_xreal_0(X5)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
fof(c_0_85, plain, ![X60, X61]:(~m1_subset_1(X60,k5_numbers)|~v7_ordinal1(X61)|k2_nat_1(X60,X61)=k2_xcmplx_0(X60,X61)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_1])])).
fof(c_0_86, plain, ![X20, X21]:(~m1_subset_1(X20,k5_numbers)|~v7_ordinal1(X21)|k2_nat_1(X20,X21)=k2_nat_1(X21,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_nat_1])])).
fof(c_0_87, plain, ![X64, X65]:(~m1_subset_1(X64,k1_numbers)|~m1_subset_1(X65,k1_numbers)|k4_power(X64,X65)=k3_power(X64,X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_power])])).
fof(c_0_88, plain, ![X87, X88, X89]:(~v7_ordinal1(X87)|(~v7_ordinal1(X88)|(~v1_xcmplx_0(X89)|k1_newton(X89,k2_xcmplx_0(X87,X88))=k3_xcmplx_0(k1_newton(X89,X87),k1_newton(X89,X88))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_newton])])])).
fof(c_0_89, plain, ![X85]:(~v1_xcmplx_0(X85)|k1_newton(X85,np__1)=X85), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_newton])])).
fof(c_0_90, plain, ![X15]:(~v3_ordinal1(X15)|~v1_finset_1(X15)|v7_ordinal1(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_card_1])])).
cnf(c_0_91, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_92, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_93, plain, (v3_ordinal1(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_94, plain, (k5_numbers=k4_ordinal1), inference(evalgc,[status(thm)],[c_0_79])).
fof(c_0_95, plain, ![X12]:(~m1_subset_1(X12,k4_ordinal1)|v1_finset_1(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_card_1])])).
cnf(c_0_96, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_97, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_98, plain, ![X73, X74]:(~v1_xcmplx_0(X73)|~v1_xcmplx_0(X74)|k2_xcmplx_0(X73,k4_xcmplx_0(X74))=k6_xcmplx_0(X73,X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc1_arithm])])).
cnf(c_0_99, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
cnf(c_0_100, plain, (k6_xcmplx_0(np__2,np__3)=k4_xcmplx_0(np__1)), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r2_r3_rm1])).
cnf(c_0_101, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_102, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_103, plain, (k7_xcmplx_0(np__2,np__2)=np__1), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_104, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_105, plain, (m1_subset_1(np__2,k1_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
fof(c_0_106, plain, ![X42]:((v1_xcmplx_0(k4_xcmplx_0(X42))|~v1_int_1(X42))&(v1_int_1(k4_xcmplx_0(X42))|~v1_int_1(X42))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_int_1])])])).
cnf(c_0_107, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_108, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_109, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
fof(c_0_110, plain, ![X17]:(~m1_subset_1(X17,k4_ordinal1)|v7_ordinal1(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_111, plain, (k4_power(X1,X2)=k3_power(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_112, plain, (k4_power(np__2,np__4)=np__16), inference(split_conjunct,[status(thm)],[t62_power])).
cnf(c_0_113, plain, (m1_subset_1(np__4,k5_numbers)), inference(split_conjunct,[status(thm)],[spc4_numerals])).
fof(c_0_114, plain, ![X50, X51]:(~v1_xreal_0(X50)|~v1_xreal_0(X51)|v1_xreal_0(k6_xcmplx_0(X50,X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_xreal_0])])).
cnf(c_0_115, plain, (k1_newton(X3,k2_xcmplx_0(X1,X2))=k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_116, plain, (k1_newton(X1,np__1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_117, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_118, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_119, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_120, plain, (v3_ordinal1(k5_numbers)), inference(rw,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_121, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_95])).
fof(c_0_122, plain, ![X84]:(~v7_ordinal1(X84)|(~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X84),np__1),np__2)),k4_xcmplx_0(np__1),k4_pepin(X84))|v1_int_2(k4_pepin(X84)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t58_pepin])])).
fof(c_0_123, negated_conjecture, ~(v1_int_2(np__257)), inference(assume_negation,[status(cth)],[t61_pepin])).
fof(c_0_124, plain, ![X24, X25]:((~r1_xxreal_0(k6_numbers,k6_xcmplx_0(X24,X25))|k1_xreal_0(X24,X25)=k6_xcmplx_0(X24,X25)|~v1_xreal_0(X25)|~v1_xreal_0(X24))&(r1_xxreal_0(k6_numbers,k6_xcmplx_0(X24,X25))|k1_xreal_0(X24,X25)=k6_numbers|~v1_xreal_0(X25)|~v1_xreal_0(X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_xreal_0])])])])])).
cnf(c_0_125, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_126, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_127, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_128, plain, (m1_subset_1(np__257,k5_numbers)), inference(split_conjunct,[status(thm)],[spc257_numerals])).
cnf(c_0_129, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_130, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_131, plain, (k6_xcmplx_0(np__2,np__3)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_132, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_133, plain, (v1_xreal_0(np__1)|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_134, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_135, plain, (m1_subset_1(np__2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_136, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_137, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
fof(c_0_138, plain, ![X46, X47]:(~v1_int_1(X46)|~v1_int_1(X47)|v1_int_1(k6_xcmplx_0(X46,X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_int_1])])).
cnf(c_0_139, plain, (m1_subset_1(np__3,k5_numbers)), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_140, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_141, plain, (k2_xcmplx_0(X1,X2)=k2_nat_1(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_142, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_143, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_110])).
fof(c_0_144, plain, ![X54, X55]:(~v1_xreal_0(X54)|~v7_ordinal1(X55)|k3_power(X54,X55)=k1_newton(X54,X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ie1_power])])).
cnf(c_0_145, plain, (k4_power(X1,X2)=k3_power(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_146, plain, (k4_power(np__2,np__4)=np__16), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_147, plain, (m1_subset_1(np__4,k1_numbers)), inference(split_conjunct,[status(thm)],[spc4_numerals])).
cnf(c_0_148, plain, (m1_subset_1(np__4,k5_numbers)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_149, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_150, plain, (k6_xcmplx_0(np__3,np__1)=np__2), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r3_r1_r2])).
fof(c_0_151, plain, ![X81, X82]:(~v7_ordinal1(X81)|(~v7_ordinal1(X82)|(~r1_xxreal_0(X81,X82)|k3_nat_d(k3_power(np__2,X82),k3_power(np__2,X81))=k3_power(np__2,k7_nat_d(X82,X81))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_nat_2])])])).
cnf(c_0_152, plain, (k4_power(np__2,np__3)=np__8), inference(split_conjunct,[status(thm)],[t61_power])).
cnf(c_0_153, plain, (m1_subset_1(np__3,k1_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_154, plain, (k1_newton(X3,k2_xcmplx_0(X1,X2))=k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X3)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_155, plain, (k1_newton(X1,np__1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_156, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r1_r2])).
cnf(c_0_157, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_158, plain, (v3_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120])])])).
cnf(c_0_159, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_160, plain, (v1_int_2(k4_pepin(X1))|~v7_ordinal1(X1)|~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k4_xcmplx_0(np__1),k4_pepin(X1))), inference(split_conjunct,[status(thm)],[c_0_122])).
fof(c_0_161, negated_conjecture, ~v1_int_2(np__257), inference(fof_simplification,[status(thm)],[c_0_123])).
cnf(c_0_162, plain, (k1_xreal_0(X1,X2)=k6_xcmplx_0(X1,X2)|~r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_163, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_164, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_165, plain, (m1_subset_1(np__257,k5_numbers)), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_166, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_167, plain, (k4_xcmplx_0(k6_xcmplx_0(np__2,np__3))=np__1), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_168, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__3))|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_131])])).
cnf(c_0_169, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_135])])])).
cnf(c_0_170, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_171, plain, (k6_xcmplx_0(np__257,np__1)=np__256), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r257_r1_r256])).
cnf(c_0_172, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_173, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X1)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_138])).
fof(c_0_174, plain, ![X8]:(~v7_ordinal1(X8)|v1_int_1(X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_int_1])])).
cnf(c_0_175, plain, (v3_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_139]), c_0_120])])])).
cnf(c_0_176, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_177, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_178, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_179, plain, (k3_power(X1,X2)=k1_newton(X1,X2)|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_180, plain, (k3_power(np__2,np__4)=np__16|~m1_subset_1(np__4,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_135])])])).
cnf(c_0_181, plain, (m1_subset_1(np__4,k1_numbers)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_182, plain, (v3_ordinal1(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_148]), c_0_120])])])).
cnf(c_0_183, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_184, plain, (k6_xcmplx_0(np__3,np__1)=np__2), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_185, plain, (k3_nat_d(k3_power(np__2,X2),k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_151])).
fof(c_0_186, plain, ![X38, X39]:(~v7_ordinal1(X38)|~v7_ordinal1(X39)|v7_ordinal1(k2_xcmplx_0(X38,X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_nat_1])])).
cnf(c_0_187, plain, (k4_power(np__2,np__3)=np__8), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_188, plain, (m1_subset_1(np__3,k1_numbers)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_189, plain, (k3_xcmplx_0(k1_newton(X1,X2),X1)=k1_newton(X1,k2_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155])])).
cnf(c_0_190, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_191, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158])])).
cnf(c_0_192, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_159, c_0_94])).
cnf(c_0_193, plain, (v1_int_2(k4_pepin(X1))|~v7_ordinal1(X1)|~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k4_xcmplx_0(np__1),k4_pepin(X1))), inference(evalgc,[status(thm)],[c_0_160])).
cnf(c_0_194, plain, (k4_pepin(np__3)=np__257), inference(split_conjunct,[status(thm)],[t53_pepin])).
cnf(c_0_195, negated_conjecture, (~v1_int_2(np__257)), inference(split_conjunct,[status(thm)],[c_0_161])).
fof(c_0_196, plain, ![X66, X67]:(~v7_ordinal1(X66)|~v7_ordinal1(X67)|k7_nat_d(X66,X67)=k1_xreal_0(X66,X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_nat_d])])).
cnf(c_0_197, plain, (k1_xreal_0(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_198, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_199, plain, (r1_xxreal_0(np__0,np__256)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r0_r256])).
cnf(c_0_200, plain, (v3_ordinal1(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_165]), c_0_120])])])).
cnf(c_0_201, plain, (k6_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k6_xcmplx_0(np__2,np__3))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167])])).
cnf(c_0_202, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_168, c_0_169])])])).
cnf(c_0_203, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__3))|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_131])])).
fof(c_0_204, plain, ![X11]:(~v1_xreal_0(X11)|v1_xcmplx_0(X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_205, plain, (k6_xcmplx_0(np__257,np__1)=np__256), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_206, plain, (v1_int_1(np__1)|~v1_int_1(k6_xcmplx_0(np__2,np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_167])])).
cnf(c_0_207, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_208, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_209, plain, (v7_ordinal1(np__3)|~v1_finset_1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_175])])).
cnf(c_0_210, plain, (v3_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_176]), c_0_120])])])).
cnf(c_0_211, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_177])])).
cnf(c_0_212, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_178, c_0_94])).
fof(c_0_213, plain, ![X43, X44]:(~v1_xcmplx_0(X43)|~v7_ordinal1(X44)|v1_xcmplx_0(k1_newton(X43,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_newton])])).
cnf(c_0_214, plain, (k3_power(X1,X2)=k1_newton(X1,X2)|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_215, plain, (k3_power(np__2,np__4)=np__16), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_180, c_0_181])])).
cnf(c_0_216, plain, (v7_ordinal1(np__4)|~v1_finset_1(np__4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_182])])).
cnf(c_0_217, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__1)|~v1_xreal_0(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184])])).
cnf(c_0_218, plain, (k3_nat_d(k3_power(np__2,X2),k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(X2,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_219, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_220, plain, (k2_xcmplx_0(np__4,np__4)=np__8), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r4_r4_r8])).
cnf(c_0_221, plain, (k3_power(np__2,np__3)=np__8), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_187]), c_0_188]), c_0_135])])])).
cnf(c_0_222, plain, (k3_xcmplx_0(X1,X1)=k1_newton(X1,np__2)|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_155]), c_0_190])])).
cnf(c_0_223, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_119])])])).
cnf(c_0_224, plain, (v1_int_2(k4_pepin(X1))|~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_193, c_0_131])])).
cnf(c_0_225, plain, (k4_pepin(np__3)=np__257), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_226, negated_conjecture, (~v1_int_2(np__257)), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_227, plain, (k7_nat_d(X1,X2)=k1_xreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_196])).
cnf(c_0_228, plain, (k1_xreal_0(X1,X2)=k6_xcmplx_0(X1,X2)|~r1_xxreal_0(np__0,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_229, plain, (r1_xxreal_0(np__0,np__256)), inference(evalgc,[status(thm)],[c_0_199])).
cnf(c_0_230, plain, (v7_ordinal1(np__257)|~v1_finset_1(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_200])])).
cnf(c_0_231, plain, (k6_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_201, c_0_202])])])).
cnf(c_0_232, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_203, c_0_169])])])).
cnf(c_0_233, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_234, plain, (v1_xreal_0(np__256)|~v1_xreal_0(np__1)|~v1_xreal_0(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_205])])).
cnf(c_0_235, plain, (m1_subset_1(np__257,k1_numbers)), inference(split_conjunct,[status(thm)],[spc257_numerals])).
fof(c_0_236, plain, ![X78, X79, X80]:((~r2_int_1(X78,X79,X80)|r1_int_1(X80,k6_xcmplx_0(X78,X79))|~v1_int_1(X80)|~v1_int_1(X79)|~v1_int_1(X78))&(~r1_int_1(X80,k6_xcmplx_0(X78,X79))|r2_int_1(X78,X79,X80)|~v1_int_1(X80)|~v1_int_1(X79)|~v1_int_1(X78))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_int_2])])])])).
cnf(c_0_237, plain, (v1_int_1(np__1)|~v1_int_1(np__3)|~v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207])])).
cnf(c_0_238, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_239, plain, (v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_192]), c_0_139])])])).
cnf(c_0_240, plain, (v7_ordinal1(np__2)|~v1_finset_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_210])])).
cnf(c_0_241, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_211, c_0_212]), c_0_212])).
cnf(c_0_242, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_213])).
fof(c_0_243, plain, ![X40, X41]:(~v1_xreal_0(X40)|~v7_ordinal1(X41)|v1_xreal_0(k1_newton(X40,X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_newton])])).
cnf(c_0_244, plain, (k1_newton(np__2,np__4)=np__16|~v7_ordinal1(np__4)|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_215])])).
cnf(c_0_245, plain, (v7_ordinal1(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_192]), c_0_148])])])).
cnf(c_0_246, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_217, c_0_169])])])).
cnf(c_0_247, plain, (k3_nat_d(k1_newton(np__2,X1),k3_power(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_214])])).
cnf(c_0_248, plain, (k6_xcmplx_0(np__8,np__1)=np__7), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r8_r1_r7])).
cnf(c_0_249, plain, (r1_xxreal_0(np__0,np__7)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r0_r7])).
cnf(c_0_250, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_251, plain, (k2_xcmplx_0(np__4,np__4)=np__8), inference(evalgc,[status(thm)],[c_0_220])).
cnf(c_0_252, plain, (k1_newton(np__2,np__3)=np__8|~v7_ordinal1(np__3)|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_221])])).
cnf(c_0_253, plain, (k3_xcmplx_0(X1,X1)=k1_newton(X1,np__2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_222, c_0_223])])])).
cnf(c_0_254, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225]), c_0_226])])).
cnf(c_0_255, plain, (k7_nat_d(X1,X2)=k1_xreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_256, plain, (k1_xreal_0(np__257,np__1)=np__256|~v1_xreal_0(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_205]), c_0_229]), c_0_169])])])).
cnf(c_0_257, plain, (v7_ordinal1(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_192]), c_0_165])])])).
cnf(c_0_258, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_231]), c_0_232])])])).
cnf(c_0_259, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_260, plain, (k2_xcmplx_0(np__256,np__1)=np__257), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r256_r1_r257])).
cnf(c_0_261, plain, (v1_xreal_0(np__256)|~v1_xreal_0(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_234, c_0_169])])])).
cnf(c_0_262, plain, (m1_subset_1(np__257,k1_numbers)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_263, plain, (r2_int_1(X2,X3,X1)|~r1_int_1(X1,k6_xcmplx_0(X2,X3))|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_264, plain, (v1_int_1(np__1)|~v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_239])])])).
cnf(c_0_265, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_192]), c_0_176])])])).
cnf(c_0_266, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_241])])).
cnf(c_0_267, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_268, plain, (k3_xcmplx_0(np__16,np__16)=np__256), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r16_r16_r256])).
cnf(c_0_269, plain, (v1_xreal_0(k1_newton(X1,X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_243])).
cnf(c_0_270, plain, (k1_newton(np__2,np__4)=np__16|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_244, c_0_245])])])).
cnf(c_0_271, plain, (v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_134]), c_0_188])])])).
cnf(c_0_272, plain, (v1_int_1(np__256)|~v1_int_1(np__1)|~v1_int_1(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_205])])).
cnf(c_0_273, plain, (k3_nat_d(k1_newton(np__2,X1),k1_newton(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_214])])).
cnf(c_0_274, plain, (k6_xcmplx_0(np__8,np__1)=np__7), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_275, plain, (r1_xxreal_0(np__0,np__7)), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_276, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251])])).
cnf(c_0_277, plain, (k1_newton(np__2,np__3)=np__8|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_252, c_0_239])])])).
cnf(c_0_278, plain, (k1_newton(k1_newton(X1,X2),np__2)=k1_newton(X1,k2_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(k1_newton(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_253])])).
cnf(c_0_279, plain, (r1_nat_d(k4_pepin(np__3),k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,k6_numbers),np__128)),np__1))), inference(split_conjunct,[status(thm)],[l109_pepin])).
cnf(c_0_280, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)),k6_xcmplx_0(np__2,np__3),np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_254, c_0_239])])])).
cnf(c_0_281, plain, (k7_nat_d(np__257,np__1)=np__256|~v1_xreal_0(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256]), c_0_223]), c_0_257])])])).
cnf(c_0_282, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_258, c_0_259])).
cnf(c_0_283, plain, (k2_xcmplx_0(np__256,np__1)=np__257), inference(evalgc,[status(thm)],[c_0_260])).
cnf(c_0_284, plain, (v1_xreal_0(np__256)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_134]), c_0_262])])])).
cnf(c_0_285, plain, (r2_int_1(X2,X3,X1)|~v1_int_1(X3)|~v1_int_1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k6_xcmplx_0(X2,X3))), inference(evalgc,[status(thm)],[c_0_263])).
cnf(c_0_286, plain, (v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_131])])).
cnf(c_0_287, plain, (v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_238]), c_0_265])])])).
cnf(c_0_288, plain, (k6_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_167]), c_0_202]), c_0_119])])])).
cnf(c_0_289, plain, (m1_subset_1(np__256,k5_numbers)), inference(split_conjunct,[status(thm)],[spc256_numerals])).
cnf(c_0_290, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_253]), c_0_265])])])).
cnf(c_0_291, plain, (k3_xcmplx_0(np__16,np__16)=np__256), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_292, plain, (v1_xreal_0(k1_newton(X1,X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_293, plain, (k1_newton(np__2,np__4)=np__16), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_270, c_0_271])])).
cnf(c_0_294, plain, (v1_int_1(np__256)|~v1_int_1(np__1)|~v7_ordinal1(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_238])])).
cnf(c_0_295, plain, (k3_power(np__2,k7_nat_d(X1,np__1))=k3_nat_d(k1_newton(np__2,X1),np__2)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__1)|~v1_xreal_0(np__2)|~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_155])])).
cnf(c_0_296, plain, (k1_xreal_0(np__8,np__1)=np__7|~v1_xreal_0(np__8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_274]), c_0_275]), c_0_169])])])).
cnf(c_0_297, plain, (v7_ordinal1(np__8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_212]), c_0_148])])])).
cnf(c_0_298, plain, (k1_newton(np__2,np__3)=np__8), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_277, c_0_271])])).
cnf(c_0_299, plain, (k1_newton(k1_newton(X1,X2),np__2)=k1_newton(X1,k2_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_278, c_0_267])).
fof(c_0_300, plain, ![X22, X23]:(~v1_xcmplx_0(X22)|~v1_xcmplx_0(X23)|k3_xcmplx_0(X22,X23)=k3_xcmplx_0(X23,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_xcmplx_0])])).
fof(c_0_301, plain, ![X71, X72]:((~r1_nat_d(X71,X72)|r1_int_1(X71,X72)|(~v7_ordinal1(X71)|~v7_ordinal1(X72)))&(~r1_int_1(X71,X72)|r1_nat_d(X71,X72)|(~v7_ordinal1(X71)|~v7_ordinal1(X72)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_nat_d])])])).
cnf(c_0_302, plain, (r1_nat_d(k4_pepin(np__3),k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,k6_numbers),np__128)),np__1))), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_303, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xreal_0(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281])])).
cnf(c_0_304, plain, (v1_xreal_0(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283]), c_0_284])])])).
cnf(c_0_305, plain, (r2_int_1(X1,k6_xcmplx_0(np__2,np__3),X2)|~r1_int_1(X2,k2_xcmplx_0(X1,np__1))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_231])])).
cnf(c_0_306, plain, (v1_int_1(k6_xcmplx_0(np__2,np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_286, c_0_287])])])).
cnf(c_0_307, plain, (v1_int_1(k2_xcmplx_0(np__1,X1))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_288])])).
cnf(c_0_308, plain, (m1_subset_1(np__256,k5_numbers)), inference(evalgc,[status(thm)],[c_0_289])).
cnf(c_0_309, plain, (v1_xcmplx_0(np__256)|~v1_xcmplx_0(np__16)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_291])])).
cnf(c_0_310, plain, (v1_xreal_0(np__16)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_245]), c_0_271])])])).
cnf(c_0_311, plain, (v1_int_1(np__256)|~v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_294, c_0_257])])])).
cnf(c_0_312, plain, (k3_power(np__2,k7_nat_d(X1,np__1))=k3_nat_d(k1_newton(np__2,X1),np__2)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v1_xreal_0(np__2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_295, c_0_223])]), c_0_259])).
cnf(c_0_313, plain, (k7_nat_d(np__8,np__1)=np__7|~v1_xreal_0(np__8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_296]), c_0_223]), c_0_297])])])).
cnf(c_0_314, plain, (v1_xreal_0(np__8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_298]), c_0_239]), c_0_271])])])).
cnf(c_0_315, plain, (r1_xxreal_0(np__1,np__8)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r8])).
cnf(c_0_316, plain, (k1_newton(np__16,np__2)=k1_newton(np__2,np__8)|~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_293]), c_0_251]), c_0_245])])])).
cnf(c_0_317, plain, (m1_subset_1(np__7,k5_numbers)), inference(split_conjunct,[status(thm)],[spc7_numerals])).
cnf(c_0_318, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_300])).
cnf(c_0_319, plain, (r1_int_1(X1,X2)|~r1_nat_d(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_301])).
cnf(c_0_320, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,k6_numbers),np__128)),np__1))), inference(rw,[status(thm)],[c_0_302, c_0_225])).
fof(c_0_321, plain, ![X16]:(~v1_xboole_0(X16)|v7_ordinal1(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_ordinal1])])).
cnf(c_0_322, plain, (m1_subset_1(np__128,k5_numbers)), inference(split_conjunct,[status(thm)],[spc128_numerals])).
cnf(c_0_323, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_303, c_0_304])])])).
cnf(c_0_324, plain, (r2_int_1(X1,k6_xcmplx_0(np__2,np__3),X2)|~r1_int_1(X2,k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~v1_int_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_305, c_0_306])])])).
cnf(c_0_325, plain, (v1_int_1(k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_307, c_0_306])])])).
cnf(c_0_326, plain, (k2_xcmplx_0(np__1,np__256)=np__257), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_241]), c_0_308]), c_0_119])])])).
cnf(c_0_327, plain, (v1_xcmplx_0(np__256)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_259]), c_0_310])])])).
cnf(c_0_328, plain, (v1_int_1(np__256)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_311, c_0_287])])])).
cnf(c_0_329, plain, (k3_power(np__2,k7_nat_d(X1,np__1))=k3_nat_d(k1_newton(np__2,X1),np__2)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_312, c_0_271])])])).
cnf(c_0_330, plain, (k7_nat_d(np__8,np__1)=np__7), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_313, c_0_314])])])).
cnf(c_0_331, plain, (r1_xxreal_0(np__1,np__8)), inference(evalgc,[status(thm)],[c_0_315])).
cnf(c_0_332, plain, (k1_newton(np__2,np__8)=np__256|~v1_xcmplx_0(np__16)|~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_316]), c_0_291])])).
cnf(c_0_333, plain, (v1_xcmplx_0(np__16)|~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_293]), c_0_245])])])).
cnf(c_0_334, plain, (m1_subset_1(np__7,k5_numbers)), inference(evalgc,[status(thm)],[c_0_317])).
cnf(c_0_335, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_318])).
cnf(c_0_336, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X1,X2)), inference(evalgc,[status(thm)],[c_0_319])).
cnf(c_0_337, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_320, c_0_198])])).
fof(c_0_338, plain, ![X62, X63]:(~m1_subset_1(X62,k5_numbers)|~v7_ordinal1(X63)|k4_nat_1(X62,X63)=k3_xcmplx_0(X62,X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_nat_1])])).
cnf(c_0_339, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_321])).
cnf(c_0_340, plain, (m1_subset_1(np__128,k5_numbers)), inference(evalgc,[status(thm)],[c_0_322])).
cnf(c_0_341, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_nat_d(np__256,np__2)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324])])).
cnf(c_0_342, plain, (v1_int_1(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_327]), c_0_328]), c_0_308])])])).
cnf(c_0_343, plain, (k3_nat_d(k1_newton(np__2,np__8),np__2)=k3_power(np__2,np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_330]), c_0_331]), c_0_297])])])).
cnf(c_0_344, plain, (k1_newton(np__2,np__8)=np__256|~v1_xcmplx_0(np__2)), inference(csr,[status(thm)],[c_0_332, c_0_333])).
cnf(c_0_345, plain, (v3_ordinal1(np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_334]), c_0_120])])])).
cnf(c_0_346, plain, (k3_xcmplx_0(k1_newton(X1,X3),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(k1_newton(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_335])])).
cnf(c_0_347, plain, (k2_xcmplx_0(np__4,np__3)=np__7), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r4_r3_r7])).
cnf(c_0_348, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_337])])).
cnf(c_0_349, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_338])).
cnf(c_0_350, plain, (k3_xcmplx_0(np__32,np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r32_r0_r0])).
cnf(c_0_351, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_339])).
cnf(c_0_352, plain, (m1_subset_1(np__32,k5_numbers)), inference(split_conjunct,[status(thm)],[spc32_numerals])).
cnf(c_0_353, plain, (v3_ordinal1(np__128)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_340]), c_0_120])])])).
cnf(c_0_354, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_nat_d(np__256,np__2)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_341, c_0_342])])).
cnf(c_0_355, plain, (k3_nat_d(np__256,np__2)=k3_power(np__2,np__7)|~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_344])])).
cnf(c_0_356, plain, (v7_ordinal1(np__7)|~v1_finset_1(np__7)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_345])])).
cnf(c_0_357, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,X3))=k1_newton(X1,k2_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_346, c_0_267]), c_0_267])).
cnf(c_0_358, plain, (k3_xcmplx_0(np__16,np__8)=np__128), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r16_r8_r128])).
cnf(c_0_359, plain, (k2_xcmplx_0(np__4,np__3)=np__7), inference(evalgc,[status(thm)],[c_0_347])).
cnf(c_0_360, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_348, c_0_257])])).
cnf(c_0_361, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_349])).
cnf(c_0_362, plain, (k3_xcmplx_0(np__32,np__0)=np__0), inference(evalgc,[status(thm)],[c_0_350])).
cnf(c_0_363, plain, (v7_ordinal1(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_164])])).
cnf(c_0_364, plain, (m1_subset_1(np__32,k5_numbers)), inference(evalgc,[status(thm)],[c_0_352])).
cnf(c_0_365, plain, (k2_xcmplx_0(np__0,np__128)=np__128), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r0_r128_r128])).
cnf(c_0_366, plain, (v7_ordinal1(np__128)|~v1_finset_1(np__128)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_353])])).
cnf(c_0_367, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_368, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_power(np__2,np__7)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_power(np__2,np__7)))|~v1_int_1(k13_newton(np__3,k3_power(np__2,np__7)))|~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_355])])).
cnf(c_0_369, plain, (v7_ordinal1(np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_192]), c_0_334])])])).
cnf(c_0_370, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__8)=k1_newton(np__2,k2_xcmplx_0(np__3,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_298]), c_0_239])])])).
cnf(c_0_371, plain, (k3_xcmplx_0(np__16,np__8)=np__128), inference(evalgc,[status(thm)],[c_0_358])).
cnf(c_0_372, plain, (k2_xcmplx_0(np__3,np__4)=np__7), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_241]), c_0_148]), c_0_139])])])).
cnf(c_0_373, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(np__0,np__128)),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,k2_nat_1(np__0,np__128)),np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_361]), c_0_362]), c_0_362]), c_0_363]), c_0_364])])])).
cnf(c_0_374, plain, (k2_xcmplx_0(np__0,np__128)=np__128), inference(evalgc,[status(thm)],[c_0_365])).
cnf(c_0_375, plain, (v7_ordinal1(np__128)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_192]), c_0_340])])])).
cnf(c_0_376, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_367])).
cnf(c_0_377, plain, (k2_xcmplx_0(np__1,X1)=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_288])])).
cnf(c_0_378, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k1_newton(np__2,np__7)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k1_newton(np__2,np__7)))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,k1_newton(np__2,np__7)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_214]), c_0_369]), c_0_271])])])).
cnf(c_0_379, plain, (k1_newton(np__2,np__7)=np__128|~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_370, c_0_293]), c_0_371]), c_0_372]), c_0_245])])])).
cnf(c_0_380, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,np__128),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,np__128),np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_141]), c_0_374]), c_0_374]), c_0_375]), c_0_376])])])).
cnf(c_0_381, plain, (v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_377]), c_0_223])])])).
cnf(c_0_382, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,np__128))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379])])).
cnf(c_0_383, plain, (r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~v7_ordinal1(k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_141]), c_0_223])])])).
cnf(c_0_384, plain, (v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[c_0_381, c_0_212])).
cnf(c_0_385, plain, (~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,np__128))|~v7_ordinal1(k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_383])])).
cnf(c_0_386, plain, (v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_241]), c_0_119])])])).
fof(c_0_387, plain, ~v1_finset_1(k4_ordinal1), inference(fof_simplification,[status(thm)],[fc7_card_1])).
cnf(c_0_388, plain, (~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,np__128))|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)), inference(csr,[status(thm)],[c_0_385, c_0_386])).
fof(c_0_389, plain, ![X68, X69, X70]:((~m2_subset_1(X70,X68,X69)|m1_subset_1(X70,X69)|(v1_xboole_0(X68)|v1_xboole_0(X69)|~m1_subset_1(X69,k1_zfmisc_1(X68))))&(~m1_subset_1(X70,X69)|m2_subset_1(X70,X68,X69)|(v1_xboole_0(X68)|v1_xboole_0(X69)|~m1_subset_1(X69,k1_zfmisc_1(X68))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_390, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_391, plain, ![X6, X7]:(~v1_finset_1(X6)|(~m1_subset_1(X7,k1_zfmisc_1(X6))|v1_finset_1(X7))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_finset_1])])])).
cnf(c_0_392, plain, (~v1_finset_1(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_387])).
cnf(c_0_393, plain, (~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)|~v7_ordinal1(k13_newton(np__3,np__128))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_238])])).
fof(c_0_394, plain, ![X9]:(~v7_ordinal1(X9)|v1_xreal_0(X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xreal_0])])).
cnf(c_0_395, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_389])).
cnf(c_0_396, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_397, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_390])).
fof(c_0_398, plain, ![X26, X27]:(~m1_subset_1(X26,k5_numbers)|~m1_subset_1(X27,k5_numbers)|m2_subset_1(k13_newton(X26,X27),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k13_newton])])).
cnf(c_0_399, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_391])).
cnf(c_0_400, plain, (~v1_finset_1(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_392])).
fof(c_0_401, plain, ![X4]:(~v1_xboole_0(X4)|v1_finset_1(X4)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_finset_1])])).
cnf(c_0_402, plain, (~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)), inference(csr,[status(thm)],[c_0_393, c_0_212])).
cnf(c_0_403, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_394])).
cnf(c_0_404, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_395])).
cnf(c_0_405, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_396])).
cnf(c_0_406, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_397, c_0_94])).
cnf(c_0_407, plain, (m2_subset_1(k13_newton(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_398])).
cnf(c_0_408, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_399])).
cnf(c_0_409, plain, (~v1_finset_1(k5_numbers)), inference(rw,[status(thm)],[c_0_400, c_0_94])).
cnf(c_0_410, plain, (v1_finset_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_401])).
cnf(c_0_411, plain, (~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)|~v1_xreal_0(k13_newton(np__3,np__128))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_259])])).
cnf(c_0_412, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_403])).
cnf(c_0_413, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_405]), c_0_406])])).
cnf(c_0_414, plain, (m2_subset_1(k13_newton(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_407])).
cnf(c_0_415, plain, (~v1_finset_1(k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_405]), c_0_409])])).
cnf(c_0_416, plain, (v1_finset_1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_410])).
cnf(c_0_417, plain, (~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)|~v7_ordinal1(k13_newton(np__3,np__128))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_411, c_0_412])])).
cnf(c_0_418, plain, (m1_subset_1(k13_newton(X1,X2),k5_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_414])])).
cnf(c_0_419, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_416])])).
cnf(c_0_420, plain, (~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)), inference(csr,[status(thm)],[c_0_417, c_0_212])).
cnf(c_0_421, plain, (m1_subset_1(k13_newton(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(sr,[status(thm)],[c_0_418, c_0_419])).
cnf(c_0_422, plain, (~v1_xcmplx_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_421]), c_0_340]), c_0_139])])])).
cnf(c_0_423, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_259]), c_0_271])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 424
# Proof object clause steps            : 308
# Proof object formula steps           : 116
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 81
# Proof object initial formulas used   : 75
# Proof object generating inferences   : 102
# Proof object simplifying inferences  : 197
# Training examples: 203 positive, 972 negative
# Training: Positive examples begin
cnf(c_0_94, plain, (k5_numbers=k4_ordinal1)).# trainpos
cnf(c_0_226, negated_conjecture, (~v1_int_2(np__257))).# trainpos
cnf(c_0_193, plain, (v1_int_2(k4_pepin(X1))|~v7_ordinal1(X1)|~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k4_xcmplx_0(np__1),k4_pepin(X1)))).# trainpos
cnf(c_0_164, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_145, plain, (k4_power(X1,X2)=k3_power(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_225, plain, (k4_pepin(np__3)=np__257)).# trainpos
cnf(c_0_404, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_176, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_397, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_165, plain, (m1_subset_1(np__257,k5_numbers))).# trainpos
cnf(c_0_285, plain, (r2_int_1(X2,X3,X1)|~v1_int_1(X3)|~v1_int_1(X2)|~v1_int_1(X1)|~r1_int_1(X1,k6_xcmplx_0(X2,X3)))).# trainpos
cnf(c_0_139, plain, (m1_subset_1(np__3,k5_numbers))).# trainpos
cnf(c_0_119, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_205, plain, (k6_xcmplx_0(np__257,np__1)=np__256)).# trainpos
cnf(c_0_405, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_400, plain, (~v1_finset_1(k4_ordinal1))).# trainpos
cnf(c_0_184, plain, (k6_xcmplx_0(np__3,np__1)=np__2)).# trainpos
cnf(c_0_336, plain, (r1_int_1(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_nat_d(X1,X2))).# trainpos
cnf(c_0_93, plain, (v3_ordinal1(k4_ordinal1))).# trainpos
cnf(c_0_126, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_190, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).# trainpos
cnf(c_0_238, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_125, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_416, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_376, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_283, plain, (k2_xcmplx_0(np__256,np__1)=np__257)).# trainpos
cnf(c_0_130, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_308, plain, (m1_subset_1(np__256,k5_numbers))).# trainpos
cnf(c_0_302, plain, (r1_nat_d(k4_pepin(np__3),k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,k6_numbers),np__128)),np__1)))).# trainpos
cnf(c_0_159, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_362, plain, (k3_xcmplx_0(np__32,np__0)=np__0)).# trainpos
cnf(c_0_218, plain, (k3_nat_d(k3_power(np__2,X2),k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(X2,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_157, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1))).# trainpos
cnf(c_0_155, plain, (k1_newton(X1,np__1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_131, plain, (k6_xcmplx_0(np__2,np__3)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_361, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_414, plain, (m2_subset_1(k13_newton(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_141, plain, (k2_xcmplx_0(X1,X2)=k2_nat_1(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_178, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_351, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_255, plain, (k7_nat_d(X1,X2)=k1_xreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_408, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_103, plain, (k7_xcmplx_0(np__2,np__2)=np__1)).# trainpos
cnf(c_0_172, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_188, plain, (m1_subset_1(np__3,k1_numbers))).# trainpos
cnf(c_0_118, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_207, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_135, plain, (m1_subset_1(np__2,k1_numbers))).# trainpos
cnf(c_0_229, plain, (r1_xxreal_0(np__0,np__256))).# trainpos
cnf(c_0_166, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_291, plain, (k3_xcmplx_0(np__16,np__16)=np__256)).# trainpos
cnf(c_0_214, plain, (k3_power(X1,X2)=k1_newton(X1,X2)|~v1_xreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_154, plain, (k1_newton(X3,k2_xcmplx_0(X1,X2))=k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X3))).# trainpos
cnf(c_0_364, plain, (m1_subset_1(np__32,k5_numbers))).# trainpos
cnf(c_0_148, plain, (m1_subset_1(np__4,k5_numbers))).# trainpos
cnf(c_0_187, plain, (k4_power(np__2,np__3)=np__8)).# trainpos
cnf(c_0_259, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_142, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_134, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_412, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_262, plain, (m1_subset_1(np__257,k1_numbers))).# trainpos
cnf(c_0_146, plain, (k4_power(np__2,np__4)=np__16)).# trainpos
cnf(c_0_340, plain, (m1_subset_1(np__128,k5_numbers))).# trainpos
cnf(c_0_374, plain, (k2_xcmplx_0(np__0,np__128)=np__128)).# trainpos
cnf(c_0_275, plain, (r1_xxreal_0(np__0,np__7))).# trainpos
cnf(c_0_331, plain, (r1_xxreal_0(np__1,np__8))).# trainpos
cnf(c_0_250, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_359, plain, (k2_xcmplx_0(np__4,np__3)=np__7)).# trainpos
cnf(c_0_181, plain, (m1_subset_1(np__4,k1_numbers))).# trainpos
cnf(c_0_251, plain, (k2_xcmplx_0(np__4,np__4)=np__8)).# trainpos
cnf(c_0_335, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_371, plain, (k3_xcmplx_0(np__16,np__8)=np__128)).# trainpos
cnf(c_0_334, plain, (m1_subset_1(np__7,k5_numbers))).# trainpos
cnf(c_0_274, plain, (k6_xcmplx_0(np__8,np__1)=np__7)).# trainpos
cnf(c_0_197, plain, (k1_xreal_0(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2)))).# trainpos
cnf(c_0_267, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_132, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_170, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_292, plain, (v1_xreal_0(k1_newton(X1,X2))|~v1_xreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_183, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_102, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_198, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_224, plain, (v1_int_2(k4_pepin(X1))|~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_363, plain, (v7_ordinal1(np__0))).# trainpos
cnf(c_0_221, plain, (k3_power(np__2,np__3)=np__8)).# trainpos
cnf(c_0_247, plain, (k3_nat_d(k1_newton(np__2,X1),k3_power(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_413, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_133, plain, (v1_xreal_0(np__1)|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_177, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_167, plain, (k4_xcmplx_0(k6_xcmplx_0(np__2,np__3))=np__1)).# trainpos
cnf(c_0_200, plain, (v3_ordinal1(np__257))).# trainpos
cnf(c_0_217, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__1)|~v1_xreal_0(np__3))).# trainpos
cnf(c_0_158, plain, (v3_ordinal1(np__1))).# trainpos
cnf(c_0_175, plain, (v3_ordinal1(np__3))).# trainpos
cnf(c_0_272, plain, (v1_int_1(np__256)|~v1_int_1(np__1)|~v1_int_1(np__257))).# trainpos
cnf(c_0_210, plain, (v3_ordinal1(np__2))).# trainpos
cnf(c_0_286, plain, (v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(np__1))).# trainpos
cnf(c_0_276, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__4))).# trainpos
cnf(c_0_203, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__3))|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_234, plain, (v1_xreal_0(np__256)|~v1_xreal_0(np__1)|~v1_xreal_0(np__257))).# trainpos
cnf(c_0_415, plain, (~v1_finset_1(k1_numbers))).# trainpos
cnf(c_0_180, plain, (k3_power(np__2,np__4)=np__16|~m1_subset_1(np__4,k1_numbers))).# trainpos
cnf(c_0_353, plain, (v3_ordinal1(np__128))).# trainpos
cnf(c_0_189, plain, (k3_xcmplx_0(k1_newton(X1,X2),X1)=k1_newton(X1,k2_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_346, plain, (k3_xcmplx_0(k1_newton(X1,X3),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(k1_newton(X1,X3)))).# trainpos
cnf(c_0_168, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__3))|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_345, plain, (v3_ordinal1(np__7))).# trainpos
cnf(c_0_182, plain, (v3_ordinal1(np__4))).# trainpos
cnf(c_0_169, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_419, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_206, plain, (v1_int_1(np__1)|~v1_int_1(k6_xcmplx_0(np__2,np__3)))).# trainpos
cnf(c_0_297, plain, (v7_ordinal1(np__8))).# trainpos
cnf(c_0_273, plain, (k3_nat_d(k1_newton(np__2,X1),k1_newton(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_191, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1))).# trainpos
cnf(c_0_240, plain, (v7_ordinal1(np__2)|~v1_finset_1(np__2))).# trainpos
cnf(c_0_252, plain, (k1_newton(np__2,np__3)=np__8|~v7_ordinal1(np__3)|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_244, plain, (k1_newton(np__2,np__4)=np__16|~v7_ordinal1(np__4)|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_294, plain, (v1_int_1(np__256)|~v1_int_1(np__1)|~v7_ordinal1(np__257))).# trainpos
cnf(c_0_254, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_230, plain, (v7_ordinal1(np__257)|~v1_finset_1(np__257))).# trainpos
cnf(c_0_337, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1)))).# trainpos
cnf(c_0_211, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_228, plain, (k1_xreal_0(X1,X2)=k6_xcmplx_0(X1,X2)|~r1_xxreal_0(np__0,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_222, plain, (k3_xcmplx_0(X1,X1)=k1_newton(X1,np__2)|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1))).# trainpos
cnf(c_0_201, plain, (k6_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k6_xcmplx_0(np__2,np__3))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_356, plain, (v7_ordinal1(np__7)|~v1_finset_1(np__7))).# trainpos
cnf(c_0_209, plain, (v7_ordinal1(np__3)|~v1_finset_1(np__3))).# trainpos
cnf(c_0_366, plain, (v7_ordinal1(np__128)|~v1_finset_1(np__128))).# trainpos
cnf(c_0_418, plain, (m1_subset_1(k13_newton(X1,X2),k5_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_216, plain, (v7_ordinal1(np__4)|~v1_finset_1(np__4))).# trainpos
cnf(c_0_265, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_223, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_239, plain, (v7_ordinal1(np__3))).# trainpos
cnf(c_0_257, plain, (v7_ordinal1(np__257))).# trainpos
cnf(c_0_237, plain, (v1_int_1(np__1)|~v1_int_1(np__3)|~v1_int_1(np__2))).# trainpos
cnf(c_0_246, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__3))).# trainpos
cnf(c_0_245, plain, (v7_ordinal1(np__4))).# trainpos
cnf(c_0_256, plain, (k1_xreal_0(np__257,np__1)=np__256|~v1_xreal_0(np__257))).# trainpos
cnf(c_0_326, plain, (k2_xcmplx_0(np__1,np__256)=np__257)).# trainpos
cnf(c_0_375, plain, (v7_ordinal1(np__128))).# trainpos
cnf(c_0_372, plain, (k2_xcmplx_0(np__3,np__4)=np__7)).# trainpos
cnf(c_0_348, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(np__257))).# trainpos
cnf(c_0_202, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__3)))).# trainpos
cnf(c_0_369, plain, (v7_ordinal1(np__7))).# trainpos
cnf(c_0_261, plain, (v1_xreal_0(np__256)|~v1_xreal_0(np__257))).# trainpos
cnf(c_0_232, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__3)))).# trainpos
cnf(c_0_296, plain, (k1_xreal_0(np__8,np__1)=np__7|~v1_xreal_0(np__8))).# trainpos
cnf(c_0_266, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers))).# trainpos
cnf(c_0_295, plain, (k3_power(np__2,k7_nat_d(X1,np__1))=k3_nat_d(k1_newton(np__2,X1),np__2)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__1)|~v1_xreal_0(np__2)|~v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_264, plain, (v1_int_1(np__1)|~v1_int_1(np__2))).# trainpos
cnf(c_0_271, plain, (v1_xreal_0(np__2))).# trainpos
cnf(c_0_281, plain, (k7_nat_d(np__257,np__1)=np__256|~v1_xreal_0(np__257))).# trainpos
cnf(c_0_311, plain, (v1_int_1(np__256)|~v1_int_1(np__1))).# trainpos
cnf(c_0_284, plain, (v1_xreal_0(np__256))).# trainpos
cnf(c_0_277, plain, (k1_newton(np__2,np__3)=np__8|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_253, plain, (k3_xcmplx_0(X1,X1)=k1_newton(X1,np__2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_280, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)),k6_xcmplx_0(np__2,np__3),np__257))).# trainpos
cnf(c_0_270, plain, (k1_newton(np__2,np__4)=np__16|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_231, plain, (k6_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_373, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(np__0,np__128)),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,k2_nat_1(np__0,np__128)),np__1)))).# trainpos
cnf(c_0_313, plain, (k7_nat_d(np__8,np__1)=np__7|~v1_xreal_0(np__8))).# trainpos
cnf(c_0_288, plain, (k6_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_287, plain, (v1_int_1(np__1))).# trainpos
cnf(c_0_314, plain, (v1_xreal_0(np__8))).# trainpos
cnf(c_0_303, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xreal_0(np__257))).# trainpos
cnf(c_0_329, plain, (k3_power(np__2,k7_nat_d(X1,np__1))=k3_nat_d(k1_newton(np__2,X1),np__2)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_305, plain, (r2_int_1(X1,k6_xcmplx_0(np__2,np__3),X2)|~r1_int_1(X2,k2_xcmplx_0(X1,np__1))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_258, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_310, plain, (v1_xreal_0(np__16))).# trainpos
cnf(c_0_278, plain, (k1_newton(k1_newton(X1,X2),np__2)=k1_newton(X1,k2_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(k1_newton(X1,X2)))).# trainpos
cnf(c_0_377, plain, (k2_xcmplx_0(np__1,X1)=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_333, plain, (v1_xcmplx_0(np__16)|~v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_307, plain, (v1_int_1(k2_xcmplx_0(np__1,X1))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_380, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,np__128),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,np__128),np__1)))).# trainpos
cnf(c_0_370, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__8)=k1_newton(np__2,k2_xcmplx_0(np__3,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_290, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_304, plain, (v1_xreal_0(np__257))).# trainpos
cnf(c_0_306, plain, (v1_int_1(k6_xcmplx_0(np__2,np__3)))).# trainpos
cnf(c_0_328, plain, (v1_int_1(np__256))).# trainpos
cnf(c_0_309, plain, (v1_xcmplx_0(np__256)|~v1_xcmplx_0(np__16))).# trainpos
cnf(c_0_379, plain, (k1_newton(np__2,np__7)=np__128|~v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_316, plain, (k1_newton(np__16,np__2)=k1_newton(np__2,np__8)|~v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_330, plain, (k7_nat_d(np__8,np__1)=np__7)).# trainpos
cnf(c_0_383, plain, (r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~v7_ordinal1(k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers))).# trainpos
cnf(c_0_381, plain, (v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_327, plain, (v1_xcmplx_0(np__256))).# trainpos
cnf(c_0_323, plain, (~r2_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257))).# trainpos
cnf(c_0_332, plain, (k1_newton(np__2,np__8)=np__256|~v1_xcmplx_0(np__16)|~v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_324, plain, (r2_int_1(X1,k6_xcmplx_0(np__2,np__3),X2)|~r1_int_1(X2,k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~v1_int_1(X2))).# trainpos
cnf(c_0_325, plain, (v1_int_1(k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_342, plain, (v1_int_1(np__257))).# trainpos
cnf(c_0_341, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_nat_d(np__256,np__2)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(np__257))).# trainpos
cnf(c_0_386, plain, (v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_343, plain, (k3_nat_d(k1_newton(np__2,np__8),np__2)=k3_power(np__2,np__7))).# trainpos
cnf(c_0_355, plain, (k3_nat_d(np__256,np__2)=k3_power(np__2,np__7)|~v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_368, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_power(np__2,np__7)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_power(np__2,np__7)))|~v1_int_1(k13_newton(np__3,k3_power(np__2,np__7)))|~v1_xcmplx_0(np__2))).# trainpos
cnf(c_0_378, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k1_newton(np__2,np__7)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k1_newton(np__2,np__7)))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,k1_newton(np__2,np__7))))).# trainpos
cnf(c_0_382, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,np__128)))).# trainpos
cnf(c_0_385, plain, (~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,np__128))|~v7_ordinal1(k2_xcmplx_0(k13_newton(np__3,np__128),np__1))|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers))).# trainpos
cnf(c_0_393, plain, (~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)|~v7_ordinal1(k13_newton(np__3,np__128)))).# trainpos
cnf(c_0_411, plain, (~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)|~v1_xreal_0(k13_newton(np__3,np__128)))).# trainpos
cnf(c_0_417, plain, (~v1_xcmplx_0(np__2)|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers)|~v7_ordinal1(k13_newton(np__3,np__128)))).# trainpos
cnf(c_0_422, plain, (~v1_xcmplx_0(np__2))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_42, plain, (k1_newton(X1,X2)=k13_newton(X1,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_128, plain, (r1_int_1(X3,k6_xcmplx_0(X1,X2))|~v1_int_1(X3)|~v1_int_1(X2)|~v1_int_1(X1)|~r2_int_1(X1,X2,X3))).#trainneg
cnf(i_0_138, plain, (k4_pepin(X1)=k1_nat_1(k2_newton(np__2,k2_newton(np__2,X1)),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_139, plain, (k2_newton(X1,X2)=k1_newton(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_43, plain, (k2_xcmplx_0(X1,X2)=k1_nat_1(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_25, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_58, plain, (k2_xcmplx_0(np__0,np__256)=np__256)).#trainneg
cnf(i_0_90, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_22, plain, (m1_subset_1(k4_pepin(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_59, plain, (k2_xcmplx_0(np__128,np__128)=np__256)).#trainneg
cnf(i_0_76, plain, (k4_xcmplx_0(k4_xcmplx_0(np__3))=np__3)).#trainneg
cnf(i_0_78, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_7, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_130, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_94, plain, (m1_subset_1(np__256,k1_numbers))).#trainneg
cnf(i_0_50, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_34, plain, (v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_52, plain, (r1_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_int_1(X1,X2))).#trainneg
cnf(i_0_17, plain, (k1_xreal_0(X1,X2)=k6_numbers|r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_140, plain, (k3_xcmplx_0(np__8,np__4)=np__32)).#trainneg
cnf(i_0_24, plain, (m1_subset_1(k7_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_100, plain, (m2_subset_1(np__257,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_64, plain, (k2_xcmplx_0(np__7,np__0)=np__7)).#trainneg
cnf(i_0_124, plain, (m1_subset_1(np__8,k5_numbers))).#trainneg
cnf(i_0_112, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_21, plain, (m1_subset_1(k3_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_81, plain, (m1_subset_1(np__128,k1_numbers))).#trainneg
cnf(i_0_65, plain, (k2_xcmplx_0(np__7,np__1)=np__8)).#trainneg
cnf(i_0_104, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_80, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_106, plain, (m1_subset_1(np__32,k1_numbers))).#trainneg
cnf(i_0_70, plain, (k7_xcmplx_0(np__256,np__2)=np__128)).#trainneg
cnf(i_0_86, plain, (m1_subset_1(np__16,k5_numbers))).#trainneg
cnf(i_0_92, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_123, plain, (m1_subset_1(np__8,k1_numbers))).#trainneg
cnf(i_0_96, plain, (m2_subset_1(np__256,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_119, plain, (m1_subset_1(np__7,k1_numbers))).#trainneg
cnf(i_0_85, plain, (m1_subset_1(np__16,k1_numbers))).#trainneg
cnf(i_0_114, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_142, plain, (m1_subset_1(np__126,k5_numbers))).#trainneg
cnf(i_0_146, plain, (m1_subset_1(np__224,k5_numbers))).#trainneg
cnf(i_0_150, plain, (m1_subset_1(np__225,k5_numbers))).#trainneg
cnf(i_0_154, plain, (m1_subset_1(np__29,k5_numbers))).#trainneg
cnf(i_0_158, plain, (m1_subset_1(np__30,k5_numbers))).#trainneg
cnf(i_0_162, plain, (m1_subset_1(np__33,k5_numbers))).#trainneg
cnf(i_0_20, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_117, plain, (m2_subset_1(np__4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_141, plain, (m1_subset_1(np__126,k1_numbers))).#trainneg
cnf(i_0_145, plain, (m1_subset_1(np__224,k1_numbers))).#trainneg
cnf(i_0_149, plain, (m1_subset_1(np__225,k1_numbers))).#trainneg
cnf(i_0_153, plain, (m1_subset_1(np__29,k1_numbers))).#trainneg
cnf(i_0_157, plain, (m1_subset_1(np__30,k1_numbers))).#trainneg
cnf(i_0_161, plain, (m1_subset_1(np__33,k1_numbers))).#trainneg
cnf(i_0_108, plain, (m2_subset_1(np__32,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_83, plain, (m2_subset_1(np__128,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_121, plain, (m2_subset_1(np__7,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_125, plain, (m2_subset_1(np__8,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_143, plain, (m2_subset_1(np__126,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_147, plain, (m2_subset_1(np__224,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_151, plain, (m2_subset_1(np__225,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_155, plain, (m2_subset_1(np__29,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_159, plain, (m2_subset_1(np__30,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_163, plain, (m2_subset_1(np__33,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_87, plain, (m2_subset_1(np__16,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_113, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_93, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_105, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_84, plain, (v2_xxreal_0(np__128))).#trainneg
cnf(i_0_97, plain, (v2_xxreal_0(np__256))).#trainneg
cnf(i_0_101, plain, (v2_xxreal_0(np__257))).#trainneg
cnf(i_0_144, plain, (v2_xxreal_0(np__126))).#trainneg
cnf(i_0_148, plain, (v2_xxreal_0(np__224))).#trainneg
cnf(i_0_152, plain, (v2_xxreal_0(np__225))).#trainneg
cnf(i_0_156, plain, (v2_xxreal_0(np__29))).#trainneg
cnf(i_0_160, plain, (v2_xxreal_0(np__30))).#trainneg
cnf(i_0_164, plain, (v2_xxreal_0(np__33))).#trainneg
cnf(i_0_118, plain, (v2_xxreal_0(np__4))).#trainneg
cnf(i_0_109, plain, (v2_xxreal_0(np__32))).#trainneg
cnf(i_0_126, plain, (v2_xxreal_0(np__8))).#trainneg
cnf(i_0_122, plain, (v2_xxreal_0(np__7))).#trainneg
cnf(i_0_88, plain, (v2_xxreal_0(np__16))).#trainneg
cnf(i_0_407, plain, (v1_int_1(np__2)|~v1_int_1(np__1)|~v1_int_1(np__3))).#trainneg
cnf(i_0_441, plain, (v7_ordinal1(np__2)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_438, plain, (v7_ordinal1(np__257)|~v7_ordinal1(np__1)|~v7_ordinal1(np__256))).#trainneg
cnf(i_0_416, plain, (k3_nat_d(k3_power(np__2,X1),k1_newton(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_412, plain, (v1_int_1(np__3)|~v1_int_1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_508, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_477, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_382, plain, (k1_nat_1(k2_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_423, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_389, plain, (k2_xcmplx_0(k2_newton(np__2,k2_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(X1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1))))).#trainneg
cnf(i_0_442, plain, (v7_ordinal1(np__256)|~v7_ordinal1(np__128))).#trainneg
cnf(i_0_473, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__1)|~v7_ordinal1(np__7))).#trainneg
cnf(i_0_400, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_448, plain, (v7_ordinal1(np__7)|~v7_ordinal1(np__3)|~v7_ordinal1(np__4))).#trainneg
cnf(i_0_521, plain, (v1_xreal_0(np__128)|~v1_xreal_0(np__2)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_470, plain, (v3_ordinal1(k3_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_491, plain, (v1_xreal_0(np__3)|~v1_xreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_455, plain, (k3_xcmplx_0(np__0,np__32)=np__0|~v1_xcmplx_0(np__32)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_391, plain, (m2_subset_1(k1_newton(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_377, plain, (r1_int_1(X1,np__256)|~r2_int_1(np__257,np__1,X1)|~v1_int_1(X1)|~v1_int_1(np__1)|~v1_int_1(np__257))).#trainneg
cnf(i_0_379, plain, (r2_int_1(np__3,np__1,X1)|~r1_int_1(X1,np__2)|~v1_int_1(np__1)|~v1_int_1(np__3)|~v1_int_1(X1))).#trainneg
cnf(i_0_466, plain, (v1_int_1(np__7)|~v1_int_1(np__1)|~v1_int_1(np__8))).#trainneg
cnf(i_0_433, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_381, plain, (k1_nat_1(k1_newton(np__2,k2_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(X1)|~v7_ordinal1(k2_newton(np__2,X1))|~m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_474, plain, (k1_nat_1(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_424, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_409, plain, (v3_ordinal1(k4_pepin(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_476, plain, (k1_nat_1(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1))))).#trainneg
cnf(i_0_380, plain, (r1_int_1(X1,np__2)|~r2_int_1(np__3,np__1,X1)|~v1_int_1(X1)|~v1_int_1(np__1)|~v1_int_1(np__3))).#trainneg
cnf(i_0_454, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_399, plain, (v3_ordinal1(np__256))).#trainneg
cnf(i_0_488, plain, (v1_xcmplx_0(np__3)|~v1_xreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_509, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_413, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k6_xcmplx_0(X1,np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_420, plain, (v3_ordinal1(np__32))).#trainneg
cnf(i_0_452, plain, (v1_xcmplx_0(np__3)|~v1_int_1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_469, plain, (r1_int_1(X1,np__7)|~r2_int_1(np__8,np__1,X1)|~v1_int_1(X1)|~v1_int_1(np__1)|~v1_int_1(np__8))).#trainneg
cnf(i_0_517, plain, (v1_xreal_0(np__7)|~v1_xreal_0(np__1)|~v1_xreal_0(np__8))).#trainneg
cnf(i_0_394, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_462, plain, (v3_ordinal1(np__8))).#trainneg
cnf(i_0_419, plain, (k3_xcmplx_0(X1,k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_376, plain, (r2_int_1(np__257,np__1,X1)|~r1_int_1(X1,np__256)|~v1_int_1(np__1)|~v1_int_1(np__257)|~v1_int_1(X1))).#trainneg
cnf(i_0_447, plain, (v3_ordinal1(k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_414, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__3))=k2_xcmplx_0(X1,np__3)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_451, plain, (v1_xcmplx_0(k6_xcmplx_0(np__2,np__3))|~v1_int_1(np__1))).#trainneg
cnf(i_0_468, plain, (r2_int_1(np__8,np__1,X1)|~r1_int_1(X1,np__7)|~v1_int_1(np__1)|~v1_int_1(np__8)|~v1_int_1(X1))).#trainneg
cnf(i_0_387, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_511, plain, (v3_ordinal1(np__126)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_512, plain, (v3_ordinal1(np__224)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_513, plain, (v3_ordinal1(np__225)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_514, plain, (v3_ordinal1(np__29)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_515, plain, (v3_ordinal1(np__30)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_516, plain, (v3_ordinal1(np__33)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_489, plain, (v3_ordinal1(np__16)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_463, plain, (k3_xcmplx_0(np__8,np__16)=np__128|~v1_xcmplx_0(np__8)|~v1_xcmplx_0(np__16))).#trainneg
cnf(i_0_479, plain, (v3_ordinal1(np__16))).#trainneg
cnf(i_0_500, plain, (v3_ordinal1(np__126))).#trainneg
cnf(i_0_501, plain, (v3_ordinal1(np__224))).#trainneg
cnf(i_0_502, plain, (v3_ordinal1(np__225))).#trainneg
cnf(i_0_503, plain, (v3_ordinal1(np__29))).#trainneg
cnf(i_0_504, plain, (v3_ordinal1(np__30))).#trainneg
cnf(i_0_505, plain, (v3_ordinal1(np__33))).#trainneg
cnf(i_0_496, plain, (k3_xcmplx_0(np__32,k7_xcmplx_0(np__0,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__32))).#trainneg
cnf(i_0_497, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(k1_newton(X1,X3),X4))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k1_newton(X1,X3))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_493, plain, (k3_xcmplx_0(np__16,k7_xcmplx_0(np__8,X1))=k7_xcmplx_0(np__128,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__8)|~v1_xcmplx_0(np__16))).#trainneg
cnf(i_0_492, plain, (k3_xcmplx_0(np__16,k7_xcmplx_0(np__16,X1))=k7_xcmplx_0(np__256,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__16))).#trainneg
cnf(i_0_495, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(X1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_494, plain, (k3_xcmplx_0(np__8,k7_xcmplx_0(np__4,X1))=k7_xcmplx_0(np__32,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(np__8))).#trainneg
cnf(i_0_498, plain, (k7_xcmplx_0(k3_xcmplx_0(X2,X1),X3)=k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_507, plain, (m2_subset_1(k4_pepin(X1),k1_numbers,k5_numbers)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_443, plain, (v7_ordinal1(k6_xcmplx_0(X1,X2))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_471, plain, (m1_subset_1(k3_power(np__2,k7_nat_d(X1,X2)),k5_numbers)|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_564, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k6_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_529, plain, (v1_int_1(np__2)|~v1_int_1(np__1)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_683, plain, (v1_xcmplx_0(np__3)|~v1_int_1(np__3))).#trainneg
cnf(i_0_603, plain, (v7_ordinal1(np__256))).#trainneg
cnf(i_0_655, plain, (v1_xreal_0(np__2)|~v1_int_1(np__1)|~v1_int_1(np__3))).#trainneg
cnf(i_0_566, plain, (v1_xcmplx_0(np__1)|~v1_int_1(k6_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_608, plain, (v7_ordinal1(k3_nat_d(X1,X2))|~v1_finset_1(k3_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_681, plain, (v1_xcmplx_0(np__3)|~v7_ordinal1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_569, plain, (v1_int_1(np__3)|~v7_ordinal1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_525, plain, (v1_int_2(k4_pepin(X1))|~r2_int_1(k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_613, plain, (v1_xreal_0(np__3)|~v7_ordinal1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_555, plain, (k3_nat_d(k3_power(np__2,X1),np__2)=k3_power(np__2,k7_nat_d(X1,np__1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__1)|~v1_xreal_0(np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_639, plain, (m2_subset_1(k3_nat_d(X1,X2),k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_607, plain, (v7_ordinal1(np__7)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_527, plain, (k3_nat_d(k3_power(np__2,X1),np__8)=k3_power(np__2,k7_nat_d(X1,np__3))|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_671, plain, (k3_nat_d(k3_power(np__2,X1),np__16)=k3_power(np__2,k7_nat_d(X1,np__4))|~r1_xxreal_0(np__4,X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__4))).#trainneg
cnf(i_0_17, plain, (k1_xreal_0(X1,X2)=np__0|r1_xxreal_0(np__0,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_595, plain, (v1_xreal_0(np__3)|~v1_int_1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_136, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_592, plain, (k1_nat_1(k2_newton(np__2,np__2),np__1)=k4_pepin(np__1)|~v7_ordinal1(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_735, plain, (k3_xcmplx_0(X1,k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_587, plain, (k2_xcmplx_0(k2_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_656, plain, (v1_xreal_0(np__256)|~v1_int_1(np__1)|~v1_int_1(np__257))).#trainneg
cnf(i_0_594, plain, (v1_xreal_0(k6_xcmplx_0(np__2,np__3))|~v1_int_1(np__1))).#trainneg
cnf(i_0_684, plain, (k1_newton(np__1,k2_xcmplx_0(X1,np__1))=k1_newton(np__1,X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(k1_newton(np__1,X1))|~v7_ordinal1(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_736, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,X1))=k1_newton(np__1,X1)|~v1_xcmplx_0(k1_newton(np__1,X1))|~v1_xcmplx_0(np__1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_774, plain, (k3_xcmplx_0(np__1,np__128)=np__128|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__256))).#trainneg
cnf(i_0_528, plain, (k3_nat_d(np__8,k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(np__3,X1))|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_602, plain, (k2_xcmplx_0(k1_newton(np__2,k2_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1)|~v7_ordinal1(k2_newton(np__2,X1)))).#trainneg
cnf(i_0_734, plain, (k3_xcmplx_0(k1_newton(X1,X2),X1)=k1_newton(X1,k2_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_673, plain, (k3_nat_d(np__16,k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(np__4,X1))|~r1_xxreal_0(X1,np__4)|~v7_ordinal1(np__4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_593, plain, (k1_nat_1(k1_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(np__2,X1)))).#trainneg
cnf(i_0_644, plain, (k2_xcmplx_0(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_657, plain, (v7_ordinal1(k4_pepin(X1))|~v1_finset_1(k4_pepin(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_578, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_589, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_617, plain, (m1_subset_1(k1_newton(X1,X2),k5_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_586, plain, (k1_nat_1(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_619, plain, (v1_xreal_0(np__7)|~v1_int_1(np__1)|~v1_int_1(np__8))).#trainneg
cnf(i_0_646, plain, (k1_nat_1(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_591, plain, (k1_nat_1(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_638, plain, (m2_subset_1(k4_pepin(X1),k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_806, plain, (v7_ordinal1(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers))).#trainneg
cnf(i_0_641, plain, (k1_nat_1(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_618, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_733, plain, (k1_newton(X1,k2_xcmplx_0(X3,X2))=k1_newton(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_571, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_745, plain, (r2_int_1(X1,k4_xcmplx_0(np__3),X2)|~r1_int_1(X2,k2_xcmplx_0(X1,np__3))|~v1_int_1(k4_xcmplx_0(np__3))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_552, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_742, plain, (v7_ordinal1(k7_nat_d(X1,X2))|~v1_finset_1(k7_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_689, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_723, plain, (m1_subset_1(k1_newton(X1,X2),k1_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_757, plain, (v7_ordinal1(np__16)|~v1_finset_1(np__16))).#trainneg
cnf(i_0_637, plain, (m2_subset_1(k7_nat_d(X1,X2),k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_667, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_747, plain, (r1_int_1(X1,k2_xcmplx_0(X2,np__3))|~r2_int_1(X2,k4_xcmplx_0(np__3),X1)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(np__3))|~v1_int_1(X2)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_727, plain, (m1_subset_1(k1_nat_1(X1,X2),k1_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_779, plain, (k3_xcmplx_0(X2,k7_xcmplx_0(X1,X3))=k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_731, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(k1_newton(X1,X3),X4))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,X2)),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k1_newton(X1,X3))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_703, plain, (v7_ordinal1(np__32)|~v1_finset_1(np__32))).#trainneg
cnf(i_0_726, plain, (m1_subset_1(k1_nat_1(X1,X2),k1_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_725, plain, (m1_subset_1(k13_newton(X1,X2),k1_numbers)|v1_xboole_0(k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_724, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_781, plain, (k3_xcmplx_0(np__8,k7_xcmplx_0(np__16,X1))=k7_xcmplx_0(np__128,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__16)|~v1_xcmplx_0(np__8))).#trainneg
cnf(i_0_676, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_785, plain, (k3_xcmplx_0(np__4,k7_xcmplx_0(np__8,X1))=k7_xcmplx_0(np__32,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__8)|~v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_758, plain, (v7_ordinal1(np__126)|~v1_finset_1(np__126))).#trainneg
cnf(i_0_759, plain, (v7_ordinal1(np__224)|~v1_finset_1(np__224))).#trainneg
cnf(i_0_760, plain, (v7_ordinal1(np__225)|~v1_finset_1(np__225))).#trainneg
cnf(i_0_761, plain, (v7_ordinal1(np__29)|~v1_finset_1(np__29))).#trainneg
cnf(i_0_762, plain, (v7_ordinal1(np__30)|~v1_finset_1(np__30))).#trainneg
cnf(i_0_763, plain, (v7_ordinal1(np__33)|~v1_finset_1(np__33))).#trainneg
cnf(i_0_48, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_616, plain, (k3_xcmplx_0(np__0,k7_xcmplx_0(np__32,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__32)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_765, plain, (k3_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4),k1_newton(X1,X2))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4)|~v1_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_612, plain, (v1_xreal_0(np__3)|~m1_subset_1(k4_xcmplx_0(np__3),k1_numbers))).#trainneg
cnf(i_0_680, plain, (v1_xcmplx_0(np__3)|~m1_subset_1(k4_xcmplx_0(np__3),k1_numbers))).#trainneg
cnf(i_0_600, plain, (v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(np__1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_803, plain, (v7_ordinal1(k6_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k6_xcmplx_0(np__2,np__3))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_804, plain, (v7_ordinal1(k6_xcmplx_0(X1,k4_xcmplx_0(np__3)))|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_700, plain, (m2_subset_1(k6_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_805, plain, (v7_ordinal1(k6_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3)))|~v1_xcmplx_0(k6_xcmplx_0(np__2,np__3))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_702, plain, (m2_subset_1(k6_xcmplx_0(X1,np__1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_532, plain, (k3_power(np__2,k7_nat_d(X1,np__3))=k3_nat_d(k1_newton(np__2,X1),np__8)|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__3)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_557, plain, (k3_power(np__2,k7_nat_d(np__3,X1))=k3_nat_d(np__8,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(np__3)|~v7_ordinal1(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_672, plain, (k3_power(np__2,k7_nat_d(X1,np__4))=k3_nat_d(k1_newton(np__2,X1),np__16)|~r1_xxreal_0(np__4,X1)|~v7_ordinal1(X1)|~v7_ordinal1(np__4)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_674, plain, (k3_power(np__2,k7_nat_d(np__4,X1))=k3_nat_d(np__16,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__4)|~v7_ordinal1(np__4)|~v7_ordinal1(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_533, plain, (k3_nat_d(np__2,k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(np__1,X1))|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~v1_xreal_0(np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_768, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(X1,X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,np__1)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_769, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k1_newton(X1,X2),X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(np__1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_609, plain, (v3_ordinal1(k3_power(np__2,k7_nat_d(X1,X2)))|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_610, plain, (v3_ordinal1(k3_power(np__2,k7_nat_d(X1,X2)))|~v7_ordinal1(k1_newton(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_611, plain, (v3_ordinal1(k3_power(np__2,k7_nat_d(X1,X2)))|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k1_newton(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_795, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k1_newton(X1,X2),X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,np__1)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_797, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(X1,X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(np__1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_808, plain, (m2_subset_1(k3_power(np__2,k7_nat_d(X1,X2)),k1_numbers,k5_numbers)|v1_xboole_0(k1_numbers)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_809, plain, (m1_subset_1(k1_newton(np__2,k7_nat_d(X1,X2)),k5_numbers)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k7_nat_d(X1,X2))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_766, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4),X5))=k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4),X5)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_792, plain, (k3_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4),k7_xcmplx_0(k1_newton(X1,X2),X5))=k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4),X5)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_819, plain, (v1_xreal_0(np__2)|~v1_int_1(np__1)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1320, plain, (v7_ordinal1(np__16))).#trainneg
cnf(i_0_812, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(np__3)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_977, plain, (k1_xreal_0(X1,X2)=k6_xcmplx_0(X1,X2)|k1_xreal_0(X1,X2)=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_813, plain, (v1_xcmplx_0(np__1)|~v1_int_1(np__3)|~v1_int_1(np__2))).#trainneg
cnf(i_0_822, plain, (v1_xcmplx_0(np__1)|~v1_int_1(np__1))).#trainneg
cnf(i_0_974, plain, (k1_xreal_0(np__3,np__1)=np__2|~r1_xxreal_0(np__0,np__2)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_937, plain, (k2_xcmplx_0(np__256,np__0)=np__256)).#trainneg
cnf(i_0_818, plain, (v1_xcmplx_0(np__3)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_972, plain, (v1_xreal_0(np__256)|~v1_int_1(np__1)|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_935, plain, (k2_xcmplx_0(np__1,np__7)=np__8)).#trainneg
cnf(i_0_934, plain, (k2_xcmplx_0(np__0,np__7)=np__7)).#trainneg
cnf(i_0_1398, plain, (v7_ordinal1(np__32))).#trainneg
cnf(i_0_908, plain, (k1_xreal_0(np__3,np__1)=np__0|r1_xxreal_0(np__0,np__2)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_1130, plain, (v1_xreal_0(np__7)|~v1_int_1(np__1))).#trainneg
cnf(i_0_517, plain, (v1_xreal_0(np__7)|~v1_xreal_0(np__8))).#trainneg
cnf(i_0_875, plain, (~r2_int_1(k1_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~v7_ordinal1(np__3)|~m1_subset_1(k3_nat_d(k7_nat_d(np__257,np__1),np__2),k5_numbers))).#trainneg
cnf(i_0_914, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(np__0,np__128)),np__1)))).#trainneg
cnf(i_0_1453, plain, (v7_ordinal1(np__126))).#trainneg
cnf(i_0_1455, plain, (v7_ordinal1(np__224))).#trainneg
cnf(i_0_1457, plain, (v7_ordinal1(np__225))).#trainneg
cnf(i_0_1459, plain, (v7_ordinal1(np__29))).#trainneg
cnf(i_0_1461, plain, (v7_ordinal1(np__30))).#trainneg
cnf(i_0_1463, plain, (v7_ordinal1(np__33))).#trainneg
cnf(i_0_971, plain, (k2_xcmplx_0(k1_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1)|~v7_ordinal1(k1_newton(np__2,X1)))).#trainneg
cnf(i_0_1164, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_965, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_943, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_931, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_1160, plain, (m2_subset_1(k4_xcmplx_0(X1),k1_numbers,k5_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1079, plain, (k2_xcmplx_0(np__1,k1_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k1_newton(np__2,k1_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(k1_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1162, plain, (m2_subset_1(k1_newton(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(k1_newton(X1,X2),k5_numbers)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1309, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1052, plain, (k2_xcmplx_0(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_885, plain, (m1_subset_1(k3_nat_d(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1161, plain, (m2_subset_1(k4_xcmplx_0(X1),k1_numbers,k5_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_922, plain, (r1_nat_d(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))).#trainneg
cnf(i_0_921, plain, (r1_nat_d(np__257,k2_xcmplx_0(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(np__1)|~m1_subset_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),k5_numbers))).#trainneg
cnf(i_0_1367, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(X1,np__2))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1081, plain, (k1_nat_1(np__1,k1_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k1_newton(np__2,k1_newton(np__2,X1)),k5_numbers)|~v7_ordinal1(k1_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1491, plain, (v7_ordinal1(k6_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(np__2)|~v7_ordinal1(np__2)|~m1_subset_1(k4_xcmplx_0(np__3),k5_numbers))).#trainneg
cnf(i_0_1266, plain, (v3_ordinal1(k13_newton(X1,X2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1076, plain, (k1_nat_1(np__1,k1_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(k1_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1323, plain, (m1_subset_1(k7_nat_d(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_916, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_xcmplx_0(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(np__128)|~m1_subset_1(k4_nat_1(np__32,np__0),k5_numbers))).#trainneg
cnf(i_0_923, plain, (r1_nat_d(np__257,k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))).#trainneg
cnf(i_0_910, plain, (k1_xreal_0(X1,k4_xcmplx_0(np__3))=np__0|r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__3))|~v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1182, plain, (k3_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,X2)),k1_newton(X1,X4))=k1_newton(X1,k2_xcmplx_0(X4,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_976, plain, (k1_xreal_0(X1,k4_xcmplx_0(np__3))=k2_xcmplx_0(X1,np__3)|~r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__3))|~v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1113, plain, (v3_ordinal1(k1_newton(X1,X2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_918, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_915, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(np__128)|~m1_subset_1(k4_nat_1(np__32,np__0),k5_numbers))).#trainneg
cnf(i_0_1230, plain, (k1_newton(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k1_newton(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1317, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1290, plain, (v3_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1148, plain, (m1_subset_1(k4_pepin(X1),k1_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_920, plain, (r1_nat_d(np__257,k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(np__1)|~m1_subset_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),k5_numbers))).#trainneg
cnf(i_0_1166, plain, (k1_nat_1(k2_newton(np__2,k1_newton(np__2,k2_xcmplx_0(X2,X1))),np__1)=k4_pepin(k2_xcmplx_0(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1180, plain, (k3_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,X2)),k1_newton(X1,X4))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),X4))|~v1_xcmplx_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_917, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_1189, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,k2_xcmplx_0(X4,X3)))=k1_newton(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X3,X4))|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1349, plain, (k3_xcmplx_0(k7_xcmplx_0(X2,X3),X1)=k3_xcmplx_0(X2,k7_xcmplx_0(X1,X3))|~v1_xcmplx_0(k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1190, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,k2_xcmplx_0(X4,X3)))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X3,X4),X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k2_xcmplx_0(X3,X4))|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1269, plain, (v3_ordinal1(k1_nat_1(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1307, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1169, plain, (k2_xcmplx_0(k2_newton(np__2,k1_newton(np__2,k2_xcmplx_0(X2,X1))),np__1)=k4_pepin(k2_xcmplx_0(X1,X2))|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,k2_xcmplx_0(X2,X1))))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1083, plain, (k1_nat_1(k1_newton(np__2,np__2),np__1)=k4_pepin(np__1)|~v7_ordinal1(np__2)|~v7_ordinal1(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_968, plain, (k2_xcmplx_0(k2_newton(np__2,np__2),np__1)=k4_pepin(np__1)|~v7_ordinal1(k2_newton(np__2,np__2))|~v7_ordinal1(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1368, plain, (k3_xcmplx_0(np__256,k7_xcmplx_0(X1,np__2))=k3_xcmplx_0(X1,np__128)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__256)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1249, plain, (v3_ordinal1(k1_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1181, plain, (m1_subset_1(k1_newton(X1,k2_xcmplx_0(X3,X2)),k5_numbers)|~m1_subset_1(k2_xcmplx_0(X2,X3),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_919, plain, (r1_nat_d(np__257,k2_nat_1(k1_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~m1_subset_1(k2_nat_1(k4_nat_1(np__32,np__0),np__128),k5_numbers))).#trainneg
cnf(i_0_1436, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1381, plain, (k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,X2)),X4)=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1126, plain, (k1_nat_1(np__1,k2_newton(np__2,np__2))=k4_pepin(np__1)|~v7_ordinal1(k2_newton(np__2,np__2))|~v7_ordinal1(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1442, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X2),X1)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1396, plain, (k3_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,X2)),k7_xcmplx_0(k1_newton(X1,X4),X5))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X4,k2_xcmplx_0(X2,X3))),X5)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1184, plain, (k3_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,X2)),k7_xcmplx_0(k1_newton(X1,X4),X5))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),X4)),X5)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X1)|~v7_ordinal1(X4)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1434, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1375, plain, (k3_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4),k1_newton(X1,X2))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,X2)),X4)|~v1_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1191, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X4,X3)),X5))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))),X5)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X3,X4))|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1388, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X4,X3)),X5))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X3,X4),X2)),X5)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k2_xcmplx_0(X3,X4))|~v7_ordinal1(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1698, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(k1_newton(X1,X3),X4))=k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,np__1)),X2),X4)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1667, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(k7_xcmplx_0(X1,X3),X4))=k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,np__1)),X3),X4)|~v1_xcmplx_0(k7_xcmplx_0(X1,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1656, plain, (k3_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4),X5),k1_newton(X1,X2))=k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4),X5)|~v1_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4),X5))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X3),X4))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_873, plain, (v1_int_1(np__3)|~m1_subset_1(k4_xcmplx_0(np__3),k5_numbers))).#trainneg
cnf(i_0_877, plain, (v1_xreal_0(np__3)|~m1_subset_1(k4_xcmplx_0(np__3),k5_numbers))).#trainneg
cnf(i_0_964, plain, (v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(np__1)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1498, plain, (m2_subset_1(k2_xcmplx_0(X1,np__3),k1_numbers,k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1490, plain, (v7_ordinal1(k6_xcmplx_0(X1,np__1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers))).#trainneg
cnf(i_0_1308, plain, (m1_subset_1(k6_xcmplx_0(X1,np__1),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1435, plain, (m1_subset_1(k6_xcmplx_0(X1,np__1),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1051, plain, (v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(np__1)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_826, plain, (k3_power(np__2,k7_nat_d(np__1,X1))=k3_nat_d(np__2,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~v1_xreal_0(np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1387, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(X1,X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(np__1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1395, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k1_newton(X1,X2),X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,np__1)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1183, plain, (m2_subset_1(k1_newton(X1,k2_xcmplx_0(X3,X2)),k1_numbers,k5_numbers)|~m1_subset_1(k2_xcmplx_0(X2,X3),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1648, plain, (m1_subset_1(k3_power(np__2,k7_nat_d(X1,X2)),k1_numbers)|~r1_xxreal_0(X2,X1)|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1219, plain, (k1_newton(X1,k2_xcmplx_0(np__1,k2_newton(np__2,k2_newton(np__2,X2))))=k1_newton(X1,k4_pepin(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1220, plain, (k1_newton(X1,k2_xcmplx_0(np__1,k2_newton(np__2,k1_newton(np__2,X2))))=k1_newton(X1,k4_pepin(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_824, plain, (v3_ordinal1(k3_power(np__2,k7_nat_d(X1,X2)))|~v7_ordinal1(k1_newton(np__2,X2))|~v7_ordinal1(k1_newton(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1350, plain, (k7_xcmplx_0(k3_xcmplx_0(X2,k7_xcmplx_0(X1,X3)),X4)=k3_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(X2,X3),X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1351, plain, (k7_xcmplx_0(k3_xcmplx_0(X2,k7_xcmplx_0(X1,X3)),X4)=k3_xcmplx_0(k7_xcmplx_0(X2,X3),k7_xcmplx_0(X1,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1369, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),k7_xcmplx_0(X3,X4))=k3_xcmplx_0(X3,k3_xcmplx_0(X1,k7_xcmplx_0(X2,X4)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1370, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),k7_xcmplx_0(X3,X4))=k3_xcmplx_0(X3,k3_xcmplx_0(X2,k7_xcmplx_0(X1,X4)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1484, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),k1_newton(X1,X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X3,np__1)),X2)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_888, plain, (m2_subset_1(k3_power(np__2,k7_nat_d(X1,X2)),k1_numbers,k5_numbers)|~v7_ordinal1(k1_newton(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_889, plain, (m2_subset_1(k3_power(np__2,k7_nat_d(X1,X2)),k1_numbers,k5_numbers)|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k1_newton(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_890, plain, (m2_subset_1(k3_power(np__2,k7_nat_d(X1,X2)),k1_numbers,k5_numbers)|~v7_ordinal1(k1_newton(np__2,X2))|~v7_ordinal1(k1_newton(np__2,X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1476, plain, (k3_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X2),X3),X1)=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(np__1,X2)),X3)|~v1_xcmplx_0(k7_xcmplx_0(k1_newton(X1,X2),X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_865, plain, (v7_ordinal1(k3_power(np__2,k7_nat_d(X1,X2)))|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~v1_finset_1(k3_power(np__2,k7_nat_d(X1,X2)))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1172, plain, (k1_nat_1(k1_newton(np__2,k1_newton(np__2,k2_xcmplx_0(X2,X1))),np__1)=k4_pepin(k2_xcmplx_0(X1,X2))|~v7_ordinal1(k1_newton(np__2,k2_xcmplx_0(X2,X1)))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1229, plain, (k1_newton(X1,k2_xcmplx_0(np__1,k1_newton(np__2,k2_newton(np__2,X2))))=k1_newton(X1,k4_pepin(X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X2)))|~v7_ordinal1(k2_newton(np__2,X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1597, plain, (v3_ordinal1(k1_newton(np__2,k7_nat_d(X1,X2)))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k7_nat_d(X1,X2))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1603, plain, (v3_ordinal1(k1_newton(np__2,k7_nat_d(X1,X2)))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(k1_newton(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2)|~v7_ordinal1(k7_nat_d(X1,X2)))).#trainneg
cnf(i_0_1609, plain, (v3_ordinal1(k1_newton(np__2,k7_nat_d(X1,X2)))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(k3_power(np__2,X2))|~v7_ordinal1(k1_newton(np__2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2)|~v7_ordinal1(k7_nat_d(X1,X2)))).#trainneg
cnf(i_0_1168, plain, (k1_nat_1(np__1,k2_newton(np__2,k1_newton(np__2,k2_xcmplx_0(X2,X1))))=k4_pepin(k2_xcmplx_0(X1,X2))|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,k2_xcmplx_0(X2,X1))))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_866, plain, (v7_ordinal1(k3_power(np__2,k7_nat_d(X1,X2)))|~v7_ordinal1(k1_newton(np__2,X2))|~v7_ordinal1(k3_power(np__2,X1))|~v1_finset_1(k3_power(np__2,k7_nat_d(X1,X2)))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1709, plain, (k7_nat_d(X1,X2)=k6_xcmplx_0(X1,X2)|k1_xreal_0(X1,X2)=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_529, plain, (v1_int_1(np__2)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1715, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1707, plain, (k1_xreal_0(X3,X4)=np__0|k6_xcmplx_0(X3,X4)!=np__0|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_418, plain, (k3_xcmplx_0(k1_newton(X1,X2),X1)=k1_newton(X1,k2_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1731, plain, (k7_nat_d(np__3,np__1)=np__2|~r1_xxreal_0(np__0,np__2)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_735, plain, (k3_xcmplx_0(X1,k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_684, plain, (k1_newton(np__1,k2_xcmplx_0(X1,np__1))=k1_newton(np__1,X1)|~v1_xcmplx_0(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2015, plain, (v3_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_736, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,X1))=k1_newton(np__1,X1)|~v1_xcmplx_0(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_527, plain, (k3_nat_d(k3_power(np__2,X1),np__8)=k3_power(np__2,k7_nat_d(X1,np__3))|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_734, plain, (k3_xcmplx_0(k1_newton(X1,X2),X1)=k1_newton(X1,k2_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1751, plain, (k7_nat_d(np__3,np__1)=np__0|r1_xxreal_0(np__0,np__2)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_2901, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k1_newton(X1,np__2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2233, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(X2,np__257))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2241, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__2))=k1_newton(X1,k2_xcmplx_0(X2,np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_419, plain, (k3_xcmplx_0(X1,k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1940, plain, (v7_ordinal1(k13_newton(X1,X2))|~v1_finset_1(k13_newton(X1,X2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_592, plain, (k1_nat_1(k2_newton(np__2,np__2),np__1)=k4_pepin(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_671, plain, (k3_nat_d(k3_power(np__2,X1),np__16)=k3_power(np__2,k7_nat_d(X1,np__4))|~r1_xxreal_0(np__4,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1797, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2290, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(np__257,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2298, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__2))=k1_newton(X1,k2_xcmplx_0(np__2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_528, plain, (k3_nat_d(np__8,k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(np__3,X1))|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1972, plain, (k7_nat_d(X1,k4_xcmplx_0(np__3))=np__0|r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__3))|~v7_ordinal1(k4_xcmplx_0(np__3))|~v7_ordinal1(X1)|~v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1724, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_673, plain, (k3_nat_d(np__16,k3_power(np__2,X1))=k3_power(np__2,k7_nat_d(np__4,X1))|~r1_xxreal_0(X1,np__4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2014, plain, (v7_ordinal1(k1_newton(X1,X2))|~v1_finset_1(k1_newton(X1,X2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2012, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__3))=k2_xcmplx_0(X1,np__3)|k2_xcmplx_0(X1,np__3)=np__0|~v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(X1)|~r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__3)))).#trainneg
cnf(i_0_557, plain, (k3_power(np__2,k7_nat_d(np__3,X1))=k3_nat_d(np__8,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1753, plain, (v1_xreal_0(np__7))).#trainneg
cnf(i_0_2011, plain, (k7_nat_d(X1,k4_xcmplx_0(np__3))=k2_xcmplx_0(X1,np__3)|~v7_ordinal1(k4_xcmplx_0(np__3))|~v7_ordinal1(X1)|~r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__3))|~v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1727, plain, (r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v7_ordinal1(k2_xcmplx_0(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~m1_subset_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),k5_numbers))).#trainneg
cnf(i_0_1762, plain, (k6_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers))).#trainneg
cnf(i_0_2196, plain, (k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(np__2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2088, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_674, plain, (k3_power(np__2,k7_nat_d(np__4,X1))=k3_nat_d(np__16,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__4)|~v7_ordinal1(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1723, plain, (r1_int_1(np__257,k2_nat_1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k2_nat_1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_591, plain, (k1_nat_1(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_1759, plain, (r1_nat_d(np__257,k2_xcmplx_0(k13_newton(np__3,k2_nat_1(np__0,np__128)),np__1))|~m1_subset_1(k13_newton(np__3,k2_nat_1(np__0,np__128)),k5_numbers))).#trainneg
cnf(i_0_1728, plain, (r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))).#trainneg
cnf(i_0_1989, plain, (k3_xcmplx_0(k1_newton(X1,np__257),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__257))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2232, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__256))=k1_newton(X1,k2_xcmplx_0(X2,np__256))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1729, plain, (r1_int_1(np__257,k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))).#trainneg
cnf(i_0_1997, plain, (k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_646, plain, (k1_nat_1(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_2154, plain, (k1_nat_1(k2_newton(np__2,k1_newton(np__2,np__2)),np__1)=k4_pepin(np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_474, plain, (k1_nat_1(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_1755, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,np__128),np__1)))).#trainneg
cnf(i_0_1761, plain, (r1_nat_d(np__257,k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(np__0,np__128))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(np__0,np__128))))).#trainneg
cnf(i_0_2138, plain, (k4_pepin(k2_xcmplx_0(X2,X1))=k4_pepin(k2_xcmplx_0(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1914, plain, (r1_nat_d(np__257,k2_xcmplx_0(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),np__1))|~m1_subset_1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),k5_numbers)|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_2188, plain, (k3_xcmplx_0(k1_newton(X1,np__257),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(np__257,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1907, plain, (r1_nat_d(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0)))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_2019, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_nat_1(np__128,np__0)),np__1)))).#trainneg
cnf(i_0_1968, plain, (r1_nat_d(np__257,k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0)))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_2146, plain, (k1_nat_1(k2_newton(np__2,k1_newton(np__2,np__257)),np__1)=k4_pepin(np__257)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1760, plain, (r1_nat_d(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(np__0,np__128))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(np__0,np__128))))).#trainneg
cnf(i_0_2228, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__8))=k1_newton(X1,k2_xcmplx_0(X2,np__8))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2270, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(X1,np__2))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_2456, plain, (m1_subset_1(k1_newton(X1,np__2),k5_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2860, plain, (m2_subset_1(k2_xcmplx_0(X1,np__3),k1_numbers,k5_numbers)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~v1_int_1(np__3))).#trainneg
cnf(i_0_2448, plain, (m1_subset_1(k1_newton(X1,np__257),k5_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1913, plain, (r1_nat_d(np__257,k2_xcmplx_0(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))),np__1))|~m1_subset_1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))),k5_numbers)|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_768, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(X1,X3))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,np__1)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1988, plain, (k3_xcmplx_0(k1_newton(X1,np__256),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__256))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2421, plain, (v7_ordinal1(k1_nat_1(X1,X2))|~v1_finset_1(k1_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1956, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))),np__1))|~m1_subset_1(k4_nat_1(np__32,np__0),k5_numbers))).#trainneg
cnf(i_0_1878, plain, (m1_subset_1(k6_xcmplx_0(X2,X1),k5_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2859, plain, (m2_subset_1(k2_xcmplx_0(X1,np__3),k1_numbers,k5_numbers)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_2317, plain, (v7_ordinal1(k1_nat_1(X1,X2))|~v1_finset_1(k1_nat_1(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1726, plain, (r1_int_1(np__257,k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k2_nat_1(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~m1_subset_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),k5_numbers))).#trainneg
cnf(i_0_2368, plain, (k2_xcmplx_0(k2_newton(np__2,k1_newton(np__2,np__2)),np__1)=k4_pepin(np__2)|~v1_xcmplx_0(np__2)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,np__2))))).#trainneg
cnf(i_0_3672, plain, (k7_nat_d(X1,X2)=np__0|k7_nat_d(X1,X2)=k6_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_454, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3693, plain, (k7_nat_d(X1,X2)=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X1)|k6_xcmplx_0(X1,X2)!=np__0|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_417, plain, (k3_nat_d(k1_newton(np__2,X1),k3_power(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_416, plain, (k3_nat_d(k3_power(np__2,X1),k1_newton(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_521, plain, (v1_xreal_0(np__128)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_3736, plain, (k1_nat_1(k2_newton(np__2,np__8),np__1)=np__257)).#trainneg
cnf(i_0_531, plain, (k3_nat_d(k1_newton(np__2,X1),k1_newton(np__2,X2))=k3_power(np__2,k7_nat_d(X1,X2))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_774, plain, (k3_xcmplx_0(np__1,np__128)=np__128|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__256))).#trainneg
cnf(i_0_413, plain, (k2_xcmplx_0(X1,k6_xcmplx_0(np__2,np__3))=k6_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3739, plain, (k1_nat_1(k1_newton(np__2,np__8),np__1)=np__257)).#trainneg
cnf(i_0_3698, plain, (v1_xcmplx_0(np__8)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_4068, plain, (k3_power(np__2,k7_nat_d(X1,np__3))=k3_nat_d(k1_newton(np__2,X1),np__8)|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3791, plain, (v1_xreal_0(k3_xcmplx_0(X1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3738, plain, (k2_xcmplx_0(k2_newton(np__2,np__8),np__1)=np__257|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_3987, plain, (k1_newton(np__1,np__2)=k1_newton(np__1,np__1))).#trainneg
cnf(i_0_3792, plain, (v7_ordinal1(k3_xcmplx_0(X1,X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4276, plain, (k3_power(np__2,k7_nat_d(X1,np__4))=k3_nat_d(k1_newton(np__2,X1),np__16)|~r1_xxreal_0(np__4,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4471, plain, (r2_int_1(X1,k6_xcmplx_0(np__2,np__3),X2)|~r1_int_1(X2,k2_xcmplx_0(np__1,X1))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X1)|~v1_int_1(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_826, plain, (k3_power(np__2,k7_nat_d(np__1,X1))=k3_nat_d(np__2,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4217, plain, (v7_ordinal1(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3986, plain, (k1_newton(np__1,np__257)=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_4759, plain, (r1_nat_d(np__257,k2_nat_1(k13_newton(np__3,k2_xcmplx_0(np__128,np__0)),np__1)))).#trainneg
cnf(i_0_4218, plain, (r1_int_1(X1,k2_xcmplx_0(X2,np__1))|~r2_int_1(X2,k6_xcmplx_0(np__2,np__3),X1)|~v1_int_1(X1)|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4213, plain, (k1_xreal_0(X1,k6_xcmplx_0(np__2,np__3))=np__0|r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4585, plain, (k3_xcmplx_0(X1,k1_newton(X1,np__256))=k1_newton(X1,np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4419, plain, (r2_int_1(k4_xcmplx_0(X1),X2,X3)|~r1_int_1(X3,k6_xcmplx_0(k4_xcmplx_0(X2),X1))|~v1_int_1(X2)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X3)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers))).#trainneg
cnf(i_0_3985, plain, (k1_newton(np__1,np__7)=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_3794, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(np__2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4491, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3795, plain, (m1_subset_1(k3_xcmplx_0(X1,X1),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4754, plain, (r1_nat_d(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0)))))|~v7_ordinal1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_3796, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4297, plain, (k3_xcmplx_0(X1,k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(np__257,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3688, plain, (~r2_int_1(k1_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~m1_subset_1(k3_nat_d(np__256,np__2),k5_numbers)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_4472, plain, (k1_xreal_0(X1,k6_xcmplx_0(np__2,np__3))=np__0|r1_xxreal_0(np__0,k2_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4838, plain, (m1_subset_1(k2_xcmplx_0(X1,np__3),k5_numbers)|~v1_xcmplx_0(X1)|~v1_int_1(np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4945, plain, (m1_subset_1(k2_xcmplx_0(X1,np__3),k5_numbers)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_3737, plain, (k1_nat_1(np__1,k2_newton(np__2,np__8))=np__257|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_4827, plain, (k3_xcmplx_0(np__2,np__1)=k3_xcmplx_0(np__1,np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3743, plain, (k2_xcmplx_0(k1_newton(np__2,np__8),np__1)=np__257|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_4469, plain, (v1_xreal_0(k2_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3692, plain, (k2_xcmplx_0(X1,np__3)=np__0|~r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__3))|~v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(X1)|k6_xcmplx_0(X1,k4_xcmplx_0(np__3))!=np__0)).#trainneg
cnf(i_0_4769, plain, (r1_nat_d(np__257,k2_nat_1(np__1,k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0)))))|~v7_ordinal1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_4290, plain, (k1_newton(X1,k2_xcmplx_0(np__257,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__257))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3818, plain, (k3_xcmplx_0(k1_newton(X1,X2),k3_xcmplx_0(X1,X1))=k1_newton(X1,k2_xcmplx_0(np__2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4902, plain, (k3_xcmplx_0(k1_newton(X1,np__256),X1)=k1_newton(X1,np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3761, plain, (k1_newton(np__2,k2_xcmplx_0(np__3,np__1))=k3_xcmplx_0(np__8,np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_529, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_5248, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)))|~v1_int_1(np__257))).#trainneg
cnf(i_0_376, plain, (r2_int_1(np__257,np__1,X1)|~r1_int_1(X1,np__256)|~v1_int_1(np__257)|~v1_int_1(X1))).#trainneg
cnf(i_0_379, plain, (r2_int_1(np__3,np__1,X1)|~r1_int_1(X1,np__2)|~v1_int_1(np__3)|~v1_int_1(X1))).#trainneg
cnf(i_0_5326, plain, (k1_newton(np__1,k2_newton(np__2,np__8))=k1_newton(np__1,np__257)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_5027, plain, (k7_nat_d(np__257,np__1)=np__0|~r2_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257))).#trainneg
cnf(i_0_466, plain, (v1_int_1(np__7)|~v1_int_1(np__8))).#trainneg
cnf(i_0_5294, plain, (v1_xreal_0(k1_newton(X1,np__2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6087, plain, (v1_xcmplx_0(k1_newton(X1,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5652, plain, (v7_ordinal1(k1_newton(X1,np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5805, plain, (r2_int_1(k4_xcmplx_0(X1),X2,X3)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X2)|~v1_int_1(X3)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~r2_int_1(k4_xcmplx_0(X2),X1,X3)|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_5031, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k5_numbers)|k7_nat_d(X1,X2)=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5024, plain, (v1_int_2(k4_pepin(X1))|k7_nat_d(k4_pepin(X1),np__1)=np__0|~r2_int_1(k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(X1))|~v7_ordinal1(X1)|~v7_ordinal1(k4_pepin(X1)))).#trainneg
cnf(i_0_380, plain, (r1_int_1(X1,np__2)|~r2_int_1(np__3,np__1,X1)|~v1_int_1(np__3)|~v1_int_1(X1))).#trainneg
cnf(i_0_377, plain, (r1_int_1(X1,np__256)|~r2_int_1(np__257,np__1,X1)|~v1_int_1(np__257)|~v1_int_1(X1))).#trainneg
cnf(i_0_468, plain, (r2_int_1(np__8,np__1,X1)|~r1_int_1(X1,np__7)|~v1_int_1(np__8)|~v1_int_1(X1))).#trainneg
cnf(i_0_5800, plain, (r2_int_1(np__3,X1,X2)|~r1_int_1(X2,k2_xcmplx_0(k4_xcmplx_0(X1),np__3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_int_1(np__3)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5035, plain, (m2_subset_1(k6_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|k7_nat_d(X1,X2)=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5618, plain, (k1_newton(np__1,np__1)=k3_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_5801, plain, (r2_int_1(np__1,X1,X2)|~r1_int_1(X2,k2_xcmplx_0(np__1,k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5056, plain, (v1_int_2(k4_pepin(X1))|~r2_int_1(k13_newton(np__3,k3_nat_d(np__0,np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(X1))|~v7_ordinal1(X1)|k6_xcmplx_0(k4_pepin(X1),np__1)!=np__0|~v7_ordinal1(k4_pepin(X1)))).#trainneg
cnf(i_0_6339, plain, (k3_xcmplx_0(k1_newton(np__1,X1),np__1)=k1_newton(np__1,k2_xcmplx_0(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5026, plain, (k7_nat_d(np__257,np__1)=np__0|~r2_int_1(k1_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~m1_subset_1(k3_nat_d(np__256,np__2),k5_numbers))).#trainneg
cnf(i_0_5802, plain, (r2_int_1(np__1,X1,X2)|~r1_int_1(X2,k2_xcmplx_0(k4_xcmplx_0(X1),np__1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5920, plain, (r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,np__128)))|~v7_ordinal1(k2_xcmplx_0(np__1,k13_newton(np__3,np__128)))|~v7_ordinal1(k13_newton(np__3,np__128)))).#trainneg
cnf(i_0_5799, plain, (r2_int_1(k4_xcmplx_0(X1),k6_xcmplx_0(np__2,np__3),X2)|~r1_int_1(X2,k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_5225, plain, (k3_nat_d(k1_newton(np__2,X1),np__2)=k3_power(np__2,k6_xcmplx_0(X1,np__1))|k7_nat_d(X1,np__1)=np__0|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5980, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,X1))=k1_newton(np__1,k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_973, plain, (k1_xreal_0(np__8,np__1)=np__7)).#trainneg
cnf(i_0_5036, plain, (v7_ordinal1(k6_xcmplx_0(X1,X2))|k7_nat_d(X1,X2)=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_finset_1(k6_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_5502, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(X1))=k6_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5119, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,np__8))=np__257|~m1_subset_1(k2_newton(np__2,np__8),k5_numbers))).#trainneg
cnf(i_0_5786, plain, (k7_nat_d(X1,k6_xcmplx_0(np__2,np__3))=np__0|r1_xxreal_0(np__0,k2_xcmplx_0(X1,np__1))|~v7_ordinal1(k6_xcmplx_0(np__2,np__3))|~v7_ordinal1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5345, plain, (k1_nat_1(k2_newton(np__2,k1_newton(k1_newton(np__2,X1),np__2)),np__1)=k4_pepin(k2_xcmplx_0(X1,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5887, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,X1))=k1_newton(np__1,k2_xcmplx_0(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5057, plain, (v1_int_2(k4_pepin(X1))|~r2_int_1(k1_newton(np__3,k3_nat_d(np__0,np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(np__0,np__2),k5_numbers)|k6_xcmplx_0(k4_pepin(X1),np__1)!=np__0|~v7_ordinal1(k4_pepin(X1)))).#trainneg
cnf(i_0_5025, plain, (v1_int_2(k4_pepin(X1))|k7_nat_d(k4_pepin(X1),np__1)=np__0|~r2_int_1(k1_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2),k5_numbers)|~v7_ordinal1(k4_pepin(X1)))).#trainneg
cnf(i_0_5955, plain, (k3_xcmplx_0(k1_newton(X1,X2),k3_xcmplx_0(X1,X1))=k1_newton(X1,k2_xcmplx_0(X2,np__2))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5244, plain, (k3_nat_d(k1_newton(np__2,X1),np__8)=k3_power(np__2,k6_xcmplx_0(X1,np__3))|k7_nat_d(X1,np__3)=np__0|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5506, plain, (r1_nat_d(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~m1_subset_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),k5_numbers)|~v1_xcmplx_0(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))).#trainneg
cnf(i_0_5097, plain, (k3_nat_d(k3_power(np__2,X1),k3_xcmplx_0(np__2,np__2))=k3_power(np__2,k7_nat_d(X1,np__2))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_5755, plain, (k1_newton(k1_newton(np__1,np__256),np__2)=k1_newton(np__1,k2_xcmplx_0(np__257,np__257)))).#trainneg
cnf(i_0_5309, plain, (k4_pepin(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))=k4_pepin(np__257)|~v1_xcmplx_0(np__2)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_5916, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5810, plain, (m1_subset_1(k1_newton(np__1,np__8),k5_numbers))).#trainneg
cnf(i_0_6044, plain, (m1_subset_1(k2_xcmplx_0(np__3,X1),k5_numbers)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1)|~v1_xreal_0(np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4425, plain, (r1_int_1(X1,k6_xcmplx_0(k4_xcmplx_0(X3),X2))|~r2_int_1(k4_xcmplx_0(X2),X3,X1)|~v1_int_1(X1)|~v1_int_1(X3)|~v1_int_1(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~m1_subset_1(k4_xcmplx_0(X3),k5_numbers))).#trainneg
cnf(i_0_4473, plain, (v1_xreal_0(k2_xcmplx_0(np__1,X1))|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3862, plain, (k1_newton(np__2,k2_xcmplx_0(np__3,np__1))=k3_xcmplx_0(np__2,np__8)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3797, plain, (m2_subset_1(k3_xcmplx_0(X1,X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5715, plain, (v7_ordinal1(k1_newton(np__1,np__256)))).#trainneg
cnf(i_0_6039, plain, (m1_subset_1(k2_xcmplx_0(np__3,X1),k5_numbers)|~v1_xcmplx_0(X1)|~v1_int_1(np__3)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3721, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__8)=k1_newton(np__2,k2_xcmplx_0(X1,np__3))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6036, plain, (v3_ordinal1(k2_xcmplx_0(X1,np__3))|~v1_xcmplx_0(X1)|~v1_int_1(np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3699, plain, (k3_xcmplx_0(np__8,k1_newton(np__2,X1))=k1_newton(np__2,k2_xcmplx_0(np__3,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4745, plain, (k3_xcmplx_0(k1_newton(X1,np__257),X1)=k1_newton(X1,k2_xcmplx_0(np__257,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4171, plain, (k1_newton(k1_newton(X1,np__257),np__2)=k1_newton(X1,k2_xcmplx_0(np__257,np__257))|~v1_xcmplx_0(k1_newton(X1,np__257))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4477, plain, (r1_int_1(X1,k2_xcmplx_0(np__1,X2))|~r2_int_1(X2,k6_xcmplx_0(np__2,np__3),X1)|~v1_int_1(X1)|~v1_int_1(k6_xcmplx_0(np__2,np__3))|~v1_int_1(X2)|~v1_xcmplx_0(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5718, plain, (m1_subset_1(k1_newton(np__1,np__256),k5_numbers))).#trainneg
cnf(i_0_5807, plain, (v7_ordinal1(k1_newton(np__1,np__8)))).#trainneg
cnf(i_0_6035, plain, (k7_nat_d(X1,k6_xcmplx_0(np__2,np__3))=np__0|r1_xxreal_0(np__0,k2_xcmplx_0(np__1,X1))|~v7_ordinal1(k6_xcmplx_0(np__2,np__3))|~v7_ordinal1(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4085, plain, (k1_newton(X1,k2_xcmplx_0(np__1,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3994, plain, (v1_xreal_0(k1_newton(np__1,X1))|~v7_ordinal1(k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4741, plain, (k3_xcmplx_0(k1_newton(np__2,np__257),np__8)=k1_newton(np__2,k2_xcmplx_0(np__257,np__3))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_4558, plain, (k3_xcmplx_0(k1_newton(np__2,np__257),np__8)=k1_newton(np__2,k2_xcmplx_0(np__3,np__257))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3978, plain, (k3_xcmplx_0(k1_newton(np__1,X1),k1_newton(np__1,X2))=k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,X2),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4243, plain, (r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(np__0,np__128))))|~v7_ordinal1(k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(np__0,np__128))))|~v7_ordinal1(k13_newton(np__3,k2_nat_1(np__0,np__128))))).#trainneg
cnf(i_0_4296, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(np__257,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5716, plain, (v1_xcmplx_0(k1_newton(np__1,np__256)))).#trainneg
cnf(i_0_5921, plain, (r1_int_1(np__257,k2_nat_1(np__1,k13_newton(np__3,np__128)))|~v7_ordinal1(k2_nat_1(np__1,k13_newton(np__3,np__128)))|~v7_ordinal1(k13_newton(np__3,np__128)))).#trainneg
cnf(i_0_5134, plain, (k3_nat_d(k1_newton(np__2,X1),k3_xcmplx_0(np__2,np__2))=k3_power(np__2,k7_nat_d(X1,np__2))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_4931, plain, (v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v1_finset_1(k2_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_1704, plain, (k7_nat_d(np__257,np__1)=np__256)).#trainneg
cnf(i_0_6867, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_7002, plain, (k6_xcmplx_0(np__1,k6_xcmplx_0(np__2,np__3))=np__2)).#trainneg
cnf(i_0_975, plain, (k1_xreal_0(np__257,np__1)=np__256)).#trainneg
cnf(i_0_3987, plain, (k1_newton(np__1,np__2)=k3_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_5248, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(np__257,np__1),np__2)))|~v1_int_1(np__257))).#trainneg
cnf(i_0_6887, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,np__1))=k3_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_6945, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,np__2))=k3_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_6597, plain, (k1_newton(np__2,np__8)=k3_xcmplx_0(np__128,np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6651, plain, (v7_ordinal1(k1_newton(np__2,np__8))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6598, plain, (k1_newton(np__2,np__8)=k3_xcmplx_0(np__2,np__128)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6574, plain, (v1_xcmplx_0(np__128)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6652, plain, (v1_xcmplx_0(k1_newton(np__2,np__8))|~v1_xcmplx_0(np__16)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6654, plain, (m1_subset_1(k1_newton(np__2,np__8),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6914, plain, (r2_int_1(np__1,k6_xcmplx_0(np__2,np__3),X1)|~r1_int_1(X1,k6_xcmplx_0(np__1,k6_xcmplx_0(np__2,np__3)))|~v1_int_1(X1))).#trainneg
cnf(i_0_6944, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__1,np__1))=k1_newton(np__1,k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_6780, plain, (r2_int_1(np__3,k6_xcmplx_0(np__2,np__3),X1)|~r1_int_1(X1,k2_xcmplx_0(np__1,np__3))|~v1_int_1(np__3)|~v1_int_1(X1))).#trainneg
cnf(i_0_7131, plain, (k1_newton(X1,k2_xcmplx_0(X2,np__2))=k1_newton(X1,k2_xcmplx_0(np__2,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6943, plain, (k1_newton(np__1,k2_xcmplx_0(X1,np__2))=k1_newton(np__1,X1)|~v1_xcmplx_0(k1_newton(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4218, plain, (r1_int_1(X1,k2_xcmplx_0(X2,np__1))|~r2_int_1(X2,k6_xcmplx_0(np__2,np__3),X1)|~v1_xcmplx_0(X2)|~v1_int_1(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_6882, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,X1))=k1_newton(np__1,X1)|~v1_xcmplx_0(k1_newton(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6650, plain, (v1_xreal_0(k1_newton(np__2,np__8))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6941, plain, (k1_newton(np__1,k2_xcmplx_0(X1,np__2))=k1_newton(np__1,k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7027, plain, (k2_xcmplx_0(np__1,k1_newton(np__2,np__8))=np__257|~m1_subset_1(k1_newton(np__2,np__8),k5_numbers))).#trainneg
cnf(i_0_6999, plain, (r2_int_1(np__1,X1,X2)|~r1_int_1(X2,k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6570, plain, (k1_newton(np__1,k1_newton(np__2,np__8))=k1_newton(np__1,np__256)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_6781, plain, (r2_int_1(np__3,X1,X2)|~r1_int_1(X2,k2_xcmplx_0(np__3,k4_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_int_1(np__3)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_4471, plain, (r2_int_1(X1,k6_xcmplx_0(np__2,np__3),X2)|~r1_int_1(X2,k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7316, plain, (m1_subset_1(k2_xcmplx_0(np__1,np__3),k5_numbers)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_7481, plain, (m1_subset_1(k2_xcmplx_0(np__1,np__3),k5_numbers)|~v1_int_1(np__3))).#trainneg
cnf(i_0_3688, plain, (~r2_int_1(k1_newton(np__3,k3_nat_d(np__256,np__2)),k6_xcmplx_0(np__2,np__3),np__257)|~m1_subset_1(k3_nat_d(np__256,np__2),k5_numbers))).#trainneg
cnf(i_0_6885, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,X1))=k1_newton(np__1,k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6884, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,X1))=k1_newton(np__1,k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6942, plain, (k1_newton(np__1,k2_xcmplx_0(X1,np__2))=k1_newton(np__1,k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7933, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,X1))=k1_newton(np__1,k2_xcmplx_0(X1,np__256))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7861, plain, (k1_newton(X1,k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))=k1_newton(X1,np__257)|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_7133, plain, (k3_xcmplx_0(k1_newton(np__1,X1),np__1)=k1_newton(np__1,k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9829, plain, (v1_int_1(np__8)|~v1_xcmplx_0(np__7)|~v1_int_1(np__7))).#trainneg
cnf(i_0_8642, plain, (k3_xcmplx_0(np__2,np__128)=np__256|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8550, plain, (k3_xcmplx_0(np__128,np__2)=np__256|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_5618, plain, (k1_newton(np__1,np__1)=np__1)).#trainneg
cnf(i_0_8641, plain, (k3_xcmplx_0(np__128,np__2)=k3_xcmplx_0(np__2,np__128)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8121, plain, (k1_xreal_0(np__1,k6_xcmplx_0(np__2,np__3))=np__2|~r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_774, plain, (k3_xcmplx_0(np__1,np__128)=np__128|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9824, plain, (v1_int_1(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8638, plain, (k1_nat_1(k3_xcmplx_0(np__2,np__128),np__1)=np__257|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8461, plain, (k1_nat_1(np__256,np__1)=np__257|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8548, plain, (k1_nat_1(k3_xcmplx_0(np__128,np__2),np__1)=np__257|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8639, plain, (k2_xcmplx_0(k3_xcmplx_0(np__2,np__128),np__1)=np__257|~v7_ordinal1(k3_xcmplx_0(np__2,np__128))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8464, plain, (v7_ordinal1(k3_xcmplx_0(np__128,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9347, plain, (~r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k3_nat_d(np__256,np__2))))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(np__257)|~m1_subset_1(k13_newton(np__3,k3_nat_d(np__256,np__2)),k5_numbers))).#trainneg
cnf(i_0_8554, plain, (v7_ordinal1(k3_xcmplx_0(np__2,np__128))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9340, plain, (r2_int_1(np__3,k6_xcmplx_0(np__2,np__3),X1)|~r1_int_1(X1,k2_xcmplx_0(np__3,np__1))|~v1_int_1(np__3)|~v1_int_1(X1))).#trainneg
cnf(i_0_8422, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__256)=k1_newton(np__2,k2_xcmplx_0(X1,np__8))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9008, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,np__257))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_8549, plain, (k2_xcmplx_0(k3_xcmplx_0(np__128,np__2),np__1)=np__257|~v7_ordinal1(k3_xcmplx_0(np__128,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9253, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__2,np__128))=np__257|~m1_subset_1(k3_xcmplx_0(np__2,np__128),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_7335, plain, (v1_xreal_0(k2_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8555, plain, (v1_xcmplx_0(k3_xcmplx_0(np__2,np__128))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9252, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__128,np__2))=np__257|~m1_subset_1(k3_xcmplx_0(np__128,np__2),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9346, plain, (v1_int_2(k4_pepin(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(np__1,k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2))))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k4_pepin(X1))|~m1_subset_1(k13_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k5_numbers))).#trainneg
cnf(i_0_8465, plain, (v1_xcmplx_0(k3_xcmplx_0(np__128,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9351, plain, (r1_int_1(X1,k2_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X2)|~v1_int_1(X1)|~v1_int_1(X2)|~r1_int_1(X1,k2_xcmplx_0(np__1,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_9535, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,np__256))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_8125, plain, (r1_int_1(X1,np__2)|~r2_int_1(np__1,k6_xcmplx_0(np__2,np__3),X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_7550, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,np__1))=k1_newton(np__1,np__256)|~v1_xcmplx_0(k1_newton(np__1,np__256)))).#trainneg
cnf(i_0_8450, plain, (k3_power(np__2,k7_nat_d(np__8,X1))=k3_nat_d(np__256,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__8)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9336, plain, (r1_int_1(X1,k2_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X2)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(X2,k5_numbers)|~r1_int_1(X1,k2_xcmplx_0(X2,np__1)))).#trainneg
cnf(i_0_6744, plain, (r2_int_1(k4_xcmplx_0(X1),k6_xcmplx_0(np__2,np__3),X2)|~r2_int_1(np__1,X1,X2)|~v1_xcmplx_0(X1)|~v1_int_1(X2)|~v1_int_1(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_8557, plain, (m1_subset_1(k3_xcmplx_0(np__2,np__128),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8381, plain, (k3_xcmplx_0(np__256,k1_newton(np__2,X1))=k1_newton(np__2,k2_xcmplx_0(X1,np__8))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6539, plain, (k1_newton(np__1,k2_xcmplx_0(k2_newton(np__2,np__8),np__1))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_8208, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),X1))=k3_xcmplx_0(k1_newton(np__1,X1),np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8423, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__256)=k1_newton(np__2,k2_xcmplx_0(np__8,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8558, plain, (k3_xcmplx_0(k3_xcmplx_0(np__2,np__128),k1_newton(np__2,X1))=k1_newton(np__2,k2_xcmplx_0(X1,np__8))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8203, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,X1),np__2))=k3_xcmplx_0(np__1,k1_newton(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6656, plain, (m2_subset_1(k1_newton(np__2,np__8),k1_numbers,k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8637, plain, (k3_nat_d(k1_newton(np__2,X1),k3_xcmplx_0(np__2,np__128))=k3_power(np__2,k7_nat_d(X1,np__8))|~r1_xxreal_0(np__8,X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6617, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__128)=k1_newton(np__2,k2_xcmplx_0(X1,np__7))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8467, plain, (m1_subset_1(k3_xcmplx_0(np__128,np__2),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6553, plain, (k1_newton(np__1,k2_xcmplx_0(k2_newton(np__2,np__8),np__2))=k3_xcmplx_0(k3_xcmplx_0(np__1,np__1),k1_newton(np__1,np__256))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_6663, plain, (v3_ordinal1(k1_newton(np__2,np__8))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_7525, plain, (v7_ordinal1(k2_xcmplx_0(X1,np__3))|~v1_finset_1(k2_xcmplx_0(X1,np__3))|~v1_xcmplx_0(X1)|~v1_int_1(np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9335, plain, (v1_int_2(k4_pepin(X1))|k6_xcmplx_0(k4_pepin(X1),np__1)!=np__0|~v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(k13_newton(np__3,k3_nat_d(np__0,np__2)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(np__0,np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(np__0,np__2)))|~v1_int_1(k4_pepin(X1)))).#trainneg
cnf(i_0_7639, plain, (r1_int_1(X1,k2_xcmplx_0(np__1,k4_xcmplx_0(X2)))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~r1_int_1(X1,k6_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9981, plain, (k1_newton(np__1,k2_xcmplx_0(X1,np__257))=k1_newton(np__1,k2_xcmplx_0(X1,np__256))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9334, plain, (k7_nat_d(k4_pepin(X1),np__1)=np__0|v1_int_2(k4_pepin(X1))|~v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k4_pepin(X1)))).#trainneg
cnf(i_0_8468, plain, (k3_xcmplx_0(k3_xcmplx_0(np__128,np__2),k1_newton(np__2,X1))=k1_newton(np__2,k2_xcmplx_0(X1,np__8))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6618, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__128)=k1_newton(np__2,k2_xcmplx_0(np__7,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7407, plain, (k1_newton(np__2,k2_xcmplx_0(np__1,np__3))=k3_xcmplx_0(np__2,np__8)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_10126, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,np__2))=k1_newton(np__1,k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_8547, plain, (k3_nat_d(k1_newton(np__2,X1),k3_xcmplx_0(np__128,np__2))=k3_power(np__2,k7_nat_d(X1,np__8))|~r1_xxreal_0(np__8,X1)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6549, plain, (k1_newton(np__1,k2_xcmplx_0(k2_newton(np__2,np__8),np__2))=k3_xcmplx_0(k1_newton(np__1,np__256),k1_newton(np__1,np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_10167, plain, (v1_int_1(np__8)|~v1_xcmplx_0(np__7))).#trainneg
cnf(i_0_10212, plain, (k3_xcmplx_0(k1_newton(np__1,np__256),np__1)=k1_newton(np__1,k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_10218, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,np__256))=k1_newton(np__1,k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_10209, plain, (k3_xcmplx_0(k1_newton(np__1,np__256),np__1)=k1_newton(np__1,k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_376, plain, (r2_int_1(np__257,np__1,X1)|~r1_int_1(X1,np__256)|~v1_int_1(X1))).#trainneg
cnf(i_0_10225, plain, (k3_xcmplx_0(k1_newton(np__1,X1),np__1)=k1_newton(np__1,k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10226, plain, (k3_xcmplx_0(k1_newton(np__1,X1),np__1)=k1_newton(np__1,k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10184, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,X1))=k1_newton(np__1,k2_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10182, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,X1))=k1_newton(np__1,k2_xcmplx_0(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10778, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__257,np__1),np__1))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_10788, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__257,np__1),np__2))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_11760, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__257,np__2),np__1))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__2)))).#trainneg
cnf(i_0_10256, plain, (k7_nat_d(np__1,k6_xcmplx_0(np__2,np__3))=np__2|~v7_ordinal1(k6_xcmplx_0(np__2,np__3))|~r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_11770, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__257,np__2),np__2))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__2)))).#trainneg
cnf(i_0_10528, plain, (k2_xcmplx_0(k3_xcmplx_0(np__2,np__128),np__1)=np__257|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_377, plain, (r1_int_1(X1,np__256)|~r2_int_1(np__257,np__1,X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_10376, plain, (r1_int_1(X1,k2_xcmplx_0(np__1,np__3))|~v1_int_1(X1)|~v1_int_1(np__3)|~r1_int_1(X1,k2_xcmplx_0(np__3,np__1)))).#trainneg
cnf(i_0_11757, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(np__257,np__2)))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__2)))).#trainneg
cnf(i_0_10775, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(np__257,np__1)))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_10790, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__257,np__1)))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_10247, plain, (~r1_int_1(np__257,k2_xcmplx_0(k1_newton(np__3,k3_nat_d(np__256,np__2)),np__1))|~v1_xcmplx_0(k1_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(k1_newton(np__3,k3_nat_d(np__256,np__2)))|~m1_subset_1(k3_nat_d(np__256,np__2),k5_numbers))).#trainneg
cnf(i_0_10890, plain, (r2_int_1(k6_xcmplx_0(np__2,np__3),k6_xcmplx_0(np__2,np__3),X1)|~r2_int_1(np__1,np__1,X1)|~v1_int_1(X1)|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers))).#trainneg
cnf(i_0_10436, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__2,np__257)))).#trainneg
cnf(i_0_10299, plain, (k1_nat_1(np__1,np__256)=np__257|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_11772, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__257,np__2)))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__2)))).#trainneg
cnf(i_0_10446, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__2))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__2,np__257)))).#trainneg
cnf(i_0_7445, plain, (k3_xcmplx_0(np__8,np__2)=k3_xcmplx_0(np__2,np__8)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_7329, plain, (r1_int_1(X1,k2_xcmplx_0(k4_xcmplx_0(X2),np__1))|~r2_int_1(np__1,X2,X1)|~v1_xcmplx_0(X2)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_7479, plain, (m2_subset_1(k1_newton(X1,np__2),k1_numbers,k5_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_9329, plain, (v1_int_2(k4_pepin(X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2),k5_numbers)|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),np__1))|~v1_xcmplx_0(k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k4_pepin(X1)))).#trainneg
cnf(i_0_7308, plain, (k4_pepin(k2_xcmplx_0(np__1,k1_newton(np__2,np__8)))=k4_pepin(np__257)|~v1_xcmplx_0(np__2)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_8455, plain, (k3_nat_d(np__256,np__2)=k3_power(np__2,np__7)|np__7=np__0|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8627, plain, (k3_nat_d(k3_xcmplx_0(np__2,np__128),k1_newton(np__2,X1))=k3_power(np__2,k7_nat_d(np__8,X1))|~r1_xxreal_0(X1,np__8)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9823, plain, (v1_int_1(k6_xcmplx_0(np__1,np__1))|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers))).#trainneg
cnf(i_0_7317, plain, (m1_subset_1(k6_xcmplx_0(np__3,X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(np__3)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7067, plain, (k1_nat_1(k2_newton(np__2,k1_newton(k1_newton(np__2,k4_xcmplx_0(X1)),np__2)),np__1)=k4_pepin(k6_xcmplx_0(k4_xcmplx_0(X1),X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_9348, plain, (k7_nat_d(k4_pepin(X1),np__1)=np__0|v1_int_2(k4_pepin(X1))|~v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(X1)|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(np__1,k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2))))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k4_pepin(X1))|~m1_subset_1(k13_newton(np__3,k3_nat_d(k6_xcmplx_0(k4_pepin(X1),np__1),np__2)),k5_numbers))).#trainneg
cnf(i_0_7482, plain, (m1_subset_1(k6_xcmplx_0(np__3,X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(np__3)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7484, plain, (m1_subset_1(k2_xcmplx_0(X1,np__3),k5_numbers)|~v1_xcmplx_0(X1)|~v1_int_1(np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_10630, plain, (r1_int_1(X1,k2_xcmplx_0(k4_xcmplx_0(X2),np__1))|~r1_int_1(X1,k6_xcmplx_0(np__1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X2))|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8537, plain, (k3_nat_d(k3_xcmplx_0(np__128,np__2),k1_newton(np__2,X1))=k3_power(np__2,k7_nat_d(np__8,X1))|~r1_xxreal_0(X1,np__8)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_7319, plain, (m1_subset_1(k2_xcmplx_0(X1,np__3),k5_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7327, plain, (r1_int_1(X1,k2_xcmplx_0(k4_xcmplx_0(X2),np__3))|~r2_int_1(np__3,X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_int_1(np__3)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_7676, plain, (k1_newton(np__2,k2_xcmplx_0(np__1,np__3))=k3_xcmplx_0(np__8,np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_7326, plain, (r1_int_1(X1,k6_xcmplx_0(np__1,X2))|~r2_int_1(k4_xcmplx_0(X2),k6_xcmplx_0(np__2,np__3),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers))).#trainneg
cnf(i_0_8946, plain, (v1_xreal_0(k1_newton(np__1,X1))|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~v1_xcmplx_0(k1_newton(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1368, plain, (k3_xcmplx_0(np__256,k7_xcmplx_0(X1,np__2))=k3_xcmplx_0(X1,np__128)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8553, plain, (v1_xreal_0(k3_xcmplx_0(np__2,np__128))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_6644, plain, (k3_power(np__2,k7_nat_d(np__7,X1))=k3_nat_d(np__128,k1_newton(np__2,X1))|~r1_xxreal_0(X1,np__7)|~v7_ordinal1(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_7174, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X1))=k1_newton(X1,k2_xcmplx_0(np__1,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10433, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(np__2,np__257)))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__2,np__257)))).#trainneg
cnf(i_0_6903, plain, (r2_int_1(np__1,np__1,X1)|~r1_int_1(X1,k2_xcmplx_0(k6_xcmplx_0(np__2,np__3),np__1))|~v1_int_1(X1)|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers))).#trainneg
cnf(i_0_7977, plain, (k3_xcmplx_0(k1_newton(np__1,np__256),k1_newton(np__1,X1))=k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,X1),np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9288, plain, (k1_newton(np__1,k2_xcmplx_0(k1_newton(np__2,np__8),np__1))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_9343, plain, (v1_int_2(k4_pepin(X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2),k5_numbers)|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(np__1,k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2))))|~v1_xcmplx_0(k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)))|~v1_int_1(k4_pepin(X1))|~m1_subset_1(k1_newton(np__3,k3_nat_d(k7_nat_d(k4_pepin(X1),np__1),np__2)),k5_numbers))).#trainneg
cnf(i_0_6552, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_newton(np__2,np__8)))=k3_xcmplx_0(k3_xcmplx_0(np__1,np__1),k1_newton(np__1,np__256))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_13840, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),np__257))=k1_newton(np__1,k2_xcmplx_0(np__256,np__2)))).#trainneg
cnf(i_0_13841, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),np__257))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_13651, plain, (v1_xreal_0(k1_newton(np__1,X1))|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13847, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__257))=k1_newton(np__1,k2_xcmplx_0(np__257,np__256)))).#trainneg
cnf(i_0_12188, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,np__8))=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_13735, plain, (k3_xcmplx_0(X1,k1_newton(X1,np__2))=k1_newton(X1,k2_xcmplx_0(np__1,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13260, plain, (r1_int_1(X1,k6_xcmplx_0(np__1,X2))|~r2_int_1(np__1,X2,X1)|~v1_xcmplx_0(X2)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers))).#trainneg
cnf(i_0_12074, plain, (k3_xcmplx_0(k1_newton(np__1,np__8),np__1)=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_13846, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__257))=k1_newton(np__1,k2_xcmplx_0(np__256,np__257)))).#trainneg
cnf(i_0_12079, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__1),np__1))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_13423, plain, (m1_subset_1(k3_power(np__2,np__7),k5_numbers)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_12077, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),np__1))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_12084, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),np__1))=np__1|~v7_ordinal1(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_13664, plain, (k3_xcmplx_0(np__256,np__1)=k3_xcmplx_0(np__2,np__128)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8379, plain, (k3_xcmplx_0(np__256,k1_newton(np__2,X1))=k1_newton(np__2,k2_xcmplx_0(np__8,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7488, plain, (v7_ordinal1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8682, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),np__2))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_13838, plain, (k3_xcmplx_0(k1_newton(k1_newton(np__1,np__256),X1),k1_newton(np__1,k2_xcmplx_0(np__257,np__257)))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8463, plain, (v1_xreal_0(k3_xcmplx_0(np__128,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_13834, plain, (k3_xcmplx_0(k1_newton(np__1,k2_xcmplx_0(np__257,np__257)),k1_newton(k1_newton(np__1,np__256),X1))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9822, plain, (v1_int_1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11443, plain, (v1_int_2(k4_pepin(X1))|k6_xcmplx_0(k4_pepin(X1),np__1)!=np__0|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(np__1,k13_newton(np__3,k3_nat_d(np__0,np__2))))|~v1_xcmplx_0(k13_newton(np__3,k3_nat_d(np__0,np__2)))|~v1_int_1(k13_newton(np__3,k3_nat_d(np__0,np__2)))|~v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(X1)|~m1_subset_1(k13_newton(np__3,k3_nat_d(np__0,np__2)),k5_numbers))).#trainneg
cnf(i_0_8245, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__1),np__2))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_13350, plain, (m1_subset_1(k6_xcmplx_0(np__3,X1),k5_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__3)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6539, plain, (k1_newton(np__1,k2_xcmplx_0(k2_newton(np__2,np__8),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_8851, plain, (k3_nat_d(k1_newton(np__2,k2_xcmplx_0(np__2,X1)),k3_power(np__2,X2))=k3_power(np__2,k7_nat_d(k2_xcmplx_0(X1,np__2),X2))|~r1_xxreal_0(X2,k2_xcmplx_0(X1,np__2))|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X2)|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7328, plain, (r1_int_1(X1,k2_xcmplx_0(np__1,k4_xcmplx_0(X2)))|~r2_int_1(np__1,X2,X1)|~v1_xcmplx_0(X2)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_8152, plain, (k7_xcmplx_0(k1_newton(np__1,k2_xcmplx_0(np__2,np__1)),X1)=k3_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11255, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,k4_xcmplx_0(X1)))=k1_newton(np__1,k2_xcmplx_0(k6_xcmplx_0(np__1,X1),np__2))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10448, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__2,np__257)))=k3_xcmplx_0(k1_newton(np__1,np__256),np__1)|~v7_ordinal1(k2_xcmplx_0(np__2,np__257)))).#trainneg
cnf(i_0_8018, plain, (k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(np__257,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8852, plain, (k3_nat_d(k1_newton(np__2,k2_xcmplx_0(np__2,X1)),k1_newton(np__2,X2))=k3_power(np__2,k7_nat_d(k2_xcmplx_0(X1,np__2),X2))|~r1_xxreal_0(X2,k2_xcmplx_0(X1,np__2))|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X2)|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9532, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_newton(np__2,k1_newton(np__2,np__2))))=k1_newton(np__1,k4_pepin(np__2))|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,np__2)))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8599, plain, (k3_xcmplx_0(k1_newton(np__2,X1),k3_xcmplx_0(np__2,np__128))=k1_newton(np__2,k2_xcmplx_0(X1,np__8))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8009, plain, (k3_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X2)),k1_newton(k1_newton(X1,X2),np__257))=k1_newton(k1_newton(X1,X2),k2_xcmplx_0(np__257,np__2))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8330, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),np__2))=np__1|~v7_ordinal1(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_9898, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,X1))=k1_newton(np__1,k2_xcmplx_0(np__256,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13843, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,np__257))=k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))).#trainneg
cnf(i_0_13590, plain, (r1_int_1(X1,k6_xcmplx_0(np__1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~m1_subset_1(k4_xcmplx_0(X2),k5_numbers)|~r1_int_1(X1,k2_xcmplx_0(np__1,k4_xcmplx_0(X2))))).#trainneg
cnf(i_0_6977, plain, (m1_subset_1(k6_xcmplx_0(np__1,X1),k5_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10474, plain, (k3_xcmplx_0(k1_newton(np__1,X1),k1_newton(np__1,np__256))=k1_newton(np__1,k2_xcmplx_0(X1,k2_xcmplx_0(np__257,np__2)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13351, plain, (m1_subset_1(k6_xcmplx_0(np__3,X1),k5_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__3)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_3899, plain, (k3_xcmplx_0(k1_newton(np__2,X1),np__16)=k1_newton(np__2,k2_xcmplx_0(np__4,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10374, plain, (~r1_int_1(np__257,k2_xcmplx_0(np__1,k1_newton(np__3,k3_nat_d(np__256,np__2))))|~v1_xcmplx_0(k1_newton(np__3,k3_nat_d(np__256,np__2)))|~v1_int_1(k1_newton(np__3,k3_nat_d(np__256,np__2)))|~m1_subset_1(k1_newton(np__3,k3_nat_d(np__256,np__2)),k5_numbers)|~m1_subset_1(k3_nat_d(np__256,np__2),k5_numbers))).#trainneg
cnf(i_0_9037, plain, (r1_int_1(X1,k2_xcmplx_0(np__3,np__1))|~v1_int_1(X1)|~v1_int_1(np__3)|~r1_int_1(X1,k2_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_3966, plain, (k1_newton(np__1,k4_xcmplx_0(X1))=k1_newton(np__1,k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5163, plain, (k2_xcmplx_0(k6_xcmplx_0(np__2,np__3),X1)=k6_xcmplx_0(X1,np__1)|~m1_subset_1(k6_xcmplx_0(np__2,np__3),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7979, plain, (r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_xcmplx_0(np__128,np__0))))|~v7_ordinal1(k2_xcmplx_0(np__1,k13_newton(np__3,k2_xcmplx_0(np__128,np__0))))|~v7_ordinal1(k13_newton(np__3,k2_xcmplx_0(np__128,np__0))))).#trainneg
cnf(i_0_4180, plain, (k3_xcmplx_0(X1,k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(np__1,np__257))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5047, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13847, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__257))=k1_newton(np__1,k2_xcmplx_0(np__256,np__257)))).#trainneg
cnf(i_0_13840, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),np__257))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_14115, plain, (~r2_int_1(k13_newton(np__3,k3_power(np__2,np__7)),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_16056, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,k2_xcmplx_0(np__257,np__2)))=k1_newton(np__1,k2_xcmplx_0(np__256,np__257)))).#trainneg
cnf(i_0_15425, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__2))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_14460, plain, (k1_newton(np__1,k2_xcmplx_0(np__8,np__2))=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_14457, plain, (k1_newton(np__1,k2_xcmplx_0(np__8,np__1))=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_14459, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,np__8))=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_14458, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,np__8))=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_14580, plain, (k1_newton(np__1,k2_xcmplx_0(np__256,np__257))=k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))).#trainneg
cnf(i_0_5755, plain, (k1_newton(k1_newton(np__1,np__256),np__2)=k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))).#trainneg
cnf(i_0_14114, plain, (~r2_int_1(k1_newton(np__3,k3_power(np__2,np__7)),k6_xcmplx_0(np__2,np__3),np__257)|~m1_subset_1(k3_power(np__2,np__7),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_15878, plain, (k1_newton(np__1,k2_xcmplx_0(np__256,X1))=k1_newton(np__1,k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16049, plain, (k3_xcmplx_0(k1_newton(k1_newton(np__1,np__256),X1),k1_newton(np__1,k2_xcmplx_0(np__256,k2_xcmplx_0(np__257,np__2))))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16133, plain, (k3_xcmplx_0(k1_newton(np__1,np__256),k1_newton(np__1,np__256))=k1_newton(np__1,k2_xcmplx_0(np__257,k2_xcmplx_0(np__257,np__2))))).#trainneg
cnf(i_0_15214, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_14868, plain, (k3_xcmplx_0(np__2,np__128)=k3_xcmplx_0(np__1,np__256)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_14109, plain, (m1_subset_1(k3_power(np__2,np__7),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_16047, plain, (k3_xcmplx_0(k1_newton(k1_newton(np__1,np__256),X1),k1_newton(np__1,k2_xcmplx_0(np__256,k2_xcmplx_0(np__257,np__2))))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16064, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(np__257,np__2)))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_15921, plain, (k1_newton(np__1,k2_xcmplx_0(np__256,k6_xcmplx_0(np__2,np__3)))=k1_newton(np__1,np__256)|~v7_ordinal1(k6_xcmplx_0(np__2,np__3))|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_16051, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__257,np__2)))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_15213, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__257))=k3_xcmplx_0(k1_newton(np__1,np__256),k1_newton(np__1,np__256))|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_16043, plain, (k3_xcmplx_0(k1_newton(np__1,k2_xcmplx_0(np__256,k2_xcmplx_0(np__257,np__2))),k1_newton(k1_newton(np__1,np__256),X1))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15209, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_16045, plain, (k3_xcmplx_0(k1_newton(np__1,k2_xcmplx_0(np__256,k2_xcmplx_0(np__257,np__2))),k1_newton(k1_newton(np__1,np__256),X1))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(X1,np__2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15212, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__256))=k3_xcmplx_0(k1_newton(np__1,np__256),k1_newton(np__1,np__256))|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_14910, plain, (v7_ordinal1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15075, plain, (v1_int_1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_16039, plain, (k1_newton(np__1,k2_xcmplx_0(np__256,k2_xcmplx_0(np__257,np__2)))=k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))).#trainneg
cnf(i_0_15208, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_newton(np__2,np__8),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_14911, plain, (v7_ordinal1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_14909, plain, (v7_ordinal1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_15074, plain, (v1_int_1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_15217, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(k2_newton(np__2,np__8),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_15216, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,k2_xcmplx_0(k2_newton(np__2,np__8),np__1)))=k3_xcmplx_0(k1_newton(np__1,np__256),k1_newton(np__1,np__256))|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_14117, plain, (~r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k3_power(np__2,np__7))))|~v1_xcmplx_0(k13_newton(np__3,k3_power(np__2,np__7)))|~v1_int_1(k13_newton(np__3,k3_power(np__2,np__7)))|~m1_subset_1(k13_newton(np__3,k3_power(np__2,np__7)),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_14118, plain, (~r1_int_1(np__257,k2_xcmplx_0(k1_newton(np__3,k3_power(np__2,np__7)),np__1))|~v1_xcmplx_0(k1_newton(np__3,k3_power(np__2,np__7)))|~v1_int_1(k1_newton(np__3,k3_power(np__2,np__7)))|~m1_subset_1(k3_power(np__2,np__7),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_16031, plain, (m1_subset_1(k6_xcmplx_0(np__1,X1),k5_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14678, plain, (m1_subset_1(k1_newton(np__2,np__7),k5_numbers)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_16221, plain, (k3_xcmplx_0(k1_newton(np__1,X1),k1_newton(np__1,k6_xcmplx_0(np__1,X2)))=k1_newton(np__1,k2_xcmplx_0(X1,k4_xcmplx_0(X2)))|~v7_ordinal1(k4_xcmplx_0(X2))|~v7_ordinal1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_16030, plain, (m1_subset_1(k6_xcmplx_0(np__1,X1),k5_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_14881, plain, (k1_newton(np__2,k2_xcmplx_0(np__8,np__1))=k3_xcmplx_0(np__256,np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_18671, plain, (m1_subset_1(k1_newton(np__2,np__7),k5_numbers))).#trainneg
cnf(i_0_16556, plain, (~r2_int_1(k13_newton(np__3,k1_newton(np__2,np__7)),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_17631, plain, (k3_xcmplx_0(np__1,np__256)=np__256|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_17835, plain, (k1_newton(np__1,k6_xcmplx_0(np__256,np__1))=k1_newton(np__1,np__256)|~v1_xcmplx_0(np__257)|~v7_ordinal1(k6_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_13846, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,np__256))=k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))).#trainneg
cnf(i_0_13847, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__257))=k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))).#trainneg
cnf(i_0_17261, plain, (~r2_int_1(k1_newton(np__3,k1_newton(np__2,np__7)),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xcmplx_0(np__2)|~m1_subset_1(k1_newton(np__2,np__7),k5_numbers))).#trainneg
cnf(i_0_17189, plain, (v7_ordinal1(k1_newton(np__1,k2_xcmplx_0(np__256,np__256))))).#trainneg
cnf(i_0_16056, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,k2_xcmplx_0(np__257,np__2)))=k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))).#trainneg
cnf(i_0_18141, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k1_newton(np__2,np__8),np__1),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k1_newton(np__2,np__8),np__1))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_17606, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k1_newton(np__2,np__8),np__1),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k1_newton(np__2,np__8),np__1))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_17190, plain, (v1_xcmplx_0(k1_newton(np__1,k2_xcmplx_0(np__256,np__256))))).#trainneg
cnf(i_0_18660, plain, (~r1_int_1(np__257,k2_xcmplx_0(k1_newton(np__3,k1_newton(np__2,np__7)),np__1))|~v1_xcmplx_0(k1_newton(np__3,k1_newton(np__2,np__7)))|~v1_xcmplx_0(np__2)|~v1_int_1(k1_newton(np__3,k1_newton(np__2,np__7))))).#trainneg
cnf(i_0_17236, plain, (k3_xcmplx_0(k1_newton(k1_newton(np__1,np__256),X1),k1_newton(np__1,k2_xcmplx_0(np__256,np__256)))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18454, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k1_newton(np__2,np__8),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k1_newton(np__2,np__8),np__1))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_17633, plain, (k2_xcmplx_0(k3_xcmplx_0(np__1,np__256),np__1)=np__257|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_17192, plain, (m1_subset_1(k1_newton(np__1,k2_xcmplx_0(np__256,np__256)),k5_numbers))).#trainneg
cnf(i_0_18658, plain, (~r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k1_newton(np__2,np__7))))|~v1_xcmplx_0(k13_newton(np__3,k1_newton(np__2,np__7)))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,k1_newton(np__2,np__7)))|~m1_subset_1(k13_newton(np__3,k1_newton(np__2,np__7)),k5_numbers))).#trainneg
cnf(i_0_17748, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(np__2,np__257)))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_18133, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__1))|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_18562, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(k1_newton(np__2,np__8),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k1_newton(np__2,np__8),np__1))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_17923, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__2,np__257)))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_18467, plain, (v7_ordinal1(k6_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_17325, plain, (k1_newton(np__1,k2_xcmplx_0(np__256,k4_xcmplx_0(X1)))=k1_newton(np__1,k6_xcmplx_0(np__257,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__257)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_17623, plain, (v7_ordinal1(k3_xcmplx_0(np__1,np__256))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_18017, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k1_newton(np__2,np__8),np__1),np__257))=k1_newton(np__1,k2_xcmplx_0(np__256,np__256))|~v7_ordinal1(k2_xcmplx_0(k1_newton(np__2,np__8),np__1))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_17597, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__2)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__1),np__2))|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_16492, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),np__257),np__257))=k3_xcmplx_0(k1_newton(np__1,np__256),k1_newton(np__1,np__256))|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),np__257)))).#trainneg
cnf(i_0_16251, plain, (k1_newton(np__1,k6_xcmplx_0(np__2,np__3))=k1_newton(np__1,k6_xcmplx_0(np__1,np__1))|~v7_ordinal1(k6_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_16058, plain, (k1_newton(np__1,k2_xcmplx_0(X1,k2_xcmplx_0(np__257,np__2)))=k1_newton(np__1,k2_xcmplx_0(X1,np__256))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17191, plain, (k3_xcmplx_0(k1_newton(np__1,k2_xcmplx_0(np__256,np__256)),k1_newton(k1_newton(np__1,np__256),X1))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16164, plain, (k1_newton(np__2,np__7)=k3_xcmplx_0(np__8,np__16)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_11259, plain, (k3_xcmplx_0(np__1,k1_newton(np__1,k2_newton(np__2,k2_newton(np__2,X1))))=k1_newton(np__1,k2_xcmplx_0(k4_pepin(X1),np__2))|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_13407, plain, (m1_subset_1(k6_xcmplx_0(np__3,X1),k5_numbers)|~v1_xcmplx_0(X1)|~v1_int_1(np__3)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5952, plain, (r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0)))))|~v7_ordinal1(k2_xcmplx_0(np__1,k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0)))))|~v7_ordinal1(k13_newton(np__3,k2_xcmplx_0(np__128,k4_nat_1(np__32,np__0))))|~v7_ordinal1(k4_nat_1(np__32,np__0)))).#trainneg
cnf(i_0_7306, plain, (v1_int_2(k4_pepin(np__257))|k6_xcmplx_0(k4_pepin(np__257),np__1)!=np__0|~r2_int_1(k13_newton(np__3,k3_nat_d(np__0,np__2)),k6_xcmplx_0(np__2,np__3),k4_pepin(np__257))|~v7_ordinal1(k4_pepin(np__257))|~v7_ordinal1(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))|~v1_xcmplx_0(np__2)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_6577, plain, (k3_xcmplx_0(np__128,k1_newton(np__2,X1))=k1_newton(np__2,k2_xcmplx_0(X1,np__7))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9531, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_newton(np__2,k1_newton(np__2,X1))))=k1_newton(np__1,k4_pepin(X1))|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4351, plain, (r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k2_xcmplx_0(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),np__1))|~v7_ordinal1(k4_nat_1(np__32,np__0))|~m1_subset_1(k13_newton(np__3,k2_nat_1(np__128,k4_nat_1(np__32,np__0))),k5_numbers))).#trainneg
cnf(i_0_11442, plain, (v1_int_2(k4_pepin(X1))|k6_xcmplx_0(k4_pepin(X1),np__1)!=np__0|~r1_int_1(k4_pepin(X1),k2_xcmplx_0(k1_newton(np__3,k3_nat_d(np__0,np__2)),np__1))|~v1_xcmplx_0(k1_newton(np__3,k3_nat_d(np__0,np__2)))|~v1_int_1(k1_newton(np__3,k3_nat_d(np__0,np__2)))|~v7_ordinal1(k4_pepin(X1))|~v7_ordinal1(X1)|~m1_subset_1(k3_nat_d(np__0,np__2),k5_numbers))).#trainneg
cnf(i_0_7166, plain, (k1_newton(np__1,k2_xcmplx_0(k2_newton(np__2,np__8),np__2))=k3_xcmplx_0(k1_newton(np__1,np__256),k3_xcmplx_0(np__1,np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_5808, plain, (v1_xcmplx_0(k1_newton(np__1,np__8)))).#trainneg
cnf(i_0_9039, plain, (v1_xreal_0(k1_newton(np__1,X1))|~v7_ordinal1(k2_xcmplx_0(np__2,X1))|~v1_xcmplx_0(k1_newton(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7204, plain, (r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v7_ordinal1(k2_xcmplx_0(np__1,k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128))))|~v1_xcmplx_0(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)))|~m1_subset_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),k5_numbers))).#trainneg
cnf(i_0_15666, plain, (k3_xcmplx_0(k1_newton(np__1,k2_xcmplx_0(np__257,np__257)),k1_newton(k1_newton(np__1,np__256),np__257))=k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(np__257,np__2)))).#trainneg
cnf(i_0_16032, plain, (m1_subset_1(k6_xcmplx_0(np__1,X1),k5_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_8509, plain, (k3_xcmplx_0(k1_newton(np__2,X1),k3_xcmplx_0(np__128,np__2))=k1_newton(np__2,k2_xcmplx_0(X1,np__8))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11518, plain, (k1_newton(np__1,k2_xcmplx_0(k4_xcmplx_0(X1),np__256))=k1_newton(np__1,k6_xcmplx_0(np__257,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__257)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_13386, plain, (k1_nat_1(k2_newton(np__2,k1_newton(k1_newton(np__2,np__1),np__2)),np__1)=k4_pepin(np__2)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_8863, plain, (k1_nat_1(k2_newton(np__2,k1_newton(np__2,k2_xcmplx_0(np__2,X1))),np__1)=k4_pepin(k2_xcmplx_0(np__2,X1))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4179, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(np__2,np__257))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5417, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(k1_newton(X1,X3),np__2))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X3,X3),X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_20746, plain, (k3_xcmplx_0(np__8,np__16)=np__128|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_18925, plain, (~r2_int_1(k13_newton(np__3,np__128),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_20751, plain, (~r2_int_1(k13_newton(np__3,k3_xcmplx_0(np__8,np__16)),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_19202, plain, (~r2_int_1(k1_newton(np__3,np__128),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_20140, plain, (k1_newton(np__1,k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)))=k1_newton(np__1,np__256)|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_21625, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(k3_xcmplx_0(X1,X1),np__2))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_19681, plain, (~r1_int_1(np__257,k2_xcmplx_0(np__1,k13_newton(np__3,np__128)))|~v1_xcmplx_0(k13_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,np__128))|~m1_subset_1(k13_newton(np__3,np__128),k5_numbers))).#trainneg
cnf(i_0_20671, plain, (k3_xcmplx_0(k3_xcmplx_0(np__8,np__16),np__2)=k1_newton(np__2,np__8)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_21339, plain, (k4_pepin(k2_xcmplx_0(np__2,X1))=k4_pepin(k2_xcmplx_0(X1,np__2))|~v1_xcmplx_0(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21151, plain, (r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)))|~v1_int_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)))|~m1_subset_1(k13_newton(np__3,k2_nat_1(k4_nat_1(np__32,np__0),np__128)),k5_numbers))).#trainneg
cnf(i_0_21034, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_21036, plain, (k1_newton(np__1,k2_xcmplx_0(k1_newton(np__2,np__8),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_21752, plain, (k3_xcmplx_0(X1,k1_newton(k1_newton(X1,X2),np__2))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X2),np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_19634, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__1,np__256))=np__257|~m1_subset_1(k3_xcmplx_0(np__1,np__256),k5_numbers)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_20672, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__8,np__16))=k1_newton(np__2,np__8)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_19857, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_xcmplx_0(k1_newton(np__2,np__8),np__1),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(k1_newton(np__2,np__8),np__1),np__1))|~v7_ordinal1(k2_xcmplx_0(k1_newton(np__2,np__8),np__1))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_20973, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__2),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__2))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_21544, plain, (k3_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X2)),k1_newton(k1_newton(X1,X2),np__2))=k1_newton(k1_newton(X1,k2_xcmplx_0(X2,X2)),np__2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20968, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__2),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__2))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_20333, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_xcmplx_0(k1_newton(np__2,np__8),np__1),np__2)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(k1_newton(np__2,np__8),np__1),np__2))|~v7_ordinal1(k2_xcmplx_0(k1_newton(np__2,np__8),np__1))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_21643, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,k2_xcmplx_0(X3,X3)))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X3,X3),X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_20967, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_newton(np__2,np__8),np__2)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__2))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_20749, plain, (~r2_int_1(k1_newton(np__3,k3_xcmplx_0(np__8,np__16)),k6_xcmplx_0(np__2,np__3),np__257)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_20976, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(k2_newton(np__2,np__8),np__2)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__2))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_21751, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(k1_newton(X1,X2),np__2))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X2),np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21395, plain, (k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(X1,np__257))=k1_newton(X1,k2_xcmplx_0(np__2,np__257))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21556, plain, (k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X2),X2))=k1_newton(k1_newton(X1,X2),k2_xcmplx_0(np__2,np__1))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21557, plain, (k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X2),X2))=k1_newton(k1_newton(X1,X2),k2_xcmplx_0(np__1,np__2))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21924, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),np__1))=np__1)).#trainneg
cnf(i_0_21922, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),np__2))=np__1)).#trainneg
cnf(i_0_23137, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__256))=k1_newton(X1,k2_xcmplx_0(np__256,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_23134, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,np__8))=k1_newton(X1,k2_xcmplx_0(np__8,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22012, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),np__257))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_23774, plain, (k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(k1_newton(X1,X2),np__2))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X2),np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_23679, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k2_xcmplx_0(np__2,np__2)),np__2))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_21907, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(k1_newton(X1,np__2),np__2))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_24266, plain, (k1_newton(k1_newton(X1,X2),k2_xcmplx_0(np__1,np__2))=k1_newton(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X2,X2)))|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22184, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k1_newton(np__2,np__8)))=k1_newton(np__1,np__256)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_22426, plain, (k1_newton(X1,k2_xcmplx_0(np__1,k2_xcmplx_0(X2,X2)))=k3_xcmplx_0(X1,k1_newton(k1_newton(X1,X2),np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_24545, plain, (k1_newton(X1,k2_xcmplx_0(np__256,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__256))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25354, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k2_newton(np__2,np__8)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_23780, plain, (k1_newton(X1,k2_xcmplx_0(np__257,np__2))=k1_newton(X1,k2_xcmplx_0(np__2,np__257))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24009, plain, (k1_newton(k1_newton(X1,X2),k2_xcmplx_0(np__2,np__1))=k1_newton(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X2,X2)))|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25158, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),k2_xcmplx_0(np__1,np__2)),np__2))=np__1|~v7_ordinal1(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_22866, plain, (k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(k1_newton(X1,np__1),np__2))=k1_newton(k1_newton(X1,np__2),np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25148, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(np__2,np__1),k2_xcmplx_0(np__2,np__1)),np__2))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_22119, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_25356, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),np__7))=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_25359, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k1_newton(np__2,np__8)))=k1_newton(np__1,np__256)|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_25361, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k2_xcmplx_0(k2_newton(np__2,np__8),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__1))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_22124, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_21812, plain, (k1_newton(np__1,k2_xcmplx_0(k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)))|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_22868, plain, (k3_xcmplx_0(k1_newton(X1,np__256),k1_newton(k1_newton(X1,np__128),np__2))=k1_newton(k1_newton(X1,np__256),np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24038, plain, (k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(np__2,np__1))=k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__257,np__257),np__257)))).#trainneg
cnf(i_0_24713, plain, (k1_newton(X1,k2_xcmplx_0(np__8,X2))=k1_newton(X1,k2_xcmplx_0(X2,np__8))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22465, plain, (k3_xcmplx_0(X1,k1_newton(X1,np__2))=k1_newton(X1,k2_xcmplx_0(np__2,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21817, plain, (k1_newton(np__1,k2_xcmplx_0(k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)))|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_24293, plain, (k1_newton(k1_newton(np__1,np__256),k2_xcmplx_0(np__1,np__2))=k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__257,np__257),np__257)))).#trainneg
cnf(i_0_22990, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__2),np__1)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__2),np__1))|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__2))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_25454, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3))))=k1_newton(np__1,np__256)|~v7_ordinal1(k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)))|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_22483, plain, (k3_xcmplx_0(X1,k1_newton(X1,k2_xcmplx_0(X2,X2)))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X2),np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25340, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,k2_xcmplx_0(np__2,np__2)))=k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_26044, plain, (~v1_xcmplx_0(k1_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~m1_subset_1(k1_newton(np__3,np__128),k5_numbers))).#trainneg
cnf(i_0_28602, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),k2_xcmplx_0(np__1,np__2)),np__1))=np__1|~v7_ordinal1(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_22865, plain, (k3_xcmplx_0(k1_newton(X1,np__8),k1_newton(k1_newton(X1,np__4),np__2))=k1_newton(k1_newton(X1,np__8),np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22708, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__2),np__2)))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(k2_newton(np__2,np__8),np__2),np__2))|~v7_ordinal1(k2_xcmplx_0(k2_newton(np__2,np__8),np__2))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_25837, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k3_xcmplx_0(np__2,np__128)))=k1_newton(np__1,np__256)|~v7_ordinal1(k3_xcmplx_0(np__2,np__128))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_22118, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(np__1,k2_newton(np__2,np__8))))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_22127, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__1,k2_newton(np__2,np__8))))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(np__1,k2_newton(np__2,np__8)))|~v7_ordinal1(k2_newton(np__2,np__8)))).#trainneg
cnf(i_0_20748, plain, (~r1_int_1(np__257,k2_xcmplx_0(k13_newton(np__3,k3_xcmplx_0(np__8,np__16)),np__1))|~v1_xcmplx_0(k13_newton(np__3,k3_xcmplx_0(np__8,np__16)))|~v1_xcmplx_0(np__2)|~v1_int_1(k13_newton(np__3,k3_xcmplx_0(np__8,np__16))))).#trainneg
cnf(i_0_21301, plain, (k1_newton(np__1,k2_xcmplx_0(k6_xcmplx_0(np__2,np__3),np__256))=k1_newton(np__1,np__256)|~v1_xcmplx_0(np__257)|~v7_ordinal1(k6_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_23675, plain, (k3_xcmplx_0(np__1,k1_newton(k1_newton(np__1,np__8),np__2))=k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__7,np__7),np__2)))).#trainneg
cnf(i_0_21385, plain, (k3_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X2)),k1_newton(k1_newton(X1,X2),np__257))=k1_newton(k1_newton(X1,X2),k2_xcmplx_0(np__2,np__257))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_28183, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X1))=k1_newton(X1,k2_xcmplx_0(np__2,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27008, plain, (k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(X1,np__2))=k1_newton(k1_newton(X1,np__2),np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19508, plain, (~r1_int_1(np__257,k2_xcmplx_0(k1_newton(np__3,np__128),np__1))|~v1_xcmplx_0(k1_newton(np__3,np__128))|~v1_xcmplx_0(np__2)|~v1_int_1(k1_newton(np__3,np__128)))).#trainneg
cnf(i_0_27744, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__1))=k1_newton(np__1,np__256)|~v1_xcmplx_0(np__257)|~v7_ordinal1(k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_25113, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(k1_newton(X1,X2),np__2))=k3_xcmplx_0(k1_newton(X1,np__2),k1_newton(k1_newton(X1,X2),np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_24544, plain, (k1_newton(X1,k2_xcmplx_0(np__257,np__256))=k1_newton(X1,k2_xcmplx_0(np__256,np__257))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25313, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k2_xcmplx_0(np__2,np__2)),np__1))=np__1|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k2_xcmplx_0(np__2,np__2)))|~v7_ordinal1(k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_27006, plain, (k1_newton(k1_newton(X1,np__2),np__2)=k1_newton(X1,k2_xcmplx_0(np__2,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22458, plain, (k3_xcmplx_0(np__1,k1_newton(k1_newton(np__1,np__8),np__2))=k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__7,np__7),np__1)))).#trainneg
cnf(i_0_25365, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k2_xcmplx_0(np__2,np__2)))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_25767, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,k1_newton(np__2,np__8)),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(np__1,k1_newton(np__2,np__8)))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_25772, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,k1_newton(np__2,np__8)),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(np__1,k1_newton(np__2,np__8)))|~v7_ordinal1(k1_newton(np__2,np__8)))).#trainneg
cnf(i_0_19391, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(np__128,np__2),np__1),np__1))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(np__128,np__2),np__1))|~v7_ordinal1(k3_xcmplx_0(np__128,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_23141, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,k6_xcmplx_0(k4_xcmplx_0(X3),X3)))=k1_newton(X1,k2_xcmplx_0(k6_xcmplx_0(k4_xcmplx_0(X3),X3),X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(k4_xcmplx_0(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(X3)))).#trainneg
cnf(i_0_23718, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(X1,np__2))=k1_newton(X1,k2_xcmplx_0(np__2,np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21811, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3))))=k1_newton(np__1,np__256)|~v7_ordinal1(k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)))|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_25836, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k3_xcmplx_0(np__128,np__2)))=k1_newton(np__1,np__256)|~v7_ordinal1(k3_xcmplx_0(np__128,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_21346, plain, (k1_nat_1(k2_newton(np__2,k1_newton(np__2,k2_xcmplx_0(np__1,np__2))),np__1)=k4_pepin(k2_xcmplx_0(np__2,np__1))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_19495, plain, (k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(np__128,np__2),np__1),np__2))=k1_newton(np__1,np__256)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(np__128,np__2),np__1))|~v7_ordinal1(k3_xcmplx_0(np__128,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_26886, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__2,np__2)))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_21820, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3))))=k1_newton(np__1,np__256)|~v7_ordinal1(k6_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)))|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_21147, plain, (v1_xreal_0(k1_newton(np__1,X1))|~v7_ordinal1(k2_xcmplx_0(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26621, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(X1,k2_xcmplx_0(np__2,np__257)))=k1_newton(X1,k2_xcmplx_0(X2,k2_xcmplx_0(np__2,np__257)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_23133, plain, (k1_newton(X1,k2_xcmplx_0(k2_xcmplx_0(X3,X3),X2))=k1_newton(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_28841, plain, (~v1_xcmplx_0(np__2)|~m1_subset_1(k1_newton(np__3,np__128),k5_numbers))).#trainneg
cnf(i_0_32543, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(np__2,np__2)))=np__1)).#trainneg
cnf(i_0_32547, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,k2_xcmplx_0(np__2,np__2)))=k1_newton(np__1,np__256))).#trainneg
cnf(i_0_32537, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(np__2,np__2)))=np__1)).#trainneg
cnf(i_0_32541, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),k2_xcmplx_0(np__1,np__2))))=np__1|~v7_ordinal1(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_29463, plain, (k1_newton(np__1,k2_xcmplx_0(np__257,k6_xcmplx_0(np__2,np__3)))=k1_newton(np__1,np__256)|~v7_ordinal1(k6_xcmplx_0(np__2,np__3))|~v1_xcmplx_0(np__257))).#trainneg
cnf(i_0_32545, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k2_xcmplx_0(k2_xcmplx_0(np__1,np__2),k2_xcmplx_0(np__1,np__2))))=np__1|~v7_ordinal1(k2_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_32539, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(k2_xcmplx_0(np__2,np__2),k2_xcmplx_0(np__2,np__2))))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__2)))).#trainneg
cnf(i_0_32542, plain, (k1_newton(np__1,k2_xcmplx_0(np__7,k2_xcmplx_0(np__2,np__2)))=k1_newton(np__1,np__8))).#trainneg
cnf(i_0_32647, plain, (k3_xcmplx_0(k1_newton(X1,X2),k1_newton(k1_newton(X1,X3),np__2))=k1_newton(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X3)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_32540, plain, (k1_newton(np__1,k2_xcmplx_0(np__2,k2_xcmplx_0(k2_xcmplx_0(np__2,np__1),k2_xcmplx_0(np__2,np__1))))=np__1|~v7_ordinal1(k2_xcmplx_0(np__2,np__1)))).#trainneg
cnf(i_0_29778, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k3_xcmplx_0(X1,X1))=k1_newton(k3_xcmplx_0(X1,X1),np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_29158, plain, (k1_newton(np__1,k2_xcmplx_0(np__1,k3_xcmplx_0(np__1,np__256)))=k1_newton(np__1,np__256)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_32644, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X1),k1_newton(k1_newton(X1,X2),np__2))=k1_newton(X1,k2_xcmplx_0(np__2,k2_xcmplx_0(X2,X2)))|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_23675, plain, (k3_xcmplx_0(np__1,k1_newton(k1_newton(np__1,np__8),np__2))=k1_newton(np__1,k2_xcmplx_0(k2_xcmplx_0(np__7,np__7),np__1)))).#trainneg
cnf(i_0_30538, plain, (k1_newton(k3_xcmplx_0(X1,X1),np__2)=k1_newton(X1,k2_xcmplx_0(np__2,np__2))|~v1_xcmplx_0(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 104
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 164
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 164
# Processed clauses                    : 2312
# ...of these trivial                  : 149
# ...subsumed                          : 988
# ...remaining for further processing  : 1175
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 49
# Backward-rewritten                   : 165
# Generated clauses                    : 34432
# ...of the previous two non-trivial   : 17183
# ...frozen by parental guidance       : 15672
# ...of these subsequently unfrozen    : 0
# Contextual simplify-reflections      : 112
# Paramodulations                      : 34430
# Factorizations                       : 2
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 961
#    Positive orientable unit clauses  : 239
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 713
# Current number of unprocessed clauses: 14603
# ...number of literals in the above   : 57926
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1389
# Proof object given clauses           : 203
# Proof search given clauses           : 1175
# Clause-clause subsumption calls (NU) : 93818
# Rec. Clause-clause subsumption calls : 46530
# Non-unit clause-clause subsumptions  : 1124
# Unit Clause-clause subsumption calls : 4207
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 167
# BW rewrite match successes           : 59
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 886204

# -------------------------------------------------
# User time                : 13.613 s
# System time              : 0.435 s
# Total time               : 14.049 s
# Maximum resident set size: 212588 pages
13.61user 0.44system 0:43.12elapsed 32%CPU (0avgtext+0avgdata 212588maxresident)k
24inputs+496outputs (0major+114572minor)pagefaults 0swaps
