# ENIGMA: LightGBM model 'Enigma/mizar40-all-T10/mega-VHSLCAXPh2e15/lgb-t150-d60-l3600-e0.15/model.lgb' loaded. (hash_base: 32768; conj_feats: 84; version: 991; iters: 150)
# Preprocessing time       : 0.668 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_l2_normsp_0, axiom, ![X1]:(l2_normsp_0(X1)=>(l1_normsp_0(X1)&l2_struct_0(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', dt_l2_normsp_0)).
fof(dt_l2_clvect_1, axiom, ![X1]:(l2_clvect_1(X1)=>(l1_clvect_1(X1)&l2_normsp_0(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', dt_l2_clvect_1)).
fof(t6_clopban3, conjecture, ![X1]:(((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_clvect_1(X1))&v3_clvect_1(X1))&v4_clvect_1(X1))&v5_clvect_1(X1))&v8_clvect_1(X1))&l2_clvect_1(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>(![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_normsp_1(X1,X2,X3)=k4_struct_0(X1))=>![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k6_numbers)))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', t6_clopban3)).
fof(dt_l1_normsp_0, axiom, ![X1]:(l1_normsp_0(X1)=>l1_struct_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', dt_l1_normsp_0)).
fof(redefinition_k1_normsp_1, axiom, ![X1, X2, X3]:((((((~(v2_struct_0(X1))&l1_struct_0(X1))&v1_funct_1(X2))&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))&m1_subset_1(X3,k5_numbers))=>k1_normsp_1(X1,X2,X3)=k1_funct_1(X2,X3)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_k1_normsp_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', fc6_ordinal1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_m2_subset_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', fc1_numbers)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_k5_numbers)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', dt_k5_numbers)).
fof(commutativity_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_nat_1(X2,X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', commutativity_k2_nat_1)).
fof(redefinition_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_k2_nat_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc8_ordinal1)).
fof(d4_normsp_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_normsp_0(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,k5_numbers,k1_numbers))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(X3=k4_normsp_0(X1,X2)<=>![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>k3_funct_2(k5_numbers,k1_numbers,X3,X4)=k1_normsp_0(X1,k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X4))))))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', d4_normsp_0)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', t1_arithm)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc3_xreal_0)).
fof(dt_k4_normsp_0, axiom, ![X1, X2]:(((((~(v2_struct_0(X1))&l1_normsp_0(X1))&v1_funct_1(X2))&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>((v1_funct_1(k4_normsp_0(X1,X2))&v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers))&m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', dt_k4_normsp_0)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', spc1_numerals)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', t6_boole)).
fof(t1_numerals, axiom, m1_subset_1(k1_xboole_0,k4_ordinal1), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', t1_numerals)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_k6_numbers)).
fof(cc2_xreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc2_xreal_0)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:(((((~(v1_xboole_0(X1))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&m1_subset_1(X4,X1))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_k3_funct_2)).
fof(fc2_normsp_0, axiom, ![X1]:(((~(v2_struct_0(X1))&v4_normsp_0(X1))&l2_normsp_0(X1))=>v1_xboole_0(k1_normsp_0(X1,k4_struct_0(X1)))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', fc2_normsp_0)).
fof(dt_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', dt_k2_nat_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc1_relset_1)).
fof(d1_series_1, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_partfun1(X1,k5_numbers))&v1_valued_0(X1))=>![X2]:(((((v1_relat_1(X2)&v4_relat_1(X2,k5_numbers))&v1_funct_1(X2))&v1_partfun1(X2,k5_numbers))&v1_valued_0(X2))=>(X2=k2_series_1(X1)<=>(k1_funct_1(X2,k6_numbers)=k1_funct_1(X1,k6_numbers)&![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_funct_1(X2,k2_nat_1(X3,np__1))=k2_xcmplx_0(k1_funct_1(X2,X3),k1_funct_1(X1,k2_nat_1(X3,np__1)))))))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', d1_series_1)).
fof(rc4_xreal_0, axiom, ?[X1]:(((v1_xboole_0(X1)&v1_xcmplx_0(X1))&v1_xxreal_0(X1))&v1_xreal_0(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', rc4_xreal_0)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_partfun1(X3,X1)=>v1_funct_2(X3,X1,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc1_funct_2)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc2_relset_1)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_k1_seq_1)).
fof(cc30_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v5_relat_1(X1,k1_numbers))=>(v1_relat_1(X1)&v3_valued_0(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc30_valued_0)).
fof(redefinition_k3_series_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>k3_series_1(X1)=k2_series_1(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', redefinition_k3_series_1)).
fof(fc85_valued_0, axiom, ![X1, X2]:(v1_membered(X2)=>v1_valued_0(k2_zfmisc_1(X1,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', fc85_valued_0)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc5_membered)).
fof(cc3_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_funct_2(X3,X1,X2)=>v1_partfun1(X3,X1)))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc3_funct_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', reflexivity_r1_tarski)).
fof(cc10_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v1_valued_0(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_valued_0(X2))), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', cc10_valued_0)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/home/yan/atp/benchmarks/mizar40/all/t6_clopban3', fc3_membered)).
fof(c_0_40, plain, ![X207]:((l1_normsp_0(X207)|~l2_normsp_0(X207))&(l2_struct_0(X207)|~l2_normsp_0(X207))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_normsp_0])])])).
fof(c_0_41, plain, ![X206]:((l1_clvect_1(X206)|~l2_clvect_1(X206))&(l2_normsp_0(X206)|~l2_clvect_1(X206))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_clvect_1])])])).
fof(c_0_42, negated_conjecture, ~(![X1]:(((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v3_normsp_0(X1))&v4_normsp_0(X1))&v2_clvect_1(X1))&v3_clvect_1(X1))&v4_clvect_1(X1))&v5_clvect_1(X1))&v8_clvect_1(X1))&l2_clvect_1(X1))=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))=>(![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_normsp_1(X1,X2,X3)=k4_struct_0(X1))=>![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k6_numbers))))), inference(assume_negation,[status(cth)],[t6_clopban3])).
cnf(c_0_43, plain, (l1_normsp_0(X1)|~l2_normsp_0(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, plain, (l2_normsp_0(X1)|~l2_clvect_1(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_45, negated_conjecture, ![X7]:(((((((((((((~v2_struct_0(esk1_0)&v13_algstr_0(esk1_0))&v2_rlvect_1(esk1_0))&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&v3_normsp_0(esk1_0))&v4_normsp_0(esk1_0))&v2_clvect_1(esk1_0))&v3_clvect_1(esk1_0))&v4_clvect_1(esk1_0))&v5_clvect_1(esk1_0))&v8_clvect_1(esk1_0))&l2_clvect_1(esk1_0))&(((v1_funct_1(esk2_0)&v1_funct_2(esk2_0,k5_numbers,u1_struct_0(esk1_0)))&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))&((~m2_subset_1(X7,k1_numbers,k5_numbers)|k1_normsp_1(esk1_0,esk2_0,X7)=k4_struct_0(esk1_0))&(m2_subset_1(esk3_0,k1_numbers,k5_numbers)&k1_seq_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_42])])])])])])])).
fof(c_0_46, plain, ![X204]:(~l1_normsp_0(X204)|l1_struct_0(X204)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_normsp_0])])).
cnf(c_0_47, plain, (l1_normsp_0(X1)|~l2_normsp_0(X1)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_48, plain, (l2_normsp_0(X1)|~l2_clvect_1(X1)), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_49, negated_conjecture, (l2_clvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_50, plain, ![X338, X339, X340]:(v2_struct_0(X338)|~l1_struct_0(X338)|~v1_funct_1(X339)|~v1_funct_2(X339,k5_numbers,u1_struct_0(X338))|~m1_subset_1(X339,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X338))))|~m1_subset_1(X340,k5_numbers)|k1_normsp_1(X338,X339,X340)=k1_funct_1(X339,X340)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k1_normsp_1])])])).
cnf(c_0_51, plain, (l1_struct_0(X1)|~l1_normsp_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (l1_normsp_0(X1)|~l2_clvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_53, negated_conjecture, (l2_clvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_49])).
fof(c_0_54, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_55, negated_conjecture, (k1_normsp_1(esk1_0,esk2_0,X1)=k4_struct_0(esk1_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (v2_struct_0(X1)|k1_normsp_1(X1,X2,X3)=k1_funct_1(X2,X3)|~l1_struct_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, plain, (l1_struct_0(X1)|~l1_normsp_0(X1)), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_58, negated_conjecture, (l1_normsp_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53])])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_60, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_61, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_62, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_63, plain, ![X354, X355, X356, X356]:((~m2_subset_1(X356,X354,X355)|m1_subset_1(X356,X355)|(v1_xboole_0(X354)|v1_xboole_0(X355)|~m1_subset_1(X355,k1_zfmisc_1(X354))))&(~m1_subset_1(X356,X355)|m2_subset_1(X356,X354,X355)|(v1_xboole_0(X354)|v1_xboole_0(X355)|~m1_subset_1(X355,k1_zfmisc_1(X354))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])])])).
fof(c_0_64, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
cnf(c_0_65, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_67, negated_conjecture, (k1_normsp_1(esk1_0,esk2_0,X1)=k4_struct_0(esk1_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_68, plain, (k1_normsp_1(X1,X2,X3)=k1_funct_1(X2,X3)|v2_struct_0(X1)|~v1_funct_1(X2)|~l1_struct_0(X1)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_69, negated_conjecture, (l1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58])])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_71, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_72, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_73, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_74, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_75, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_77, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_78, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_79, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_66])).
fof(c_0_80, plain, ![X165, X166]:(~m1_subset_1(X165,k5_numbers)|~v7_ordinal1(X166)|k2_nat_1(X165,X166)=k2_nat_1(X166,X165)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_nat_1])])).
fof(c_0_81, plain, ![X343, X344]:(~m1_subset_1(X343,k5_numbers)|~v7_ordinal1(X344)|k2_nat_1(X343,X344)=k2_xcmplx_0(X343,X344)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_1])])).
cnf(c_0_82, negated_conjecture, (k1_funct_1(esk2_0,X1)=k4_struct_0(esk1_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70]), c_0_71]), c_0_72])]), c_0_73])])).
cnf(c_0_83, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_84, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_85, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_86, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_87, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_88, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_89, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_90, plain, ![X157]:(~m1_subset_1(X157,k4_ordinal1)|v7_ordinal1(X157)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_91, plain, ![X175, X176, X177, X178]:((X177!=k4_normsp_0(X175,X176)|(~m2_subset_1(X178,k1_numbers,k5_numbers)|k3_funct_2(k5_numbers,k1_numbers,X177,X178)=k1_normsp_0(X175,k3_funct_2(k5_numbers,u1_struct_0(X175),X176,X178)))|(~v1_funct_1(X177)|~v1_funct_2(X177,k5_numbers,k1_numbers)|~m1_subset_1(X177,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X176)|~v1_funct_2(X176,k5_numbers,u1_struct_0(X175))|~m1_subset_1(X176,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X175)))))|(v2_struct_0(X175)|~l1_normsp_0(X175)))&((m2_subset_1(esk5_3(X175,X176,X177),k1_numbers,k5_numbers)|X177=k4_normsp_0(X175,X176)|(~v1_funct_1(X177)|~v1_funct_2(X177,k5_numbers,k1_numbers)|~m1_subset_1(X177,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X176)|~v1_funct_2(X176,k5_numbers,u1_struct_0(X175))|~m1_subset_1(X176,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X175)))))|(v2_struct_0(X175)|~l1_normsp_0(X175)))&(k3_funct_2(k5_numbers,k1_numbers,X177,esk5_3(X175,X176,X177))!=k1_normsp_0(X175,k3_funct_2(k5_numbers,u1_struct_0(X175),X176,esk5_3(X175,X176,X177)))|X177=k4_normsp_0(X175,X176)|(~v1_funct_1(X177)|~v1_funct_2(X177,k5_numbers,k1_numbers)|~m1_subset_1(X177,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X176)|~v1_funct_2(X176,k5_numbers,u1_struct_0(X175))|~m1_subset_1(X176,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X175)))))|(v2_struct_0(X175)|~l1_normsp_0(X175))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_normsp_0])])])])])])])])).
cnf(c_0_92, negated_conjecture, (k4_struct_0(esk1_0)=k1_funct_1(esk2_0,esk3_0)|~m1_subset_1(esk3_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_93, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_83]), c_0_85])]), c_0_86]), c_0_87])])).
fof(c_0_94, plain, ![X366]:(~v1_xcmplx_0(X366)|k2_xcmplx_0(X366,k6_numbers)=X366), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
fof(c_0_95, plain, ![X122]:(~v1_xreal_0(X122)|v1_xcmplx_0(X122)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_96, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_97, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_98, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_99, plain, (k3_funct_2(k5_numbers,k1_numbers,X1,X4)=k1_normsp_0(X2,k3_funct_2(k5_numbers,u1_struct_0(X2),X3,X4))|v2_struct_0(X2)|X1!=k4_normsp_0(X2,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~l1_normsp_0(X2)), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_100, plain, ![X197, X198]:(((v1_funct_1(k4_normsp_0(X197,X198))|(v2_struct_0(X197)|~l1_normsp_0(X197)|~v1_funct_1(X198)|~v1_funct_2(X198,k5_numbers,u1_struct_0(X197))|~m1_subset_1(X198,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X197))))))&(v1_funct_2(k4_normsp_0(X197,X198),k5_numbers,k1_numbers)|(v2_struct_0(X197)|~l1_normsp_0(X197)|~v1_funct_1(X198)|~v1_funct_2(X198,k5_numbers,u1_struct_0(X197))|~m1_subset_1(X198,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X197)))))))&(m1_subset_1(k4_normsp_0(X197,X198),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|(v2_struct_0(X197)|~l1_normsp_0(X197)|~v1_funct_1(X198)|~v1_funct_2(X198,k5_numbers,u1_struct_0(X197))|~m1_subset_1(X198,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X197))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_normsp_0])])])])).
cnf(c_0_101, negated_conjecture, (k4_struct_0(esk1_0)=k1_funct_1(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_102, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_103, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_104, plain, ![X379]:(~v1_xboole_0(X379)|X379=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_105, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_106, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_107, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97])])).
cnf(c_0_108, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_109, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(split_conjunct,[status(thm)],[t1_numerals])).
cnf(c_0_110, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_111, plain, ![X106]:(~v7_ordinal1(X106)|v1_xreal_0(X106)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xreal_0])])).
cnf(c_0_112, plain, (k3_funct_2(k5_numbers,k1_numbers,X1,X4)=k1_normsp_0(X2,k3_funct_2(k5_numbers,u1_struct_0(X2),X3,X4))|v2_struct_0(X2)|X1!=k4_normsp_0(X2,X3)|~v1_funct_1(X3)|~v1_funct_1(X1)|~l1_normsp_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_113, plain, (v1_funct_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_114, plain, (v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_115, plain, (m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_100])).
fof(c_0_116, plain, ![X345, X346, X347, X348]:(v1_xboole_0(X345)|~v1_funct_1(X347)|~v1_funct_2(X347,X345,X346)|~m1_subset_1(X347,k1_zfmisc_1(k2_zfmisc_1(X345,X346)))|~m1_subset_1(X348,X345)|k3_funct_2(X345,X346,X347,X348)=k1_funct_1(X347,X348)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_117, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,esk3_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_101])])).
cnf(c_0_118, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_119, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_120, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
fof(c_0_121, plain, ![X242]:(v2_struct_0(X242)|~v4_normsp_0(X242)|~l2_normsp_0(X242)|v1_xboole_0(k1_normsp_0(X242,k4_struct_0(X242)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_normsp_0])])])).
fof(c_0_122, plain, ![X187, X188]:(~m1_subset_1(X187,k5_numbers)|~v7_ordinal1(X188)|m2_subset_1(k2_nat_1(X187,X188),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_nat_1])])).
cnf(c_0_123, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_124, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_125, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_107])])).
cnf(c_0_126, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_108, c_0_79])).
cnf(c_0_127, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_128, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_129, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_130, plain, (k1_normsp_0(X1,k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(k4_normsp_0(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_112])])).
cnf(c_0_131, plain, (v2_struct_0(X1)|v1_funct_1(k4_normsp_0(X1,X2))|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_132, plain, (v2_struct_0(X1)|v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_133, plain, (v2_struct_0(X1)|m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_134, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_135, negated_conjecture, (k1_funct_1(esk2_0,esk3_0)=k1_funct_1(esk2_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])])])).
cnf(c_0_136, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_137, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_138, plain, (v2_struct_0(X1)|v1_xboole_0(k1_normsp_0(X1,k4_struct_0(X1)))|~v4_normsp_0(X1)|~l2_normsp_0(X1)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_139, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_140, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_141, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_125, c_0_126]), c_0_126])).
cnf(c_0_142, plain, (m1_subset_1(k6_numbers,k5_numbers)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_128]), c_0_79])).
cnf(c_0_143, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_144, plain, (k1_normsp_0(X1,k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_130, c_0_131]), c_0_132]), c_0_133])).
cnf(c_0_145, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_146, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,np__1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_117, c_0_135])])).
cnf(c_0_147, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_148, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_137, c_0_128])).
cnf(c_0_149, plain, (v2_struct_0(X1)|v1_xboole_0(k1_normsp_0(X1,k4_struct_0(X1)))|~v4_normsp_0(X1)|~l2_normsp_0(X1)), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_150, negated_conjecture, (v4_normsp_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_151, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_152, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_142])])])).
cnf(c_0_153, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_126])])).
cnf(c_0_154, plain, (k1_normsp_0(X1,k1_funct_1(X2,X3))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_87])])).
cnf(c_0_155, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,np__1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_85])]), c_0_87]), c_0_86])])).
cnf(c_0_156, plain, (k1_normsp_0(X1,k4_struct_0(X1))=k6_numbers|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_157, negated_conjecture, (k4_struct_0(esk1_0)=k1_funct_1(esk2_0,np__1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_135])])).
cnf(c_0_158, negated_conjecture, (v4_normsp_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_159, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_97])])).
cnf(c_0_160, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[c_0_152, c_0_153])).
cnf(c_0_161, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,np__1))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_72])])])).
cnf(c_0_162, negated_conjecture, (k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1))=k6_numbers|~l2_normsp_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_158])]), c_0_73])])).
cnf(c_0_163, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_142])])])).
fof(c_0_164, plain, ![X64, X65, X66]:(~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))|v1_relat_1(X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_165, plain, ![X171, X172, X173]:(((k1_funct_1(X172,k6_numbers)=k1_funct_1(X171,k6_numbers)|X172!=k2_series_1(X171)|(~v1_relat_1(X172)|~v4_relat_1(X172,k5_numbers)|~v1_funct_1(X172)|~v1_partfun1(X172,k5_numbers)|~v1_valued_0(X172))|(~v1_relat_1(X171)|~v4_relat_1(X171,k5_numbers)|~v1_funct_1(X171)|~v1_partfun1(X171,k5_numbers)|~v1_valued_0(X171)))&(~m2_subset_1(X173,k1_numbers,k5_numbers)|k1_funct_1(X172,k2_nat_1(X173,np__1))=k2_xcmplx_0(k1_funct_1(X172,X173),k1_funct_1(X171,k2_nat_1(X173,np__1)))|X172!=k2_series_1(X171)|(~v1_relat_1(X172)|~v4_relat_1(X172,k5_numbers)|~v1_funct_1(X172)|~v1_partfun1(X172,k5_numbers)|~v1_valued_0(X172))|(~v1_relat_1(X171)|~v4_relat_1(X171,k5_numbers)|~v1_funct_1(X171)|~v1_partfun1(X171,k5_numbers)|~v1_valued_0(X171))))&((m2_subset_1(esk4_2(X171,X172),k1_numbers,k5_numbers)|k1_funct_1(X172,k6_numbers)!=k1_funct_1(X171,k6_numbers)|X172=k2_series_1(X171)|(~v1_relat_1(X172)|~v4_relat_1(X172,k5_numbers)|~v1_funct_1(X172)|~v1_partfun1(X172,k5_numbers)|~v1_valued_0(X172))|(~v1_relat_1(X171)|~v4_relat_1(X171,k5_numbers)|~v1_funct_1(X171)|~v1_partfun1(X171,k5_numbers)|~v1_valued_0(X171)))&(k1_funct_1(X172,k2_nat_1(esk4_2(X171,X172),np__1))!=k2_xcmplx_0(k1_funct_1(X172,esk4_2(X171,X172)),k1_funct_1(X171,k2_nat_1(esk4_2(X171,X172),np__1)))|k1_funct_1(X172,k6_numbers)!=k1_funct_1(X171,k6_numbers)|X172=k2_series_1(X171)|(~v1_relat_1(X172)|~v4_relat_1(X172,k5_numbers)|~v1_funct_1(X172)|~v1_partfun1(X172,k5_numbers)|~v1_valued_0(X172))|(~v1_relat_1(X171)|~v4_relat_1(X171,k5_numbers)|~v1_funct_1(X171)|~v1_partfun1(X171,k5_numbers)|~v1_valued_0(X171))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_series_1])])])])])])])).
cnf(c_0_166, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~l2_normsp_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162]), c_0_58]), c_0_70]), c_0_71])]), c_0_73])])).
cnf(c_0_167, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[c_0_163, c_0_126])).
cnf(c_0_168, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_164])).
fof(c_0_169, plain, (((v1_xboole_0(esk49_0)&v1_xcmplx_0(esk49_0))&v1_xxreal_0(esk49_0))&v1_xreal_0(esk49_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
cnf(c_0_170, plain, (X1=k2_series_1(X2)|k1_funct_1(X1,k2_nat_1(esk4_2(X2,X1),np__1))!=k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk4_2(X2,X1),np__1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_funct_1(X2)|~v1_partfun1(X2,k5_numbers)|~v1_valued_0(X2)), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_171, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~l2_normsp_0(esk1_0)|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_172, negated_conjecture, (v1_funct_1(k4_normsp_0(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_70]), c_0_58]), c_0_71]), c_0_72])]), c_0_73])])).
cnf(c_0_173, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_174, plain, (v1_xboole_0(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_175, plain, (X1=k2_series_1(X2)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk4_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk4_2(X2,X1),np__1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_176, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~l2_normsp_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_171]), c_0_172])]), c_0_87])])).
cnf(c_0_177, plain, (v1_relat_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_133])])).
fof(c_0_178, plain, ![X59, X60, X61]:(~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60)))|(~v1_partfun1(X61,X59)|v1_funct_2(X61,X59,X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])).
fof(c_0_179, plain, ![X99, X100, X101]:((v4_relat_1(X101,X99)|~m1_subset_1(X101,k1_zfmisc_1(k2_zfmisc_1(X99,X100))))&(v5_relat_1(X101,X100)|~m1_subset_1(X101,k1_zfmisc_1(k2_zfmisc_1(X99,X100))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_180, plain, (v1_xcmplx_0(esk49_0)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_181, plain, (v1_xboole_0(esk49_0)), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_182, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk4_2(X2,X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_96]), c_0_119])])])).
cnf(c_0_183, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_48]), c_0_53])])])).
cnf(c_0_184, negated_conjecture, (v1_relat_1(k4_normsp_0(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_70]), c_0_58]), c_0_71]), c_0_72])]), c_0_73])])).
cnf(c_0_185, plain, (v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_partfun1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_178])).
cnf(c_0_186, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_179])).
cnf(c_0_187, plain, (v1_xcmplx_0(esk49_0)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_188, plain, (esk49_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_181])])).
cnf(c_0_189, negated_conjecture, (k2_series_1(X1)=k4_normsp_0(esk1_0,esk2_0)|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)))))!=k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(np__1,esk4_2(X1,k4_normsp_0(esk1_0,esk2_0))))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v7_ordinal1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)))|~v1_valued_0(X1)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183]), c_0_184]), c_0_172])])])).
cnf(c_0_190, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_191, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_192, plain, (v1_xcmplx_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_193, negated_conjecture, (k2_series_1(X1)=k4_normsp_0(esk1_0,esk2_0)|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)))))!=k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(np__1,esk4_2(X1,k4_normsp_0(esk1_0,esk2_0))))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_189, c_0_190]), c_0_126]), c_0_191])).
cnf(c_0_194, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_192])])).
cnf(c_0_195, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))))!=k6_numbers|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_183]), c_0_194]), c_0_184]), c_0_172])])])).
cnf(c_0_196, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))))!=k6_numbers|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_195, c_0_190]), c_0_191])).
cnf(c_0_197, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_183])])).
cnf(c_0_198, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_197, c_0_190])).
cnf(c_0_199, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_151]), c_0_85])]), c_0_86]), c_0_87])])).
cnf(c_0_200, plain, (m2_subset_1(esk4_2(X1,X2),k1_numbers,k5_numbers)|X2=k2_series_1(X1)|k1_funct_1(X2,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_relat_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_funct_1(X2)|~v1_partfun1(X2,k5_numbers)|~v1_valued_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_201, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199]), c_0_119])])])).
cnf(c_0_202, plain, (X2=k2_series_1(X1)|m2_subset_1(esk4_2(X1,X2),k1_numbers,k5_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_200])).
fof(c_0_203, plain, ![X341, X342]:(~v1_relat_1(X341)|~v1_funct_1(X341)|~v3_valued_0(X341)|k1_seq_1(X341,X342)=k1_funct_1(X341,X342)), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])])])).
fof(c_0_204, plain, ![X108]:((v1_relat_1(X108)|(~v1_relat_1(X108)|~v5_relat_1(X108,k1_numbers)))&(v3_valued_0(X108)|(~v1_relat_1(X108)|~v5_relat_1(X108,k1_numbers)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc30_valued_0])])])).
fof(c_0_205, plain, ![X349]:(~v1_funct_1(X349)|~v1_funct_2(X349,k5_numbers,k1_numbers)|~m1_subset_1(X349,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k3_series_1(X349)=k2_series_1(X349)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_series_1])])).
cnf(c_0_206, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_201, c_0_126])).
cnf(c_0_207, plain, (m1_subset_1(esk4_2(X1,X2),k5_numbers)|X2=k2_series_1(X1)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_202]), c_0_85])]), c_0_86]), c_0_87])])).
cnf(c_0_208, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_203])).
cnf(c_0_209, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_210, plain, (k3_series_1(X1)=k2_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_211, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207]), c_0_184]), c_0_172])])])).
fof(c_0_212, plain, ![X260, X261]:(~v1_membered(X261)|v1_valued_0(k2_zfmisc_1(X260,X261))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc85_valued_0])])).
fof(c_0_213, plain, ![X137]:(~v3_membered(X137)|v1_membered(X137)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_214, negated_conjecture, (k1_seq_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_215, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_216, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_217, plain, (k2_series_1(X1)=k3_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_218, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_211, c_0_191])).
fof(c_0_219, plain, ![X111, X112, X113]:(v1_xboole_0(X112)|(~m1_subset_1(X113,k1_zfmisc_1(k2_zfmisc_1(X111,X112)))|(~v1_funct_2(X113,X111,X112)|v1_partfun1(X113,X111)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_funct_2])])])])])])).
cnf(c_0_220, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_212])).
cnf(c_0_221, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_213])).
fof(c_0_222, plain, ![X371, X372, X371, X372]:((~m1_subset_1(X371,k1_zfmisc_1(X372))|r1_tarski(X371,X372))&(~r1_tarski(X371,X372)|m1_subset_1(X371,k1_zfmisc_1(X372)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
fof(c_0_223, plain, ![X357]:r1_tarski(X357,X357), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_224, negated_conjecture, (k1_seq_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_225, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216])])).
cnf(c_0_226, negated_conjecture, (k3_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218]), c_0_172])])])).
cnf(c_0_227, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_219])).
fof(c_0_228, plain, ![X15, X16]:(~v1_relat_1(X15)|~v1_valued_0(X15)|(~m1_subset_1(X16,k1_zfmisc_1(X15))|v1_valued_0(X16))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc10_valued_0])])])])])).
cnf(c_0_229, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1)), inference(evalgc,[status(thm)],[c_0_220])).
cnf(c_0_230, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_221])).
cnf(c_0_231, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_232, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_222])).
cnf(c_0_233, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_223])).
cnf(c_0_234, negated_conjecture, (k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers|~v5_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_225])])).
cnf(c_0_235, negated_conjecture, (k3_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_226, c_0_190])).
cnf(c_0_236, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_179])).
cnf(c_0_237, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_238, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_239, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230])])).
cnf(c_0_240, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_241, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_232])).
cnf(c_0_242, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_243, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~v5_relat_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235]), c_0_184]), c_0_172])])])).
cnf(c_0_244, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_245, plain, (v1_partfun1(k4_normsp_0(X1,X2),k5_numbers)|v2_struct_0(X1)|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_133]), c_0_86])])).
cnf(c_0_246, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_247, plain, (v1_valued_0(k2_zfmisc_1(X1,k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_240])])).
cnf(c_0_248, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_242])])).
cnf(c_0_249, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_243, c_0_244])).
cnf(c_0_250, plain, (v1_partfun1(k4_normsp_0(X1,X2),k5_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(csr,[status(thm)],[c_0_245, c_0_132])).
cnf(c_0_251, plain, (v1_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_133]), c_0_247])])])).
cnf(c_0_252, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_248])])).
cnf(c_0_253, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_250]), c_0_58]), c_0_70]), c_0_71]), c_0_72])]), c_0_73])])).
cnf(c_0_254, plain, (v1_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_251, c_0_252])])).
cnf(c_0_255, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254]), c_0_58]), c_0_70]), c_0_71]), c_0_72])]), c_0_73])])).
cnf(c_0_256, negated_conjecture, (~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_183]), c_0_93])])])).
cnf(c_0_257, negated_conjecture, (~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_133]), c_0_58]), c_0_70]), c_0_71]), c_0_72])]), c_0_73])])).
cnf(c_0_258, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_132]), c_0_58]), c_0_70]), c_0_71]), c_0_72])]), c_0_73]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 259
# Proof object clause steps            : 184
# Proof object formula steps           : 75
# Proof object conjectures             : 60
# Proof object clause conjectures      : 57
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 55
# Proof object initial formulas used   : 40
# Proof object generating inferences   : 51
# Proof object simplifying inferences  : 134
# Training examples: 109 positive, 1737 negative
# Training: Positive examples begin
cnf(c_0_79, plain, (k4_ordinal1=k5_numbers)).#trainpos
cnf(c_0_128, plain, (k1_xboole_0=k6_numbers)).#trainpos
cnf(c_0_158, negated_conjecture, (v4_normsp_0(esk1_0))).#trainpos
cnf(c_0_73, negated_conjecture, (~v2_struct_0(esk1_0))).#trainpos
cnf(c_0_224, negated_conjecture, (k1_seq_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers)).#trainpos
cnf(c_0_53, negated_conjecture, (l2_clvect_1(esk1_0))).#trainpos
cnf(c_0_86, plain, (~v1_xboole_0(k1_numbers))).#trainpos
cnf(c_0_72, negated_conjecture, (v1_funct_1(esk2_0))).#trainpos
cnf(c_0_240, plain, (v3_membered(k1_numbers))).#trainpos
cnf(c_0_83, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers))).#trainpos
cnf(c_0_71, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_175, plain, (X1=k2_series_1(X2)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk4_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk4_2(X2,X1),np__1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainpos
cnf(c_0_187, plain, (v1_xcmplx_0(esk49_0))).#trainpos
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainpos
cnf(c_0_67, negated_conjecture, (k1_normsp_1(esk1_0,esk2_0,X1)=k4_struct_0(esk1_0)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainpos
cnf(c_0_118, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainpos
cnf(c_0_78, plain, (~v1_xboole_0(k4_ordinal1))).#trainpos
cnf(c_0_48, plain, (l2_normsp_0(X1)|~l2_clvect_1(X1))).#trainpos
cnf(c_0_181, plain, (v1_xboole_0(esk49_0))).#trainpos
cnf(c_0_119, plain, (m1_subset_1(np__1,k5_numbers))).#trainpos
cnf(c_0_230, plain, (v1_membered(X1)|~v3_membered(X1))).#trainpos
cnf(c_0_124, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainpos
cnf(c_0_137, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainpos
cnf(c_0_127, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainpos
cnf(c_0_57, plain, (l1_struct_0(X1)|~l1_normsp_0(X1))).#trainpos
cnf(c_0_188, plain, (esk49_0=k6_numbers)).#trainpos
cnf(c_0_192, plain, (v1_xcmplx_0(k6_numbers))).#trainpos
cnf(c_0_143, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainpos
cnf(c_0_47, plain, (l1_normsp_0(X1)|~l2_normsp_0(X1))).#trainpos
cnf(c_0_85, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainpos
cnf(c_0_123, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainpos
cnf(c_0_52, plain, (l1_normsp_0(X1)|~l2_clvect_1(X1))).#trainpos
cnf(c_0_58, negated_conjecture, (l1_normsp_0(esk1_0))).#trainpos
cnf(c_0_215, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainpos
cnf(c_0_96, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_97, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_182, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk4_2(X2,X1)))).#trainpos
cnf(c_0_69, negated_conjecture, (l1_struct_0(esk1_0))).#trainpos
cnf(c_0_151, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_173, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_194, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).#trainpos
cnf(c_0_140, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xreal_0(X1))).#trainpos
cnf(c_0_147, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainpos
cnf(c_0_84, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainpos
cnf(c_0_237, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainpos
cnf(c_0_202, plain, (X2=k2_series_1(X1)|m2_subset_1(esk4_2(X1,X2),k1_numbers,k5_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainpos
cnf(c_0_131, plain, (v2_struct_0(X1)|v1_funct_1(k4_normsp_0(X1,X2))|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainpos
cnf(c_0_132, plain, (v2_struct_0(X1)|v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainpos
cnf(c_0_133, plain, (v2_struct_0(X1)|m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainpos
cnf(c_0_68, plain, (k1_normsp_1(X1,X2,X3)=k1_funct_1(X2,X3)|v2_struct_0(X1)|~v1_funct_1(X2)|~l1_struct_0(X1)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainpos
cnf(c_0_145, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainpos
cnf(c_0_112, plain, (k3_funct_2(k5_numbers,k1_numbers,X1,X4)=k1_normsp_0(X2,k3_funct_2(k5_numbers,u1_struct_0(X2),X3,X4))|v2_struct_0(X2)|X1!=k4_normsp_0(X2,X3)|~v1_funct_1(X3)|~v1_funct_1(X1)|~l1_normsp_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainpos
cnf(c_0_107, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_82, negated_conjecture, (k1_funct_1(esk2_0,X1)=k4_struct_0(esk1_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_159, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_93, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers))).#trainpos
cnf(c_0_172, negated_conjecture, (v1_funct_1(k4_normsp_0(esk1_0,esk2_0)))).#trainpos
cnf(c_0_177, plain, (v1_relat_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainpos
cnf(c_0_130, plain, (k1_normsp_0(X1,k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(k4_normsp_0(X1,X2)))).#trainpos
cnf(c_0_92, negated_conjecture, (k4_struct_0(esk1_0)=k1_funct_1(esk2_0,esk3_0)|~m1_subset_1(esk3_0,k5_numbers))).#trainpos
cnf(c_0_117, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,esk3_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_184, negated_conjecture, (v1_relat_1(k4_normsp_0(esk1_0,esk2_0)))).#trainpos
cnf(c_0_154, plain, (k1_normsp_0(X1,k1_funct_1(X2,X3))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainpos
cnf(c_0_135, negated_conjecture, (k1_funct_1(esk2_0,esk3_0)=k1_funct_1(esk2_0,np__1))).#trainpos
cnf(c_0_146, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,np__1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_157, negated_conjecture, (k4_struct_0(esk1_0)=k1_funct_1(esk2_0,np__1))).#trainpos
cnf(c_0_155, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,np__1)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_161, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,np__1))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainpos
cnf(c_0_242, plain, (r1_tarski(X1,X1))).#trainpos
cnf(c_0_108, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainpos
cnf(c_0_153, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_125, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainpos
cnf(c_0_152, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_163, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_216, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainpos
cnf(c_0_225, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainpos
cnf(c_0_234, negated_conjecture, (k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers|~v5_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0))))).#trainpos
cnf(c_0_229, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainpos
cnf(c_0_239, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainpos
cnf(c_0_247, plain, (v1_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainpos
cnf(c_0_241, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainpos
cnf(c_0_248, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainpos
cnf(c_0_246, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainpos
cnf(c_0_149, plain, (v2_struct_0(X1)|v1_xboole_0(k1_normsp_0(X1,k4_struct_0(X1)))|~v4_normsp_0(X1)|~l2_normsp_0(X1))).#trainpos
cnf(c_0_156, plain, (k1_normsp_0(X1,k4_struct_0(X1))=k6_numbers|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainpos
cnf(c_0_162, negated_conjecture, (k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1))=k6_numbers|~l2_normsp_0(esk1_0))).#trainpos
cnf(c_0_166, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~l2_normsp_0(esk1_0))).#trainpos
cnf(c_0_252, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainpos
cnf(c_0_176, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~l2_normsp_0(esk1_0))).#trainpos
cnf(c_0_183, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainpos
cnf(c_0_244, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainpos
cnf(c_0_191, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_190, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainpos
cnf(c_0_217, plain, (k2_series_1(X1)=k3_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_199, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainpos
cnf(c_0_207, plain, (m1_subset_1(esk4_2(X1,X2),k5_numbers)|X2=k2_series_1(X1)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainpos
cnf(c_0_245, plain, (v1_partfun1(k4_normsp_0(X1,X2),k5_numbers)|v2_struct_0(X1)|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainpos
cnf(c_0_251, plain, (v1_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainpos
cnf(c_0_189, negated_conjecture, (k2_series_1(X1)=k4_normsp_0(esk1_0,esk2_0)|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)))))!=k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(np__1,esk4_2(X1,k4_normsp_0(esk1_0,esk2_0))))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v7_ordinal1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)))|~v1_valued_0(X1)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainpos
cnf(c_0_195, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))))!=k6_numbers|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainpos
cnf(c_0_197, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainpos
cnf(c_0_201, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0))))).#trainpos
cnf(c_0_211, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers))).#trainpos
cnf(c_0_226, negated_conjecture, (k3_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainpos
cnf(c_0_243, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~v5_relat_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_253, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_255, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainpos
cnf(c_0_256, negated_conjecture, (~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainpos
cnf(c_0_257, negated_conjecture, (~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_18, negated_conjecture, (v13_algstr_0(esk1_0))).#trainneg
cnf(i_0_17, negated_conjecture, (v2_rlvect_1(esk1_0))).#trainneg
cnf(i_0_16, negated_conjecture, (v3_rlvect_1(esk1_0))).#trainneg
cnf(i_0_15, negated_conjecture, (v4_rlvect_1(esk1_0))).#trainneg
cnf(i_0_14, negated_conjecture, (v3_normsp_0(esk1_0))).#trainneg
cnf(i_0_12, negated_conjecture, (v2_clvect_1(esk1_0))).#trainneg
cnf(i_0_11, negated_conjecture, (v3_clvect_1(esk1_0))).#trainneg
cnf(i_0_10, negated_conjecture, (v4_clvect_1(esk1_0))).#trainneg
cnf(i_0_9, negated_conjecture, (v5_clvect_1(esk1_0))).#trainneg
cnf(i_0_8, negated_conjecture, (v8_clvect_1(esk1_0))).#trainneg
cnf(i_0_307, plain, (~v2_struct_0(esk28_0))).#trainneg
cnf(i_0_353, plain, (v1_funct_2(esk42_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_245, plain, (~v1_xboole_0(k4_numbers))).#trainneg
cnf(i_0_221, plain, (l2_clvect_1(esk11_0))).#trainneg
cnf(i_0_282, plain, (v1_funct_1(esk18_0))).#trainneg
cnf(i_0_301, plain, (v1_funct_1(esk24_0))).#trainneg
cnf(i_0_323, plain, (v1_funct_1(esk34_0))).#trainneg
cnf(i_0_355, plain, (v1_funct_1(esk42_0))).#trainneg
cnf(i_0_363, plain, (v1_funct_1(esk43_0))).#trainneg
cnf(i_0_250, plain, (v5_membered(k4_numbers))).#trainneg
cnf(i_0_219, plain, (l1_struct_0(esk9_0))).#trainneg
cnf(i_0_306, plain, (v7_struct_0(esk28_0))).#trainneg
cnf(i_0_283, plain, (v1_relat_1(esk18_0))).#trainneg
cnf(i_0_302, plain, (v1_relat_1(esk24_0))).#trainneg
cnf(i_0_325, plain, (v1_relat_1(esk34_0))).#trainneg
cnf(i_0_358, plain, (v1_relat_1(esk42_0))).#trainneg
cnf(i_0_365, plain, (v1_relat_1(esk43_0))).#trainneg
cnf(i_0_352, plain, (v1_valued_0(esk42_0))).#trainneg
cnf(i_0_258, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_200, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_419, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_243, plain, (~v1_xboole_0(k3_numbers))).#trainneg
cnf(i_0_246, plain, (v7_membered(k1_numbers))).#trainneg
cnf(i_0_209, plain, (l1_clvect_1(X1)|~l2_clvect_1(X1))).#trainneg
cnf(i_0_290, plain, (v6_membered(esk20_0))).#trainneg
cnf(i_0_314, plain, (v6_membered(esk31_0))).#trainneg
cnf(i_0_340, plain, (v6_membered(esk39_0))).#trainneg
cnf(i_0_351, plain, (v2_valued_0(esk42_0))).#trainneg
cnf(i_0_220, plain, (l2_algstr_0(esk10_0))).#trainneg
cnf(i_0_330, plain, (v1_xcmplx_0(esk36_0))).#trainneg
cnf(i_0_369, plain, (v1_xcmplx_0(esk44_0))).#trainneg
cnf(i_0_350, plain, (v3_valued_0(esk42_0))).#trainneg
cnf(i_0_305, plain, (v1_xxreal_0(esk27_0))).#trainneg
cnf(i_0_329, plain, (v1_xxreal_0(esk36_0))).#trainneg
cnf(i_0_332, plain, (v1_xxreal_0(esk37_0))).#trainneg
cnf(i_0_368, plain, (v1_xxreal_0(esk44_0))).#trainneg
cnf(i_0_371, plain, (v1_xxreal_0(esk45_0))).#trainneg
cnf(i_0_378, plain, (v1_xxreal_0(esk49_0))).#trainneg
cnf(i_0_272, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_381, plain, (v1_xxreal_0(esk50_0))).#trainneg
cnf(i_0_291, plain, (~v1_xboole_0(esk20_0))).#trainneg
cnf(i_0_48, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_304, plain, (v1_xreal_0(esk26_0))).#trainneg
cnf(i_0_327, plain, (v1_xreal_0(esk36_0))).#trainneg
cnf(i_0_366, plain, (v1_xreal_0(esk44_0))).#trainneg
cnf(i_0_377, plain, (v1_xreal_0(esk49_0))).#trainneg
cnf(i_0_244, plain, (v4_membered(k3_numbers))).#trainneg
cnf(i_0_300, plain, (v4_valued_0(esk24_0))).#trainneg
cnf(i_0_349, plain, (v4_valued_0(esk42_0))).#trainneg
cnf(i_0_372, plain, (v7_ordinal1(esk46_0))).#trainneg
cnf(i_0_383, plain, (v7_ordinal1(esk51_0))).#trainneg
cnf(i_0_239, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_303, plain, (v1_xboole_0(esk25_0))).#trainneg
cnf(i_0_382, plain, (v1_xboole_0(esk50_0))).#trainneg
cnf(i_0_260, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_315, plain, (~v1_xboole_0(esk31_0))).#trainneg
cnf(i_0_319, plain, (~v1_xboole_0(esk32_0))).#trainneg
cnf(i_0_292, plain, (v3_ordinal1(esk21_0))).#trainneg
cnf(i_0_322, plain, (~v1_xboole_0(esk34_0))).#trainneg
cnf(i_0_316, plain, (v3_ordinal1(esk32_0))).#trainneg
cnf(i_0_326, plain, (~v1_xboole_0(esk35_0))).#trainneg
cnf(i_0_341, plain, (~v1_xboole_0(esk39_0))).#trainneg
cnf(i_0_157, plain, (v6_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_43, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_318, plain, (v1_ordinal1(esk32_0))).#trainneg
cnf(i_0_317, plain, (v2_ordinal1(esk32_0))).#trainneg
cnf(i_0_348, plain, (v5_valued_0(esk42_0))).#trainneg
cnf(i_0_248, plain, (v7_membered(k4_numbers))).#trainneg
cnf(i_0_247, plain, (v7_membered(k3_numbers))).#trainneg
cnf(i_0_249, plain, (v7_membered(k4_ordinal1))).#trainneg
cnf(i_0_339, plain, (v7_membered(esk39_0))).#trainneg
cnf(i_0_281, plain, (v1_funcop_1(esk18_0))).#trainneg
cnf(i_0_423, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_328, plain, (v2_xxreal_0(esk36_0))).#trainneg
cnf(i_0_331, plain, (v2_xxreal_0(esk37_0))).#trainneg
cnf(i_0_367, plain, (v3_xxreal_0(esk44_0))).#trainneg
cnf(i_0_370, plain, (v3_xxreal_0(esk45_0))).#trainneg
cnf(i_0_218, plain, (l1_normsp_0(esk8_0))).#trainneg
cnf(i_0_359, plain, (~v1_xboole_0(esk42_0))).#trainneg
cnf(i_0_362, plain, (~v1_xboole_0(esk43_0))).#trainneg
cnf(i_0_223, plain, (l2_struct_0(esk13_0))).#trainneg
cnf(i_0_259, plain, (~v1_finset_1(k4_numbers))).#trainneg
cnf(i_0_308, plain, (l2_struct_0(esk28_0))).#trainneg
cnf(i_0_420, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_384, plain, (~v1_xboole_0(esk51_0))).#trainneg
cnf(i_0_263, plain, (~v1_finset_1(k3_numbers))).#trainneg
cnf(i_0_216, plain, (l1_algstr_0(esk6_0))).#trainneg
cnf(i_0_217, plain, (l1_clvect_1(esk7_0))).#trainneg
cnf(i_0_222, plain, (l2_normsp_0(esk12_0))).#trainneg
cnf(i_0_361, plain, (v1_finset_1(esk43_0))).#trainneg
cnf(i_0_356, plain, (v5_relat_1(esk42_0,k5_numbers))).#trainneg
cnf(i_0_354, plain, (v1_partfun1(esk42_0,k5_numbers))).#trainneg
cnf(i_0_324, plain, (v4_relat_1(esk34_0,k5_numbers))).#trainneg
cnf(i_0_357, plain, (v4_relat_1(esk42_0,k5_numbers))).#trainneg
cnf(i_0_364, plain, (v4_relat_1(esk43_0,k5_numbers))).#trainneg
cnf(i_0_212, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_202, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_436, plain, (esk50_0=k6_numbers)).#trainneg
cnf(i_0_438, plain, (esk25_0=k6_numbers)).#trainneg
cnf(i_0_381, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_439, plain, (l1_struct_0(esk8_0))).#trainneg
cnf(i_0_440, plain, (l1_struct_0(esk28_0))).#trainneg
cnf(i_0_441, plain, (l1_struct_0(esk13_0))).#trainneg
cnf(i_0_442, plain, (l1_struct_0(esk6_0))).#trainneg
cnf(i_0_377, plain, (v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_343, plain, (v1_funct_1(esk40_1(X1)))).#trainneg
cnf(i_0_203, plain, (l2_algstr_0(X1)|~l1_clvect_1(X1))).#trainneg
cnf(i_0_112, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_120, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_92, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_101, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_83, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_139, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_133, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_102, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_50, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_49, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_443, plain, (l2_algstr_0(esk7_0))).#trainneg
cnf(i_0_445, plain, (v1_xxreal_0(esk51_0))).#trainneg
cnf(i_0_446, plain, (v1_xxreal_0(esk46_0))).#trainneg
cnf(i_0_238, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_444, plain, (l2_algstr_0(X1)|~l2_clvect_1(X1))).#trainneg
cnf(i_0_447, plain, (v3_membered(k3_numbers))).#trainneg
cnf(i_0_448, plain, (v1_xreal_0(esk51_0))).#trainneg
cnf(i_0_449, plain, (v1_xreal_0(esk46_0))).#trainneg
cnf(i_0_451, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_452, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_450, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_453, plain, (v3_membered(k4_numbers))).#trainneg
cnf(i_0_78, plain, (v7_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_113, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_207, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_210, plain, (l2_struct_0(X1)|~l2_normsp_0(X1))).#trainneg
cnf(i_0_206, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_375, plain, (v1_xboole_0(X1)|v1_subset_1(esk48_1(X1),X1))).#trainneg
cnf(i_0_241, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_262, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_275, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_455, plain, (l1_normsp_0(esk12_0))).#trainneg
cnf(i_0_457, plain, (l2_struct_0(esk10_0))).#trainneg
cnf(i_0_458, plain, (l2_struct_0(esk7_0))).#trainneg
cnf(i_0_27, plain, (~v2_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_454, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_460, plain, (l2_struct_0(esk12_0))).#trainneg
cnf(i_0_467, plain, (l1_struct_0(esk12_0))).#trainneg
cnf(i_0_468, plain, (l1_struct_0(esk10_0))).#trainneg
cnf(i_0_459, plain, (l2_struct_0(X1)|~l2_clvect_1(X1))).#trainneg
cnf(i_0_198, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_30, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_32, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_34, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_36, plain, (v1_rat_1(X2)|~v4_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_38, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_40, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_427, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_428, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_469, plain, (l1_struct_0(esk7_0))).#trainneg
cnf(i_0_471, plain, (l1_normsp_0(esk11_0))).#trainneg
cnf(i_0_462, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_474, plain, (l2_struct_0(esk11_0))).#trainneg
cnf(i_0_475, negated_conjecture, (l2_struct_0(esk1_0))).#trainneg
cnf(i_0_509, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_497, plain, (v1_rat_1(k6_numbers)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_463, plain, (k2_xcmplx_0(esk44_0,k6_numbers)=esk44_0)).#trainneg
cnf(i_0_464, plain, (k2_xcmplx_0(esk36_0,k6_numbers)=esk36_0)).#trainneg
cnf(i_0_412, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_163, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_183, plain, (v1_funct_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_185, plain, (v1_relat_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_181, plain, (v1_valued_0(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_172, plain, (v2_struct_0(X1)|m1_subset_1(k1_normsp_0(X1,X2),k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_179, plain, (v2_struct_0(X1)|v1_funct_1(k2_normsp_0(X1,X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~l1_normsp_0(X1)|~v5_relat_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_180, plain, (v2_struct_0(X1)|v1_relat_1(k2_normsp_0(X1,X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~l1_normsp_0(X1)|~v5_relat_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_489, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_508, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_519, plain, (l1_struct_0(esk11_0))).#trainneg
cnf(i_0_520, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(np__1,esk4_2(X2,X1))))!=k1_funct_1(X1,k2_xcmplx_0(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v7_ordinal1(esk4_2(X2,X1))|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_539, negated_conjecture, (v1_relat_1(esk2_0))).#trainneg
cnf(i_0_540, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_503, plain, (v1_int_1(k6_numbers)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_523, plain, (l1_struct_0(X1)|~l2_clvect_1(X1))).#trainneg
cnf(i_0_476, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_214, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_147, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_191, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_190, plain, (m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_409, plain, (k2_normsp_0(X1,X2)=k4_normsp_0(X1,X2)|v2_struct_0(X1)|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_477, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_580, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_495, plain, (v1_rat_1(np__1)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_565, plain, (~v1_xboole_0(k3_series_1(X1))|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_499, plain, (v1_rat_1(k5_numbers)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_505, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_480, plain, (v1_xcmplx_0(k4_struct_0(X1))|~v1_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_481, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_173, plain, (v2_struct_0(X1)|m1_subset_1(k1_normsp_1(X1,X2,X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_struct_0(X1)|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_169, plain, (X3=k4_normsp_0(X1,X2)|v2_struct_0(X1)|m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_normsp_0(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_418, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_clvect_1(X1)|~v3_clvect_1(X1)|~v4_clvect_1(X1)|~v5_clvect_1(X1)|~v8_clvect_1(X1)|~l2_clvect_1(X1)|~v1_funct_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_416, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(esk56_2(X1,X2),np__1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_clvect_1(X1)|~v3_clvect_1(X1)|~v4_clvect_1(X1)|~v5_clvect_1(X1)|~v8_clvect_1(X1)|~l2_clvect_1(X1)|~v1_funct_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_513, plain, (k1_seq_1(esk42_0,X1)=k1_funct_1(esk42_0,X1))).#trainneg
cnf(i_0_528, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_529, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_496, plain, (v1_rat_1(np__1)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_501, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_502, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_482, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_487, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_493, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_555, negated_conjecture, (v1_partfun1(esk2_0,k5_numbers)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_486, plain, (v1_xxreal_0(k4_struct_0(X1))|~v2_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_488, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_492, plain, (v1_xreal_0(k4_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_532, plain, (v1_xreal_0(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_536, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_550, negated_conjecture, (m1_subset_1(esk3_0,k1_numbers))).#trainneg
cnf(i_0_553, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_581, plain, (k2_xcmplx_0(np__1,k6_numbers)=np__1)).#trainneg
cnf(i_0_592, plain, (m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_595, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(esk56_2(X1,X2),k5_numbers))).#trainneg
cnf(i_0_636, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_633, plain, (m2_subset_1(esk44_0,k1_numbers,k5_numbers)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_642, negated_conjecture, (v7_ordinal1(esk3_0))).#trainneg
cnf(i_0_591, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_494, negated_conjecture, (v1_rat_1(esk2_0)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_634, plain, (m2_subset_1(esk36_0,k1_numbers,k5_numbers)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_645, negated_conjecture, (v1_xreal_0(esk3_0))).#trainneg
cnf(i_0_671, negated_conjecture, (v1_xxreal_0(esk3_0))).#trainneg
cnf(i_0_498, plain, (v1_rat_1(k4_struct_0(X1))|~v4_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_504, plain, (v1_int_1(k4_struct_0(X1))|~v5_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_500, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_506, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_596, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(esk56_2(X1,X2),k5_numbers))).#trainneg
cnf(i_0_612, negated_conjecture, (k1_funct_1(esk2_0,esk5_3(X1,X2,X3))=k4_struct_0(esk1_0)|X3=k4_normsp_0(X1,X2)|v2_struct_0(X1)|~m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_511, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_566, plain, (v1_relat_1(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_510, plain, (v7_ordinal1(k4_struct_0(X1))|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_637, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_713, plain, (v1_xxreal_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_714, plain, (v1_xreal_0(k4_struct_0(X1))|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_515, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_516, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(esk4_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_xcmplx_0(esk4_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_3, negated_conjecture, (k1_normsp_1(esk1_0,esk2_0,X1)=k1_funct_1(esk2_0,esk3_0)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_591, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,esk3_0),u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_700, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_702, negated_conjecture, (k1_funct_1(esk2_0,esk3_0)=k1_funct_1(esk2_0,esk44_0)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_703, negated_conjecture, (k1_funct_1(esk2_0,esk3_0)=k1_funct_1(esk2_0,esk36_0)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_749, negated_conjecture, (k1_funct_1(esk2_0,esk44_0)=k1_funct_1(esk2_0,esk36_0)|~m1_subset_1(esk44_0,k5_numbers)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_583, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_640, negated_conjecture, (v1_rat_1(esk3_0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_530, plain, (v1_xcmplx_0(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_620, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_641, negated_conjecture, (v1_int_1(esk3_0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_621, plain, (v1_xreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_531, plain, (v1_xxreal_0(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_533, plain, (v1_rat_1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_700, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3, negated_conjecture, (k1_normsp_1(esk1_0,esk2_0,X1)=k1_funct_1(esk2_0,np__1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_750, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,np__1))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_752, negated_conjecture, (k1_funct_1(esk2_0,esk36_0)=k1_funct_1(esk2_0,np__1)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_753, negated_conjecture, (k1_funct_1(esk2_0,esk44_0)=k1_funct_1(esk2_0,np__1)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_810, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_767, plain, (v1_xcmplx_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_646, negated_conjecture, (v1_rat_1(esk3_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_534, plain, (v1_int_1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v5_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_780, plain, (v1_xxreal_0(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_643, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_846, negated_conjecture, (v1_xcmplx_0(esk3_0))).#trainneg
cnf(i_0_554, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_535, plain, (v7_ordinal1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_membered(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_544, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_867, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_619, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_582, plain, (k2_xcmplx_0(k4_struct_0(X1),k6_numbers)=k4_struct_0(X1)|~l2_struct_0(X1)|~v1_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_647, negated_conjecture, (v1_int_1(esk3_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_812, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X1,esk2_0))))!=k1_funct_1(esk2_0,k2_nat_1(np__1,esk4_2(X1,esk2_0)))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(X1)|~v1_valued_0(esk2_0)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_559, plain, (v1_xcmplx_0(k3_series_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_622, negated_conjecture, (v1_xreal_0(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_848, negated_conjecture, (v1_xboole_0(esk2_0)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|m1_subset_1(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_878, plain, (m2_subset_1(k4_struct_0(X1),k1_numbers,k5_numbers)|~m1_subset_1(k4_struct_0(X1),k5_numbers)|~l2_struct_0(X1)|~v1_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_560, plain, (v1_xxreal_0(k3_series_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_561, plain, (v1_xreal_0(k3_series_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_393, plain, (v1_funct_1(esk53_1(X1)))).#trainneg
cnf(i_0_401, plain, (v1_funct_1(esk55_1(X1)))).#trainneg
cnf(i_0_345, plain, (v1_relat_1(esk40_1(X1)))).#trainneg
cnf(i_0_396, plain, (v1_relat_1(esk53_1(X1)))).#trainneg
cnf(i_0_403, plain, (v1_relat_1(esk55_1(X1)))).#trainneg
cnf(i_0_391, plain, (v1_valued_0(esk53_1(X1)))).#trainneg
cnf(i_0_390, plain, (v2_valued_0(esk53_1(X1)))).#trainneg
cnf(i_0_389, plain, (v3_valued_0(esk53_1(X1)))).#trainneg
cnf(i_0_388, plain, (v4_valued_0(esk53_1(X1)))).#trainneg
cnf(i_0_684, negated_conjecture, (v1_rat_1(esk2_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_689, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_652, plain, (v1_rat_1(k6_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_882, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_nat_1(np__1,esk4_2(esk2_0,esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers)|~m1_subset_1(k2_nat_1(np__1,esk4_2(esk2_0,esk2_0)),k5_numbers))).#trainneg
cnf(i_0_562, plain, (v1_rat_1(k3_series_1(X1))|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_653, plain, (v1_int_1(k6_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_556, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_857, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_563, plain, (v1_int_1(k3_series_1(X1))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_564, plain, (v7_ordinal1(k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_898, negated_conjecture, (u1_struct_0(esk1_0)=k6_numbers|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_899, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_399, plain, (v4_valued_0(esk55_1(X1)))).#trainneg
cnf(i_0_320, plain, (v1_xboole_0(esk33_1(X1)))).#trainneg
cnf(i_0_342, plain, (v5_ordinal1(esk40_1(X1)))).#trainneg
cnf(i_0_385, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk52_1(X1)))).#trainneg
cnf(i_0_435, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_41, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_929, plain, (esk33_1(X1)=k6_numbers)).#trainneg
cnf(i_0_877, negated_conjecture, (v1_xxreal_0(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_895, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_893, negated_conjecture, (v1_rat_1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_571, plain, (v1_xcmplx_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_135, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_128, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_152, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_944, plain, (v3_valued_0(esk24_0))).#trainneg
cnf(i_0_945, plain, (k1_seq_1(esk24_0,X1)=k1_funct_1(esk24_0,X1))).#trainneg
cnf(i_0_946, plain, (v1_valued_0(esk24_0))).#trainneg
cnf(i_0_159, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_84, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_115, plain, (v8_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_941, plain, (v3_valued_0(esk55_1(X1)))).#trainneg
cnf(i_0_951, plain, (v1_valued_0(esk55_1(X1)))).#trainneg
cnf(i_0_847, negated_conjecture, (k2_xcmplx_0(esk3_0,k6_numbers)=esk3_0)).#trainneg
cnf(i_0_896, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers)|~v7_ordinal1(esk4_2(esk2_0,esk2_0)))).#trainneg
cnf(i_0_572, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_947, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_669, plain, (m1_subset_1(esk44_0,k1_numbers)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_960, plain, (v1_rat_1(esk44_0)|~v4_membered(k1_numbers)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_956, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_890, plain, (v1_xxreal_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_573, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_574, plain, (v1_rat_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_955, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_964, negated_conjecture, (k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers|~v1_xboole_0(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_134, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_394, plain, (v5_relat_1(esk53_1(X1),k3_numbers))).#trainneg
cnf(i_0_224, plain, (m1_subset_1(esk14_1(X1),X1))).#trainneg
cnf(i_0_344, plain, (v5_relat_1(esk40_1(X1),X1))).#trainneg
cnf(i_0_392, plain, (v1_partfun1(esk53_1(X1),X1))).#trainneg
cnf(i_0_400, plain, (v1_partfun1(esk55_1(X1),X1))).#trainneg
cnf(i_0_395, plain, (v4_relat_1(esk53_1(X1),X1))).#trainneg
cnf(i_0_402, plain, (v4_relat_1(esk55_1(X1),X1))).#trainneg
cnf(i_0_121, plain, (v1_xboole_0(X1)|v2_xxreal_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_976, plain, (v1_xcmplx_0(esk14_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_977, plain, (v1_xxreal_0(esk14_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_961, plain, (v1_int_1(esk44_0)|~v5_membered(k1_numbers)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_813, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(X1,k2_xcmplx_0(esk4_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(X1)|~v1_valued_0(esk2_0)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_575, plain, (v1_int_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_978, plain, (v1_xreal_0(esk14_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_962, plain, (v7_ordinal1(esk44_0)|~v6_membered(k1_numbers)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_686, plain, (m1_subset_1(esk36_0,k1_numbers)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_901, plain, (v1_xxreal_0(k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_576, plain, (v7_ordinal1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_585, plain, (v1_xcmplx_0(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_586, plain, (v1_xxreal_0(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_298, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk23_1(X1)))).#trainneg
cnf(i_0_386, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk52_1(X1)))).#trainneg
cnf(i_0_397, plain, (v1_zfmisc_1(X1)|~v1_zfmisc_1(esk54_1(X1)))).#trainneg
cnf(i_0_110, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_1001, plain, (~v1_xboole_0(esk37_0))).#trainneg
cnf(i_0_1002, plain, (~v1_xboole_0(esk36_0))).#trainneg
cnf(i_0_132, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_1004, plain, (~v1_xboole_0(esk45_0))).#trainneg
cnf(i_0_979, plain, (v1_rat_1(esk14_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_980, plain, (v1_int_1(esk14_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_1005, plain, (~v1_xboole_0(esk44_0))).#trainneg
cnf(i_0_587, plain, (v1_xreal_0(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_981, plain, (v7_ordinal1(esk14_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_991, plain, (v1_rat_1(esk36_0)|~v4_membered(k1_numbers)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_992, plain, (v1_int_1(esk36_0)|~v5_membered(k1_numbers)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_849, plain, (v1_xboole_0(k3_series_1(X1))|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X2,k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_891, plain, (v1_xreal_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_588, plain, (v1_rat_1(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_589, plain, (v1_int_1(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_108, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_1017, plain, (~v2_xxreal_0(esk45_0))).#trainneg
cnf(i_0_1018, plain, (~v2_xxreal_0(esk44_0))).#trainneg
cnf(i_0_309, plain, (v2_struct_0(X1)|v1_zfmisc_1(esk29_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_77, plain, (v6_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_56, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_237, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_264, plain, (v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_273, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v8_struct_0(X1))).#trainneg
cnf(i_0_1033, plain, (u1_struct_0(X1)=k6_numbers|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1020, plain, (v7_ordinal1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_815, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(X1,k2_nat_1(esk4_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_nat_1(esk4_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(X1)|~v1_valued_0(esk2_0)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_590, plain, (v7_ordinal1(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1021, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1022, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1023, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1031, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_597, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk56_2(X1,X2)))).#trainneg
cnf(i_0_599, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v7_ordinal1(esk56_2(X1,X2)))).#trainneg
cnf(i_0_97, plain, (v1_funcop_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_79, plain, (v2_funcop_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funcop_1(X1))).#trainneg
cnf(i_0_321, plain, (m1_subset_1(esk33_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1078, plain, (v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_347, plain, (m1_subset_1(esk41_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_151, plain, (v13_struct_0(X1,k1_xboole_0)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_54, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_1095, plain, (v5_relat_1(esk42_0,k4_numbers))).#trainneg
cnf(i_0_1096, plain, (v5_relat_1(esk24_0,k4_numbers))).#trainneg
cnf(i_0_141, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_1100, plain, (v5_relat_1(esk42_0,k3_numbers))).#trainneg
cnf(i_0_1101, plain, (v5_relat_1(esk24_0,k3_numbers))).#trainneg
cnf(i_0_1093, plain, (v5_relat_1(esk55_1(X1),k4_numbers))).#trainneg
cnf(i_0_1057, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk4_2(esk2_0,esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers)|~m1_subset_1(k2_nat_1(esk4_2(esk2_0,esk2_0),np__1),k5_numbers))).#trainneg
cnf(i_0_1072, plain, (v1_rat_1(k6_numbers)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1073, plain, (v1_int_1(k6_numbers)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1110, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1122, plain, (k2_xcmplx_0(k6_numbers,np__1)=np__1)).#trainneg
cnf(i_0_1126, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1042, plain, (u1_struct_0(esk13_0)=k6_numbers|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_1094, plain, (v5_relat_1(esk53_1(X1),k4_numbers))).#trainneg
cnf(i_0_1043, plain, (u1_struct_0(esk12_0)=k6_numbers|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_609, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1108, plain, (v1_rat_1(k6_numbers)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1109, plain, (v1_int_1(k6_numbers)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_657, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|v2_struct_0(X3)|~v1_membered(u1_struct_0(X3))|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X3))|~v1_funct_1(X1))).#trainneg
cnf(i_0_658, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|v2_struct_0(X3)|~v2_membered(u1_struct_0(X3))|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X3))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1180, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1181, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1215, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1044, plain, (u1_struct_0(esk11_0)=k6_numbers|~v2_struct_0(esk11_0))).#trainneg
cnf(i_0_1098, plain, (v5_relat_1(esk55_1(X1),k3_numbers))).#trainneg
cnf(i_0_1045, plain, (u1_struct_0(esk10_0)=k6_numbers|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_659, plain, (v1_xreal_0(k1_funct_1(X1,X2))|v2_struct_0(X3)|~v3_membered(u1_struct_0(X3))|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X3))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1202, plain, (v1_rat_1(k6_numbers)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_660, plain, (v1_rat_1(k1_funct_1(X1,X2))|v2_struct_0(X3)|~v4_membered(u1_struct_0(X3))|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X3))|~v1_funct_1(X1))).#trainneg
cnf(i_0_661, plain, (v1_int_1(k1_funct_1(X1,X2))|v2_struct_0(X3)|~v5_membered(u1_struct_0(X3))|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X3))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1249, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_226, plain, (v9_struct_0(k4_struct_0(X1),X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_1270, negated_conjecture, (v9_struct_0(k1_funct_1(esk2_0,np__1),esk1_0))).#trainneg
cnf(i_0_1272, negated_conjecture, (v9_struct_0(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_277, plain, (v9_struct_0(esk16_1(X1),X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_1046, plain, (u1_struct_0(esk9_0)=k6_numbers|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_1090, plain, (v1_relat_1(esk41_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1047, plain, (u1_struct_0(esk8_0)=k6_numbers|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_1113, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_662, plain, (v7_ordinal1(k1_funct_1(X1,X2))|v2_struct_0(X3)|~v6_membered(u1_struct_0(X3))|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X3))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1048, plain, (u1_struct_0(esk7_0)=k6_numbers|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_1034, plain, (X1=u1_struct_0(X2)|~v1_xboole_0(X1)|~l1_struct_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_691, plain, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_310, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk29_1(X1)))).#trainneg
cnf(i_0_373, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk47_1(X1)))).#trainneg
cnf(i_0_312, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(esk30_1(X1)))).#trainneg
cnf(i_0_434, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_285, plain, (v1_funct_1(esk19_2(X1,X2)))).#trainneg
cnf(i_0_288, plain, (v1_relat_1(esk19_2(X1,X2)))).#trainneg
cnf(i_0_295, plain, (v1_relat_1(esk22_2(X1,X2)))).#trainneg
cnf(i_0_296, plain, (v1_xboole_0(esk22_2(X1,X2)))).#trainneg
cnf(i_0_299, plain, (v1_xboole_0(X1)|m1_subset_1(esk23_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_376, plain, (v1_xboole_0(X1)|m1_subset_1(esk48_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_387, plain, (v1_xboole_0(X1)|m1_subset_1(esk52_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_398, plain, (v1_zfmisc_1(X1)|m1_subset_1(esk54_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1292, plain, (esk22_2(X1,X2)=k6_numbers)).#trainneg
cnf(i_0_1049, plain, (u1_struct_0(esk6_0)=k6_numbers|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_692, plain, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3),k1_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1025, negated_conjecture, (v7_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_1027, negated_conjecture, (v1_int_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_720, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(np__1,esk4_2(X2,X1))))!=k1_funct_1(X1,k2_xcmplx_0(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk4_2(X2,X1)))).#trainneg
cnf(i_0_23, plain, (v2_struct_0(X1)|v13_struct_0(X1,np__1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_21, plain, (v5_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_25, plain, (v6_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_149, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_156, plain, (v4_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_158, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,k1_xboole_0))).#trainneg
cnf(i_0_26, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_105, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_1421, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_118, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_1384, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1408, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1116, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk3_0)=esk3_0)).#trainneg
cnf(i_0_777, plain, (k2_xcmplx_0(k1_normsp_0(X1,X2),k6_numbers)=k1_normsp_0(X1,X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1385, plain, (v7_ordinal1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1386, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1412, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_1387, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_778, plain, (v1_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_782, plain, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1431, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_95, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_346, plain, (~v1_subset_1(esk41_1(X1),X1))).#trainneg
cnf(i_0_63, plain, (v5_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_62, plain, (v6_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_68, plain, (v7_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v5_valued_0(X1))).#trainneg
cnf(i_0_73, plain, (v8_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v6_valued_0(X1))).#trainneg
cnf(i_0_278, plain, (m1_subset_1(esk16_1(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_124, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1501, negated_conjecture, (v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_1508, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_1511, plain, (v3_ordinal1(k6_numbers))).#trainneg
cnf(i_0_1287, plain, (u1_struct_0(X1)=u1_struct_0(X2)|~l1_struct_0(X2)|~v2_struct_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1423, negated_conjecture, (k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers|~v5_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_784, plain, (v1_rat_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v4_membered(k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1388, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1483, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_1531, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_842, plain, (v1_int_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v5_membered(k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_850, plain, (v1_xboole_0(k4_normsp_0(X1,X2))|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|m1_subset_1(X3,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~m1_subset_1(X3,k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1530, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_334, plain, (v1_xboole_0(X2)|v1_funct_1(esk38_2(X1,X2)))).#trainneg
cnf(i_0_337, plain, (v1_xboole_0(X2)|v1_relat_1(esk38_2(X1,X2)))).#trainneg
cnf(i_0_266, plain, (v2_valued_0(k2_zfmisc_1(X2,X1))|~v2_membered(X1))).#trainneg
cnf(i_0_267, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_274, plain, (v4_valued_0(k2_zfmisc_1(X2,X1))|~v6_membered(X1))).#trainneg
cnf(i_0_280, plain, (v7_struct_0(X1)|m1_subset_1(esk17_1(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_286, plain, (v5_relat_1(esk19_2(X1,X2),X2))).#trainneg
cnf(i_0_293, plain, (v5_relat_1(esk22_2(X1,X2),X2))).#trainneg
cnf(i_0_287, plain, (v4_relat_1(esk19_2(X1,X2),X1))).#trainneg
cnf(i_0_294, plain, (v4_relat_1(esk22_2(X1,X2),X1))).#trainneg
cnf(i_0_1009, plain, (v1_xxreal_0(esk14_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1544, negated_conjecture, (k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers|~v5_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_852, plain, (v1_xxreal_0(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1398, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1399, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_854, plain, (v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_89, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_162, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1595, plain, (k2_xcmplx_0(X1,k3_series_1(X2))=k2_xcmplx_0(k3_series_1(X2),X1)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_67, plain, (v5_membered(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_22, plain, (v6_ordinal1(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_72, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_45, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_57, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_59, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1675, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_61, plain, (v4_membered(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1571, plain, (v1_valued_0(k2_zfmisc_1(X1,k3_numbers)))).#trainneg
cnf(i_0_1620, plain, (v5_membered(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1473, negated_conjecture, (k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers|~v5_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_885, plain, (k2_xcmplx_0(k3_series_1(X1),k6_numbers)=k3_series_1(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1631, plain, (v6_ordinal1(k5_numbers)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1653, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1726, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_1715, plain, (v5_membered(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_892, plain, (k1_seq_1(esk53_1(X1),X2)=k1_funct_1(esk53_1(X1),X2))).#trainneg
cnf(i_0_53, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_127, plain, (v1_zfmisc_1(X2)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_279, plain, (v7_struct_0(X1)|~l2_struct_0(X1)|~v9_struct_0(esk17_1(X1),X1))).#trainneg
cnf(i_0_251, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_335, plain, (v1_xboole_0(X2)|v5_relat_1(esk38_2(X1,X2),X2))).#trainneg
cnf(i_0_333, plain, (v1_xboole_0(X2)|v1_partfun1(esk38_2(X1,X2),X1))).#trainneg
cnf(i_0_336, plain, (v1_xboole_0(X2)|v4_relat_1(esk38_2(X1,X2),X1))).#trainneg
cnf(i_0_20, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_269, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k4_numbers)|~v5_membered(X1))).#trainneg
cnf(i_0_268, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k3_numbers)|~v4_membered(X1))).#trainneg
cnf(i_0_430, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1518, plain, (u1_struct_0(X1)=u1_struct_0(esk13_0)|~l1_struct_0(X1)|~v2_struct_0(esk13_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1572, plain, (v1_valued_0(k2_zfmisc_1(X1,k4_numbers)))).#trainneg
cnf(i_0_1664, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_938, plain, (v1_xcmplx_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1791, plain, (v2_membered(k5_numbers))).#trainneg
cnf(i_0_1735, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1692, plain, (v1_rat_1(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1693, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_950, plain, (k1_seq_1(esk55_1(X1),X2)=k1_funct_1(esk55_1(X1),X2))).#trainneg
cnf(i_0_954, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_311, plain, (v2_struct_0(X1)|m1_subset_1(esk29_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_374, plain, (v2_struct_0(X1)|m1_subset_1(esk47_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_313, plain, (v7_struct_0(X1)|m1_subset_1(esk30_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_28, plain, (v2_valued_0(X2)|~v1_relat_1(X1)|~v2_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_31, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_37, plain, (v4_valued_0(X2)|~v1_relat_1(X1)|~v4_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_270, plain, (v1_funct_1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funcop_1(X1))).#trainneg
cnf(i_0_271, plain, (v1_relat_1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funcop_1(X1))).#trainneg
cnf(i_0_252, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_253, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_254, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1915, plain, (v1_funct_1(k1_funct_1(esk18_0,X1)))).#trainneg
cnf(i_0_1687, plain, (v4_membered(k5_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_963, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1751, plain, (v1_zfmisc_1(k5_numbers)|~v1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_1939, plain, (v4_membered(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_1716, plain, (v5_membered(k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1617, plain, (v5_membered(k6_numbers)|~v5_membered(X1))).#trainneg
cnf(i_0_1948, plain, (v5_membered(k6_numbers))).#trainneg
cnf(i_0_965, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_966, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1953, plain, (v3_membered(k6_numbers))).#trainneg
cnf(i_0_257, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_232, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_231, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_230, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_229, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_227, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_88, plain, (v1_xboole_0(X1)|v1_subset_1(X2,X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_360, plain, (m1_subset_1(esk42_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_175, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1519, plain, (u1_struct_0(X1)=u1_struct_0(esk12_0)|~l1_struct_0(X1)|~v2_struct_0(esk12_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1584, plain, (v5_relat_1(esk40_1(k4_numbers),k3_numbers))).#trainneg
cnf(i_0_1977, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(esk3_0)|~v3_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_970, plain, (v1_relat_1(k4_normsp_0(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1868, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2073, plain, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1883, plain, (v2_valued_0(k5_numbers)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1628, plain, (v6_ordinal1(k6_numbers)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1639, plain, (v6_membered(k6_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_967, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_968, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_2108, plain, (v6_membered(k6_numbers))).#trainneg
cnf(i_0_29, plain, (v10_algstr_0(X2,X1)|~v13_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2131, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X2)|~v2_struct_0(X2)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2132, plain, (X1=k1_normsp_0(X2,k4_struct_0(X2))|v2_struct_0(X2)|~v1_xboole_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2155, plain, (k1_normsp_0(X1,k4_struct_0(X1))=k1_normsp_0(X2,k4_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X2)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2168, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),np__1)=k6_numbers|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_2169, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),esk3_0)=k6_numbers|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_1676, plain, (v1_valued_0(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_971, plain, (v1_funct_1(k4_normsp_0(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_2104, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k1_numbers))).#trainneg
cnf(i_0_1898, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1913, plain, (v4_valued_0(k5_numbers)|~v4_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1650, plain, (v1_membered(k6_numbers)|~v1_membered(X1))).#trainneg
cnf(i_0_2214, plain, (v1_membered(k6_numbers))).#trainneg
cnf(i_0_972, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_973, plain, (v1_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_2172, negated_conjecture, (k1_normsp_0(esk1_0,k1_funct_1(esk2_0,X1))=k6_numbers|~l2_normsp_0(esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2199, negated_conjecture, (k1_normsp_0(X1,k4_struct_0(X1))=k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1))|v2_struct_0(X1)|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2202, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_1520, plain, (u1_struct_0(X1)=u1_struct_0(esk11_0)|~l1_struct_0(X1)|~v2_struct_0(esk11_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2018, negated_conjecture, (v3_xxreal_0(esk3_0)|v2_xxreal_0(esk3_0)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_974, plain, (v1_partfun1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1)|v1_xboole_0(X2)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2))).#trainneg
cnf(i_0_2105, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1940, plain, (v4_membered(k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_2064, plain, (v5_membered(esk42_0)|~v5_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1661, plain, (v2_membered(k6_numbers)|~v2_membered(X1))).#trainneg
cnf(i_0_2268, plain, (v2_membered(k6_numbers))).#trainneg
cnf(i_0_975, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(X1)))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X2,esk14_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_982, plain, (k2_xcmplx_0(esk14_1(X1),k6_numbers)=esk14_1(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_994, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2207, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k1_numbers))).#trainneg
cnf(i_0_2065, plain, (v6_ordinal1(esk42_0)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2066, plain, (v6_membered(esk42_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1684, plain, (v4_membered(k6_numbers)|~v4_membered(X1))).#trainneg
cnf(i_0_996, plain, (k2_xcmplx_0(k1_normsp_1(X1,X2,X3),k6_numbers)=k1_normsp_1(X1,X2,X3)|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2252, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2294, plain, (v4_membered(k6_numbers))).#trainneg
cnf(i_0_432, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_117, plain, (~v1_xboole_0(X1)|~v1_subset_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_255, plain, (v1_rat_1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_2360, negated_conjecture, (v1_rat_1(k1_funct_1(esk2_0,esk36_0))|~v5_relat_1(esk2_0,k3_numbers)|~m1_subset_1(esk44_0,k5_numbers)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_256, plain, (v1_int_1(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_2369, negated_conjecture, (v1_int_1(k1_funct_1(esk2_0,esk36_0))|~v5_relat_1(esk2_0,k4_numbers)|~m1_subset_1(esk44_0,k5_numbers)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_201, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1521, plain, (u1_struct_0(X1)=u1_struct_0(esk10_0)|~l1_struct_0(X1)|~v2_struct_0(esk10_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1917, plain, (v1_relat_1(k1_funct_1(esk18_0,X1)))).#trainneg
cnf(i_0_2326, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k2_xcmplx_0(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2208, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k4_numbers))).#trainneg
cnf(i_0_2069, plain, (v1_membered(esk42_0)|~v1_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2070, plain, (v2_membered(esk42_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1748, plain, (v1_zfmisc_1(k6_numbers)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_2408, plain, (v1_zfmisc_1(k6_numbers)|v1_xboole_0(X1))).#trainneg
cnf(i_0_2415, plain, (v1_zfmisc_1(k6_numbers))).#trainneg
cnf(i_0_431, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_228, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_276, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_35, plain, (v5_relat_1(X2,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_33, plain, (v5_relat_1(X2,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_284, plain, (v1_funct_2(esk19_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_182, plain, (v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_184, plain, (v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_234, plain, (v2_struct_0(X1)|v5_relat_1(k2_normsp_0(X1,X2),k1_numbers)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l1_normsp_0(X1)|~v5_relat_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_289, plain, (m1_subset_1(esk19_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2434, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_2428, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk41_1(X2)))).#trainneg
cnf(i_0_2106, plain, (v2_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v3_valued_0(k1_numbers))).#trainneg
cnf(i_0_2330, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k2_nat_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1010, plain, (v1_xreal_0(esk14_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_2500, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2071, plain, (v3_membered(esk42_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2072, plain, (v4_membered(esk42_0)|~v4_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2107, plain, (v6_ordinal1(k6_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2570, plain, (v6_ordinal1(k6_numbers))).#trainneg
cnf(i_0_1011, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k3_series_1(X1))|m1_subset_1(esk14_1(k3_series_1(X1)),k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_424, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_338, plain, (v1_xboole_0(X2)|m1_subset_1(esk38_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_225, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk15_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_39, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_42, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_44, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_2697, plain, (v3_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2700, plain, (k1_seq_1(k4_normsp_0(X1,X2),X3)=k1_funct_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~v1_relat_1(k4_normsp_0(X1,X2))|~v1_funct_1(k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_60, plain, (v4_valued_0(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_114, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_2432, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk14_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1963, plain, (v1_valued_0(k2_zfmisc_1(X1,k6_numbers)))).#trainneg
cnf(i_0_2333, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=esk2_0|k1_funct_1(esk2_0,k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0)),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2516, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2725, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1024, plain, (v1_xreal_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_555, negated_conjecture, (v1_partfun1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_103, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_2760, negated_conjecture, (v5_relat_1(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2775, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_58, plain, (v5_relat_1(X2,k4_numbers)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_46, plain, (v5_relat_1(X2,k3_numbers)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_1113, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_1026, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_2209, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k3_numbers))).#trainneg
cnf(i_0_2666, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2667, plain, (v1_valued_0(k6_numbers)|~v1_membered(X1))).#trainneg
cnf(i_0_2815, plain, (v1_valued_0(k6_numbers))).#trainneg
cnf(i_0_1028, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_1029, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_167, plain, (k1_funct_1(X1,k6_numbers)=k1_funct_1(X2,k6_numbers)|X1!=k2_series_1(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_192, plain, (v1_funct_1(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_111, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_82, plain, (v1_partfun1(X2,X1)|~v1_xboole_0(X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_166, plain, (k2_xcmplx_0(k1_funct_1(X2,X1),k1_funct_1(X3,k2_nat_1(X1,np__1)))=k1_funct_1(X2,k2_nat_1(X1,np__1))|X2!=k2_series_1(X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_valued_0(X3)|~v1_valued_0(X2)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_188, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_417, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),esk56_2(X1,X2))=k1_seq_1(k4_normsp_0(X1,X2),esk56_2(X1,X2))|v2_struct_0(X1)|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v3_normsp_0(X1)|~v4_normsp_0(X1)|~v2_clvect_1(X1)|~v3_clvect_1(X1)|~v4_clvect_1(X1)|~v5_clvect_1(X1)|~v8_clvect_1(X1)|~l2_clvect_1(X1)|~v1_funct_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_168, plain, (X1=k4_normsp_0(X2,X3)|v2_struct_0(X2)|k3_funct_2(k5_numbers,k1_numbers,X1,esk5_3(X2,X3,X1))!=k1_normsp_0(X2,k3_funct_2(k5_numbers,u1_struct_0(X2),X3,esk5_3(X2,X3,X1)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~l1_normsp_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2)))))).#trainneg
cnf(i_0_551, plain, (m1_subset_1(k2_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1522, plain, (u1_struct_0(X1)=u1_struct_0(esk9_0)|~l1_struct_0(X1)|~v2_struct_0(esk9_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2082, plain, (m1_subset_1(k1_funct_1(esk42_0,X1),k1_numbers))).#trainneg
cnf(i_0_2262, negated_conjecture, (v2_xxreal_0(esk3_0)|~v1_xboole_0(esk3_0)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_1057, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk4_2(esk2_0,esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(k2_nat_1(esk4_2(esk2_0,esk2_0),np__1),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_1030, negated_conjecture, (v1_rat_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_2211, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2680, negated_conjecture, (v2_valued_0(esk2_0)|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2681, plain, (v2_valued_0(k6_numbers)|~v2_membered(X1))).#trainneg
cnf(i_0_3036, plain, (v2_valued_0(k6_numbers))).#trainneg
cnf(i_0_1060, plain, (v1_xxreal_0(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1061, plain, (v1_xreal_0(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_558, plain, (m1_subset_1(esk4_2(X1,X2),k1_numbers)|X2=k2_series_1(X1)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_567, plain, (v1_partfun1(k3_series_1(X1),k5_numbers)|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_577, plain, (v2_struct_0(X1)|~v1_xboole_0(k4_normsp_0(X1,X2))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3067, negated_conjecture, (~v1_xboole_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_593, plain, (m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)|X3=k4_normsp_0(X1,X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_594, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_numbers)|X3=k4_normsp_0(X1,X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_623, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_624, plain, (m1_subset_1(k2_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_630, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_631, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3022, plain, (v1_xreal_0(k1_funct_1(esk42_0,X1)))).#trainneg
cnf(i_0_895, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_2212, plain, (v4_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2694, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2695, plain, (v3_valued_0(k6_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_3155, plain, (v3_valued_0(k6_numbers))).#trainneg
cnf(i_0_1064, plain, (k2_xcmplx_0(k5_numbers,k6_numbers)=k5_numbers|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_721, plain, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_722, plain, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3),k1_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_743, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,esk44_0))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_744, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,esk44_0),u1_struct_0(X1))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_746, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,esk36_0))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_747, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,esk36_0),u1_struct_0(X1))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_751, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_754, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,esk36_0))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk44_0)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(esk44_0,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk44_0,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_768, negated_conjecture, (v1_xcmplx_0(k1_funct_1(esk2_0,np__1))|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1523, plain, (u1_struct_0(X1)=u1_struct_0(esk8_0)|~l1_struct_0(X1)|~v2_struct_0(esk8_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2083, plain, (m1_subset_1(k1_funct_1(esk24_0,X1),k1_numbers))).#trainneg
cnf(i_0_882, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_nat_1(np__1,esk4_2(esk2_0,esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1))|~v4_relat_1(esk2_0,k5_numbers)|~v1_valued_0(esk2_0)|~m1_subset_1(k2_nat_1(np__1,esk4_2(esk2_0,esk2_0)),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_1075, plain, (v1_relat_1(k4_normsp_0(X1,k6_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3147, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v1_relat_1(k5_numbers)|~v5_relat_1(k1_numbers,k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2711, negated_conjecture, (v4_valued_0(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2712, plain, (v4_valued_0(k6_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_1067, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_3294, plain, (v4_valued_0(k6_numbers))).#trainneg
cnf(i_0_769, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,np__1))|~v2_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_770, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,np__1))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_771, negated_conjecture, (v1_rat_1(k1_funct_1(esk2_0,np__1))|~v4_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_772, negated_conjecture, (v1_int_1(k1_funct_1(esk2_0,np__1))|~v5_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_773, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,np__1))|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_774, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,np__1))|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_775, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,np__1))|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_779, plain, (v1_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_783, plain, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_785, plain, (v1_rat_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v4_membered(k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3282, plain, (v1_xreal_0(k1_funct_1(esk24_0,X1)))).#trainneg
cnf(i_0_3292, negated_conjecture, (v4_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_2333, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=esk2_0|k1_funct_1(esk2_0,k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(k2_nat_1(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1076, plain, (v1_funct_1(k4_normsp_0(X1,k6_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3148, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v1_relat_1(k5_numbers)|~v5_relat_1(k1_numbers,k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2814, negated_conjecture, (v1_valued_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3358, negated_conjecture, (v3_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3370, negated_conjecture, (v1_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1068, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1077, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_2(k6_numbers,X1,X2)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3369, negated_conjecture, (k1_seq_1(esk2_0,X1)=k1_funct_1(esk2_0,X1)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_794, negated_conjecture, (v1_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_795, negated_conjecture, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_796, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_797, negated_conjecture, (v1_rat_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v4_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_798, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0),k1_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_799, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),np__1)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1524, plain, (u1_struct_0(X1)=u1_struct_0(esk7_0)|~l1_struct_0(X1)|~v2_struct_0(esk7_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1586, plain, (v5_relat_1(esk19_2(X1,k4_numbers),k3_numbers))).#trainneg
cnf(i_0_2330, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(k2_nat_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1079, plain, (v1_partfun1(k6_numbers,X1)|v1_xboole_0(X2)|~v1_funct_2(k6_numbers,X1,X2))).#trainneg
cnf(i_0_3035, negated_conjecture, (v2_valued_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2029, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v2_xxreal_0(k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1574, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1081, plain, (v1_xcmplx_0(esk41_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1082, plain, (v1_xxreal_0(esk41_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_805, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,X2))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),np__1)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3473, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,np__1))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),np__1)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_806, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_808, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,X2))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3443, plain, (v1_valued_0(k2_zfmisc_1(X1,esk39_0)))).#trainneg
cnf(i_0_2326, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),np__1),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1083, plain, (v1_xreal_0(esk41_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2405, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k5_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3151, negated_conjecture, (v3_valued_0(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1084, plain, (v1_rat_1(esk41_1(X1))|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1085, plain, (v1_int_1(esk41_1(X1))|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3516, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_809, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk2_0,X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1525, plain, (u1_struct_0(X1)=u1_struct_0(esk6_0)|~l1_struct_0(X1)|~v2_struct_0(esk6_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_3444, plain, (v1_valued_0(k2_zfmisc_1(X1,esk31_0)))).#trainneg
cnf(i_0_812, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X1,esk2_0))))!=k1_funct_1(esk2_0,k2_nat_1(np__1,esk4_2(X1,esk2_0)))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1086, plain, (v7_ordinal1(esk41_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2812, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k5_numbers)|~v5_relat_1(k1_numbers,k3_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3364, negated_conjecture, (v1_valued_0(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1087, plain, (v1_relat_1(k4_normsp_0(X1,esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1088, plain, (v1_funct_1(k4_normsp_0(X1,esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_814, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(X1,k2_xcmplx_0(np__1,esk4_2(X1,esk2_0))))!=k1_funct_1(esk2_0,k2_xcmplx_0(np__1,esk4_2(X1,esk2_0)))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v7_ordinal1(esk4_2(X1,esk2_0))|~v1_valued_0(X1)|~v1_valued_0(esk2_0)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_816, negated_conjecture, (X1=k2_series_1(esk2_0)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(esk2_0,X1)),k1_funct_1(esk2_0,np__1))!=k1_funct_1(X1,k2_xcmplx_0(esk4_2(esk2_0,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(esk2_0,X1),k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk2_0,X1),np__1),k5_numbers))).#trainneg
cnf(i_0_3445, plain, (v1_valued_0(k2_zfmisc_1(X1,esk20_0)))).#trainneg
cnf(i_0_813, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(X1,k2_xcmplx_0(esk4_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1089, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(esk41_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(esk41_1(k2_zfmisc_1(X1,X2)),X1,X2)|~v1_funct_1(esk41_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3436, negated_conjecture, (v2_valued_0(esk2_0)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3749, negated_conjecture, (v2_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1091, plain, (v1_partfun1(esk41_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(X2)|~v1_funct_2(esk41_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_1092, plain, (v1_xboole_0(esk41_1(X1))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X2,esk41_1(X1)))).#trainneg
cnf(i_0_817, negated_conjecture, (X1=k2_series_1(esk2_0)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(esk2_0,X1)),k1_funct_1(esk2_0,np__1))!=k1_funct_1(X1,k2_xcmplx_0(np__1,esk4_2(esk2_0,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v7_ordinal1(esk4_2(esk2_0,X1))|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k2_xcmplx_0(np__1,esk4_2(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_818, negated_conjecture, (X1=k2_series_1(esk2_0)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(esk2_0,X1)),k1_funct_1(esk2_0,np__1))!=k1_funct_1(X1,k2_nat_1(esk4_2(esk2_0,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k2_nat_1(esk4_2(esk2_0,X1),np__1),k5_numbers))).#trainneg
cnf(i_0_819, negated_conjecture, (X1=k2_series_1(esk2_0)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(esk2_0,X1)),k1_funct_1(esk2_0,np__1))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(esk2_0,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v7_ordinal1(esk4_2(esk2_0,X1))|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k2_nat_1(np__1,esk4_2(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_820, negated_conjecture, (X1=k2_series_1(esk2_0)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(esk2_0,X1)),k1_funct_1(esk2_0,np__1))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(esk2_0,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(esk2_0,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(esk2_0,X1),k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(k2_nat_1(np__1,esk4_2(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_829, negated_conjecture, (v1_xcmplx_0(k1_funct_1(esk2_0,X1))|~v1_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_830, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,X1))|~v2_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_831, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,X1))|~v3_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_832, negated_conjecture, (v1_rat_1(k1_funct_1(esk2_0,X1))|~v4_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_833, negated_conjecture, (v1_int_1(k1_funct_1(esk2_0,X1))|~v5_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_834, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,X1))|~v6_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3835, negated_conjecture, (m2_subset_1(k1_funct_1(esk2_0,X1),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(esk2_0,X1),k1_numbers)|~v6_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2437, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1774, plain, (v5_relat_1(k2_zfmisc_1(X1,k4_numbers),k4_numbers))).#trainneg
cnf(i_0_815, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(X1,k2_nat_1(esk4_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_nat_1(esk4_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(esk2_0,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1097, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2074, plain, (v1_zfmisc_1(esk42_0)|~v1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2099, negated_conjecture, (v2_xxreal_0(esk3_0)|~v1_xboole_0(esk3_0)|~v3_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_1400, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1102, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1115, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk4_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(np__1,esk4_2(X2,X1))))!=k1_funct_1(X1,k2_xcmplx_0(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_843, negated_conjecture, (v1_int_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v5_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_844, plain, (v1_int_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v5_membered(k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_851, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_855, negated_conjecture, (v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_856, plain, (v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_858, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,np__1))=k1_funct_1(k4_normsp_0(X1,esk2_0),X2)|v2_struct_0(X1)|~m1_subset_1(k4_normsp_0(X1,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(k4_normsp_0(X1,esk2_0),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,esk2_0))|~l1_normsp_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_3957, negated_conjecture, (k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1))=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_3980, negated_conjecture, (k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1))=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1778, plain, (v5_relat_1(k2_zfmisc_1(X1,k3_numbers),k3_numbers))).#trainneg
cnf(i_0_1118, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_xcmplx_0(esk4_2(X2,X1),np__1)),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_xcmplx_0(esk4_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k5_numbers)|~m1_subset_1(k1_funct_1(X2,k2_xcmplx_0(esk4_2(X2,X1),np__1)),k5_numbers))).#trainneg
cnf(i_0_2266, plain, (v5_membered(esk42_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2406, plain, (v1_membered(esk42_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1119, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(esk4_2(X2,X1),np__1)),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(esk4_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k5_numbers)|~m1_subset_1(k1_funct_1(X2,k2_nat_1(esk4_2(X2,X1),np__1)),k5_numbers))).#trainneg
cnf(i_0_1120, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v7_ordinal1(esk4_2(X2,X1))|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k5_numbers)|~m1_subset_1(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k5_numbers))).#trainneg
cnf(i_0_3998, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),np__1)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4000, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),esk3_0)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4002, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_4005, negated_conjecture, (k1_normsp_0(esk1_0,k1_funct_1(esk2_0,X1))=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1688, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1954, plain, (v5_relat_1(k2_zfmisc_1(X1,k6_numbers),k4_numbers))).#trainneg
cnf(i_0_1121, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k5_numbers)|~m1_subset_1(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k5_numbers))).#trainneg
cnf(i_0_2407, plain, (v2_membered(esk42_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2563, plain, (v3_membered(esk42_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3152, negated_conjecture, (v3_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1689, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2341, plain, (v5_relat_1(k2_zfmisc_1(X1,k6_numbers),k3_numbers))).#trainneg
cnf(i_0_2567, plain, (v4_membered(esk42_0)|~v5_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3293, negated_conjecture, (v4_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3365, negated_conjecture, (v1_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1690, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2764, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_4030, plain, (v1_membered(esk42_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_4159, plain, (v2_membered(esk42_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3437, negated_conjecture, (v2_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3535, negated_conjecture, (v3_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_859, negated_conjecture, (v1_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_860, negated_conjecture, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_861, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_862, negated_conjecture, (v1_rat_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v4_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1691, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2779, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_1147, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3356, negated_conjecture, (v5_relat_1(esk2_0,k4_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_4170, plain, (v4_membered(esk42_0)|~v6_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_3710, negated_conjecture, (v1_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3750, negated_conjecture, (v2_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1144, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_863, negated_conjecture, (v1_int_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v5_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_864, negated_conjecture, (v7_ordinal1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_865, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2),k1_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_866, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),np__1)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1694, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2854, plain, (v1_partfun1(esk19_2(X1,X1),X1))).#trainneg
cnf(i_0_3357, negated_conjecture, (v5_relat_1(esk2_0,k3_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_4359, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1172, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_xcmplx_0(esk4_2(X2,X1),np__1)),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_xcmplx_0(esk4_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k1_numbers)|~m1_subset_1(k1_funct_1(X2,k2_xcmplx_0(esk4_2(X2,X1),np__1)),k1_numbers))).#trainneg
cnf(i_0_868, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X3)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4140, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1173, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(esk4_2(X2,X1),np__1)),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(esk4_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k1_numbers)|~m1_subset_1(k1_funct_1(X2,k2_nat_1(esk4_2(X2,X1),np__1)),k1_numbers))).#trainneg
cnf(i_0_4364, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1174, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v7_ordinal1(esk4_2(X2,X1))|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k1_numbers)|~m1_subset_1(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k1_numbers))).#trainneg
cnf(i_0_1175, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k1_funct_1(X1,esk4_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(np__1,esk4_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X2,X1),k5_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_2(X2,X1)),k1_numbers)|~m1_subset_1(k1_funct_1(X2,k2_nat_1(np__1,esk4_2(X2,X1))),k1_numbers))).#trainneg
cnf(i_0_870, negated_conjecture, (k1_normsp_0(X1,k1_funct_1(esk2_0,X3))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_879, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)=k1_funct_1(esk2_0,np__1)|~v1_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4169, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4176, plain, (v5_relat_1(k2_zfmisc_1(X1,k4_numbers),k3_numbers))).#trainneg
cnf(i_0_2292, plain, (v6_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1433, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1434, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_4365, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_4462, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_4175, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2080, plain, (m1_subset_1(k1_funct_1(esk55_1(X1),X2),k1_numbers))).#trainneg
cnf(i_0_4028, plain, (v5_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1435, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1441, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1204, plain, (v1_int_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_969, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|m1_subset_1(esk14_1(esk2_0),k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1149, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1150, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4232, plain, (v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4584, plain, (v1_xreal_0(k1_funct_1(esk55_1(X1),X2)))).#trainneg
cnf(i_0_1442, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1206, plain, (v1_xcmplx_0(k1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_4164, plain, (v3_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1443, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1616, negated_conjecture, (v5_membered(esk2_0)|~v5_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1207, plain, (v1_xcmplx_0(k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1209, plain, (v1_xcmplx_0(k1_funct_1(k6_numbers,X1))|v2_struct_0(X2)|~v1_membered(u1_struct_0(X2))|~l1_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4569, plain, (v1_xcmplx_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2081, plain, (m1_subset_1(k1_funct_1(esk53_1(X1),X2),k1_numbers))).#trainneg
cnf(i_0_3153, negated_conjecture, (v3_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1211, plain, (v1_xxreal_0(k1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_4178, plain, (v1_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1627, negated_conjecture, (v6_ordinal1(esk2_0)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1638, negated_conjecture, (v6_membered(esk2_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4678, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1212, plain, (v1_xxreal_0(k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1214, plain, (v1_xxreal_0(k1_funct_1(k6_numbers,X1))|v2_struct_0(X2)|~v2_membered(u1_struct_0(X2))|~l1_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4699, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X2),k5_numbers))).#trainneg
cnf(i_0_4720, plain, (k2_xcmplx_0(k3_series_1(X1),X2)=k2_xcmplx_0(X2,k3_series_1(X1))|~m1_subset_1(k1_zfmisc_1(X2),k5_numbers)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4728, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(k1_zfmisc_1(X2),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4798, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1217, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4866, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_3154, negated_conjecture, (v3_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1252, plain, (v1_xreal_0(k1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_4181, plain, (v2_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1649, negated_conjecture, (v1_membered(esk2_0)|~v1_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1660, negated_conjecture, (v2_membered(esk2_0)|~v2_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4233, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1253, plain, (v1_xreal_0(k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1255, plain, (v1_xreal_0(k1_funct_1(k6_numbers,X1))|v2_struct_0(X2)|~v3_membered(u1_struct_0(X2))|~l1_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1221, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1222, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4577, plain, (v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4683, plain, (v1_xreal_0(k1_funct_1(esk53_1(X1),X2)))).#trainneg
cnf(i_0_3366, negated_conjecture, (v1_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1257, plain, (v1_rat_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_4255, plain, (v4_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1671, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1683, negated_conjecture, (v4_membered(esk2_0)|~v4_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1259, plain, (v1_rat_1(k1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1260, plain, (v1_rat_1(k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1237, plain, (m2_subset_1(esk51_0,k1_numbers,k5_numbers)|~m1_subset_1(esk51_0,k1_numbers))).#trainneg
cnf(i_0_1238, plain, (m2_subset_1(esk46_0,k1_numbers,k5_numbers)|~m1_subset_1(esk46_0,k1_numbers))).#trainneg
cnf(i_0_1239, plain, (m2_subset_1(esk14_1(X1),k1_numbers,k5_numbers)|~m1_subset_1(esk14_1(X1),k1_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_1241, plain, (m2_subset_1(k1_normsp_0(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|~m1_subset_1(k1_normsp_0(X1,X2),k1_numbers)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1244, plain, (m2_subset_1(k3_series_1(X1),k1_numbers,k5_numbers)|~m1_subset_1(k3_series_1(X1),k1_numbers)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1245, plain, (m2_subset_1(k4_normsp_0(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|~m1_subset_1(k4_normsp_0(X1,X2),k1_numbers)|~l1_normsp_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1246, plain, (m2_subset_1(k4_struct_0(X1),k1_numbers,k5_numbers)|~m1_subset_1(k4_struct_0(X1),k1_numbers)|~l2_struct_0(X1)|~v6_membered(u1_struct_0(X1)))).#trainneg
cnf(i_0_1247, plain, (m2_subset_1(k1_normsp_1(X1,X2,X3),k1_numbers,k5_numbers)|v2_struct_0(X1)|~m1_subset_1(k1_normsp_1(X1,X2,X3),k1_numbers)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1248, plain, (m2_subset_1(k5_numbers,k1_numbers,k5_numbers)|~m1_subset_1(k5_numbers,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1382, plain, (v1_rat_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3164, plain, (m1_subset_1(k1_seq_1(k6_numbers,X1),k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3367, negated_conjecture, (v1_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1262, plain, (v1_rat_1(k1_funct_1(k6_numbers,X1))|v2_struct_0(X2)|~v4_membered(u1_struct_0(X2))|~l1_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1747, negated_conjecture, (v1_zfmisc_1(esk2_0)|~v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2056, plain, (v3_ordinal1(esk42_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_4360, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4463, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_1264, plain, (v1_int_1(k1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|~l1_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1265, plain, (v1_int_1(k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2))|v2_struct_0(X1)|~l1_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1267, plain, (v1_int_1(k1_funct_1(k6_numbers,X1))|v2_struct_0(X2)|~l1_struct_0(X2)|~v5_membered(u1_struct_0(X2))|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1268, plain, (m1_subset_1(X1,k5_numbers)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1281, plain, (v7_ordinal1(k1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1282, plain, (v7_ordinal1(k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),X2))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1284, plain, (v7_ordinal1(k1_funct_1(k6_numbers,X1))|v2_struct_0(X2)|~v6_membered(u1_struct_0(X2))|~l1_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X2))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1383, plain, (v1_int_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_5032, plain, (v1_xreal_0(k1_seq_1(k6_numbers,X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3438, negated_conjecture, (v2_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1295, plain, (v1_xcmplx_0(esk23_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2057, plain, (v1_xcmplx_0(esk42_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2058, plain, (v1_xxreal_0(esk42_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_4570, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1296, plain, (v1_xxreal_0(esk23_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1297, plain, (v1_xreal_0(esk23_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1298, plain, (v1_rat_1(esk23_1(X1))|v1_xboole_0(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1299, plain, (v1_int_1(esk23_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1300, plain, (v7_ordinal1(esk23_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1301, plain, (v1_relat_1(k4_normsp_0(X1,esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1302, plain, (v1_funct_1(k4_normsp_0(X1,esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1303, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(esk23_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(esk23_1(k2_zfmisc_1(X1,X2)),X1,X2)|~v1_funct_1(esk23_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1304, plain, (v1_relat_1(esk23_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1305, plain, (v1_partfun1(esk23_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(esk23_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_1306, plain, (v1_xcmplx_0(k1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1396, plain, (v1_int_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5077, plain, (v1_xreal_0(k1_funct_1(k6_numbers,X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_3439, negated_conjecture, (v2_valued_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1307, plain, (v1_xxreal_0(k1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_2059, plain, (v1_xreal_0(esk42_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2060, plain, (v1_rat_1(esk42_0)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_4578, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4664, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1308, plain, (v1_xreal_0(k1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1309, plain, (v1_rat_1(k1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1310, plain, (v1_int_1(k1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1311, plain, (v7_ordinal1(k1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1312, plain, (v1_xboole_0(esk23_1(X1))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X2,esk23_1(X1)))).#trainneg
cnf(i_0_1313, plain, (v1_xcmplx_0(esk48_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1314, plain, (v1_xxreal_0(esk48_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1315, plain, (v1_xreal_0(esk48_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1316, plain, (v1_rat_1(esk48_1(X1))|v1_xboole_0(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1317, plain, (v1_int_1(esk48_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1318, plain, (v7_ordinal1(esk48_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1319, plain, (v1_relat_1(k4_normsp_0(X1,esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1320, plain, (v1_funct_1(k4_normsp_0(X1,esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1321, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(esk48_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(esk48_1(k2_zfmisc_1(X1,X2)),X1,X2)|~v1_funct_1(esk48_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1397, plain, (v1_rat_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1322, plain, (v1_relat_1(esk48_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1783, plain, (u1_struct_0(esk12_0)=u1_struct_0(esk13_0)|~v2_struct_0(esk13_0)|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_2061, plain, (v1_int_1(esk42_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2062, plain, (v7_ordinal1(esk42_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_4679, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4975, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1323, plain, (v1_partfun1(esk48_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(esk48_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_1324, plain, (v1_xcmplx_0(k1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1325, plain, (v1_xxreal_0(k1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1326, plain, (v1_xreal_0(k1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1327, plain, (v1_rat_1(k1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1328, plain, (v1_int_1(k1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1329, plain, (v7_ordinal1(k1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1330, plain, (v1_xboole_0(esk48_1(X1))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X2,esk48_1(X1)))).#trainneg
cnf(i_0_1331, plain, (v1_xcmplx_0(esk52_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1332, plain, (v1_xxreal_0(esk52_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1333, plain, (v1_xreal_0(esk52_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1334, plain, (v1_rat_1(esk52_1(X1))|v1_xboole_0(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1335, plain, (v1_int_1(esk52_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1336, plain, (v7_ordinal1(esk52_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1407, plain, (v1_rat_1(k6_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1337, plain, (v1_relat_1(k4_normsp_0(X1,esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1784, plain, (u1_struct_0(esk13_0)=u1_struct_0(esk11_0)|~v2_struct_0(esk13_0)|~v2_struct_0(esk11_0))).#trainneg
cnf(i_0_4671, negated_conjecture, (v5_membered(esk2_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1338, plain, (v1_funct_1(k4_normsp_0(X1,esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1339, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(esk52_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(esk52_1(k2_zfmisc_1(X1,X2)),X1,X2)|~v1_funct_1(esk52_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1340, plain, (v1_relat_1(esk52_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1341, plain, (v1_partfun1(esk52_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(esk52_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_1342, plain, (v1_xcmplx_0(k1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1343, plain, (v1_xxreal_0(k1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1344, plain, (v1_xreal_0(k1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1345, plain, (v1_rat_1(k1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1346, plain, (v1_int_1(k1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1347, plain, (v7_ordinal1(k1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1348, plain, (v1_xboole_0(esk52_1(X1))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X2,esk52_1(X1)))).#trainneg
cnf(i_0_1349, plain, (v1_xcmplx_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1350, plain, (v1_xxreal_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1351, plain, (v1_xreal_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4234, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1352, plain, (v1_rat_1(esk54_1(X1))|v1_zfmisc_1(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1785, plain, (u1_struct_0(esk13_0)=u1_struct_0(esk10_0)|~v2_struct_0(esk13_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_4915, negated_conjecture, (v1_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4917, negated_conjecture, (v2_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1353, plain, (v1_int_1(esk54_1(X1))|v1_zfmisc_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1354, plain, (v7_ordinal1(esk54_1(X1))|v1_zfmisc_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1355, plain, (v1_relat_1(k4_normsp_0(X1,esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1356, plain, (v1_funct_1(k4_normsp_0(X1,esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1357, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_zfmisc_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(esk54_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(esk54_1(k2_zfmisc_1(X1,X2)),X1,X2)|~v1_funct_1(esk54_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1358, plain, (v1_relat_1(esk54_1(k2_zfmisc_1(X1,X2)))|v1_zfmisc_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1359, plain, (v1_partfun1(esk54_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(X2)|v1_zfmisc_1(k2_zfmisc_1(X1,X2))|~v1_funct_2(esk54_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_1360, plain, (v1_xcmplx_0(k1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1361, plain, (v1_xxreal_0(k1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1362, plain, (v1_xreal_0(k1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1363, plain, (v1_rat_1(k1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1364, plain, (v1_int_1(k1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1365, plain, (v7_ordinal1(k1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),X2))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1366, plain, (v1_xboole_0(esk54_1(X1))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|v1_zfmisc_1(X1)|~m1_subset_1(X2,esk54_1(X1)))).#trainneg
cnf(i_0_4235, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1786, plain, (u1_struct_0(esk13_0)=u1_struct_0(esk9_0)|~v2_struct_0(esk13_0)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_4979, negated_conjecture, (v3_membered(esk2_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4984, negated_conjecture, (v4_membered(esk2_0)|~v5_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1380, negated_conjecture, (v1_int_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1372, negated_conjecture, (m2_subset_1(esk2_0,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_1381, negated_conjecture, (v1_rat_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1389, plain, (v1_xreal_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1390, negated_conjecture, (v7_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1391, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1392, negated_conjecture, (v1_int_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1393, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1394, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1395, negated_conjecture, (v1_rat_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4571, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_5036, plain, (m1_subset_1(k1_funct_1(k6_numbers,X1),k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1401, plain, (v1_xreal_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1787, plain, (u1_struct_0(esk13_0)=u1_struct_0(esk8_0)|~v2_struct_0(esk13_0)|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_5045, plain, (v3_ordinal1(esk42_0)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_5085, plain, (v1_xcmplx_0(esk42_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_1402, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1403, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1404, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1405, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1406, negated_conjecture, (v1_rat_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1414, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1415, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1416, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1417, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1418, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1419, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1420, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k1_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1422, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_1424, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_4572, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1432, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_5287, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k4_numbers))).#trainneg
cnf(i_0_5086, plain, (v1_xxreal_0(esk42_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_5105, plain, (v1_xreal_0(esk42_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2729, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1425, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1426, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_5308, plain, (k2_zfmisc_1(X1,X2)=k1_normsp_0(X3,k4_struct_0(X3))|v2_struct_0(X3)|~l2_normsp_0(X3)|~v4_normsp_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5432, plain, (v1_funct_2(X1,X2,X3)|v2_struct_0(X4)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_normsp_0(X4,k4_struct_0(X4))))|~l2_normsp_0(X4)|~v1_xboole_0(X3)|~v4_normsp_0(X4))).#trainneg
cnf(i_0_5435, plain, (v1_partfun1(X1,X2)|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k1_normsp_0(X3,k4_struct_0(X3))))|~v1_funct_2(X1,X2,X2)|~l2_normsp_0(X3)|~v1_xboole_0(X2)|~v4_normsp_0(X3))).#trainneg
cnf(i_0_1427, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1428, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1429, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1430, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k4_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1436, plain, (v1_xreal_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1437, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1438, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_5305, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1439, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1440, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1444, plain, (v1_xreal_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1445, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1446, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1447, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1448, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1459, plain, (k2_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3),k6_numbers)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_5684, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1466, plain, (k2_xcmplx_0(k5_numbers,k6_numbers)=k5_numbers|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1474, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1475, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_5717, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,X2)|~v1_xboole_0(X2)|~m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k6_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1460, negated_conjecture, (k2_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0),k6_numbers)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1461, negated_conjecture, (k2_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2),k6_numbers)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1462, plain, (k2_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3),k6_numbers)=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1465, plain, (m2_subset_1(k5_numbers,k1_numbers,k5_numbers)|~m1_subset_1(k5_numbers,k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1476, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_3527, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=esk2_0|k1_funct_1(esk2_0,k2_xcmplx_0(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0)))!=k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k6_numbers)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(np__1,esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),esk2_0),k5_numbers))).#trainneg
cnf(i_0_1481, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_5843, plain, (v5_relat_1(k5_numbers,k4_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1477, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1478, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1479, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1480, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k3_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1482, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_1484, plain, (v1_xcmplx_0(esk16_1(X1))|~v1_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1485, plain, (v1_xxreal_0(esk16_1(X1))|~v2_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1486, plain, (v1_xreal_0(esk16_1(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1487, plain, (v1_rat_1(esk16_1(X1))|~v4_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1488, plain, (v1_int_1(esk16_1(X1))|~v5_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1489, plain, (v7_ordinal1(esk16_1(X1))|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1491, plain, (v3_ordinal1(esk54_1(X1))|v1_zfmisc_1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1492, plain, (v3_ordinal1(esk52_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5306, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(X3)|~l1_struct_0(X3)|~v2_struct_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1493, plain, (v3_ordinal1(esk48_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5866, plain, (v5_relat_1(k5_numbers,k3_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1495, plain, (v3_ordinal1(esk41_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1497, plain, (v3_ordinal1(esk23_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1498, plain, (v3_ordinal1(esk16_1(X1))|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1499, plain, (v3_ordinal1(esk14_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1502, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1503, plain, (v3_ordinal1(k1_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_ordinal1(k1_numbers)|~l1_normsp_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1504, negated_conjecture, (v3_ordinal1(k1_funct_1(esk2_0,np__1))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1505, negated_conjecture, (v3_ordinal1(k1_funct_1(esk2_0,X1))|~v3_ordinal1(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1506, plain, (v3_ordinal1(k1_funct_1(X1,X2))|v2_struct_0(X3)|~v3_ordinal1(u1_struct_0(X3))|~l1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X3))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,u1_struct_0(X3))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1512, plain, (v3_ordinal1(k3_series_1(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1513, plain, (v3_ordinal1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1514, plain, (v3_ordinal1(k4_struct_0(X1))|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1515, plain, (v3_ordinal1(k1_normsp_1(X1,X2,X3))|v2_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1536, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|v1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~v1_funct_2(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk54_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_5297, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v5_relat_1(k1_numbers,k3_numbers))).#trainneg
cnf(i_0_1537, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~v1_funct_2(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk52_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1538, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~v1_funct_2(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk48_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1539, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~v1_funct_2(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk41_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1540, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|v1_xboole_0(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))|~v1_funct_2(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk23_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))))).#trainneg
cnf(i_0_1541, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)))),k5_numbers,u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))))).#trainneg
cnf(i_0_1542, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(esk1_0,esk2_0))|m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_1543, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,k6_numbers))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,k6_numbers))|~v1_funct_2(k6_numbers,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1545, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_5878, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk13_0)|~v1_xboole_0(X2)|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_1553, plain, (k1_seq_1(k2_zfmisc_1(X1,X2),X3)=k1_funct_1(k2_zfmisc_1(X1,X2),X3)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_5109, plain, (v1_rat_1(esk42_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_5135, plain, (v1_int_1(esk42_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2750, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5990, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1546, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1547, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_5993, plain, (k2_zfmisc_1(X1,X2)=k1_normsp_0(X3,k4_struct_0(X3))|v2_struct_0(X3)|~l2_normsp_0(X3)|~v4_normsp_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1548, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1549, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1550, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~v7_ordinal1(esk56_2(X1,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1551, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v5_relat_1(k3_series_1(k4_normsp_0(X1,X2)),k5_numbers)|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2))))).#trainneg
cnf(i_0_1555, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k4_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1556, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1557, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_1558, plain, (v3_ordinal1(esk17_1(X1))|v7_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1559, plain, (v1_xcmplx_0(esk17_1(X1))|v7_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1560, plain, (v1_xxreal_0(esk17_1(X1))|v7_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1561, plain, (v1_xreal_0(esk17_1(X1))|v7_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1562, plain, (v1_rat_1(esk17_1(X1))|v7_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1563, plain, (v1_int_1(esk17_1(X1))|v7_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1564, plain, (v7_ordinal1(esk17_1(X1))|v7_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_1566, plain, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k3_funct_2(k5_numbers,u1_struct_0(X1),X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1567, negated_conjecture, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),esk3_0))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1568, negated_conjecture, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,esk2_0),X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k1_funct_1(esk2_0,np__1),u1_struct_0(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,u1_struct_0(X1)))).#trainneg
cnf(i_0_1569, plain, (v1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,X2),X3))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k1_funct_1(X2,X3),u1_struct_0(X1))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1570, plain, (k2_xcmplx_0(k5_numbers,k6_numbers)=k5_numbers|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5879, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk12_0)|~v1_xboole_0(X2)|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_1575, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1576, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1577, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1588, plain, (k2_xcmplx_0(X1,esk44_0)=k2_xcmplx_0(esk44_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6500, negated_conjecture, (k2_xcmplx_0(esk44_0,esk2_0)=k2_xcmplx_0(esk2_0,esk44_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_6506, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk44_0)=k2_xcmplx_0(esk44_0,k1_funct_1(esk2_0,X1))|~v1_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6528, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk44_0)=k2_xcmplx_0(esk44_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6530, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk44_0)=k2_xcmplx_0(esk44_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6531, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk44_0)=k2_xcmplx_0(esk44_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6532, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk44_0)=k2_xcmplx_0(esk44_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_5880, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk11_0)|~v1_xboole_0(X2)|~v2_struct_0(esk11_0))).#trainneg
cnf(i_0_1589, plain, (k2_xcmplx_0(X1,esk36_0)=k2_xcmplx_0(esk36_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1590, plain, (k2_xcmplx_0(X1,esk14_1(X2))=k2_xcmplx_0(esk14_1(X2),X1)|~v1_xcmplx_0(X1)|~v1_membered(X2))).#trainneg
cnf(i_0_1592, plain, (k2_xcmplx_0(X1,k1_normsp_0(X2,X3))=k2_xcmplx_0(k1_normsp_0(X2,X3),X1)|v2_struct_0(X2)|~v1_xcmplx_0(X1)|~l1_normsp_0(X2)|~v1_membered(k1_numbers)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_6533, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk44_0)=k2_xcmplx_0(esk44_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_6604, negated_conjecture, (k2_xcmplx_0(esk36_0,esk2_0)=k2_xcmplx_0(esk2_0,esk36_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_6610, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk36_0)=k2_xcmplx_0(esk36_0,k1_funct_1(esk2_0,X1))|~v1_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6752, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk36_0)=k2_xcmplx_0(esk36_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6754, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk36_0)=k2_xcmplx_0(esk36_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6755, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk36_0)=k2_xcmplx_0(esk36_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5881, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk10_0)|~v1_xboole_0(X2)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_1593, plain, (k2_xcmplx_0(X1,np__1)=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1594, plain, (k2_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1596, plain, (k2_xcmplx_0(X1,k4_struct_0(X2))=k2_xcmplx_0(k4_struct_0(X2),X1)|~v1_xcmplx_0(X1)|~l2_struct_0(X2)|~v1_membered(u1_struct_0(X2)))).#trainneg
cnf(i_0_6756, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk36_0)=k2_xcmplx_0(esk36_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_6757, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk36_0)=k2_xcmplx_0(esk36_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1591, negated_conjecture, (k2_xcmplx_0(X1,esk3_0)=k2_xcmplx_0(esk3_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6949, negated_conjecture, (k2_xcmplx_0(esk3_0,esk2_0)=k2_xcmplx_0(esk2_0,esk3_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k5_numbers))).#trainneg
cnf(i_0_6955, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk3_0)=k2_xcmplx_0(esk3_0,k1_funct_1(esk2_0,X1))|~v1_membered(u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6977, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk3_0)=k2_xcmplx_0(esk3_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5882, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk9_0)|~v1_xboole_0(X2)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_1597, plain, (k2_xcmplx_0(X1,k1_normsp_1(X2,X3,X4))=k2_xcmplx_0(k1_normsp_1(X2,X3,X4),X1)|v2_struct_0(X2)|~v1_xcmplx_0(X1)|~v1_membered(u1_struct_0(X2))|~l1_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X2))))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1598, plain, (k2_xcmplx_0(X1,k5_numbers)=k2_xcmplx_0(k5_numbers,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1599, plain, (k2_xcmplx_0(X1,k5_numbers)=k2_xcmplx_0(k5_numbers,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_6979, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk3_0)=k2_xcmplx_0(esk3_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v6_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6980, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk3_0)=k2_xcmplx_0(esk3_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6981, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk3_0)=k2_xcmplx_0(esk3_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_6982, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk3_0)=k2_xcmplx_0(esk3_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1600, plain, (k2_xcmplx_0(X1,k5_numbers)=k2_xcmplx_0(k5_numbers,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1601, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7209, plain, (k2_xcmplx_0(k3_series_1(X1),X2)=k2_xcmplx_0(X2,k3_series_1(X1))|~v1_xreal_0(X2)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5883, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk8_0)|~v1_xboole_0(X2)|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_6509, plain, (k2_xcmplx_0(k6_numbers,esk44_0)=esk44_0)).#trainneg
cnf(i_0_1602, plain, (k2_xcmplx_0(k6_numbers,k3_series_1(X1))=k3_series_1(X1)|~v1_xreal_0(k3_series_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5193, negated_conjecture, (v1_membered(esk2_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1603, plain, (m2_subset_1(k2_xcmplx_0(X2,k3_series_1(X1)),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(k3_series_1(X1),k5_numbers)|~v1_xcmplx_0(X2)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1606, plain, (m2_subset_1(k2_xcmplx_0(k3_series_1(X2),X1),k1_numbers,k5_numbers)|~v7_ordinal1(k3_series_1(X2))|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1610, plain, (v5_membered(esk54_1(X1))|v1_zfmisc_1(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_1611, plain, (v5_membered(esk52_1(X1))|v1_xboole_0(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_1612, plain, (v5_membered(esk48_1(X1))|v1_xboole_0(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_1613, plain, (v5_membered(esk41_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_1614, plain, (v5_membered(esk23_1(X1))|v1_xboole_0(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_1615, plain, (v5_membered(esk14_1(k1_zfmisc_1(X1)))|~v5_membered(X1))).#trainneg
cnf(i_0_1618, plain, (v5_membered(k3_series_1(X1))|~v5_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1619, plain, (v5_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v5_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_5884, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk7_0)|~v1_xboole_0(X2)|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_6203, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6613, plain, (k2_xcmplx_0(k6_numbers,esk36_0)=esk36_0)).#trainneg
cnf(i_0_1621, plain, (v6_ordinal1(esk54_1(X1))|v1_zfmisc_1(X1)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_5198, negated_conjecture, (v2_membered(esk2_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3448, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1622, plain, (v6_ordinal1(esk52_1(X1))|v1_xboole_0(X1)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1623, plain, (v6_ordinal1(esk48_1(X1))|v1_xboole_0(X1)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1624, plain, (v6_ordinal1(esk41_1(X1))|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1625, plain, (v6_ordinal1(esk23_1(X1))|v1_xboole_0(X1)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1626, plain, (v6_ordinal1(esk14_1(k1_zfmisc_1(X1)))|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1629, plain, (v6_ordinal1(k3_series_1(X1))|~v6_ordinal1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1630, plain, (v6_ordinal1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_ordinal1(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1632, plain, (v6_membered(esk54_1(X1))|v1_zfmisc_1(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1633, plain, (v6_membered(esk52_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1634, plain, (v6_membered(esk48_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1635, plain, (v6_membered(esk41_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1636, plain, (v6_membered(esk23_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1637, plain, (v6_membered(esk14_1(k1_zfmisc_1(X1)))|~v6_membered(X1))).#trainneg
cnf(i_0_1640, plain, (v6_membered(k3_series_1(X1))|~v6_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5885, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk6_0)|~v1_xboole_0(X2)|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_6384, plain, (v5_relat_1(k2_zfmisc_1(X1,esk39_0),k4_numbers))).#trainneg
cnf(i_0_1641, plain, (v6_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v6_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7254, plain, (m2_subset_1(esk44_0,k1_numbers,k5_numbers)|~v7_ordinal1(esk44_0))).#trainneg
cnf(i_0_5232, negated_conjecture, (v4_membered(esk2_0)|~v6_membered(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1643, plain, (v1_membered(esk54_1(X1))|v1_zfmisc_1(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_1644, plain, (v1_membered(esk52_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_1645, plain, (v1_membered(esk48_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_1646, plain, (v1_membered(esk41_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1647, plain, (v1_membered(esk23_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_1648, plain, (v1_membered(esk14_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_1651, plain, (v1_membered(k3_series_1(X1))|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1652, plain, (v1_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1654, plain, (v2_membered(esk54_1(X1))|v1_zfmisc_1(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1655, plain, (v2_membered(esk52_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1656, plain, (v2_membered(esk48_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1657, plain, (v2_membered(esk41_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_1658, plain, (v2_membered(esk23_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_5991, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(X3)|~l1_struct_0(X3)|~v2_struct_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1659, plain, (v2_membered(esk14_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_5253, plain, (v1_xcmplx_0(esk42_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_1662, plain, (v2_membered(k3_series_1(X1))|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1663, plain, (v2_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1665, plain, (v3_membered(esk54_1(X1))|v1_zfmisc_1(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1666, plain, (v3_membered(esk52_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1667, plain, (v3_membered(esk48_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1668, plain, (v3_membered(esk41_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1669, plain, (v3_membered(esk23_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1670, plain, (v3_membered(esk14_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1673, plain, (v3_membered(k3_series_1(X1))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1674, plain, (v3_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1677, plain, (v4_membered(esk54_1(X1))|v1_zfmisc_1(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_1678, plain, (v4_membered(esk52_1(X1))|v1_xboole_0(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_1679, plain, (v4_membered(esk48_1(X1))|v1_xboole_0(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_1680, plain, (v4_membered(esk41_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_6385, plain, (v5_relat_1(k2_zfmisc_1(X1,esk31_0),k4_numbers))).#trainneg
cnf(i_0_3730, negated_conjecture, (k2_series_1(esk2_0)=k4_normsp_0(esk1_0,esk2_0)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_xcmplx_0(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),np__1))!=k2_xcmplx_0(k6_numbers,k1_funct_1(esk2_0,np__1))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),np__1),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1681, plain, (v4_membered(esk23_1(X1))|v1_xboole_0(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_5298, plain, (v1_xxreal_0(esk42_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_1682, plain, (v4_membered(esk14_1(k1_zfmisc_1(X1)))|~v4_membered(X1))).#trainneg
cnf(i_0_1685, plain, (v4_membered(k3_series_1(X1))|~v4_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1686, plain, (v4_membered(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1695, plain, (v1_relat_1(k4_normsp_0(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1696, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xboole_0(k4_normsp_0(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|m1_subset_1(X2,k2_zfmisc_1(k5_numbers,k1_numbers))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k4_normsp_0(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1697, plain, (v1_funct_1(k4_normsp_0(X1,k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1700, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X1)|v1_xboole_0(X2)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2))).#trainneg
cnf(i_0_1701, plain, (v1_xcmplx_0(k1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1702, plain, (v1_xxreal_0(k1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1703, plain, (v1_xreal_0(k1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1704, plain, (v1_rat_1(k1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1705, plain, (v1_int_1(k1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2))|v2_struct_0(X1)|~l1_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_1706, plain, (v7_ordinal1(k1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),X2))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(k2_zfmisc_1(k5_numbers,u1_struct_0(X1)),k5_numbers,u1_struct_0(X1))|~v1_funct_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))).#trainneg
cnf(i_0_7427, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk13_0)|~v1_xboole_0(X1)|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_6386, plain, (v5_relat_1(k2_zfmisc_1(X1,esk20_0),k4_numbers))).#trainneg
cnf(i_0_7470, negated_conjecture, (k2_series_1(esk2_0)=k4_normsp_0(esk1_0,esk2_0)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_xcmplx_0(np__1,esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0))))!=k2_xcmplx_0(k6_numbers,k1_funct_1(esk2_0,np__1))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(np__1,esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),k5_numbers))).#trainneg
cnf(i_0_7356, plain, (m2_subset_1(esk36_0,k1_numbers,k5_numbers)|~v7_ordinal1(esk36_0))).#trainneg
cnf(i_0_1724, plain, (k2_xcmplx_0(k6_numbers,k3_series_1(X1))=k3_series_1(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1729, plain, (v1_xboole_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1732, plain, (v1_xboole_0(esk41_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1734, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1737, plain, (v1_xboole_0(k3_series_1(X1))|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7592, plain, (k3_series_1(X1)=k1_normsp_0(X2,k4_struct_0(X2))|v2_struct_0(X2)|~l2_normsp_0(X2)|~v4_normsp_0(X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1738, plain, (v1_xboole_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_7428, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk12_0)|~v1_xboole_0(X1)|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_7579, plain, (esk41_1(X1)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7710, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1988, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v3_xxreal_0(k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1743, plain, (v1_zfmisc_1(esk48_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_1744, plain, (v1_zfmisc_1(esk41_1(X1))|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_1745, plain, (v1_zfmisc_1(esk23_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_1746, plain, (v1_zfmisc_1(esk14_1(k1_zfmisc_1(X1)))|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_1749, plain, (v1_zfmisc_1(k3_series_1(X1))|~v1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1750, plain, (v1_zfmisc_1(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_6388, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k4_numbers))).#trainneg
cnf(i_0_3449, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4697, negated_conjecture, (v6_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1772, plain, (v5_relat_1(esk38_2(X1,k4_numbers),k3_numbers)|~v1_relat_1(esk38_2(X1,k4_numbers)))).#trainneg
cnf(i_0_1773, plain, (v1_xboole_0(X2)|~r2_hidden(X2,X1)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_1777, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k4_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1779, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~v5_membered(X2))).#trainneg
cnf(i_0_1780, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),k3_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1788, plain, (u1_struct_0(esk7_0)=u1_struct_0(esk13_0)|~v2_struct_0(esk13_0)|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_1789, plain, (u1_struct_0(esk6_0)=u1_struct_0(esk13_0)|~v2_struct_0(esk13_0)|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_1792, plain, (v1_xcmplx_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1793, plain, (v1_xcmplx_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1794, plain, (v1_xcmplx_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1795, plain, (v1_xcmplx_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_7429, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk11_0)|~v1_xboole_0(X1)|~v2_struct_0(esk11_0))).#trainneg
cnf(i_0_6393, plain, (v5_relat_1(k2_zfmisc_1(X1,esk39_0),k3_numbers))).#trainneg
cnf(i_0_1798, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5048, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5057, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),k5_numbers))).#trainneg
cnf(i_0_1799, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_1800, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_1801, plain, (v3_ordinal1(esk29_1(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1802, plain, (v1_xcmplx_0(esk29_1(X1))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1803, plain, (v1_xxreal_0(esk29_1(X1))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1804, plain, (v1_xreal_0(esk29_1(X1))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1805, plain, (v1_rat_1(esk29_1(X1))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1806, plain, (v1_int_1(esk29_1(X1))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1807, plain, (v7_ordinal1(esk29_1(X1))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1808, plain, (v1_xboole_0(esk29_1(X1))|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(X2,u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,esk29_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1809, plain, (v5_membered(esk29_1(X1))|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1810, plain, (v6_ordinal1(esk29_1(X1))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1811, plain, (v6_membered(esk29_1(X1))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1812, plain, (v1_membered(esk29_1(X1))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_7430, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk10_0)|~v1_xboole_0(X1)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_6394, plain, (v5_relat_1(k2_zfmisc_1(X1,esk31_0),k3_numbers))).#trainneg
cnf(i_0_1813, plain, (v2_membered(esk29_1(X1))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_5171, negated_conjecture, (v5_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_5229, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1814, plain, (v3_membered(esk29_1(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1815, plain, (v4_membered(esk29_1(X1))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1819, plain, (v3_ordinal1(esk47_1(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1820, plain, (v1_xcmplx_0(esk47_1(X1))|v2_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1821, plain, (v1_xxreal_0(esk47_1(X1))|v2_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1822, plain, (v1_xreal_0(esk47_1(X1))|v2_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1823, plain, (v1_rat_1(esk47_1(X1))|v2_struct_0(X1)|~v4_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1824, plain, (v1_int_1(esk47_1(X1))|v2_struct_0(X1)|~v5_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1825, plain, (v7_ordinal1(esk47_1(X1))|v2_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1826, plain, (v1_xboole_0(esk47_1(X1))|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(X2,u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,esk47_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1827, plain, (v5_membered(esk47_1(X1))|v2_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_7431, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk9_0)|~v1_xboole_0(X1)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_6395, plain, (v5_relat_1(k2_zfmisc_1(X1,esk20_0),k3_numbers))).#trainneg
cnf(i_0_3772, negated_conjecture, (k2_series_1(esk2_0)=k4_normsp_0(esk1_0,esk2_0)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),np__1))!=k2_xcmplx_0(k6_numbers,k1_funct_1(esk2_0,np__1))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(esk2_0)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k2_nat_1(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),np__1),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1828, plain, (v6_ordinal1(esk47_1(X1))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_5984, plain, (v1_rat_1(esk42_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_7284, negated_conjecture, (v1_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_7373, negated_conjecture, (v2_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1829, plain, (v6_membered(esk47_1(X1))|v2_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1830, plain, (v1_membered(esk47_1(X1))|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1831, plain, (v2_membered(esk47_1(X1))|v2_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1832, plain, (v3_membered(esk47_1(X1))|v2_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1833, plain, (v4_membered(esk47_1(X1))|v2_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1835, plain, (v1_zfmisc_1(esk47_1(X1))|v2_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1837, plain, (v3_ordinal1(esk30_1(X1))|v7_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1838, plain, (v1_xcmplx_0(esk30_1(X1))|v7_struct_0(X1)|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1839, plain, (v1_xxreal_0(esk30_1(X1))|v7_struct_0(X1)|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1840, plain, (v1_xreal_0(esk30_1(X1))|v7_struct_0(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1841, plain, (v1_rat_1(esk30_1(X1))|v7_struct_0(X1)|~v4_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1842, plain, (v1_int_1(esk30_1(X1))|v7_struct_0(X1)|~v5_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_7432, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk8_0)|~v1_xboole_0(X1)|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_6397, plain, (v5_relat_1(k2_zfmisc_1(X1,k5_numbers),k3_numbers))).#trainneg
cnf(i_0_7918, negated_conjecture, (k2_series_1(esk2_0)=k4_normsp_0(esk1_0,esk2_0)|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(np__1,esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0))))!=k2_xcmplx_0(k6_numbers,k1_funct_1(esk2_0,np__1))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(k2_nat_1(np__1,esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0))),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(esk2_0,k4_normsp_0(esk1_0,esk2_0)),k5_numbers))).#trainneg
cnf(i_0_1843, plain, (v7_ordinal1(esk30_1(X1))|v7_struct_0(X1)|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_7416, negated_conjecture, (v4_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1844, plain, (v1_xboole_0(esk30_1(X1))|v1_xboole_0(u1_struct_0(X1))|m1_subset_1(X2,u1_struct_0(X1))|v7_struct_0(X1)|~m1_subset_1(X2,esk30_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1845, plain, (v5_membered(esk30_1(X1))|v7_struct_0(X1)|~v5_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1846, plain, (v6_ordinal1(esk30_1(X1))|v7_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1847, plain, (v6_membered(esk30_1(X1))|v7_struct_0(X1)|~v6_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1848, plain, (v1_membered(esk30_1(X1))|v7_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1849, plain, (v2_membered(esk30_1(X1))|v7_struct_0(X1)|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1850, plain, (v3_membered(esk30_1(X1))|v7_struct_0(X1)|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1851, plain, (v4_membered(esk30_1(X1))|v7_struct_0(X1)|~v4_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1852, plain, (v1_xboole_0(esk30_1(X1))|v7_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1855, plain, (v1_valued_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1856, plain, (v1_valued_0(esk52_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1857, plain, (v1_valued_0(esk48_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7433, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk7_0)|~v1_xboole_0(X1)|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_1858, plain, (v1_valued_0(esk47_1(X1))|v2_struct_0(X1)|~v1_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_5029, plain, (v3_ordinal1(k1_seq_1(k6_numbers,X1))|~v3_ordinal1(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1859, plain, (v1_valued_0(esk41_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1860, plain, (v1_valued_0(esk30_1(X1))|v7_struct_0(X1)|~v1_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1861, plain, (v1_valued_0(esk29_1(X1))|v2_struct_0(X1)|~v1_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1862, plain, (v1_valued_0(esk23_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1863, plain, (v1_valued_0(esk14_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1864, negated_conjecture, (v1_valued_0(esk2_0)|~v1_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~v1_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1866, plain, (v1_valued_0(k3_series_1(X1))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1870, plain, (v2_valued_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1871, plain, (v2_valued_0(esk52_1(X1))|v1_xboole_0(X1)|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1872, plain, (v2_valued_0(esk48_1(X1))|v1_xboole_0(X1)|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1873, plain, (v2_valued_0(esk47_1(X1))|v2_struct_0(X1)|~v2_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1874, plain, (v2_valued_0(esk41_1(X1))|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1875, plain, (v2_valued_0(esk30_1(X1))|v7_struct_0(X1)|~v2_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_7434, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(esk6_0)|~v1_xboole_0(X1)|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_1876, plain, (v2_valued_0(esk29_1(X1))|v2_struct_0(X1)|~v2_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_5030, plain, (v1_xcmplx_0(k1_seq_1(k6_numbers,X1))|~v1_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1877, plain, (v2_valued_0(esk23_1(X1))|v1_xboole_0(X1)|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1878, plain, (v2_valued_0(esk14_1(k1_zfmisc_1(X1)))|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1879, negated_conjecture, (v2_valued_0(esk2_0)|~v2_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~v1_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1881, plain, (v2_valued_0(k3_series_1(X1))|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1882, plain, (v2_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1885, plain, (v3_valued_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1886, plain, (v3_valued_0(esk52_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1887, plain, (v3_valued_0(esk48_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1888, plain, (v3_valued_0(esk47_1(X1))|v2_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1889, plain, (v3_valued_0(esk41_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1890, plain, (v3_valued_0(esk30_1(X1))|v7_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1891, plain, (v3_valued_0(esk29_1(X1))|v2_struct_0(X1)|~v3_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1892, plain, (v3_valued_0(esk23_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1893, plain, (v3_valued_0(esk14_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1894, negated_conjecture, (v3_valued_0(esk2_0)|~v3_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~v1_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5031, plain, (v1_xxreal_0(k1_seq_1(k6_numbers,X1))|~v2_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1896, plain, (v3_valued_0(k3_series_1(X1))|~v3_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1900, plain, (v4_valued_0(esk54_1(X1))|v1_zfmisc_1(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1901, plain, (v4_valued_0(esk52_1(X1))|v1_xboole_0(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1902, plain, (v4_valued_0(esk48_1(X1))|v1_xboole_0(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1903, plain, (v4_valued_0(esk47_1(X1))|v2_struct_0(X1)|~v4_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1904, plain, (v4_valued_0(esk41_1(X1))|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1905, plain, (v4_valued_0(esk30_1(X1))|v7_struct_0(X1)|~v4_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1906, plain, (v4_valued_0(esk29_1(X1))|v2_struct_0(X1)|~v4_valued_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1907, plain, (v4_valued_0(esk23_1(X1))|v1_xboole_0(X1)|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1908, plain, (v4_valued_0(esk14_1(k1_zfmisc_1(X1)))|~v4_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1909, negated_conjecture, (v4_valued_0(esk2_0)|~v4_valued_0(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))|~v1_relat_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1911, plain, (v4_valued_0(k3_series_1(X1))|~v4_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1912, plain, (v4_valued_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~v4_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1916, plain, (v1_funct_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2092, plain, (u1_struct_0(esk12_0)=u1_struct_0(esk11_0)|~v2_struct_0(esk12_0)|~v2_struct_0(esk11_0))).#trainneg
cnf(i_0_5033, plain, (v1_rat_1(k1_seq_1(k6_numbers,X1))|~v4_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1918, plain, (v1_relat_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1919, plain, (k2_xcmplx_0(X1,k1_funct_1(X2,X3))=k2_xcmplx_0(k1_funct_1(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1920, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)=k1_funct_1(X1,X2)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1922, negated_conjecture, (v1_xcmplx_0(k1_funct_1(esk2_0,esk36_0))|~v1_valued_0(esk2_0)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_1923, negated_conjecture, (v1_xcmplx_0(k1_funct_1(esk2_0,esk44_0))|~v1_valued_0(esk2_0)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_1924, negated_conjecture, (v1_xcmplx_0(k1_funct_1(esk2_0,np__1))|~v1_valued_0(esk2_0))).#trainneg
cnf(i_0_1925, negated_conjecture, (v1_xcmplx_0(k1_funct_1(esk2_0,X1))|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8305, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk3_0)=k2_xcmplx_0(esk3_0,k1_funct_1(esk2_0,X1))|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8306, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk36_0)=k2_xcmplx_0(esk36_0,k1_funct_1(esk2_0,X1))|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8307, negated_conjecture, (k2_xcmplx_0(k1_funct_1(esk2_0,X1),esk44_0)=k2_xcmplx_0(esk44_0,k1_funct_1(esk2_0,X1))|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2093, plain, (u1_struct_0(esk12_0)=u1_struct_0(esk10_0)|~v2_struct_0(esk12_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5034, plain, (v1_int_1(k1_seq_1(k6_numbers,X1))|~v5_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_5035, plain, (v7_ordinal1(k1_seq_1(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_1942, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_1943, plain, (v1_xreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1928, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,esk36_0))|~v2_valued_0(esk2_0)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_1929, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,esk44_0))|~v2_valued_0(esk2_0)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_1930, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,np__1))|~v2_valued_0(esk2_0))).#trainneg
cnf(i_0_1931, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,X1))|~v2_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1934, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,esk36_0))|~v3_valued_0(esk2_0)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_1935, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,esk44_0))|~v3_valued_0(esk2_0)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_1936, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,np__1))|~v3_valued_0(esk2_0))).#trainneg
cnf(i_0_1937, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,X1))|~v3_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5681, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_8438, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2094, plain, (u1_struct_0(esk12_0)=u1_struct_0(esk9_0)|~v2_struct_0(esk12_0)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_1964, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1965, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1966, plain, (m2_subset_1(k1_funct_1(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v4_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1968, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,esk36_0))|~v4_valued_0(esk2_0)|~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_5692, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2095, plain, (u1_struct_0(esk12_0)=u1_struct_0(esk8_0)|~v2_struct_0(esk12_0)|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_1969, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,esk44_0))|~v4_valued_0(esk2_0)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_1970, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,np__1))|~v4_valued_0(esk2_0))).#trainneg
cnf(i_0_1971, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,X1))|~v4_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1973, plain, (v2_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v3_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1974, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v3_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2096, plain, (u1_struct_0(esk12_0)=u1_struct_0(esk7_0)|~v2_struct_0(esk12_0)|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_8564, plain, (v2_xxreal_0(k6_numbers)|~v3_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_1994, plain, (v2_xxreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1995, plain, (v2_xxreal_0(X2)|~v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2015, plain, (v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2097, plain, (u1_struct_0(esk12_0)=u1_struct_0(esk6_0)|~v2_struct_0(esk12_0)|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_2035, plain, (v3_xxreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2055, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk41_1(X1)))).#trainneg
cnf(i_0_4579, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_3694, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,X1))!=k1_funct_1(esk2_0,k2_nat_1(np__1,esk4_2(esk2_0,esk2_0)))|~v1_valued_0(esk2_0)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k2_nat_1(np__1,esk4_2(esk2_0,esk2_0)),k5_numbers))).#trainneg
cnf(i_0_2084, plain, (m1_subset_1(k1_seq_1(k2_zfmisc_1(X1,X2),X3),k1_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_8628, plain, (v2_xxreal_0(k6_numbers)|~v3_xxreal_0(esk36_0))).#trainneg
cnf(i_0_8630, plain, (v2_xxreal_0(k6_numbers)|~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_8389, plain, (v1_xxreal_0(k1_seq_1(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_2085, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_2086, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_2063, plain, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers))|m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~m1_subset_1(X1,esk42_0))).#trainneg
cnf(i_0_2087, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_2088, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2089, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_2117, plain, (k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|~v1_xboole_0(k3_series_1(k4_normsp_0(X1,X2)))|~v1_relat_1(k3_series_1(k4_normsp_0(X1,X2)))|~v1_funct_1(k3_series_1(k4_normsp_0(X1,X2)))|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4580, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_7337, plain, (v3_membered(esk41_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_2122, negated_conjecture, (m2_subset_1(esk56_2(esk1_0,esk2_0),k1_numbers,k5_numbers)|k1_seq_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),k6_numbers)!=k1_seq_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)|~v1_xboole_0(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_relat_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)))|~v1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_2133, negated_conjecture, (v1_xboole_0(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_9141, negated_conjecture, (v1_xboole_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),np__1))|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_7463, plain, (v3_membered(esk41_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_8838, plain, (v3_xxreal_0(k6_numbers)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_9170, plain, (~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_5309, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2146, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk13_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk13_0))).#trainneg
cnf(i_0_9143, negated_conjecture, (v1_xboole_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),esk3_0))|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_9156, negated_conjecture, (k1_normsp_0(X1,k4_struct_0(X1))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),np__1)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_9214, negated_conjecture, (k1_normsp_0(X1,k4_struct_0(X1))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),esk3_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_7584, plain, (esk14_1(k1_zfmisc_1(X1))=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3715, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,X1))!=k1_funct_1(esk2_0,k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)))|~v1_valued_0(esk2_0)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k2_xcmplx_0(np__1,esk4_2(esk2_0,esk2_0)),k5_numbers))).#trainneg
cnf(i_0_2147, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk12_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk12_0))).#trainneg
cnf(i_0_2256, plain, (u1_struct_0(esk10_0)=u1_struct_0(esk11_0)|~v2_struct_0(esk11_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_8736, plain, (v2_xxreal_0(k6_numbers)|~v3_xxreal_0(esk36_0))).#trainneg
cnf(i_0_5994, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2148, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk11_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk11_0))).#trainneg
cnf(i_0_2149, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk10_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_9276, negated_conjecture, (k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1))=k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),esk3_0)|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_9435, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),esk3_0))|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_9461, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(esk1_0,esk2_0),np__1))|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_2150, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk9_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_2257, plain, (u1_struct_0(esk9_0)=u1_struct_0(esk11_0)|~v2_struct_0(esk11_0)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_8764, plain, (v2_xxreal_0(k6_numbers)|~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_5693, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6202, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2151, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk8_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_2152, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk7_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_2153, plain, (k1_normsp_0(X1,k4_struct_0(X1))=u1_struct_0(esk6_0)|v2_struct_0(X1)|~l2_normsp_0(X1)|~v4_normsp_0(X1)|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_2173, plain, (v1_xcmplx_0(k1_normsp_0(X2,k4_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2174, plain, (v1_xxreal_0(k1_normsp_0(X2,k4_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2175, plain, (v1_xreal_0(k1_normsp_0(X2,k4_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2176, plain, (v1_rat_1(k1_normsp_0(X2,k4_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~v4_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2177, plain, (v1_int_1(k1_normsp_0(X2,k4_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~v5_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2178, plain, (v7_ordinal1(k1_normsp_0(X2,k4_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2179, plain, (k2_xcmplx_0(k1_normsp_0(X2,k4_struct_0(X2)),k6_numbers)=k1_normsp_0(X2,k4_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2180, plain, (v1_xxreal_0(k1_normsp_0(X2,k4_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_2181, plain, (m1_subset_1(k1_normsp_0(X2,k4_struct_0(X2)),k1_numbers)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_normsp_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(X1)|~l2_normsp_0(X2)|~v4_normsp_0(X1)|~v4_normsp_0(X2))).#trainneg
cnf(i_0_3450, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2201, negated_conjecture, (k1_funct_1(k4_normsp_0(esk1_0,esk2_0),np__1)=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~l2_normsp_0(esk1_0))).#trainneg
cnf(i_0_2696, plain, (v3_valued_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3727, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,np__1))!=k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1))|~v1_valued_0(esk2_0)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1),k5_numbers)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2258, plain, (u1_struct_0(esk8_0)=u1_struct_0(esk11_0)|~v2_struct_0(esk11_0)|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_4698, negated_conjecture, (v6_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4980, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2259, plain, (u1_struct_0(esk7_0)=u1_struct_0(esk11_0)|~v2_struct_0(esk11_0)|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_2260, plain, (u1_struct_0(esk6_0)=u1_struct_0(esk11_0)|~v2_struct_0(esk11_0)|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_9643, plain, (k1_seq_1(k3_series_1(X1),X2)=k1_funct_1(k3_series_1(X1),X2)|~v1_relat_1(k3_series_1(X1))|~v1_funct_1(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9680, negated_conjecture, (k1_funct_1(k3_series_1(k4_normsp_0(esk1_0,esk2_0)),esk3_0)!=k6_numbers|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_9681, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2)))).#trainneg
cnf(i_0_9683, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_xcmplx_0(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_xcmplx_0(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2)))).#trainneg
cnf(i_0_9684, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(esk56_2(X1,X2),np__1))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(esk56_2(X1,X2),np__1))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2)))).#trainneg
cnf(i_0_9685, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k2_nat_1(np__1,esk56_2(X1,X2)))!=k1_seq_1(k4_normsp_0(X1,X2),k2_nat_1(np__1,esk56_2(X1,X2)))|k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(esk56_2(X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2)))).#trainneg
cnf(i_0_9687, plain, (k1_seq_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2)))).#trainneg
cnf(i_0_2755, plain, (v5_relat_1(esk41_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_8629, negated_conjecture, (v2_xxreal_0(k6_numbers)|~v3_xxreal_0(esk3_0)|~v2_xxreal_0(esk3_0))).#trainneg
cnf(i_0_5172, negated_conjecture, (v5_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5194, negated_conjecture, (v1_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2267, plain, (v5_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_2270, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(X1)))|v1_xboole_0(X1)|m1_subset_1(esk14_1(esk14_1(k1_zfmisc_1(X1))),X1))).#trainneg
cnf(i_0_9728, plain, (k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|~m1_subset_1(k4_normsp_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k4_normsp_0(X1,X2))|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9820, plain, (k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),X3)=k1_seq_1(k4_normsp_0(X1,X2),X3)|m2_subset_1(esk56_2(X1,X2),k1_numbers,k5_numbers)|v2_struct_0(X1)|k1_funct_1(k3_series_1(k4_normsp_0(X1,X2)),k6_numbers)!=k1_seq_1(k4_normsp_0(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(k4_normsp_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(k4_normsp_0(X1,X2))|~v1_funct_1(X2)|~l2_clvect_1(X1)|~v8_clvect_1(X1)|~v5_clvect_1(X1)|~v4_clvect_1(X1)|~v3_clvect_1(X1)|~v2_clvect_1(X1)|~v4_normsp_0(X1)|~v3_normsp_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~l1_normsp_0(X1))).#trainneg
cnf(i_0_2770, plain, (v4_relat_1(esk41_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_5199, negated_conjecture, (v2_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5230, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2239, negated_conjecture, (v1_xcmplx_0(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2240, negated_conjecture, (v1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2241, negated_conjecture, (v1_xreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2242, negated_conjecture, (v1_rat_1(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v4_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2243, negated_conjecture, (v1_int_1(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v5_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2244, negated_conjecture, (v7_ordinal1(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2245, negated_conjecture, (k2_xcmplx_0(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)),k6_numbers)=k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v1_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2246, negated_conjecture, (v1_xxreal_0(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)))|v2_struct_0(X1)|~l1_normsp_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_2247, negated_conjecture, (m1_subset_1(k1_normsp_0(esk1_0,k1_funct_1(esk2_0,np__1)),k1_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_normsp_0(esk1_0)|~l2_normsp_0(X1)|~v4_normsp_0(X1))).#trainneg
cnf(i_0_9730, plain, (v5_relat_1(esk41_1(k2_zfmisc_1(X1,k4_numbers)),k3_numbers))).#trainneg
cnf(i_0_3739, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k2_xcmplx_0(k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,X1))!=k1_funct_1(esk2_0,k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1))|~v1_valued_0(esk2_0)|~m1_subset_1(esk4_2(esk2_0,esk2_0),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk2_0,esk2_0),np__1),k5_numbers))).#trainneg
cnf(i_0_7285, negated_conjecture, (v1_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_7374, negated_conjecture, (v2_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2285, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_2286, plain, (v1_xxreal_0(k4_normsp_0(X1,X2))|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2287, plain, (m1_subset_1(k1_seq_1(k5_numbers,X1),k1_numbers)|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v5_relat_1(k1_numbers,k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2288, plain, (k1_seq_1(k5_numbers,X1)=k1_funct_1(k5_numbers,X1)|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v5_relat_1(k1_numbers,k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2293, plain, (v6_membered(esk42_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2664, plain, (v1_valued_0(esk19_2(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_7417, negated_conjecture, (v4_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_7844, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2312, plain, (k2_xcmplx_0(k1_funct_1(X2,X3),k6_numbers)=k1_funct_1(X2,X3)|v2_struct_0(X1)|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_2313, negated_conjecture, (k3_funct_2(k5_numbers,k1_numbers,k4_normsp_0(X1,k4_normsp_0(esk1_0,esk2_0)),esk3_0)=k1_normsp_0(X1,k6_numbers)|v2_struct_0(X1)|~l1_normsp_0(X1)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,u1_struct_0(X1))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2314, negated_conjecture, (m1_subset_1(k6_numbers,u1_struct_0(X1))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,u1_struct_0(X1))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2321, negated_conjecture, (k2_series_1(X1)=k4_normsp_0(esk1_0,esk2_0)|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_xcmplx_0(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),np__1)))!=k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_xcmplx_0(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),np__1))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(X1)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2324, negated_conjecture, (k2_series_1(X1)=k4_normsp_0(esk1_0,esk2_0)|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_nat_1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),np__1)))!=k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k2_nat_1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),np__1))|k1_funct_1(k4_normsp_0(esk1_0,esk2_0),k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(X1)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X1,k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_2678, plain, (v2_valued_0(esk19_2(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_2378, plain, (v3_ordinal1(k7_real_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7847, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))),k5_numbers))).#trainneg
cnf(i_0_2379, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2380, plain, (v1_xxreal_0(k7_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10031, negated_conjecture, (k2_series_1(k4_normsp_0(esk1_0,esk2_0))=k4_normsp_0(esk1_0,esk2_0)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),np__1),k5_numbers)|~m1_subset_1(esk4_2(k4_normsp_0(esk1_0,esk2_0),k4_normsp_0(esk1_0,esk2_0)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_10064, negated_conjecture, (v1_xreal_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k5_numbers)|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10065, negated_conjecture, (m2_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers,k5_numbers)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers)|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10066, negated_conjecture, (v1_xreal_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10067, negated_conjecture, (v1_xreal_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10068, negated_conjecture, (v1_xreal_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k4_numbers))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10069, negated_conjecture, (v1_xreal_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k3_numbers))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_2692, plain, (v3_valued_0(esk19_2(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_6494, plain, (k2_xcmplx_0(esk44_0,esk36_0)=k2_xcmplx_0(esk36_0,esk44_0))).#trainneg
cnf(i_0_10092, negated_conjecture, (k1_seq_1(k4_normsp_0(esk1_0,esk2_0),esk3_0)!=k6_numbers|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2381, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5683, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2382, plain, (v1_rat_1(k7_real_1(X1,X2))|~v4_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2383, plain, (v1_int_1(k7_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10083, negated_conjecture, (v2_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10084, negated_conjecture, (v4_valued_0(k4_normsp_0(esk1_0,esk2_0))|~v4_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10085, negated_conjecture, (v3_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k4_normsp_0(esk1_0,esk2_0),k5_numbers,k1_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_9184, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10119, plain, (v1_xreal_0(k2_xcmplx_0(esk36_0,esk44_0)))).#trainneg
cnf(i_0_6499, negated_conjecture, (k2_xcmplx_0(esk3_0,esk44_0)=k2_xcmplx_0(esk44_0,esk3_0))).#trainneg
cnf(i_0_2384, plain, (v7_ordinal1(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2393, plain, (u1_struct_0(esk10_0)=u1_struct_0(esk9_0)|~v2_struct_0(esk10_0)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_10730, negated_conjecture, (~v3_xxreal_0(esk3_0)|~v1_xboole_0(k2_xcmplx_0(esk44_0,esk3_0)))).#trainneg
cnf(i_0_10371, plain, (v1_valued_0(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2385, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2386, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10103, negated_conjecture, (m2_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_10104, negated_conjecture, (m2_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers,k5_numbers)|~v1_partfun1(k4_normsp_0(esk1_0,esk2_0),k5_numbers)|~v1_valued_0(k4_normsp_0(esk1_0,esk2_0))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_normsp_0(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_10149, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_10865, plain, (m1_subset_1(esk44_0,k1_numbers))).#trainneg
cnf(i_0_2387, plain, (m1_subset_1(k7_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2394, plain, (u1_struct_0(esk8_0)=u1_struct_0(esk10_0)|~v2_struct_0(esk10_0)|~v2_struct_0(esk8_0))).#trainneg
cnf(i_0_10729, negated_conjecture, (~v3_xxreal_0(esk3_0)|~v2_xxreal_0(k2_xcmplx_0(esk44_0,esk3_0))|~v1_xxreal_0(k2_xcmplx_0(esk44_0,esk3_0)))).#trainneg
cnf(i_0_10907, plain, (v3_ordinal1(esk44_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2395, plain, (u1_struct_0(esk7_0)=u1_struct_0(esk10_0)|~v2_struct_0(esk10_0)|~v2_struct_0(esk7_0))).#trainneg
cnf(i_0_10713, negated_conjecture, (m2_subset_1(k2_xcmplx_0(esk44_0,esk3_0),k1_numbers,k5_numbers)|~v7_ordinal1(esk44_0))).#trainneg
cnf(i_0_10720, negated_conjecture, (v1_xreal_0(k2_xcmplx_0(esk44_0,esk3_0)))).#trainneg
cnf(i_0_10721, negated_conjecture, (m2_subset_1(k2_xcmplx_0(esk44_0,esk3_0),k1_numbers,k5_numbers)|~m1_subset_1(esk44_0,k5_numbers))).#trainneg
cnf(i_0_10866, plain, (m1_subset_1(esk36_0,k1_numbers))).#trainneg
cnf(i_0_10927, negated_conjecture, (m1_subset_1(k2_xcmplx_0(esk44_0,esk3_0),k5_numbers)|~v7_ordinal1(esk44_0))).#trainneg
cnf(i_0_2709, plain, (v4_valued_0(esk19_2(X1,X2))|~v6_membered(X2))).#trainneg
cnf(i_0_10911, plain, (v1_rat_1(esk44_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_2396, plain, (u1_struct_0(esk6_0)=u1_struct_0(esk10_0)|~v2_struct_0(esk10_0)|~v2_struct_0(esk6_0))).#trainneg
cnf(i_0_10912, plain, (v1_int_1(esk44_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_10913, plain, (v7_ordinal1(esk44_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2403, plain, (m1_subset_1(k1_seq_1(k5_numbers,X1),k1_numbers)|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2404, plain, (k1_seq_1(k5_numbers,X1)=k1_funct_1(k5_numbers,X1)|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v5_relat_1(k1_numbers,k4_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_10928, negated_conjecture, (m1_subset_1(k2_xcmplx_0(esk44_0,esk3_0),k1_numbers)|~v7_ordinal1(esk44_0))).#trainneg
cnf(i_0_10957, plain, (m2_subset_1(esk44_0,k1_numbers,k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_10984, negated_conjecture, (m1_subset_1(k2_xcmplx_0(esk44_0,esk3_0),k1_numbers)|~m1_subset_1(k1_zfmisc_1(esk44_0),k5_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 261
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 435
# Removed in clause preprocessing      : 46
# Initial clauses in saturation        : 389
# Processed clauses                    : 2914
# ...of these trivial                  : 41
# ...subsumed                          : 1027
# ...remaining for further processing  : 1846
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 65
# Backward-rewritten                   : 48
# Generated clauses                    : 10568
# ...of the previous two non-trivial   : 9485
# Contextual simplify-reflections      : 1020
# Paramodulations                      : 10566
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
# Propositional unsat check successes  : 0
# Current number of processed clauses  : 1732
#    Positive orientable unit clauses  : 269
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 38
#    Non-unit-clauses                  : 1425
# Current number of unprocessed clauses: 6440
# ...number of literals in the above   : 45270
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1960
# Proof object given clauses           : 109
# Proof search given clauses           : 1846
# Clause-clause subsumption calls (NU) : 2143321
# Rec. Clause-clause subsumption calls : 192984
# Non-unit clause-clause subsumptions  : 1749
# Unit Clause-clause subsumption calls : 29800
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57
# BW rewrite match successes           : 26
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 302338

# -------------------------------------------------
# User time                : 3.345 s
# System time              : 0.172 s
# Total time               : 3.517 s
# ...preprocessing         : 0.668 s
# ...main loop             : 2.849 s
# Maximum resident set size: 105040 pages
