2020-05-31 21:48:36.491198: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69
2020-05-31 21:48:36.597693: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 21:48:36.768624: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 21:48:37.338165: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 21:48:39.214302: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2723129 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69' loaded (query=128; context=256; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.823 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', d8_lattice3)).
fof(d4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v1_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r1_lattice3(X1,u1_struct_0(X1),X2)))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', d4_yellow_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', t1_subset)).
fof(d15_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(((v1_orders_2(X3)&v4_yellow_0(X3,X1))&m1_yellow_0(X3,X1))=>(X3=k5_yellow_0(X1,X2)<=>u1_struct_0(X3)=X2)))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', d15_yellow_0)).
fof(d2_yellow_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_orders_2(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', d2_yellow_2)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', t2_subset)).
fof(dt_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', dt_k2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', cc1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', cc2_relset_1)).
fof(dt_k1_yellow_2, axiom, ![X1, X2, X3]:(((((((~(v2_struct_0(X1))&l1_orders_2(X1))&~(v2_struct_0(X2)))&l1_orders_2(X2))&v1_funct_1(X3))&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>((v1_orders_2(k1_yellow_2(X1,X2,X3))&v4_yellow_0(k1_yellow_2(X1,X2,X3),X2))&m1_yellow_0(k1_yellow_2(X1,X2,X3),X2))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', dt_k1_yellow_2)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', redefinition_k2_relset_1)).
fof(t6_lattice8, conjecture, ![X1]:((((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_yellow_0(X1))&l1_orders_2(X1))=>![X2]:(((((~(v2_struct_0(X2))&v3_orders_2(X2))&v4_orders_2(X2))&v5_orders_2(X2))&l1_orders_2(X2))=>![X3]:((((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&v5_orders_3(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>v1_yellow_0(k1_yellow_2(X1,X2,X3))))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', t6_lattice8)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', reflexivity_r1_tarski)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', dt_l1_orders_2)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((X2!=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X3)))&(X2=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0)))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', d1_funct_2)).
fof(d5_orders_3, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(v5_orders_3(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_orders_2(X1,X4,X5)=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>![X7]:(m1_subset_1(X7,u1_struct_0(X2))=>((X6=k1_funct_1(X3,X4)&X7=k1_funct_1(X3,X5))=>r1_orders_2(X2,X6,X7))))))))))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', d5_orders_3)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', fc1_subset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', redefinition_k1_relset_1)).
fof(d3_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X1))&X3=k1_funct_1(X1,X4))))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', d3_funct_1)).
fof(t60_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:((v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((((X5=X3&X6=X4)&r1_orders_2(X1,X3,X4))&r2_hidden(X5,u1_struct_0(X2)))=>r1_orders_2(X2,X5,X6)))))))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', t60_yellow_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', t4_subset)).
fof(dt_k5_yellow_0, axiom, ![X1, X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>((v1_orders_2(k5_yellow_0(X1,X2))&v4_yellow_0(k5_yellow_0(X1,X2),X1))&m1_yellow_0(k5_yellow_0(X1,X2),X1))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', dt_k5_yellow_0)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', dt_m1_yellow_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t6_lattice8', fc1_xboole_0)).
fof(c_0_26, plain, ![X14, X15, X16, X17]:((~r1_lattice3(X14,X15,X16)|(~m1_subset_1(X17,u1_struct_0(X14))|(~r2_hidden(X17,X15)|r1_orders_2(X14,X16,X17)))|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&((m1_subset_1(esk4_3(X14,X15,X16),u1_struct_0(X14))|r1_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&((r2_hidden(esk4_3(X14,X15,X16),X15)|r1_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))&(~r1_orders_2(X14,X16,esk4_3(X14,X15,X16))|r1_lattice3(X14,X15,X16)|~m1_subset_1(X16,u1_struct_0(X14))|~l1_orders_2(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
fof(c_0_27, plain, ![X54, X56]:(((m1_subset_1(esk12_1(X54),u1_struct_0(X54))|~v1_yellow_0(X54)|~l1_orders_2(X54))&(r1_lattice3(X54,u1_struct_0(X54),esk12_1(X54))|~v1_yellow_0(X54)|~l1_orders_2(X54)))&(~m1_subset_1(X56,u1_struct_0(X54))|~r1_lattice3(X54,u1_struct_0(X54),X56)|v1_yellow_0(X54)|~l1_orders_2(X54))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_0])])])])])).
cnf(c_0_28, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29, plain, (r1_lattice3(X1,u1_struct_0(X1),esk12_1(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_30, plain, ![X124, X125]:(~r2_hidden(X124,X125)|m1_subset_1(X124,X125)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_31, plain, (r1_orders_2(X1,X3,X4)|~l1_orders_2(X1)|~r2_hidden(X4,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_28])).
cnf(c_0_32, plain, (r1_lattice3(X1,u1_struct_0(X1),esk12_1(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_29])).
cnf(c_0_33, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_34, plain, (m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (r1_orders_2(X1,esk12_1(X1),X2)|~r2_hidden(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32])])).
cnf(c_0_37, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_33])).
cnf(c_0_38, plain, (m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_34])).
fof(c_0_39, plain, ![X11, X12, X13]:((X13!=k5_yellow_0(X11,X12)|u1_struct_0(X13)=X12|(~v1_orders_2(X13)|~v4_yellow_0(X13,X11)|~m1_yellow_0(X13,X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_orders_2(X11))&(u1_struct_0(X13)!=X12|X13=k5_yellow_0(X11,X12)|(~v1_orders_2(X13)|~v4_yellow_0(X13,X11)|~m1_yellow_0(X13,X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_orders_2(X11))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_yellow_0])])])])).
fof(c_0_40, plain, ![X27, X28, X29]:(v2_struct_0(X27)|~l1_orders_2(X27)|(v2_struct_0(X28)|~l1_orders_2(X28)|(~v1_funct_1(X29)|~v1_funct_2(X29,u1_struct_0(X27),u1_struct_0(X28))|~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X27),u1_struct_0(X28))))|k1_yellow_2(X27,X28,X29)=k5_yellow_0(X28,k2_relset_1(u1_struct_0(X28),X29))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_yellow_2])])])])).
cnf(c_0_41, plain, (r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (r1_orders_2(X1,esk12_1(X1),X2)|~r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
fof(c_0_43, plain, ![X78, X79]:(~m1_subset_1(X78,X79)|(v1_xboole_0(X79)|r2_hidden(X78,X79))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_44, plain, (u1_struct_0(X1)=X3|X1!=k5_yellow_0(X2,X3)|~v1_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_46, plain, ![X76, X77]:(~v1_relat_1(X77)|~v5_relat_1(X77,X76)|m1_subset_1(k2_relset_1(X76,X77),k1_zfmisc_1(X76))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_relset_1])])).
cnf(c_0_47, plain, (r1_lattice3(X1,X2,esk12_1(X1))|~m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk4_3(X1,X2,esk12_1(X1)),u1_struct_0(X1))|~v1_yellow_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42])])).
cnf(c_0_48, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (u1_struct_0(X1)=X3|X1!=k5_yellow_0(X2,X3)|~l1_orders_2(X2)|~v1_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_50, plain, (k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_52, plain, ![X91, X92, X93]:(~m1_subset_1(X93,k1_zfmisc_1(k2_zfmisc_1(X91,X92)))|v1_relat_1(X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_53, plain, ![X106, X107, X108]:((v4_relat_1(X108,X106)|~m1_subset_1(X108,k1_zfmisc_1(k2_zfmisc_1(X106,X107))))&(v5_relat_1(X108,X107)|~m1_subset_1(X108,k1_zfmisc_1(k2_zfmisc_1(X106,X107))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_54, plain, (r1_lattice3(X1,X2,esk12_1(X1))|~r2_hidden(esk4_3(X1,X2,esk12_1(X1)),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)), inference(csr,[status(thm)],[c_0_47, c_0_38])).
cnf(c_0_55, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_56, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_57, plain, (u1_struct_0(X1)=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X4)|v2_struct_0(X2)|X1!=k1_yellow_2(X4,X2,X3)|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~v1_orders_2(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50])])).
cnf(c_0_58, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_59, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_61, plain, (r1_lattice3(X1,X2,esk12_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk4_3(X1,X2,esk12_1(X1)),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_62, plain, (r1_lattice3(X1,X2,X3)|m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (u1_struct_0(X1)=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k1_yellow_2(X4,X2,X3)|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~v1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~v5_relat_1(X3,u1_struct_0(X2))|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58])])).
cnf(c_0_64, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_65, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_60])).
fof(c_0_66, plain, ![X51, X52, X53]:(((v1_orders_2(k1_yellow_2(X51,X52,X53))|(v2_struct_0(X51)|~l1_orders_2(X51)|v2_struct_0(X52)|~l1_orders_2(X52)|~v1_funct_1(X53)|~v1_funct_2(X53,u1_struct_0(X51),u1_struct_0(X52))|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X51),u1_struct_0(X52))))))&(v4_yellow_0(k1_yellow_2(X51,X52,X53),X52)|(v2_struct_0(X51)|~l1_orders_2(X51)|v2_struct_0(X52)|~l1_orders_2(X52)|~v1_funct_1(X53)|~v1_funct_2(X53,u1_struct_0(X51),u1_struct_0(X52))|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X51),u1_struct_0(X52)))))))&(m1_yellow_0(k1_yellow_2(X51,X52,X53),X52)|(v2_struct_0(X51)|~l1_orders_2(X51)|v2_struct_0(X52)|~l1_orders_2(X52)|~v1_funct_1(X53)|~v1_funct_2(X53,u1_struct_0(X51),u1_struct_0(X52))|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X51),u1_struct_0(X52))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_yellow_2])])])])).
cnf(c_0_67, plain, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(X1,X2,esk12_1(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk12_1(X1),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62])])).
fof(c_0_68, plain, ![X64, X65]:(~v1_relat_1(X65)|~v5_relat_1(X65,X64)|k2_relset_1(X64,X65)=k10_xtuple_0(X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_69, plain, (u1_struct_0(X1)=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X4)|v2_struct_0(X2)|X1!=k1_yellow_2(X4,X2,X3)|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~v1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_70, plain, (v1_orders_2(k1_yellow_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_71, plain, (v4_yellow_0(k1_yellow_2(X1,X2,X3),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, plain, (m1_yellow_0(k1_yellow_2(X1,X2,X3),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_73, negated_conjecture, ~(![X1]:((((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_yellow_0(X1))&l1_orders_2(X1))=>![X2]:(((((~(v2_struct_0(X2))&v3_orders_2(X2))&v4_orders_2(X2))&v5_orders_2(X2))&l1_orders_2(X2))=>![X3]:((((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&v5_orders_3(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>v1_yellow_0(k1_yellow_2(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t6_lattice8])).
cnf(c_0_74, plain, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(X1,X2,esk12_1(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)), inference(csr,[status(thm)],[c_0_67, c_0_38])).
fof(c_0_75, plain, ![X104, X105]:((~m1_subset_1(X104,k1_zfmisc_1(X105))|r1_tarski(X104,X105))&(~r1_tarski(X104,X105)|m1_subset_1(X104,k1_zfmisc_1(X105)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_76, plain, ![X130]:r1_tarski(X130,X130), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
fof(c_0_77, plain, ![X88]:(v2_struct_0(X88)|~l1_struct_0(X88)|~v1_xboole_0(u1_struct_0(X88))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_78, plain, ![X75]:(~l1_orders_2(X75)|l1_struct_0(X75)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_79, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_80, plain, (u1_struct_0(k1_yellow_2(X1,X2,X3))=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~v4_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~v1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_69])])).
cnf(c_0_81, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_82, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v4_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_83, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_72])).
fof(c_0_84, plain, ![X68, X69, X70]:(((~v1_funct_2(X70,X68,X69)|X68=k1_relset_1(X68,X70)|X69=k1_xboole_0|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))&(X68!=k1_relset_1(X68,X70)|v1_funct_2(X70,X68,X69)|X69=k1_xboole_0|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69)))))&((~v1_funct_2(X70,X68,X69)|X70=k1_xboole_0|X69!=k1_xboole_0|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))&(X70!=k1_xboole_0|v1_funct_2(X70,X68,X69)|X69!=k1_xboole_0|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_85, negated_conjecture, ((((((~v2_struct_0(esk1_0)&v3_orders_2(esk1_0))&v4_orders_2(esk1_0))&v5_orders_2(esk1_0))&v1_yellow_0(esk1_0))&l1_orders_2(esk1_0))&(((((~v2_struct_0(esk2_0)&v3_orders_2(esk2_0))&v4_orders_2(esk2_0))&v5_orders_2(esk2_0))&l1_orders_2(esk2_0))&((((v1_funct_1(esk3_0)&v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&v5_orders_3(esk3_0,esk1_0,esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&~v1_yellow_0(k1_yellow_2(esk1_0,esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_73])])])])).
fof(c_0_86, plain, ![X40, X41, X42, X43, X44, X45, X46]:((~v5_orders_3(X42,X40,X41)|(~m1_subset_1(X43,u1_struct_0(X40))|(~m1_subset_1(X44,u1_struct_0(X40))|(~r1_orders_2(X40,X43,X44)|(~m1_subset_1(X45,u1_struct_0(X41))|(~m1_subset_1(X46,u1_struct_0(X41))|(X45!=k1_funct_1(X42,X43)|X46!=k1_funct_1(X42,X44)|r1_orders_2(X41,X45,X46)))))))|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))&((m1_subset_1(esk8_3(X40,X41,X42),u1_struct_0(X40))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))&((m1_subset_1(esk9_3(X40,X41,X42),u1_struct_0(X40))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))&((r1_orders_2(X40,esk8_3(X40,X41,X42),esk9_3(X40,X41,X42))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))&((m1_subset_1(esk10_3(X40,X41,X42),u1_struct_0(X41))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))&((m1_subset_1(esk11_3(X40,X41,X42),u1_struct_0(X41))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))&(((esk10_3(X40,X41,X42)=k1_funct_1(X42,esk8_3(X40,X41,X42))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))&(esk11_3(X40,X41,X42)=k1_funct_1(X42,esk9_3(X40,X41,X42))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40)))&(~r1_orders_2(X41,esk10_3(X40,X41,X42),esk11_3(X40,X41,X42))|v5_orders_3(X42,X40,X41)|(~v1_funct_1(X42)|~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|~l1_orders_2(X41)|~l1_orders_2(X40))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_orders_3])])])])])).
cnf(c_0_87, plain, (r1_orders_2(X1,esk12_1(X1),X2)|v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_74])])).
cnf(c_0_88, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_89, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
fof(c_0_90, plain, ![X159]:~v1_xboole_0(k1_zfmisc_1(X159)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_91, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_92, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_93, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_94, plain, (u1_struct_0(k1_yellow_2(X1,X2,X3))=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_83])).
fof(c_0_95, plain, ![X66, X67]:(~v1_relat_1(X67)|~v4_relat_1(X67,X66)|k1_relset_1(X66,X67)=k9_xtuple_0(X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_96, plain, (X2=k1_relset_1(X2,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_97, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_98, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_99, plain, (r1_orders_2(X3,X6,X7)|~v5_orders_3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~r1_orders_2(X2,X4,X5)|~m1_subset_1(X6,u1_struct_0(X3))|~m1_subset_1(X7,u1_struct_0(X3))|X6!=k1_funct_1(X1,X4)|X7!=k1_funct_1(X1,X5)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~l1_orders_2(X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_100, negated_conjecture, (v5_orders_3(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_101, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_102, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_103, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_104, plain, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk12_1(X1),X2)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)), inference(csr,[status(thm)],[c_0_87, c_0_38])).
cnf(c_0_105, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_106, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_107, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_108, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_109, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_110, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_111, plain, (u1_struct_0(k1_yellow_2(X3,X1,X2))=k10_xtuple_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~v5_relat_1(X2,u1_struct_0(X1))|~v1_relat_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94])])).
fof(c_0_112, plain, ![X30, X31, X32, X34, X35, X36, X38]:((((r2_hidden(esk5_3(X30,X31,X32),k9_xtuple_0(X30))|~r2_hidden(X32,X31)|X31!=k10_xtuple_0(X30)|(~v1_relat_1(X30)|~v1_funct_1(X30)))&(X32=k1_funct_1(X30,esk5_3(X30,X31,X32))|~r2_hidden(X32,X31)|X31!=k10_xtuple_0(X30)|(~v1_relat_1(X30)|~v1_funct_1(X30))))&(~r2_hidden(X35,k9_xtuple_0(X30))|X34!=k1_funct_1(X30,X35)|r2_hidden(X34,X31)|X31!=k10_xtuple_0(X30)|(~v1_relat_1(X30)|~v1_funct_1(X30))))&((~r2_hidden(esk6_2(X30,X36),X36)|(~r2_hidden(X38,k9_xtuple_0(X30))|esk6_2(X30,X36)!=k1_funct_1(X30,X38))|X36=k10_xtuple_0(X30)|(~v1_relat_1(X30)|~v1_funct_1(X30)))&((r2_hidden(esk7_2(X30,X36),k9_xtuple_0(X30))|r2_hidden(esk6_2(X30,X36),X36)|X36=k10_xtuple_0(X30)|(~v1_relat_1(X30)|~v1_funct_1(X30)))&(esk6_2(X30,X36)=k1_funct_1(X30,esk7_2(X30,X36))|r2_hidden(esk6_2(X30,X36),X36)|X36=k10_xtuple_0(X30)|(~v1_relat_1(X30)|~v1_funct_1(X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_1])])])])])])).
cnf(c_0_113, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_114, plain, (X3=k1_xboole_0|k1_relset_1(X2,X1)=X2|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_115, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_116, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_117, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_118, plain, (r1_orders_2(X3,X6,X7)|X7!=k1_funct_1(X1,X5)|X6!=k1_funct_1(X1,X4)|~l1_orders_2(X3)|~l1_orders_2(X2)|~v1_funct_1(X1)|~v5_orders_3(X1,X2,X3)|~r1_orders_2(X2,X4,X5)|~m1_subset_1(X7,u1_struct_0(X3))|~m1_subset_1(X6,u1_struct_0(X3))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_119, negated_conjecture, (v5_orders_3(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_120, negated_conjecture, (v1_funct_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_121, negated_conjecture, (l1_orders_2(esk2_0)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_122, negated_conjecture, (l1_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_123, plain, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk12_1(X1),X2)|v1_xboole_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_55])])).
cnf(c_0_124, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106])])).
cnf(c_0_125, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_126, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_127, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_128, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_129, plain, (u1_struct_0(k1_yellow_2(X1,X2,X3))=k10_xtuple_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_111, c_0_64]), c_0_65])).
cnf(c_0_130, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_131, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k9_xtuple_0(X2))|X3!=k1_funct_1(X2,X1)|X4!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_132, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_133, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116])])])).
cnf(c_0_134, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_117])).
fof(c_0_135, plain, ![X19, X20, X21, X22, X23, X24]:(~l1_orders_2(X19)|(~v4_yellow_0(X20,X19)|~m1_yellow_0(X20,X19)|(~m1_subset_1(X21,u1_struct_0(X19))|(~m1_subset_1(X22,u1_struct_0(X19))|(~m1_subset_1(X23,u1_struct_0(X20))|(~m1_subset_1(X24,u1_struct_0(X20))|(X23!=X21|X24!=X22|~r1_orders_2(X19,X21,X22)|~r2_hidden(X23,u1_struct_0(X20))|r1_orders_2(X20,X23,X24)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_yellow_0])])])).
cnf(c_0_136, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,X4)|~r1_orders_2(esk1_0,X4,X3)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_115]), c_0_119]), c_0_116]), c_0_120]), c_0_121]), c_0_122])])])).
cnf(c_0_137, plain, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk12_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_125])])).
cnf(c_0_138, negated_conjecture, (v1_yellow_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_139, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_122])])])).
fof(c_0_140, plain, ![X151, X152, X153]:(~r2_hidden(X151,X152)|~m1_subset_1(X152,k1_zfmisc_1(X153))|m1_subset_1(X151,X153)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_141, plain, (k2_relset_1(u1_struct_0(X2),X3)=k10_xtuple_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_129])])).
cnf(c_0_142, negated_conjecture, (~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_143, plain, (r2_hidden(X3,X4)|X4!=k10_xtuple_0(X2)|X3!=k1_funct_1(X2,X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_144, negated_conjecture, (k9_xtuple_0(esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,u1_struct_0(esk1_0))|~v1_relat_1(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_145, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_115])])).
cnf(c_0_146, negated_conjecture, (v1_relat_1(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_115])])).
cnf(c_0_147, plain, (r1_orders_2(X2,X5,X6)|~l1_orders_2(X1)|~v4_yellow_0(X2,X1)|~m1_yellow_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|X5!=X3|X6!=X4|~r1_orders_2(X1,X3,X4)|~r2_hidden(X5,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_148, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,esk12_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_122]), c_0_138])]), c_0_139])])).
cnf(c_0_149, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_150, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k10_xtuple_0(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_115]), c_0_116]), c_0_120]), c_0_121]), c_0_122])]), c_0_127]), c_0_142])])).
cnf(c_0_151, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_115])])).
cnf(c_0_152, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|X1!=k1_funct_1(X3,X4)|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X4,k9_xtuple_0(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_55])])).
cnf(c_0_153, negated_conjecture, (k9_xtuple_0(esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_145]), c_0_146])])).
cnf(c_0_154, plain, (r1_orders_2(X2,X5,X6)|X4!=X6|X3!=X5|~l1_orders_2(X1)|~v4_yellow_0(X2,X1)|~m1_yellow_0(X2,X1)|~r1_orders_2(X1,X3,X4)|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X5,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_155, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|X1!=k1_funct_1(esk3_0,X2)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_148])])).
cnf(c_0_156, plain, (r2_hidden(esk5_3(X1,X2,X3),k9_xtuple_0(X1))|~r2_hidden(X3,X2)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_157, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_158, negated_conjecture, (m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_150]), c_0_151]), c_0_146])])])).
cnf(c_0_159, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,X3)|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_146]), c_0_120])]), c_0_139])])).
cnf(c_0_160, plain, (r1_orders_2(X1,X2,X3)|~r1_orders_2(X4,X2,X3)|~r2_hidden(X2,u1_struct_0(X1))|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~l1_orders_2(X4)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_154, c_0_37])])])])).
cnf(c_0_161, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),k1_funct_1(esk3_0,X1))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_155])])).
cnf(c_0_162, plain, (X1=k1_funct_1(X2,esk5_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_163, plain, (r2_hidden(esk5_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_164, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k10_xtuple_0(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158])])).
cnf(c_0_165, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,X1),X2)|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_159])])).
cnf(c_0_166, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),k1_funct_1(esk3_0,X2))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(X1))|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161]), c_0_121])])])).
cnf(c_0_167, plain, (k1_funct_1(X2,esk5_3(X2,X3,X1))=X1|X3!=k10_xtuple_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_168, plain, (m1_subset_1(esk5_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_163])])).
cnf(c_0_169, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165])])).
cnf(c_0_170, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X3)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk3_0,X2,X3),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_146]), c_0_120])])])).
cnf(c_0_171, negated_conjecture, (m1_subset_1(esk5_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_153]), c_0_146]), c_0_120])])])).
cnf(c_0_172, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk5_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_167]), c_0_146]), c_0_120])])])).
cnf(c_0_173, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(X2,X3)|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171])])).
cnf(c_0_174, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(X1,u1_struct_0(esk2_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_171])])).
cnf(c_0_175, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(X2,X3)|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_173, c_0_174]), c_0_169])).
cnf(c_0_176, plain, (X1=k5_yellow_0(X3,X2)|u1_struct_0(X1)!=X2|~v1_orders_2(X1)|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_177, plain, (v1_yellow_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,u1_struct_0(X2),X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_178, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_175])])).
cnf(c_0_179, plain, (X1=k5_yellow_0(X3,X2)|u1_struct_0(X1)!=X2|~l1_orders_2(X3)|~v1_orders_2(X1)|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_180, plain, (v1_yellow_0(X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,u1_struct_0(X2),X1)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_181, negated_conjecture, (~v1_yellow_0(k1_yellow_2(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_182, negated_conjecture, (r1_lattice3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_178])])).
cnf(c_0_183, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_184, plain, ![X62, X63]:(((v1_orders_2(k5_yellow_0(X62,X63))|(~l1_orders_2(X62)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X62)))))&(v4_yellow_0(k5_yellow_0(X62,X63),X62)|(~l1_orders_2(X62)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X62))))))&(m1_yellow_0(k5_yellow_0(X62,X63),X62)|(~l1_orders_2(X62)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X62)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_yellow_0])])])).
cnf(c_0_185, plain, (k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,X4)|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(k1_yellow_2(X1,X2,X3))!=X4|~m1_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~v1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_82])])).
cnf(c_0_186, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X1,X2,X3),k2_relset_1(u1_struct_0(X2),X3),X4)|~m1_subset_1(X4,k2_relset_1(u1_struct_0(X2),X3))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_94])])).
cnf(c_0_187, negated_conjecture, (~v1_yellow_0(k1_yellow_2(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_188, plain, (k1_yellow_2(X1,X2,X3)=k1_yellow_2(X4,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~l1_orders_2(X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_50])])).
fof(c_0_189, plain, ![X57, X58]:(~l1_orders_2(X57)|(~m1_yellow_0(X58,X57)|l1_orders_2(X58))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
cnf(c_0_190, negated_conjecture, (k1_yellow_2(X1,esk2_0,esk3_0)=k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_150]), c_0_120]), c_0_121])]), c_0_142])])).
cnf(c_0_191, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0)))|~r2_hidden(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(X1))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_182, c_0_37])).
cnf(c_0_192, plain, (r1_lattice3(X1,X2,X3)|r2_hidden(esk4_3(X1,X2,X3),X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_193, plain, (m1_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_184])).
cnf(c_0_194, plain, (v4_yellow_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_50])])).
cnf(c_0_195, plain, (k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(k1_yellow_2(X1,X2,X3))!=X4|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_185, c_0_81]), c_0_83])).
cnf(c_0_196, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_150]), c_0_120]), c_0_121])]), c_0_142])])).
cnf(c_0_197, negated_conjecture, (v2_struct_0(X1)|~v1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_115]), c_0_116]), c_0_120]), c_0_121]), c_0_122])]), c_0_127]), c_0_142])])).
cnf(c_0_198, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_189])).
cnf(c_0_199, negated_conjecture, (m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_190]), c_0_120]), c_0_121])]), c_0_142])])).
cnf(c_0_200, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(X1,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk4_3(X1,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(X1))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~l1_orders_2(X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192])])).
cnf(c_0_201, plain, (m1_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_202, negated_conjecture, (v4_yellow_0(k5_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk2_0),esk3_0)),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_115]), c_0_116]), c_0_120]), c_0_121]), c_0_122])]), c_0_142]), c_0_127])])).
cnf(c_0_203, plain, (k5_yellow_0(X2,u1_struct_0(k1_yellow_2(X1,X2,X3)))=k1_yellow_2(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(u1_struct_0(k1_yellow_2(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_195])])).
cnf(c_0_204, negated_conjecture, (v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_196, c_0_197])).
cnf(c_0_205, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_206, negated_conjecture, (m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_115]), c_0_116]), c_0_122])]), c_0_127])])).
cnf(c_0_207, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(X1,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk4_3(X1,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(X1))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_200, c_0_37])).
cnf(c_0_208, negated_conjecture, (u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0))=k10_xtuple_0(esk3_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_150]), c_0_120]), c_0_121])]), c_0_142])])).
cnf(c_0_209, negated_conjecture, (m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_190]), c_0_158]), c_0_121])])])).
cnf(c_0_210, negated_conjecture, (v4_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_202, c_0_150])])).
cnf(c_0_211, plain, (k5_yellow_0(X1,k10_xtuple_0(X3))=k1_yellow_2(X2,X1,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_129])])).
cnf(c_0_212, negated_conjecture, (v2_struct_0(X1)|~r1_lattice3(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_190])])).
cnf(c_0_213, negated_conjecture, (l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_121])])])).
cnf(c_0_214, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk4_3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208])])).
cnf(c_0_215, negated_conjecture, (l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_209]), c_0_121])])])).
cnf(c_0_216, negated_conjecture, (v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_158]), c_0_120]), c_0_121])]), c_0_142])])).
cnf(c_0_217, negated_conjecture, (v2_struct_0(X1)|~r1_lattice3(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_212, c_0_213])])).
cnf(c_0_218, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk4_3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_214, c_0_215]), c_0_216]), c_0_209])).
cnf(c_0_219, plain, (r1_lattice3(k1_yellow_2(X1,X2,X3),X4,X5)|m1_subset_1(esk4_3(k1_yellow_2(X1,X2,X3),X4,X5),k10_xtuple_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X5,k10_xtuple_0(X3))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_129])])).
cnf(c_0_220, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(X3)|~r1_lattice3(k1_yellow_2(X3,esk2_0,esk3_0),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(esk2_0))|~l1_orders_2(X3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_211]), c_0_158]), c_0_120]), c_0_121])]), c_0_142])])).
cnf(c_0_221, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219]), c_0_120]), c_0_121])]), c_0_142])])).
cnf(c_0_222, negated_conjecture, (v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_115]), c_0_116]), c_0_122])]), c_0_127])])).
cnf(c_0_223, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_221, c_0_215]), c_0_37])).
cnf(c_0_224, negated_conjecture, (v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223])])).
cnf(c_0_225, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_224, c_0_37])).
cnf(c_0_226, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_165])])).
cnf(c_0_227, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_38]), c_0_122]), c_0_138])])])).
cnf(c_0_228, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_229, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_128]), c_0_121])])])).
cnf(c_0_230, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_115]), c_0_116]), c_0_122])]), c_0_127])])).
cnf(c_0_231, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_232, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_229, c_0_230]), c_0_231])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 233
# Proof object clause steps            : 181
# Proof object formula steps           : 52
# Proof object conjectures             : 80
# Proof object clause conjectures      : 77
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 46
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 69
# Proof object simplifying inferences  : 146
# Training examples: 117 positive, 964 negative
# Training: Positive examples begin
cnf(c_0_138, negated_conjecture, (v1_yellow_0(esk1_0))).# trainpos
cnf(c_0_122, negated_conjecture, (l1_orders_2(esk1_0))).# trainpos
cnf(c_0_121, negated_conjecture, (l1_orders_2(esk2_0))).# trainpos
cnf(c_0_120, negated_conjecture, (v1_funct_1(esk3_0))).# trainpos
cnf(c_0_127, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_142, negated_conjecture, (~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_231, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_187, negated_conjecture, (~v1_yellow_0(k1_yellow_2(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_119, negated_conjecture, (v5_orders_3(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_110, plain, (l1_struct_0(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_116, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_115, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_106, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_125, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_37, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_105, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_157, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_109, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_205, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1))).# trainpos
cnf(c_0_55, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_38, plain, (m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_64, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_65, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_134, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_93, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_132, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_32, plain, (r1_lattice3(X1,u1_struct_0(X1),esk12_1(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_180, plain, (v1_yellow_0(X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,u1_struct_0(X2),X1))).# trainpos
cnf(c_0_58, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_201, plain, (m1_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_114, plain, (X3=k1_xboole_0|k1_relset_1(X2,X1)=X2|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_143, plain, (r2_hidden(X3,X4)|X4!=k10_xtuple_0(X2)|X3!=k1_funct_1(X2,X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2)))).# trainpos
cnf(c_0_49, plain, (u1_struct_0(X1)=X3|X1!=k5_yellow_0(X2,X3)|~l1_orders_2(X2)|~v1_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_167, plain, (k1_funct_1(X2,esk5_3(X2,X3,X1))=X1|X3!=k10_xtuple_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_192, plain, (r1_lattice3(X1,X2,X3)|r2_hidden(esk4_3(X1,X2,X3),X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).# trainpos
cnf(c_0_163, plain, (r2_hidden(esk5_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).# trainpos
cnf(c_0_62, plain, (r1_lattice3(X1,X2,X3)|m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).# trainpos
cnf(c_0_41, plain, (r1_lattice3(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2)))).# trainpos
cnf(c_0_179, plain, (X1=k5_yellow_0(X3,X2)|u1_struct_0(X1)!=X2|~l1_orders_2(X3)|~v1_orders_2(X1)|~v4_yellow_0(X1,X3)|~m1_yellow_0(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3))))).# trainpos
cnf(c_0_31, plain, (r1_orders_2(X1,X3,X4)|~l1_orders_2(X1)|~r2_hidden(X4,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).# trainpos
cnf(c_0_118, plain, (r1_orders_2(X3,X6,X7)|X7!=k1_funct_1(X1,X5)|X6!=k1_funct_1(X1,X4)|~l1_orders_2(X3)|~l1_orders_2(X2)|~v1_funct_1(X1)|~v5_orders_3(X1,X2,X3)|~r1_orders_2(X2,X4,X5)|~m1_subset_1(X7,u1_struct_0(X3))|~m1_subset_1(X6,u1_struct_0(X3))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).# trainpos
cnf(c_0_81, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_154, plain, (r1_orders_2(X2,X5,X6)|X4!=X6|X3!=X5|~l1_orders_2(X1)|~v4_yellow_0(X2,X1)|~m1_yellow_0(X2,X1)|~r1_orders_2(X1,X3,X4)|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X5,u1_struct_0(X2)))).# trainpos
cnf(c_0_82, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v4_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_83, plain, (v2_struct_0(X2)|v2_struct_0(X1)|m1_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_50, plain, (k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_133, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0)).# trainpos
cnf(c_0_160, plain, (r1_orders_2(X1,X2,X3)|~r1_orders_2(X4,X2,X3)|~r2_hidden(X2,u1_struct_0(X1))|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~l1_orders_2(X4))).# trainpos
cnf(c_0_188, plain, (k1_yellow_2(X1,X2,X3)=k1_yellow_2(X4,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~l1_orders_2(X4))).# trainpos
cnf(c_0_128, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_124, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_146, negated_conjecture, (v1_relat_1(esk3_0))).# trainpos
cnf(c_0_57, plain, (u1_struct_0(X1)=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X4)|v2_struct_0(X2)|X1!=k1_yellow_2(X4,X2,X3)|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~v1_orders_2(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X4))).# trainpos
cnf(c_0_36, plain, (r1_orders_2(X1,esk12_1(X1),X2)|~r2_hidden(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1))).# trainpos
cnf(c_0_151, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_194, plain, (v4_yellow_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_185, plain, (k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,X4)|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(k1_yellow_2(X1,X2,X3))!=X4|~m1_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~v1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X1))).# trainpos
cnf(c_0_63, plain, (u1_struct_0(X1)=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k1_yellow_2(X4,X2,X3)|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~v1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~v5_relat_1(X3,u1_struct_0(X2))|~v1_relat_1(X3))).# trainpos
cnf(c_0_139, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_229, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_47, plain, (r1_lattice3(X1,X2,esk12_1(X1))|~m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk4_3(X1,X2,esk12_1(X1)),u1_struct_0(X1))|~v1_yellow_0(X1))).# trainpos
cnf(c_0_145, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_197, negated_conjecture, (v2_struct_0(X1)|~v1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_202, negated_conjecture, (v4_yellow_0(k5_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk2_0),esk3_0)),esk2_0))).# trainpos
cnf(c_0_203, plain, (k5_yellow_0(X2,u1_struct_0(k1_yellow_2(X1,X2,X3)))=k1_yellow_2(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(u1_struct_0(k1_yellow_2(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_136, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,X4)|~r1_orders_2(esk1_0,X4,X3)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_144, negated_conjecture, (k9_xtuple_0(esk3_0)=u1_struct_0(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,u1_struct_0(esk1_0))|~v1_relat_1(esk3_0))).# trainpos
cnf(c_0_80, plain, (u1_struct_0(k1_yellow_2(X1,X2,X3))=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~v4_yellow_0(k1_yellow_2(X1,X2,X3),X2)|~v1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_61, plain, (r1_lattice3(X1,X2,esk12_1(X1))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk4_3(X1,X2,esk12_1(X1)),u1_struct_0(X1)))).# trainpos
cnf(c_0_186, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X1,X2,X3),k2_relset_1(u1_struct_0(X2),X3),X4)|~m1_subset_1(X4,k2_relset_1(u1_struct_0(X2),X3))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_111, plain, (u1_struct_0(k1_yellow_2(X3,X1,X2))=k10_xtuple_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~v5_relat_1(X2,u1_struct_0(X1))|~v1_relat_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(X3))).# trainpos
cnf(c_0_211, plain, (k5_yellow_0(X1,k10_xtuple_0(X3))=k1_yellow_2(X2,X1,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X1)|~l1_orders_2(X2))).# trainpos
cnf(c_0_141, plain, (k2_relset_1(u1_struct_0(X2),X3)=k10_xtuple_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_150, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=k10_xtuple_0(esk3_0))).# trainpos
cnf(c_0_219, plain, (r1_lattice3(k1_yellow_2(X1,X2,X3),X4,X5)|m1_subset_1(esk4_3(k1_yellow_2(X1,X2,X3),X4,X5),k10_xtuple_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X5,k10_xtuple_0(X3))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).# trainpos
cnf(c_0_67, plain, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(X1,X2,esk12_1(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk12_1(X1),u1_struct_0(X1)))).# trainpos
cnf(c_0_158, negated_conjecture, (m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_196, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_208, negated_conjecture, (u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0))=k10_xtuple_0(esk3_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_87, plain, (r1_orders_2(X1,esk12_1(X1),X2)|v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1))).# trainpos
cnf(c_0_210, negated_conjecture, (v4_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0))).# trainpos
cnf(c_0_164, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k10_xtuple_0(esk3_0)))).# trainpos
cnf(c_0_190, negated_conjecture, (k1_yellow_2(X1,esk2_0,esk3_0)=k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_216, negated_conjecture, (v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_209, negated_conjecture, (m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_199, negated_conjecture, (m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_215, negated_conjecture, (l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_206, negated_conjecture, (m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0))).# trainpos
cnf(c_0_123, plain, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk12_1(X1),X2)|v1_xboole_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,X3))).# trainpos
cnf(c_0_137, plain, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk12_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1))).# trainpos
cnf(c_0_213, negated_conjecture, (l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))).# trainpos
cnf(c_0_148, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,esk12_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_155, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|X1!=k1_funct_1(esk3_0,X2)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_161, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),k1_funct_1(esk3_0,X1))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_212, negated_conjecture, (v2_struct_0(X1)|~r1_lattice3(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~l1_orders_2(X1))).# trainpos
cnf(c_0_220, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(X3)|~r1_lattice3(k1_yellow_2(X3,esk2_0,esk3_0),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(esk2_0))|~l1_orders_2(X3))).# trainpos
cnf(c_0_222, negated_conjecture, (v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_166, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),k1_funct_1(esk3_0,X2))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(X1))|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_170, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X3)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk3_0,X2,X3),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,X2))).# trainpos
cnf(c_0_152, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|X1!=k1_funct_1(X3,X4)|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X4,k9_xtuple_0(X3)))).# trainpos
cnf(c_0_159, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,X3)|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_165, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r2_hidden(k1_funct_1(esk3_0,X1),X2)|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_169, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_172, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk5_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).# trainpos
cnf(c_0_168, plain, (m1_subset_1(esk5_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).# trainpos
cnf(c_0_171, negated_conjecture, (m1_subset_1(esk5_3(esk3_0,X1,X2),u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,X1))).# trainpos
cnf(c_0_174, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|m1_subset_1(X1,u1_struct_0(esk2_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_173, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(X2,X3)|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_178, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_182, negated_conjecture, (r1_lattice3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(X1)))).# trainpos
cnf(c_0_200, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(X1,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk4_3(X1,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(X1))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~l1_orders_2(X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1)))).# trainpos
cnf(c_0_214, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk4_3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_221, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0)))).# trainpos
cnf(c_0_224, negated_conjecture, (v2_struct_0(X1)|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_226, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_227, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).# trainpos
cnf(c_0_230, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_73, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_122, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_15, negated_conjecture, (v3_orders_2(esk1_0))).#trainneg
cnf(i_0_9, negated_conjecture, (v3_orders_2(esk2_0))).#trainneg
cnf(i_0_14, negated_conjecture, (v4_orders_2(esk1_0))).#trainneg
cnf(i_0_8, negated_conjecture, (v4_orders_2(esk2_0))).#trainneg
cnf(i_0_13, negated_conjecture, (v5_orders_2(esk1_0))).#trainneg
cnf(i_0_7, negated_conjecture, (v5_orders_2(esk2_0))).#trainneg
cnf(i_0_229, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_154, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_98, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_146, plain, (v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_184, plain, (v1_xboole_0(esk23_0))).#trainneg
cnf(i_0_153, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_120, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_170, plain, (v7_ordinal1(esk21_0))).#trainneg
cnf(i_0_181, plain, (v1_xreal_0(esk23_0))).#trainneg
cnf(i_0_182, plain, (v1_xxreal_0(esk23_0))).#trainneg
cnf(i_0_183, plain, (v1_xcmplx_0(esk23_0))).#trainneg
cnf(i_0_171, plain, (~v1_xboole_0(esk21_0))).#trainneg
cnf(i_0_188, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_233, plain, (v1_funct_1(esk28_1(X1)))).#trainneg
cnf(i_0_236, plain, (v1_relat_1(esk28_1(X1)))).#trainneg
cnf(i_0_76, plain, (v1_xboole_0(esk13_1(X1)))).#trainneg
cnf(i_0_232, plain, (v1_finseq_1(esk28_1(X1)))).#trainneg
cnf(i_0_144, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_59, plain, (X1=k1_xboole_0|X3!=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_68, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_119, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_117, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_118, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_155, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_172, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_147, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_94, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_136, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_235, plain, (v4_relat_1(esk28_1(X1),k5_numbers))).#trainneg
cnf(i_0_79, plain, (m1_subset_1(esk14_1(X1),X1))).#trainneg
cnf(i_0_234, plain, (v5_relat_1(esk28_1(X1),X1))).#trainneg
cnf(i_0_238, plain, (v1_funct_1(esk29_2(X1,X2)))).#trainneg
cnf(i_0_127, plain, (v1_relat_1(esk15_2(X1,X2)))).#trainneg
cnf(i_0_241, plain, (v1_relat_1(esk29_2(X1,X2)))).#trainneg
cnf(i_0_128, plain, (v1_xboole_0(esk15_2(X1,X2)))).#trainneg
cnf(i_0_134, plain, (~v1_subset_1(esk17_1(X1),X1))).#trainneg
cnf(i_0_231, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_131, plain, (~v2_struct_0(X1)|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_152, plain, (~v2_struct_0(X1)|~l1_orders_2(X1)|~v1_lattice3(X1))).#trainneg
cnf(i_0_148, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk19_1(X1)))).#trainneg
cnf(i_0_247, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_185, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_248, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_97, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_108, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_143, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_151, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_223, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_230, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_156, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_166, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_77, plain, (m1_subset_1(esk13_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_135, plain, (m1_subset_1(esk17_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_125, plain, (v5_relat_1(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_239, plain, (v5_relat_1(esk29_2(X1,X2),X2))).#trainneg
cnf(i_0_126, plain, (v4_relat_1(esk15_2(X1,X2),X1))).#trainneg
cnf(i_0_240, plain, (v4_relat_1(esk29_2(X1,X2),X1))).#trainneg
cnf(i_0_168, plain, (v1_funct_1(esk20_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_178, plain, (v1_funct_1(esk22_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_225, plain, (v1_orders_2(esk27_1(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_169, plain, (v1_relat_1(esk20_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_211, plain, (v1_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_216, plain, (v1_xboole_0(X1)|v3_card_1(esk26_1(X1),np__1))).#trainneg
cnf(i_0_90, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_109, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_209, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_132, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk16_1(X1)))).#trainneg
cnf(i_0_149, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk19_1(X1)))).#trainneg
cnf(i_0_210, plain, (v2_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_243, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_180, plain, (v1_relat_1(esk22_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116, plain, (v1_xboole_0(k10_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_96, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_187, plain, (l1_struct_0(esk24_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_177, plain, (v1_finset_1(esk22_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_175, plain, (v1_finseq_1(esk22_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_174, plain, (v2_finseq_1(esk22_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_142, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_246, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_165, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_179, plain, (v4_relat_1(esk22_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_224, plain, (v4_yellow_0(esk27_1(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_226, plain, (m1_yellow_0(esk27_1(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_167, plain, (v3_card_1(esk20_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_212, plain, (v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_91, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_140, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_138, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_158, plain, (v1_finset_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_137, plain, (v2_struct_0(X1)|v2_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_227, plain, (v2_struct_0(X1)|v3_yellow_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_176, plain, (v3_card_1(esk22_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_249, plain, (v3_card_1(esk30_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_186, plain, (v13_struct_0(esk24_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_107, plain, (v1_xboole_0(k2_zfmisc_1(X2,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_106, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_237, plain, (m1_subset_1(esk28_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_133, plain, (v1_xboole_0(X1)|m1_subset_1(esk16_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_150, plain, (v1_xboole_0(X1)|m1_subset_1(esk19_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_217, plain, (v1_xboole_0(X1)|m1_subset_1(esk26_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_82, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_129, plain, (m1_subset_1(esk15_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_242, plain, (m1_subset_1(esk29_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_58, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_121, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_85, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_208, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_163, plain, (v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_75, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_218, plain, (v1_yellow_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v8_struct_0(X1))).#trainneg
cnf(i_0_80, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=X1|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_157, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_100, plain, (l1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_101, plain, (v1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_92, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_93, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_89, plain, (v3_orders_2(X1)|~v3_orders_2(X2)|~l1_orders_2(X2)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_105, plain, (v4_orders_2(X1)|~v4_orders_2(X2)|~l1_orders_2(X2)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_55, plain, (v1_orders_2(k5_yellow_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_52, plain, (v2_struct_0(X2)|v2_struct_0(X1)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v2_struct_0(k1_yellow_2(X1,X2,X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_145, plain, (v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_103, plain, (v5_orders_2(X1)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_130, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_113, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_192, plain, (v2_struct_0(X2)|v1_funct_1(esk25_2(X1,X2))|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_195, plain, (v2_struct_0(X2)|v1_relat_1(esk25_2(X1,X2))|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_71, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_173, plain, (v1_partfun1(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_54, plain, (v4_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_25, plain, (r1_tarski(u1_orders_2(X1),u1_orders_2(X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X1,X2))).#trainneg
cnf(i_0_123, plain, (v3_card_1(k9_xtuple_0(X2),X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_card_1(X1)|~v3_card_1(X2,X1))).#trainneg
cnf(i_0_60, plain, (X3=k1_xboole_0|v1_funct_2(X2,X1,X3)|k1_relset_1(X1,X2)!=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_193, plain, (v2_struct_0(X2)|v5_relat_1(esk25_2(X1,X2),u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_194, plain, (v2_struct_0(X2)|v4_relat_1(esk25_2(X1,X2),u1_struct_0(X1))|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_191, plain, (v2_struct_0(X2)|v1_partfun1(esk25_2(X1,X2),u1_struct_0(X1))|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_112, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_struct_0(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_78, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_34, plain, (v5_orders_3(X3,X2,X1)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~r1_orders_2(X1,esk10_3(X2,X1,X3),esk11_3(X2,X1,X3)))).#trainneg
cnf(i_0_114, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_24, plain, (m1_yellow_0(X1,X2)|~l1_orders_2(X2)|~l1_orders_2(X1)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~r1_tarski(u1_orders_2(X1),u1_orders_2(X2)))).#trainneg
cnf(i_0_141, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_66, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_29, plain, (X2=k10_xtuple_0(X1)|r2_hidden(esk6_2(X1,X2),X2)|r2_hidden(esk7_2(X1,X2),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_189, plain, (v2_struct_0(X2)|v5_orders_3(esk25_2(X1,X2),X1,X2)|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_190, plain, (v2_struct_0(X2)|v1_funct_2(esk25_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_28, plain, (X2=k10_xtuple_0(X1)|k1_funct_1(X1,esk7_2(X1,X2))=esk6_2(X1,X2)|r2_hidden(esk6_2(X1,X2),X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_196, plain, (v2_struct_0(X2)|m1_subset_1(esk25_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v3_orders_2(X2)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_30, plain, (X2=k10_xtuple_0(X1)|esk6_2(X1,X2)!=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_xtuple_0(X1))|~r2_hidden(esk6_2(X1,X2),X2))).#trainneg
cnf(i_0_67, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r3_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_62, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_69, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_41, plain, (v5_orders_3(X3,X1,X2)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_40, plain, (v5_orders_3(X3,X1,X2)|m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_38, plain, (v5_orders_3(X3,X1,X2)|m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_37, plain, (v5_orders_3(X3,X1,X2)|m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_36, plain, (k1_funct_1(X3,esk8_3(X1,X2,X3))=esk10_3(X1,X2,X3)|v5_orders_3(X3,X1,X2)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_35, plain, (k1_funct_1(X3,esk9_3(X1,X2,X3))=esk11_3(X1,X2,X3)|v5_orders_3(X3,X1,X2)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_39, plain, (v5_orders_3(X3,X1,X2)|r1_orders_2(X1,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_584, plain, (u1_struct_0(k5_yellow_0(X1,X2))=X2|~m1_yellow_0(k5_yellow_0(X1,X2),X1)|~v4_yellow_0(k5_yellow_0(X1,X2),X1)|~v1_orders_2(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_251, plain, (esk23_0=k1_xboole_0)).#trainneg
cnf(i_0_250, negated_conjecture, (esk3_0=k1_xboole_0|u1_struct_0(esk2_0)!=k1_xboole_0)).#trainneg
cnf(i_0_252, plain, (esk18_0=k1_xboole_0)).#trainneg
cnf(i_0_356, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_256, plain, (~v7_ordinal1(k5_numbers))).#trainneg
cnf(i_0_377, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_253, plain, (esk13_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_254, plain, (esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2))).#trainneg
cnf(i_0_577, plain, (r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_576, plain, (r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_575, plain, (k1_relset_1(u1_struct_0(X1),u1_orders_2(X1))=u1_struct_0(X1)|u1_struct_0(X1)=k1_xboole_0|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_255, plain, (esk15_2(X1,X2)=k1_xboole_0)).#trainneg
cnf(i_0_595, plain, (r1_lattice3(X2,X1,X3)|~v1_xboole_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_581, plain, (m1_yellow_0(X1,X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_644, plain, (r1_orders_2(X1,X2,X3)|X3!=k1_funct_1(u1_orders_2(X1),X4)|X2!=k1_funct_1(u1_orders_2(X1),X5)|~r1_orders_2(X1,X5,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~v5_orders_3(u1_orders_2(X1),X1,X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_282, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk17_1(X2)))).#trainneg
cnf(i_0_283, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk14_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_358, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_307, plain, (k5_numbers=X1|g1_orders_2(k5_numbers,esk28_1(k5_numbers))!=g1_orders_2(X1,X2))).#trainneg
cnf(i_0_263, plain, (esk17_1(k2_zfmisc_1(X1,X2))=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(esk17_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_274, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_279, plain, (~v1_xboole_0(esk20_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_360, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_266, plain, (esk17_1(k2_zfmisc_1(X1,X1))=X2|g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_267, plain, (esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=X2|g1_orders_2(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_461, plain, (u1_orders_2(X1)=k1_xboole_0|u1_struct_0(X1)!=k1_xboole_0|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_562, plain, (u1_struct_0(X1)=k1_xboole_0|v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|k1_relset_1(u1_struct_0(X1),u1_orders_2(X1))!=u1_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_703, plain, (v1_yellow_0(k5_yellow_0(X1,X2))|~r1_lattice3(k5_yellow_0(X1,X2),X2,X3)|~m1_subset_1(X3,X2)|~l1_orders_2(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_181, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_182, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_77, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_275, plain, (k10_xtuple_0(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_183, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_752, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_276, plain, (X1=k10_xtuple_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_277, plain, (k9_xtuple_0(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_717, plain, (u1_struct_0(k1_yellow_2(X3,X1,X2))=k2_relset_1(u1_struct_0(X1),X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_773, plain, (r1_orders_2(X1,esk12_1(X2),X3)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(esk12_1(X2),u1_struct_0(X2))|~l1_orders_2(X2)|~r2_hidden(X3,u1_struct_0(X2))|~v1_yellow_0(X2))).#trainneg
cnf(i_0_696, plain, (r1_tarski(X2,u1_struct_0(X3))|~m1_yellow_0(k5_yellow_0(X1,X2),X3)|~l1_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_699, plain, (r1_lattice3(k5_yellow_0(X1,X2),X2,esk12_1(k5_yellow_0(X1,X2)))|~l1_orders_2(k5_yellow_0(X1,X2))|~v1_yellow_0(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_286, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_127, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_280, plain, (v1_xboole_0(esk20_1(k1_xboole_0))|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_278, plain, (X1=k9_xtuple_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_478, plain, (v1_funct_1(k5_numbers)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers))).#trainneg
cnf(i_0_296, plain, (~v1_xboole_0(esk22_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_298, plain, (~v1_xboole_0(esk30_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_285, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_287, plain, (X1=X2|g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_709, plain, (r1_tarski(u1_struct_0(X1),X3)|~m1_yellow_0(X1,k5_yellow_0(X2,X3))|~l1_orders_2(k5_yellow_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2))).#trainneg
cnf(i_0_767, plain, (r1_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_701, plain, (m1_subset_1(u1_orders_2(k5_yellow_0(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~l1_orders_2(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_807, plain, (v4_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))),X1)|v2_struct_0(X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_841, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_894, plain, (esk20_1(k1_xboole_0)=k1_xboole_0|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_126, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_288, plain, (X1=X2|g1_orders_2(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_125, plain, (v5_relat_1(k1_xboole_0,X2))).#trainneg
cnf(i_0_294, plain, (~v1_finseq_1(k5_numbers)|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers))).#trainneg
cnf(i_0_750, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_754, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_290, plain, (v7_ordinal1(esk22_1(X1))|~v3_ordinal1(esk22_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_813, plain, (v1_yellow_0(k1_yellow_2(X4,X1,X2))|v2_struct_0(X4)|v2_struct_0(X1)|~r1_lattice3(k1_yellow_2(X4,X1,X2),k2_relset_1(u1_struct_0(X1),X2),X3)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k2_relset_1(u1_struct_0(X1),X2))|~l1_orders_2(k1_yellow_2(X4,X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_863, plain, (r1_lattice3(X1,X2,esk12_1(X3))|~m1_subset_1(esk12_1(X3),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk12_1(X3),u1_struct_0(X1))|~r2_hidden(esk4_3(X1,X2,esk12_1(X3)),u1_struct_0(X3))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~m1_subset_1(esk4_3(X1,X2,esk12_1(X3)),u1_struct_0(X1))|~l1_orders_2(X3)|~v1_yellow_0(X3))).#trainneg
cnf(i_0_936, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,esk12_1(X4))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(X4),u1_struct_0(esk1_0))|~r2_hidden(esk12_1(X4),u1_struct_0(esk1_0))|~r2_hidden(X3,u1_struct_0(X4))|~m1_yellow_0(esk1_0,X4)|~v4_yellow_0(esk1_0,X4)|~l1_orders_2(X4)|~v1_yellow_0(X4))).#trainneg
cnf(i_0_911, plain, (r1_orders_2(X1,X2,X2)|v2_struct_0(X3)|~r2_hidden(X2,u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X3))|~l1_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_860, negated_conjecture, (v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_yellow_0(k5_yellow_0(esk2_0,k2_relset_1(u1_struct_0(esk2_0),esk3_0))))).#trainneg
cnf(i_0_939, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,X3)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_865, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),X2),u1_struct_0(X3))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_yellow_0(k1_yellow_2(X4,X1,X2),X3)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X3)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_876, plain, (r1_lattice3(k1_yellow_2(X3,X1,X2),k2_relset_1(u1_struct_0(X1),X2),esk12_1(k1_yellow_2(X3,X1,X2)))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(k1_yellow_2(X3,X1,X2))|~l1_orders_2(X1)|~v1_yellow_0(k1_yellow_2(X3,X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_840, plain, (v1_funct_1(k1_xboole_0)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_366, plain, (l1_orders_2(esk27_1(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_947, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_784, plain, (u1_struct_0(X1)=k5_numbers|g1_orders_2(k5_numbers,esk28_1(k5_numbers))!=X1|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_295, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_297, plain, (v1_xboole_0(esk22_1(k1_xboole_0))|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_301, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_299, plain, (v1_xboole_0(esk30_1(k1_xboole_0))|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_300, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1031, plain, (r1_orders_2(X1,esk4_3(X2,X3,X4),esk4_3(X2,X3,X4))|v2_struct_0(X2)|r1_lattice3(X2,X3,X4)|~r2_hidden(esk4_3(X2,X3,X4),u1_struct_0(X1))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~l1_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(X2)))).#trainneg
cnf(i_0_1022, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),X2)|X2!=k1_funct_1(esk3_0,X3)|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~r2_hidden(X3,u1_struct_0(X1))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_1021, plain, (r1_lattice3(X1,u1_struct_0(X2),esk12_1(X2))|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~m1_subset_1(esk4_3(X1,u1_struct_0(X2),esk12_1(X2)),u1_struct_0(X1))|~l1_orders_2(X2)|~v1_yellow_0(X2)|~m1_subset_1(esk12_1(X2),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1039, negated_conjecture, (r1_orders_2(esk2_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk5_3(esk3_0,X2,X1),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_961, plain, (r1_tarski(k2_relset_1(u1_struct_0(X2),X3),u1_struct_0(X4))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_yellow_0(k1_yellow_2(X1,X2,X3),X4)|~l1_orders_2(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_972, plain, (r1_orders_2(k1_yellow_2(X1,X2,X3),X4,X5)|v2_struct_0(X2)|v2_struct_0(X1)|X5!=k1_funct_1(X6,X7)|X4!=k1_funct_1(X6,X8)|~r1_orders_2(X9,X8,X7)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),k2_relset_1(u1_struct_0(X2),X3))))|~m1_subset_1(X5,k2_relset_1(u1_struct_0(X2),X3))|~m1_subset_1(X4,k2_relset_1(u1_struct_0(X2),X3))|~m1_subset_1(X7,u1_struct_0(X9))|~m1_subset_1(X8,u1_struct_0(X9))|~v5_orders_3(X6,X9,k1_yellow_2(X1,X2,X3))|~v1_funct_2(X6,u1_struct_0(X9),k2_relset_1(u1_struct_0(X2),X3))|~v1_funct_1(X6)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X9)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1020, plain, (r1_lattice3(X1,X2,esk12_1(X3))|v1_xboole_0(u1_struct_0(X3))|~r2_hidden(esk12_1(X3),u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~m1_subset_1(esk4_3(X1,X2,esk12_1(X3)),u1_struct_0(X1))|~l1_orders_2(X3)|~v1_yellow_0(X3)|~m1_subset_1(esk4_3(X1,X2,esk12_1(X3)),u1_struct_0(X3)))).#trainneg
cnf(i_0_1061, plain, (v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1075, plain, (esk22_1(k1_xboole_0)=k1_xboole_0|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_306, plain, (esk28_1(k5_numbers)=X1|g1_orders_2(k5_numbers,esk28_1(k5_numbers))!=g1_orders_2(X2,X1))).#trainneg
cnf(i_0_302, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_999, plain, (v3_card_1(k1_xboole_0,k1_xboole_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_303, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_741, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_304, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,esk28_1(X2)))).#trainneg
cnf(i_0_308, plain, (esk28_1(X1)=k1_xboole_0|X1!=k1_xboole_0|~v1_funct_2(esk28_1(X1),k5_numbers,X1))).#trainneg
cnf(i_0_1148, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r1_lattice3(k1_yellow_2(X1,X2,X3),k10_xtuple_0(X3),X4)|~m1_subset_1(X4,k10_xtuple_0(X3))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1142, plain, (r1_tarski(k10_xtuple_0(X3),u1_struct_0(X4))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_yellow_0(k1_yellow_2(X1,X2,X3),X4)|~l1_orders_2(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1126, plain, (r1_orders_2(X1,esk4_3(X2,X3,X4),esk4_3(X2,X3,X4))|r1_lattice3(X2,X3,X4)|v2_struct_0(X2)|~r2_hidden(esk4_3(X2,X3,X4),u1_struct_0(X1))|~m1_yellow_0(X1,X5)|~v4_yellow_0(X1,X5)|~m1_subset_1(esk4_3(X2,X3,X4),u1_struct_0(X1))|~m1_subset_1(esk4_3(X2,X3,X4),u1_struct_0(X5))|~l1_orders_2(X5)|~r2_hidden(esk4_3(X2,X3,X4),u1_struct_0(X5))|~m1_yellow_0(X5,X2)|~v4_yellow_0(X5,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1145, plain, (r1_lattice3(k1_yellow_2(X1,X2,X3),k10_xtuple_0(X3),esk12_1(k1_yellow_2(X1,X2,X3)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~v1_yellow_0(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1178, plain, (k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|k10_xtuple_0(X3)!=X4|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1189, plain, (r1_lattice3(X1,u1_struct_0(X2),esk12_1(X2))|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~l1_orders_2(X2)|~v1_yellow_0(X2)|~m1_subset_1(esk12_1(X2),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1130, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),X2)|v1_xboole_0(u1_struct_0(X1))|X2!=k1_funct_1(esk3_0,X3)|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_825, plain, (l1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1229, plain, (u1_orders_2(X1)=esk28_1(k5_numbers)|g1_orders_2(k5_numbers,esk28_1(k5_numbers))!=X1|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_309, plain, (esk16_1(k2_zfmisc_1(X1,X1))=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk16_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_1225, plain, (v3_card_1(k1_xboole_0,k1_xboole_0)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_1222, plain, (v1_finseq_1(k1_xboole_0)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_1223, plain, (v2_finseq_1(k1_xboole_0)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_755, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_310, plain, (X1=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk16_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_311, plain, (esk16_1(k2_zfmisc_1(X1,X2))=k1_xboole_0|v1_xboole_0(k2_zfmisc_1(X1,X2))|X2!=k1_xboole_0|~v1_funct_2(esk16_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_1218, plain, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(X2,X3,esk12_1(X1))|~r2_hidden(esk12_1(X1),u1_struct_0(X2))|~m1_yellow_0(X2,X1)|~v4_yellow_0(X2,X1)|~m1_subset_1(esk4_3(X2,X3,esk12_1(X1)),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk12_1(X1),u1_struct_0(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1155, plain, (r1_lattice3(k1_yellow_2(X1,X2,X3),X4,esk12_1(k1_yellow_2(X1,X2,X3)))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk4_3(k1_yellow_2(X1,X2,X3),X4,esk12_1(k1_yellow_2(X1,X2,X3))),k10_xtuple_0(X3))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~v1_yellow_0(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1321, plain, (r1_tarski(k10_xtuple_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_orders_2(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_1303, plain, (u1_struct_0(X1)=k10_xtuple_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|X1!=k1_yellow_2(X4,X3,X2)|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~v1_orders_2(X1)|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_462, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1087, plain, (esk30_1(k1_xboole_0)=k1_xboole_0|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_827, plain, (k1_xboole_0=X1|g1_orders_2(X2,k1_xboole_0)!=g1_orders_2(X3,X1))).#trainneg
cnf(i_0_312, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk16_1(X2)))).#trainneg
cnf(i_0_835, plain, (v1_partfun1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_376, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_354, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1278, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_314, plain, (esk19_1(k2_zfmisc_1(X1,X1))=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_315, plain, (X1=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_1378, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),X2)|r1_lattice3(X1,X3,X4)|X2!=k1_funct_1(esk3_0,esk4_3(X1,X3,X4))|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk4_3(X1,X3,X4),u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_1316, plain, (r1_tarski(k10_xtuple_0(X1),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X4)|~m1_yellow_0(k5_yellow_0(X3,k2_relset_1(u1_struct_0(X3),X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_orders_2(X2)|~l1_orders_2(X3)|~l1_orders_2(X4))).#trainneg
cnf(i_0_989, plain, (k1_yellow_2(X1,X2,X3)=k5_yellow_0(X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|k2_relset_1(u1_struct_0(X2),X3)!=X4|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_910, plain, (r1_tarski(u1_struct_0(k5_yellow_0(X1,X2)),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(k5_yellow_0(X1,X2))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1111, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_826, plain, (v1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1118, plain, (esk17_1(k1_xboole_0)=X2|g1_orders_2(X1,esk17_1(k1_xboole_0))!=g1_orders_2(X3,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_316, plain, (esk19_1(k2_zfmisc_1(X1,X2))=k1_xboole_0|v1_xboole_0(k2_zfmisc_1(X1,X2))|X2!=k1_xboole_0|~v1_funct_2(esk19_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_414, plain, (l1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))).#trainneg
cnf(i_0_1277, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_493, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2))).#trainneg
cnf(i_0_317, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk19_1(X2)))).#trainneg
cnf(i_0_319, plain, (esk26_1(k2_zfmisc_1(X1,X1))=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk26_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_1026, plain, (r1_orders_2(X1,esk12_1(X2),esk12_1(X2))|v2_struct_0(X2)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~l1_orders_2(X2)|~v3_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_1018, plain, (r1_lattice3(X1,X2,esk12_1(k5_yellow_0(X3,X4)))|~r2_hidden(esk4_3(X1,X2,esk12_1(k5_yellow_0(X3,X4))),X4)|~r2_hidden(esk12_1(k5_yellow_0(X3,X4)),u1_struct_0(X1))|~m1_yellow_0(X1,k5_yellow_0(X3,X4))|~v4_yellow_0(X1,k5_yellow_0(X3,X4))|~m1_subset_1(esk4_3(X1,X2,esk12_1(k5_yellow_0(X3,X4))),u1_struct_0(X1))|~l1_orders_2(k5_yellow_0(X3,X4))|~v1_yellow_0(k5_yellow_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3))).#trainneg
cnf(i_0_1253, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_472, plain, (v1_funct_1(esk17_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_948, negated_conjecture, (v3_ordinal1(u1_struct_0(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|~v5_ordinal1(esk3_0))).#trainneg
cnf(i_0_320, plain, (X1=X2|v1_xboole_0(k2_zfmisc_1(X1,X1))|g1_orders_2(X1,esk26_1(k2_zfmisc_1(X1,X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_823, plain, (v1_xboole_0(X1)|~v1_partfun1(k1_xboole_0,X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_423, plain, (v1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))).#trainneg
cnf(i_0_1069, plain, (u1_struct_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))=k5_numbers|~v1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))|~l1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))).#trainneg
cnf(i_0_321, plain, (esk26_1(k2_zfmisc_1(X1,X2))=k1_xboole_0|v1_xboole_0(k2_zfmisc_1(X1,X2))|X2!=k1_xboole_0|~v1_funct_2(esk26_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_322, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk26_1(X2)))).#trainneg
cnf(i_0_860, negated_conjecture, (v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_1564, plain, (g1_orders_2(k5_numbers,u1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))=g1_orders_2(k5_numbers,esk28_1(k5_numbers)))).#trainneg
cnf(i_0_1548, plain, (r1_lattice3(X1,X2,esk12_1(k1_yellow_2(X5,X3,X4)))|v2_struct_0(X5)|v2_struct_0(X3)|~r2_hidden(esk4_3(X1,X2,esk12_1(k1_yellow_2(X5,X3,X4))),k2_relset_1(u1_struct_0(X3),X4))|~r2_hidden(esk12_1(k1_yellow_2(X5,X3,X4)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X5,X3,X4))|~v4_yellow_0(X1,k1_yellow_2(X5,X3,X4))|~m1_subset_1(esk4_3(X1,X2,esk12_1(k1_yellow_2(X5,X3,X4))),u1_struct_0(X1))|~m1_subset_1(k2_relset_1(u1_struct_0(X3),X4),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(k1_yellow_2(X5,X3,X4))|~l1_orders_2(X3)|~v1_yellow_0(k1_yellow_2(X5,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X5),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_orders_2(X5))).#trainneg
cnf(i_0_1547, plain, (r1_lattice3(X1,X2,esk12_1(k1_yellow_2(X5,X3,X4)))|v2_struct_0(X3)|v2_struct_0(X5)|~r2_hidden(esk4_3(X1,X2,esk12_1(k1_yellow_2(X5,X3,X4))),k10_xtuple_0(X4))|~r2_hidden(esk12_1(k1_yellow_2(X5,X3,X4)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X5,X3,X4))|~v4_yellow_0(X1,k1_yellow_2(X5,X3,X4))|~m1_subset_1(esk4_3(X1,X2,esk12_1(k1_yellow_2(X5,X3,X4))),u1_struct_0(X1))|~m1_subset_1(k10_xtuple_0(X4),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(k1_yellow_2(X5,X3,X4))|~l1_orders_2(X3)|~v1_yellow_0(k1_yellow_2(X5,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X5),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_orders_2(X5))).#trainneg
cnf(i_0_1504, plain, (r1_tarski(u1_struct_0(k1_yellow_2(X3,X1,X2)),k10_xtuple_0(X2))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(k1_yellow_2(X3,X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_1472, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),k1_funct_1(esk3_0,esk4_3(X1,X2,X3)))|r1_lattice3(X1,X2,X3)|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk4_3(X1,X2,X3)),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_385, plain, (v1_relat_1(esk17_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_330, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk29_2(X2,X3)))).#trainneg
cnf(i_0_1463, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_332, plain, (esk29_2(X1,X1)=X2|g1_orders_2(X1,esk29_2(X1,X1))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_333, plain, (X1=X2|g1_orders_2(X1,esk29_2(X1,X1))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_1624, negated_conjecture, (u1_struct_0(X1)=k10_xtuple_0(esk3_0)|v2_struct_0(X2)|X1!=k1_yellow_2(X2,esk2_0,esk3_0)|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~v1_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1501, negated_conjecture, (k1_yellow_2(X1,esk2_0,esk3_0)=k5_yellow_0(esk2_0,X2)|v2_struct_0(X1)|k10_xtuple_0(esk3_0)!=X2|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1432, negated_conjecture, (r1_tarski(k10_xtuple_0(esk3_0),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1692, plain, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X1,esk12_1(X1),esk4_3(X2,X3,X4))|r1_lattice3(X2,X3,X4)|~m1_subset_1(esk4_3(X2,X3,X4),u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1636, negated_conjecture, (r1_tarski(k10_xtuple_0(esk3_0),u1_struct_0(X2))|v2_struct_0(X1)|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),X2)|~l1_orders_2(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_413, plain, (l1_orders_2(g1_orders_2(X1,esk29_2(X1,X1))))).#trainneg
cnf(i_0_1696, plain, (u1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))=esk17_1(k1_xboole_0)|g1_orders_2(X1,esk17_1(k1_xboole_0))!=g1_orders_2(k5_numbers,esk28_1(k5_numbers))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_334, plain, (esk29_2(X1,X2)=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(esk29_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_1082, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_335, plain, (v1_funct_2(esk29_2(X1,X2),X1,X2)|X2!=k1_xboole_0|esk29_2(X1,X2)!=k1_xboole_0)).#trainneg
cnf(i_0_336, plain, (v1_funct_2(esk28_1(X1),k5_numbers,X1)|X1!=k1_xboole_0|esk28_1(X1)!=k1_xboole_0)).#trainneg
cnf(i_0_1773, negated_conjecture, (r1_tarski(u1_struct_0(X1),k10_xtuple_0(esk3_0))|v2_struct_0(X2)|~m1_yellow_0(X1,k1_yellow_2(X2,esk2_0,esk3_0))|~l1_orders_2(k1_yellow_2(X2,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_744, plain, (l1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1619, negated_conjecture, (v3_ordinal1(k10_xtuple_0(esk3_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1697, plain, (u1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))=k1_xboole_0|g1_orders_2(X1,k1_xboole_0)!=g1_orders_2(k5_numbers,esk28_1(k5_numbers)))).#trainneg
cnf(i_0_337, plain, (v1_funct_2(esk26_1(k2_zfmisc_1(X1,X2)),X1,X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|X2!=k1_xboole_0|esk26_1(k2_zfmisc_1(X1,X2))!=k1_xboole_0)).#trainneg
cnf(i_0_1742, plain, (v1_relat_1(esk17_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1459, negated_conjecture, (v3_ordinal1(esk3_0)|~v1_card_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_338, plain, (v1_funct_2(esk19_1(k2_zfmisc_1(X1,X2)),X1,X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|X2!=k1_xboole_0|esk19_1(k2_zfmisc_1(X1,X2))!=k1_xboole_0)).#trainneg
cnf(i_0_339, plain, (v1_funct_2(esk17_1(k2_zfmisc_1(X1,X2)),X1,X2)|X2!=k1_xboole_0|esk17_1(k2_zfmisc_1(X1,X2))!=k1_xboole_0)).#trainneg
cnf(i_0_1785, negated_conjecture, (k1_yellow_2(X2,esk2_0,esk3_0)=k1_yellow_2(X1,esk2_0,esk3_0)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1823, negated_conjecture, (k1_yellow_2(X1,esk2_0,esk3_0)=k5_yellow_0(esk2_0,X2)|v2_struct_0(X1)|u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0))!=X2|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),esk2_0)|~v1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1576, plain, (r1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers)),X1,X2)|X2!=k1_funct_1(X3,X4)|X1!=k1_funct_1(X3,X5)|~r1_orders_2(X6,X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),k5_numbers)))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,u1_struct_0(X6))|~m1_subset_1(X5,u1_struct_0(X6))|~v5_orders_3(X3,X6,g1_orders_2(k5_numbers,esk28_1(k5_numbers)))|~v1_funct_2(X3,u1_struct_0(X6),k5_numbers)|~v1_funct_1(X3)|~l1_orders_2(X6))).#trainneg
cnf(i_0_1780, negated_conjecture, (r1_tarski(u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)),k10_xtuple_0(esk3_0))|v2_struct_0(X1)|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1986, plain, (v1_relat_1(esk17_1(k1_xboole_0)))).#trainneg
cnf(i_0_1771, negated_conjecture, (r1_tarski(k1_xboole_0,u1_struct_0(esk2_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_1449, plain, (u1_orders_2(X1)=k1_xboole_0|g1_orders_2(X2,k1_xboole_0)!=X1|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_340, plain, (v1_funct_2(esk16_1(k2_zfmisc_1(X1,X2)),X1,X2)|v1_xboole_0(k2_zfmisc_1(X1,X2))|X2!=k1_xboole_0|esk16_1(k2_zfmisc_1(X1,X2))!=k1_xboole_0)).#trainneg
cnf(i_0_1866, plain, (v4_relat_1(k2_zfmisc_1(X3,X4),X1)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_342, plain, (v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|X2!=k1_xboole_0|esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))!=k1_xboole_0)).#trainneg
cnf(i_0_345, plain, (v3_ordinal1(esk29_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1602, plain, (k2_relset_1(u1_struct_0(X1),X2)=k10_xtuple_0(X2)|v2_struct_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~v1_funct_2(X2,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_509, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_460, plain, (v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1621, negated_conjecture, (v1_funct_1(k10_xtuple_0(esk3_0))|~v1_relat_1(u1_struct_0(esk2_0))|~v1_funct_1(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1521, plain, (esk17_1(k1_xboole_0)=u1_orders_2(X1)|g1_orders_2(X2,esk17_1(k1_xboole_0))!=X1|~v1_xboole_0(X2)|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_346, plain, (v3_ordinal1(esk28_1(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_836, plain, (v1_funct_2(k1_xboole_0,X1,X2)|X2!=k1_xboole_0)).#trainneg
cnf(i_0_1460, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_347, plain, (v3_ordinal1(esk26_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_348, plain, (v3_ordinal1(esk19_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1529, plain, (m1_subset_1(u1_struct_0(k1_yellow_2(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1543, plain, (r1_orders_2(X1,esk12_1(X2),esk12_1(X2))|v2_struct_0(X2)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~m1_subset_1(esk12_1(X2),u1_struct_0(X1))|~m1_subset_1(esk12_1(X2),u1_struct_0(X3))|~l1_orders_2(X3)|~r2_hidden(esk12_1(X2),u1_struct_0(X3))|~m1_yellow_0(X3,X2)|~v4_yellow_0(X3,X2)|~l1_orders_2(X2)|~v1_yellow_0(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1488, negated_conjecture, (r1_tarski(k10_xtuple_0(esk3_0),u1_struct_0(X1))|~m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_527, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_464, plain, (v4_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1748, negated_conjecture, (v4_yellow_0(k5_yellow_0(esk2_0,k1_xboole_0),esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_349, plain, (v3_ordinal1(esk17_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_895, plain, (X1=esk20_1(k1_xboole_0)|~v1_xboole_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_1462, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))).#trainneg
cnf(i_0_1867, plain, (v1_partfun1(k2_zfmisc_1(X3,X4),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_350, plain, (v3_ordinal1(esk16_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_352, plain, (v3_ordinal1(esk14_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_2177, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,u1_struct_0(k1_yellow_2(X3,X2,X4)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~l1_orders_2(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_2192, plain, (m1_subset_1(k2_relset_1(u1_struct_0(X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2158, plain, (m1_subset_1(X1,X2)|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k10_xtuple_0(X3)))).#trainneg
cnf(i_0_2193, plain, (m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2213, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3)))).#trainneg
cnf(i_0_455, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_359, plain, (v1_xboole_0(u1_struct_0(esk24_1(X1)))|~v2_struct_0(esk24_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_361, plain, (v2_struct_0(esk24_1(X1))|~v1_xboole_0(u1_struct_0(esk24_1(X1)))|~v1_card_1(X1))).#trainneg
cnf(i_0_2199, plain, (r1_orders_2(X1,esk12_1(X2),esk12_1(X2))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X3))|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~m1_yellow_0(X3,X2)|~v4_yellow_0(X1,X3)|~v4_yellow_0(X3,X2)|~l1_orders_2(X2)|~v1_yellow_0(X2)|~v3_orders_2(X2)|~m1_subset_1(esk12_1(X2),u1_struct_0(X3)))).#trainneg
cnf(i_0_964, plain, (r1_lattice3(k1_yellow_2(X1,X2,X3),k2_relset_1(u1_struct_0(X2),X3),esk12_1(k1_yellow_2(X1,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~v1_yellow_0(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_866, plain, (r1_tarski(X1,u1_struct_0(k5_yellow_0(X2,X1)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(k5_yellow_0(X2,X1))|~l1_orders_2(X2))).#trainneg
cnf(i_0_708, plain, (r1_orders_2(k5_yellow_0(X1,X2),X3,X4)|X4!=k1_funct_1(X5,X6)|X3!=k1_funct_1(X5,X7)|~r1_orders_2(X8,X7,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X8),X2)))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X6,u1_struct_0(X8))|~m1_subset_1(X7,u1_struct_0(X8))|~v5_orders_3(X5,X8,k5_yellow_0(X1,X2))|~v1_funct_2(X5,u1_struct_0(X8),X2)|~v1_funct_1(X5)|~l1_orders_2(k5_yellow_0(X1,X2))|~l1_orders_2(X8)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1180, plain, (u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)))=k10_xtuple_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1874, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X3)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_371, plain, (v1_xboole_0(esk17_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2225, plain, (esk30_1(k1_xboole_0)=esk20_1(k1_xboole_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_365, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2226, plain, (esk20_1(k1_xboole_0)=esk22_1(k1_xboole_0)|~v1_card_1(k1_xboole_0)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_2125, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_367, plain, (v1_xboole_0(esk29_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_368, plain, (v1_xboole_0(esk28_1(X1))|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X1)))).#trainneg
cnf(i_0_2262, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,k10_xtuple_0(X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~l1_orders_2(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_2372, plain, (r1_tarski(k10_xtuple_0(X1),u1_struct_0(k1_yellow_2(X3,X2,X1)))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(k10_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(k1_yellow_2(X3,X2,X1))|~l1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_orders_2(X3))).#trainneg
cnf(i_0_547, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(u1_orders_2(X2)))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2))).#trainneg
cnf(i_0_1817, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|r1_lattice3(X3,X4,X5)|X2!=k1_funct_1(esk3_0,esk4_3(X3,X4,X5))|X1!=k1_funct_1(esk3_0,esk12_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk4_3(X3,X4,X5),u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(X3))|~l1_orders_2(X3))).#trainneg
cnf(i_0_810, plain, (k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))=k5_yellow_0(X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)))!=X4|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2513, negated_conjecture, (r1_orders_2(esk2_0,X1,k1_funct_1(esk3_0,esk4_3(X2,X3,X4)))|r1_lattice3(X2,X3,X4)|X1!=k1_funct_1(esk3_0,esk12_1(esk1_0))|~m1_subset_1(esk4_3(X2,X3,X4),u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk4_3(X2,X3,X4)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_432, plain, (v1_xboole_0(esk28_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2535, negated_conjecture, (k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))=k5_yellow_0(esk2_0,X1)|u1_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))!=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_374, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2173, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))),k5_numbers))).#trainneg
cnf(i_0_380, plain, (v3_ordinal1(esk12_1(X1))|~v3_ordinal1(u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_383, plain, (v1_relat_1(esk26_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2496, plain, (m1_subset_1(k10_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(k1_yellow_2(X2,X3,X1))))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_orders_2(k1_yellow_2(X2,X3,X1))|~l1_orders_2(X3)|~l1_orders_2(X2))).#trainneg
cnf(i_0_1505, plain, (r1_tarski(u1_struct_0(k1_yellow_2(X3,X1,X2)),k2_relset_1(u1_struct_0(X1),X2))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(k1_yellow_2(X3,X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_1799, negated_conjecture, (u1_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))=k10_xtuple_0(esk3_0)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1119, plain, (esk14_1(k1_zfmisc_1(k1_xboole_0))=X2|g1_orders_2(X1,esk14_1(k1_zfmisc_1(k1_xboole_0)))!=g1_orders_2(X3,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_384, plain, (v1_relat_1(esk19_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_386, plain, (v1_relat_1(esk16_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_388, plain, (v1_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1764, negated_conjecture, (u1_struct_0(X1)=k10_xtuple_0(esk3_0)|v2_struct_0(X2)|X1!=k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~v1_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1494, plain, (r1_tarski(k10_xtuple_0(u1_orders_2(X1)),u1_struct_0(X2))|v2_struct_0(X1)|~m1_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))),X2)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2554, negated_conjecture, (k1_yellow_2(X2,esk2_0,esk3_0)=k5_yellow_0(esk2_0,X1)|v2_struct_0(X2)|u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0))!=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_2683, negated_conjecture, (u1_struct_0(X1)=k10_xtuple_0(esk3_0)|X1!=k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~v1_orders_2(X1))).#trainneg
cnf(i_0_2620, negated_conjecture, (u1_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))=k10_xtuple_0(esk3_0))).#trainneg
cnf(i_0_2598, plain, (r1_tarski(u1_struct_0(k1_yellow_2(X1,X2,X3)),k10_xtuple_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v5_relat_1(X3,u1_struct_0(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_475, plain, (v1_funct_1(esk14_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2654, plain, (u1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))=esk14_1(k1_zfmisc_1(k1_xboole_0))|g1_orders_2(X1,esk14_1(k1_zfmisc_1(k1_xboole_0)))!=g1_orders_2(k5_numbers,esk28_1(k5_numbers))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_395, plain, (v5_relat_1(esk26_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2237, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))),k5_numbers))).#trainneg
cnf(i_0_396, plain, (v5_relat_1(esk19_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_397, plain, (v5_relat_1(esk17_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_886, plain, (k5_yellow_0(X1,u1_struct_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X3))))=k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X1)|~l1_orders_2(X2))).#trainneg
cnf(i_0_407, plain, (v4_relat_1(esk17_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_398, plain, (v5_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1076, plain, (X1=esk22_1(k1_xboole_0)|~v1_xboole_0(X1)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_400, plain, (v5_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_405, plain, (v4_relat_1(esk26_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2797, plain, (r1_tarski(k2_relset_1(u1_struct_0(X2),X3),k10_xtuple_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2711, plain, (r1_orders_2(X1,esk12_1(X2),X3)|v1_xboole_0(u1_struct_0(X2))|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(esk12_1(X2),u1_struct_0(X2))|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_2787, plain, (m1_subset_1(u1_struct_0(k1_yellow_2(X1,X2,X3)),k1_zfmisc_1(k10_xtuple_0(X3)))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2198, plain, (r1_orders_2(X1,esk12_1(X2),esk12_1(X2))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~r2_hidden(esk12_1(X2),k10_xtuple_0(X5))|~m1_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~m1_yellow_0(k1_yellow_2(X3,X4,X5),X2)|~v4_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~v4_yellow_0(k1_yellow_2(X3,X4,X5),X2)|~l1_orders_2(X2)|~v1_yellow_0(X2)|~v3_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_orders_2(X4)|~l1_orders_2(X3))).#trainneg
cnf(i_0_422, plain, (v1_orders_2(g1_orders_2(X1,esk29_2(X1,X1))))).#trainneg
cnf(i_0_2784, negated_conjecture, (u1_struct_0(k5_yellow_0(esk2_0,k1_xboole_0))=k1_xboole_0|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_2706, negated_conjecture, (u1_struct_0(X1)=k1_xboole_0|X1!=k5_yellow_0(esk2_0,k1_xboole_0)|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~v1_orders_2(X1)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_406, plain, (v4_relat_1(esk19_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1088, plain, (X1=esk30_1(k1_xboole_0)|~v1_xboole_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_2866, negated_conjecture, (l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_408, plain, (v4_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_410, plain, (v4_relat_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_2885, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|v1_xboole_0(u1_struct_0(X4))|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,esk12_1(X4))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(X4),u1_struct_0(esk1_0))|~r2_hidden(esk12_1(X4),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X4)|~v4_yellow_0(esk1_0,X4)|~m1_subset_1(X3,u1_struct_0(X4))|~l1_orders_2(X4)|~v1_yellow_0(X4))).#trainneg
cnf(i_0_1373, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),X2)|X2!=k1_funct_1(esk3_0,esk12_1(X1))|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(X1),u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_683, plain, (m1_yellow_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)),X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1190, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),X2),u1_struct_0(X3))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)),X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3)|~l1_orders_2(X1)|~l1_orders_2(X4))).#trainneg
cnf(i_0_2767, negated_conjecture, (k2_relset_1(u1_struct_0(X1),X2)=k10_xtuple_0(X2)|v2_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),u1_struct_0(X1))))|~v1_funct_2(X2,k10_xtuple_0(esk3_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_745, plain, (v1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_2655, plain, (esk14_1(k1_zfmisc_1(k1_xboole_0))=u1_orders_2(X1)|g1_orders_2(X2,esk14_1(k1_zfmisc_1(k1_xboole_0)))!=X1|~v1_xboole_0(X2)|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_415, plain, (l1_orders_2(g1_orders_2(X1,esk26_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_2250, negated_conjecture, (m1_yellow_0(k5_yellow_0(esk2_0,k1_xboole_0),esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_1864, plain, (v5_relat_1(k2_zfmisc_1(X3,X4),X2)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_416, plain, (l1_orders_2(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_417, plain, (l1_orders_2(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_794, plain, (v4_yellow_0(k5_yellow_0(k5_yellow_0(X1,X2),k2_relset_1(X2,X3)),k5_yellow_0(X1,X2))|v2_struct_0(k5_yellow_0(X1,X2))|v2_struct_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X2)))|~v1_funct_2(X3,u1_struct_0(X4),X2)|~v1_funct_1(X3)|~l1_orders_2(k5_yellow_0(X1,X2))|~l1_orders_2(X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3008, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk5_3(esk3_0,X2,X1),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_3010, negated_conjecture, (r1_lattice3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(X2)))|v1_xboole_0(u1_struct_0(X2))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X2)),u1_struct_0(esk2_0))|esk4_3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(X2)))!=k1_funct_1(esk3_0,esk12_1(X2))|~r2_hidden(esk12_1(X2),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X2)|~v4_yellow_0(esk1_0,X2)|~m1_subset_1(esk4_3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(X2))),u1_struct_0(esk2_0))|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_456, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2854, plain, (esk30_1(k1_xboole_0)=esk22_1(k1_xboole_0)|~v7_ordinal1(k1_xboole_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_418, plain, (l1_orders_2(g1_orders_2(X1,esk16_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_420, plain, (l1_orders_2(g1_orders_2(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_424, plain, (v1_orders_2(g1_orders_2(X1,esk26_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_819, plain, (r1_lattice3(k5_yellow_0(X1,X2),X3,esk12_1(k5_yellow_0(X1,X2)))|~r2_hidden(esk4_3(k5_yellow_0(X1,X2),X3,esk12_1(k5_yellow_0(X1,X2))),X2)|~l1_orders_2(k5_yellow_0(X1,X2))|~v1_yellow_0(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2887, plain, (r1_orders_2(X1,esk12_1(X2),X3)|v1_xboole_0(u1_struct_0(X2))|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(esk12_1(X2),u1_struct_0(X4))|~l1_orders_2(X4)|~r2_hidden(esk12_1(X2),u1_struct_0(X4))|~m1_yellow_0(X4,X2)|~v4_yellow_0(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_2196, negated_conjecture, (r1_orders_2(X1,esk12_1(X2),esk12_1(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~r2_hidden(esk12_1(X2),k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,k1_yellow_2(X3,esk2_0,esk3_0))|~m1_yellow_0(k1_yellow_2(X3,esk2_0,esk3_0),X2)|~v4_yellow_0(X1,k1_yellow_2(X3,esk2_0,esk3_0))|~v4_yellow_0(k1_yellow_2(X3,esk2_0,esk3_0),X2)|~l1_orders_2(X2)|~v1_yellow_0(X2)|~v3_orders_2(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(esk2_0))|~l1_orders_2(X3))).#trainneg
cnf(i_0_864, plain, (r1_orders_2(X1,esk12_1(X2),X3)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(esk12_1(X2),u1_struct_0(X4))|~l1_orders_2(X4)|~r2_hidden(esk12_1(X2),u1_struct_0(X4))|~r2_hidden(X3,u1_struct_0(X2))|~m1_yellow_0(X4,X2)|~v4_yellow_0(X4,X2)|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_1107, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1527, plain, (m1_subset_1(u1_struct_0(esk27_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2931, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))=k1_xboole_0|~v1_funct_2(u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0)),k1_xboole_0,k1_xboole_0)|~l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_425, plain, (v1_orders_2(g1_orders_2(X1,esk19_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_2424, plain, (v1_funct_2(k1_xboole_0,X1,X2)|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2221, negated_conjecture, (v3_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~m1_yellow_0(k5_yellow_0(esk2_0,k1_xboole_0),esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_2222, negated_conjecture, (v4_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~m1_yellow_0(k5_yellow_0(esk2_0,k1_xboole_0),esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_2241, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_426, plain, (v1_orders_2(g1_orders_2(X1,esk17_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_427, plain, (v1_orders_2(g1_orders_2(X1,esk16_1(k2_zfmisc_1(X1,X1))))|v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_2197, plain, (r1_orders_2(X1,esk12_1(X2),esk12_1(X2))|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X3)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~r2_hidden(esk12_1(X2),k2_relset_1(u1_struct_0(X4),X5))|~m1_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~m1_yellow_0(k1_yellow_2(X3,X4,X5),X2)|~v4_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~v4_yellow_0(k1_yellow_2(X3,X4,X5),X2)|~l1_orders_2(X2)|~v1_yellow_0(X2)|~v3_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_orders_2(X4)|~l1_orders_2(X3))).#trainneg
cnf(i_0_1363, plain, (v1_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk12_1(X2),u1_struct_0(X2))|~l1_orders_2(k5_yellow_0(X1,u1_struct_0(X2)))|~l1_orders_2(X1)|~r2_hidden(esk12_1(X2),u1_struct_0(k5_yellow_0(X1,u1_struct_0(X2))))|~m1_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)),X2)|~v4_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)),X2)|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_3216, plain, (r1_orders_2(X1,esk12_1(X2),X3)|v1_xboole_0(u1_struct_0(X4))|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~r2_hidden(X3,u1_struct_0(X2))|~m1_yellow_0(X1,X4)|~m1_yellow_0(X4,X2)|~v4_yellow_0(X1,X4)|~v4_yellow_0(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~l1_orders_2(X2)|~v1_yellow_0(X2)|~m1_subset_1(esk12_1(X2),u1_struct_0(X4)))).#trainneg
cnf(i_0_3182, plain, (r1_lattice3(k1_yellow_2(X4,X1,X2),X3,esk12_1(k1_yellow_2(X4,X1,X2)))|v2_struct_0(X4)|v2_struct_0(X1)|~r2_hidden(esk4_3(k1_yellow_2(X4,X1,X2),X3,esk12_1(k1_yellow_2(X4,X1,X2))),k2_relset_1(u1_struct_0(X1),X2))|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(k1_yellow_2(X4,X1,X2))|~l1_orders_2(X1)|~v1_yellow_0(k1_yellow_2(X4,X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_3197, plain, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X2,esk12_1(X1),X3)|v1_xboole_0(u1_struct_0(X4))|~r2_hidden(esk12_1(X1),u1_struct_0(X2))|~m1_yellow_0(X2,X4)|~m1_yellow_0(X4,X1)|~v4_yellow_0(X2,X4)|~v4_yellow_0(X4,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk12_1(X1),u1_struct_0(X4)))).#trainneg
cnf(i_0_2883, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),X2),k10_xtuple_0(X2))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(k5_yellow_0(X1,k10_xtuple_0(X2)))|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1761, negated_conjecture, (m1_subset_1(esk4_3(X1,k10_xtuple_0(esk3_0),X2),u1_struct_0(esk2_0))|r1_lattice3(X1,k10_xtuple_0(esk3_0),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2373, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),X2),u1_struct_0(k1_yellow_2(X3,X1,X2)))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(k1_yellow_2(X3,X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_1774, negated_conjecture, (v1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3225, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3256, plain, (v1_funct_2(k1_xboole_0,X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2736, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))=k1_xboole_0|k10_xtuple_0(esk3_0)!=k1_xboole_0|~v1_funct_2(u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))),k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))|~l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_429, plain, (v1_orders_2(g1_orders_2(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_431, plain, (v1_xboole_0(esk29_2(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3378, negated_conjecture, (r1_orders_2(X1,esk4_3(X2,k10_xtuple_0(esk3_0),X3),esk4_3(X2,k10_xtuple_0(esk3_0),X3))|r1_lattice3(X2,k10_xtuple_0(esk3_0),X3)|~r2_hidden(esk4_3(X2,k10_xtuple_0(esk3_0),X3),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_3381, negated_conjecture, (r1_lattice3(X1,k10_xtuple_0(esk3_0),esk12_1(esk2_0))|~r2_hidden(esk12_1(esk2_0),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~v1_yellow_0(esk2_0)|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3379, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(esk2_0,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(X1)))|esk4_3(esk2_0,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(X1)))!=k1_funct_1(esk3_0,esk12_1(X1))|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_3402, negated_conjecture, (r1_tarski(k10_xtuple_0(esk3_0),u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3340, negated_conjecture, (v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_xboole_0(k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_905, plain, (r1_tarski(u1_struct_0(X1),k2_relset_1(u1_struct_0(X2),X3))|v2_struct_0(X4)|v2_struct_0(X2)|~m1_yellow_0(X1,k1_yellow_2(X4,X2,X3))|~m1_subset_1(k2_relset_1(u1_struct_0(X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(k1_yellow_2(X4,X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X4))).#trainneg
cnf(i_0_3411, plain, (v1_funct_1(k5_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_614, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),esk25_2(X1,X1)))|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2223, negated_conjecture, (v5_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~m1_yellow_0(k5_yellow_0(esk2_0,k1_xboole_0),esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_1744, negated_conjecture, (v3_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0))).#trainneg
cnf(i_0_435, plain, (v1_xboole_0(esk17_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3155, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(esk2_0,X2,k1_funct_1(esk3_0,esk12_1(X1)))|esk4_3(esk2_0,X2,k1_funct_1(esk3_0,esk12_1(X1)))!=k1_funct_1(esk3_0,esk12_1(X1))|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_3215, plain, (r1_orders_2(X1,esk12_1(X2),X3)|v2_struct_0(X4)|v2_struct_0(X5)|~r2_hidden(esk12_1(X2),u1_struct_0(X1))|~r2_hidden(esk12_1(X2),k10_xtuple_0(X6))|~r2_hidden(X3,u1_struct_0(X2))|~m1_yellow_0(X1,k1_yellow_2(X4,X5,X6))|~m1_yellow_0(k1_yellow_2(X4,X5,X6),X2)|~v4_yellow_0(X1,k1_yellow_2(X4,X5,X6))|~v4_yellow_0(k1_yellow_2(X4,X5,X6),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,k10_xtuple_0(X6))|~l1_orders_2(X2)|~v1_yellow_0(X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X5))|~v1_funct_1(X6)|~l1_orders_2(X5)|~l1_orders_2(X4))).#trainneg
cnf(i_0_740, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(k1_yellow_2(X1,X2,X3)))|~l1_orders_2(k1_yellow_2(X1,X2,X3)))).#trainneg
cnf(i_0_1292, plain, (v1_yellow_0(k5_yellow_0(X2,k10_xtuple_0(X3)))|v2_struct_0(X1)|v2_struct_0(X2)|~r1_lattice3(k5_yellow_0(X2,k10_xtuple_0(X3)),k2_relset_1(u1_struct_0(X2),X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k2_relset_1(u1_struct_0(X2),X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k5_yellow_0(X2,k10_xtuple_0(X3)))|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_1036, negated_conjecture, (r1_orders_2(esk2_0,X1,k1_funct_1(esk3_0,X2))|X1!=k1_funct_1(esk3_0,X2)|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2370, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k5_yellow_0(X2,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(k5_yellow_0(X2,X1))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1112, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_441, plain, (v1_xboole_0(esk29_2(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1708, plain, (u1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))=esk28_1(k5_numbers))).#trainneg
cnf(i_0_438, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3502, plain, (v1_xboole_0(esk17_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1083, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3442, negated_conjecture, (r1_orders_2(X1,esk12_1(esk2_0),X2)|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_3441, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,X2,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X1,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X1,X2,esk3_0)))|~m1_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_3613, plain, (v1_xboole_0(esk17_1(k1_xboole_0)))).#trainneg
cnf(i_0_445, plain, (v1_xboole_0(esk17_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1745, negated_conjecture, (v4_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0))).#trainneg
cnf(i_0_1746, negated_conjecture, (v5_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~m1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),esk2_0))).#trainneg
cnf(i_0_448, plain, (v1_xboole_0(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2292, plain, (m1_subset_1(k2_relset_1(u1_struct_0(X1),u1_orders_2(X1)),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1282, plain, (v1_yellow_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_lattice3(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)),k10_xtuple_0(X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)))|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_3586, plain, (k1_relset_1(k5_numbers,esk28_1(k5_numbers))=k5_numbers|k5_numbers=k1_xboole_0|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_451, plain, (v3_ordinal1(u1_orders_2(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3878, plain, (esk17_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_3656, plain, (v4_relat_1(k2_zfmisc_1(X4,X5),X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_452, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_453, plain, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~r2_hidden(X1,u1_orders_2(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_3847, negated_conjecture, (r1_lattice3(X1,X2,esk12_1(esk2_0))|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(X1))|~r2_hidden(esk4_3(X1,X2,esk12_1(esk2_0)),k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk4_3(X1,X2,esk12_1(esk2_0)),u1_struct_0(X1))|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_3870, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,X2,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~m1_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X1,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_3697, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_457, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_3801, plain, (v5_relat_1(k2_zfmisc_1(X3,X4),X2)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_458, plain, (u1_struct_0(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_459, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_814, plain, (v1_yellow_0(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,X2)|~l1_orders_2(k5_yellow_0(X1,X2))|~l1_orders_2(X1)|~v1_xboole_0(X2)|~m1_subset_1(X3,u1_struct_0(k5_yellow_0(X1,X2))))).#trainneg
cnf(i_0_4112, plain, (v1_yellow_0(k1_yellow_2(X4,X1,X2))|v2_struct_0(X1)|v2_struct_0(X4)|~v1_xboole_0(k10_xtuple_0(X2))|~m1_subset_1(X3,u1_struct_0(k1_yellow_2(X4,X1,X2)))|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k10_xtuple_0(X2))|~l1_orders_2(k1_yellow_2(X4,X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_3937, plain, (m1_subset_1(k10_xtuple_0(u1_orders_2(X1)),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~v1_relat_1(u1_orders_2(X1)))).#trainneg
cnf(i_0_1409, plain, (r1_lattice3(k1_yellow_2(X1,X2,X3),X4,esk12_1(k1_yellow_2(X1,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(k1_yellow_2(X1,X2,X3))|~m1_subset_1(esk4_3(k1_yellow_2(X1,X2,X3),X4,esk12_1(k1_yellow_2(X1,X2,X3))),k10_xtuple_0(X3)))).#trainneg
cnf(i_0_1673, negated_conjecture, (k2_relset_1(u1_struct_0(X1),X2)=k10_xtuple_0(X2)|v2_struct_0(k1_yellow_2(X3,esk2_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),u1_struct_0(X1))))|~v1_funct_2(X2,k10_xtuple_0(esk3_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(k1_yellow_2(X3,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(esk2_0))|~l1_orders_2(X3))).#trainneg
cnf(i_0_2505, plain, (r1_tarski(k10_xtuple_0(X1),k2_relset_1(u1_struct_0(X3),X1))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_orders_2(k1_yellow_2(X2,X3,X1))|~l1_orders_2(X3)|~l1_orders_2(X2))).#trainneg
cnf(i_0_1611, plain, (r1_orders_2(X1,X2,X3)|X3!=k1_funct_1(X4,X5)|X2!=k1_funct_1(X4,X6)|~r1_orders_2(g1_orders_2(k5_numbers,esk28_1(k5_numbers)),X6,X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X6,k5_numbers)|~v5_orders_3(X4,g1_orders_2(k5_numbers,esk28_1(k5_numbers)),X1)|~v1_funct_2(X4,k5_numbers,u1_struct_0(X1))|~v1_funct_1(X4)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1612, plain, (r1_orders_2(X1,esk4_3(X2,X3,X4),esk4_3(X2,X3,X4))|r1_lattice3(X2,X3,X4)|v2_struct_0(X2)|~r2_hidden(esk4_3(X2,X3,X4),u1_struct_0(X1))|~r2_hidden(esk4_3(X2,X3,X4),k5_numbers)|~m1_yellow_0(X1,g1_orders_2(k5_numbers,esk28_1(k5_numbers)))|~m1_yellow_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers)),X2)|~v4_yellow_0(X1,g1_orders_2(k5_numbers,esk28_1(k5_numbers)))|~v4_yellow_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers)),X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_3982, plain, (~v1_subset_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_465, plain, (v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|u1_struct_0(X1)!=k1_xboole_0|u1_orders_2(X1)!=k1_xboole_0|~l1_orders_2(X1))).#trainneg
cnf(i_0_466, plain, (v3_orders_2(esk27_1(X1))|~m1_yellow_0(esk27_1(X1),X1)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2607, plain, (r1_tarski(u1_struct_0(k5_yellow_0(X2,k10_xtuple_0(X3))),k2_relset_1(u1_struct_0(X2),X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k5_yellow_0(X2,k10_xtuple_0(X3)))|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_739, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_xboole_0(u1_struct_0(k5_yellow_0(X2,X1)))|~l1_orders_2(k5_yellow_0(X2,X1)))).#trainneg
cnf(i_0_1301, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),X2),u1_struct_0(X3))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_yellow_0(k5_yellow_0(X1,k10_xtuple_0(X2)),X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3846, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,esk12_1(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(esk1_0))|~r2_hidden(esk12_1(esk2_0),u1_struct_0(esk1_0))|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~m1_yellow_0(esk1_0,esk2_0)|~v4_yellow_0(esk1_0,esk2_0)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_546, plain, (l1_orders_2(k5_yellow_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3954, plain, (k9_xtuple_0(esk28_1(k5_numbers))=k5_numbers|k5_numbers=k1_xboole_0|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_467, plain, (v4_orders_2(esk27_1(X1))|~m1_yellow_0(esk27_1(X1),X1)|~l1_orders_2(X1)|~v4_orders_2(X1))).#trainneg
cnf(i_0_470, plain, (v1_funct_1(esk26_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_471, plain, (v1_funct_1(esk19_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1683, negated_conjecture, (r1_orders_2(X1,esk4_3(X2,X3,X4),esk4_3(X2,X3,X4))|r1_lattice3(X2,X3,X4)|v2_struct_0(X2)|v2_struct_0(X5)|~r2_hidden(esk4_3(X2,X3,X4),u1_struct_0(X1))|~r2_hidden(esk4_3(X2,X3,X4),k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,k1_yellow_2(X5,esk2_0,esk3_0))|~m1_yellow_0(k1_yellow_2(X5,esk2_0,esk3_0),X2)|~v4_yellow_0(X1,k1_yellow_2(X5,esk2_0,esk3_0))|~v4_yellow_0(k1_yellow_2(X5,esk2_0,esk3_0),X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X5),u1_struct_0(esk2_0))|~l1_orders_2(X5))).#trainneg
cnf(i_0_4208, plain, (r1_tarski(k10_xtuple_0(X1),k2_relset_1(u1_struct_0(X2),X1))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_orders_2(k5_yellow_0(X2,k10_xtuple_0(X1)))|~l1_orders_2(X2)|~l1_orders_2(X3)|~m1_subset_1(k10_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_2889, plain, (v1_xboole_0(u1_struct_0(k1_yellow_2(X1,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_xboole_0(k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_3011, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(X2)),X3)|v1_xboole_0(u1_struct_0(X2))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(X2)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X2)),u1_struct_0(esk2_0))|X3!=k1_funct_1(esk3_0,esk12_1(X2))|~r2_hidden(esk12_1(X2),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X2)|~v4_yellow_0(esk1_0,X2)|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_1438, plain, (r1_tarski(k10_xtuple_0(u1_orders_2(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1129, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),X2)|r1_lattice3(X3,u1_struct_0(X1),X4)|X2!=k1_funct_1(esk3_0,esk4_3(X3,u1_struct_0(X1),X4))|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk4_3(X3,u1_struct_0(X1),X4),u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3))).#trainneg
cnf(i_0_4310, plain, (l1_orders_2(k1_yellow_2(X3,X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_4098, plain, (u1_struct_0(X1)=X2|X1!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_1089, plain, (m1_subset_1(esk28_1(X1),k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_473, plain, (v1_funct_1(esk16_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4317, plain, (k5_numbers=k1_xboole_0|~v1_xboole_0(esk28_1(k5_numbers))|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_2930, negated_conjecture, (v5_relat_1(u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0)),k1_xboole_0)|~l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_479, plain, (v1_funct_1(u1_orders_2(X1))|~v1_relat_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~v1_funct_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_480, plain, (v5_orders_2(esk27_1(X1))|~m1_yellow_0(esk27_1(X1),X1)|~l1_orders_2(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_4383, plain, (r1_tarski(k10_xtuple_0(u1_orders_2(k1_yellow_2(X1,X2,X3))),k10_xtuple_0(X3))|v2_struct_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_funct_2(u1_orders_2(k1_yellow_2(X1,X2,X3)),k10_xtuple_0(X3),k10_xtuple_0(X3))|~v1_funct_1(u1_orders_2(k1_yellow_2(X1,X2,X3)))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_4364, plain, (r1_tarski(k10_xtuple_0(u1_orders_2(X1)),k2_relset_1(u1_struct_0(X1),u1_orders_2(X1)))|v2_struct_0(X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3869, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,X2,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(esk2_0),k2_relset_1(u1_struct_0(X2),esk3_0))|~m1_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X1,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_3439, negated_conjecture, (v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~r2_hidden(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_2261, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,k2_relset_1(u1_struct_0(X2),X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~l1_orders_2(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_4377, plain, (r1_tarski(k10_xtuple_0(esk28_1(k5_numbers)),k5_numbers)|v2_struct_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_482, plain, (v1_xboole_0(X1)|~v1_partfun1(esk29_2(X1,X2),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2932, negated_conjecture, (v4_relat_1(u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0)),k1_xboole_0)|~l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_4007, negated_conjecture, (v1_xboole_0(u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0)))|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_1528, plain, (m1_subset_1(u1_struct_0(k5_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4404, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),X2),k10_xtuple_0(X2))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(X3))).#trainneg
cnf(i_0_4477, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,X2,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X3,X2,esk3_0)))|~m1_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X1,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(X2))|~l1_orders_2(X3))).#trainneg
cnf(i_0_4533, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(k5_yellow_0(X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2))).#trainneg
cnf(i_0_4448, negated_conjecture, (r1_tarski(k10_xtuple_0(u1_orders_2(k1_yellow_2(X2,esk2_0,esk3_0))),k10_xtuple_0(esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(u1_orders_2(k1_yellow_2(X2,esk2_0,esk3_0)),k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(u1_orders_2(k1_yellow_2(X2,esk2_0,esk3_0)))|~l1_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_4418, plain, (m1_subset_1(X1,k1_xboole_0)|~r2_hidden(X1,esk28_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4525, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_486, plain, (v1_xboole_0(X1)|~v1_partfun1(esk17_1(k2_zfmisc_1(X1,X2)),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4008, negated_conjecture, (v1_xboole_0(u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))|~v1_xboole_0(k2_zfmisc_1(k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_489, plain, (v1_xboole_0(X1)|~v1_partfun1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3080, negated_conjecture, (k2_relset_1(k10_xtuple_0(esk3_0),X1)=k10_xtuple_0(X1)|v2_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))))|~v1_funct_2(X1,k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4617, negated_conjecture, (r1_tarski(k10_xtuple_0(u1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))),k10_xtuple_0(esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(u1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0)),k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(u1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4582, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),u1_orders_2(X1)),k10_xtuple_0(u1_orders_2(X1)))|v2_struct_0(X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4590, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,X2,esk3_0))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(k5_yellow_0(X2,k10_xtuple_0(esk3_0))))|~m1_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X1,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~v1_yellow_0(esk2_0)|~m1_subset_1(k10_xtuple_0(esk3_0),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_3811, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X5)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4440, plain, (v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_494, plain, (v3_ordinal1(k2_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_495, plain, (v1_xboole_0(X1)|~v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v1_xboole_0(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_496, plain, (l1_orders_2(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_4603, plain, (m1_subset_1(X1,u1_struct_0(X2))|v1_xboole_0(u1_struct_0(k5_yellow_0(X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(k5_yellow_0(X2,X3))))).#trainneg
cnf(i_0_4594, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,X2,esk3_0))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(k1_yellow_2(X3,X2,esk3_0)))|~m1_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X1,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~v1_yellow_0(esk2_0)|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X3,X2,esk3_0))))).#trainneg
cnf(i_0_764, negated_conjecture, (X1=k10_xtuple_0(esk3_0)|r2_hidden(esk7_2(esk3_0,X1),k9_xtuple_0(esk3_0))|r2_hidden(esk6_2(esk3_0,X1),X1))).#trainneg
cnf(i_0_1729, plain, (r1_lattice3(X1,k10_xtuple_0(X2),esk12_1(k1_yellow_2(X3,X4,X2)))|v2_struct_0(X3)|v2_struct_0(X4)|~r2_hidden(esk12_1(k1_yellow_2(X3,X4,X2)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X3,X4,X2))|~v4_yellow_0(X1,k1_yellow_2(X3,X4,X2))|~m1_subset_1(esk4_3(X1,k10_xtuple_0(X2),esk12_1(k1_yellow_2(X3,X4,X2))),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X4)))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X2)|~l1_orders_2(k1_yellow_2(X3,X4,X2))|~l1_orders_2(X4)|~l1_orders_2(X3)|~v1_yellow_0(k1_yellow_2(X3,X4,X2))|~m1_subset_1(esk12_1(k1_yellow_2(X3,X4,X2)),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4772, plain, (v1_xboole_0(u1_struct_0(k1_yellow_2(X2,X1,X3)))|m1_subset_1(X4,u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(k1_yellow_2(X2,X1,X3)))|~m1_subset_1(u1_struct_0(k1_yellow_2(X2,X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X2))).#trainneg
cnf(i_0_4777, plain, (v1_xboole_0(u1_struct_0(k5_yellow_0(X1,X2)))|m1_subset_1(esk12_1(k5_yellow_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~l1_orders_2(k5_yellow_0(X1,X2))|~v1_yellow_0(k5_yellow_0(X1,X2)))).#trainneg
cnf(i_0_4782, plain, (v1_xboole_0(u1_struct_0(k5_yellow_0(X1,X2)))|m1_subset_1(esk4_3(k5_yellow_0(X1,X2),X3,X4),u1_struct_0(X1))|r1_lattice3(k5_yellow_0(X1,X2),X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(k5_yellow_0(X1,X2)))|~l1_orders_2(k5_yellow_0(X1,X2)))).#trainneg
cnf(i_0_4774, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_497, plain, (v1_orders_2(g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_4425, negated_conjecture, (v1_funct_1(u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0)))|~v1_funct_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_4426, negated_conjecture, (v1_funct_1(u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))|~v1_funct_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_498, plain, (k2_relset_1(k2_zfmisc_1(X1,X1),X2)=X3|g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X2))!=g1_orders_2(X4,X3)|~v5_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_499, plain, (X1=X2|g1_orders_2(X1,k2_relset_1(k2_zfmisc_1(X1,X1),X3))!=g1_orders_2(X2,X4)|~v5_relat_1(X3,k2_zfmisc_1(X1,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_4308, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk2_0)),X1)|X1!=k1_funct_1(esk3_0,X2)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(esk1_0))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(esk1_0,esk2_0)|~v4_yellow_0(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk2_0)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_4113, plain, (v1_yellow_0(k1_yellow_2(X4,X1,X2))|v2_struct_0(X4)|v2_struct_0(X1)|~v1_xboole_0(k2_relset_1(u1_struct_0(X1),X2))|~m1_subset_1(X3,u1_struct_0(k1_yellow_2(X4,X1,X2)))|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k2_relset_1(u1_struct_0(X1),X2))|~l1_orders_2(k1_yellow_2(X4,X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_4595, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1324, plain, (r1_orders_2(X1,esk4_3(X2,X3,X4),esk4_3(X2,X3,X4))|r1_lattice3(X2,X3,X4)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X6)|~r2_hidden(esk4_3(X2,X3,X4),u1_struct_0(X1))|~r2_hidden(esk4_3(X2,X3,X4),k10_xtuple_0(X7))|~m1_yellow_0(X1,k1_yellow_2(X5,X6,X7))|~m1_yellow_0(k1_yellow_2(X5,X6,X7),X2)|~v4_yellow_0(X1,k1_yellow_2(X5,X6,X7))|~v4_yellow_0(k1_yellow_2(X5,X6,X7),X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X6))))|~v1_funct_2(X7,u1_struct_0(X5),u1_struct_0(X6))|~v1_funct_1(X7)|~l1_orders_2(X6)|~l1_orders_2(X5))).#trainneg
cnf(i_0_1588, plain, (r1_tarski(u1_struct_0(X1),k5_numbers)|~m1_yellow_0(X1,g1_orders_2(k5_numbers,esk28_1(k5_numbers))))).#trainneg
cnf(i_0_4793, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|v1_yellow_0(k1_yellow_2(X3,X2,esk3_0))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_yellow_0(k1_yellow_2(X3,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X3,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X3,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X3)|~l1_orders_2(X1)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_2176, plain, (v1_funct_1(u1_struct_0(k1_yellow_2(X1,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_relat_1(u1_struct_0(X2))|~v1_funct_1(u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_4894, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|m1_subset_1(esk14_1(k10_xtuple_0(esk3_0)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_4092, negated_conjecture, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0)))!=g1_orders_2(X1,X2)|~l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_500, plain, (v1_xboole_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_501, plain, (v5_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_502, plain, (k2_relset_1(k2_zfmisc_1(X1,X2),X3)=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_4961, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0))))).#trainneg
cnf(i_0_4890, plain, (v1_xboole_0(u1_struct_0(k1_yellow_2(X5,X1,X2)))|r1_lattice3(k1_yellow_2(X5,X1,X2),X3,X4)|m1_subset_1(esk4_3(k1_yellow_2(X5,X1,X2),X3,X4),u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X5)|~m1_subset_1(X4,u1_struct_0(k1_yellow_2(X5,X1,X2)))|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X5),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X5))).#trainneg
cnf(i_0_4881, plain, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(k5_yellow_0(X1,X2),X3,esk12_1(X1))|v1_xboole_0(u1_struct_0(k5_yellow_0(X1,X2)))|~r2_hidden(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,X2)))|~m1_yellow_0(k5_yellow_0(X1,X2),X1)|~v4_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4519, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))=esk22_1(k1_xboole_0)|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_503, plain, (v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_504, plain, (v1_xboole_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_505, plain, (v4_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5072, plain, (v1_xboole_0(u1_struct_0(X1))|r1_lattice3(k1_yellow_2(X4,X1,X2),X3,esk12_1(X1))|v2_struct_0(X4)|v2_struct_0(X1)|~r2_hidden(esk12_1(X1),u1_struct_0(k1_yellow_2(X4,X1,X2)))|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_5064, plain, (v1_yellow_0(k5_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,X2)))|~l1_orders_2(k5_yellow_0(X1,X2))|~r2_hidden(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_5011, plain, (v5_relat_1(k10_xtuple_0(X3),X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5070, negated_conjecture, (r1_lattice3(k1_yellow_2(X2,esk2_0,esk3_0),X1,esk12_1(esk2_0))|v2_struct_0(X2)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0)))|~v1_yellow_0(esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_5048, plain, (r1_orders_2(X1,esk4_3(k1_yellow_2(X2,X3,X4),X5,X6),esk4_3(k1_yellow_2(X2,X3,X4),X5,X6))|v2_struct_0(X3)|v1_xboole_0(u1_struct_0(k1_yellow_2(X2,X3,X4)))|r1_lattice3(k1_yellow_2(X2,X3,X4),X5,X6)|v2_struct_0(X2)|~r2_hidden(esk4_3(k1_yellow_2(X2,X3,X4),X5,X6),u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~l1_orders_2(X3)|~v3_orders_2(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X6,u1_struct_0(k1_yellow_2(X2,X3,X4)))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_orders_2(X2))).#trainneg
cnf(i_0_5065, plain, (r1_orders_2(k5_yellow_0(X1,X2),esk12_1(X1),X3)|v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X3,X4)|~m1_subset_1(X3,u1_struct_0(k5_yellow_0(X1,X2)))|~m1_subset_1(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,X2)))|~l1_orders_2(k5_yellow_0(X1,X2))|~r2_hidden(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_1113, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4986, negated_conjecture, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,k1_xboole_0)!=g1_orders_2(X1,X2)|~v1_xboole_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_506, plain, (v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|X2!=k1_xboole_0|k2_relset_1(k2_zfmisc_1(X1,X2),X3)!=k1_xboole_0|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_824, plain, (k1_relset_1(X1,k1_xboole_0)=X1|X2=k1_xboole_0|~v1_funct_2(k1_xboole_0,X1,X2))).#trainneg
cnf(i_0_5105, negated_conjecture, (v5_relat_1(esk22_1(k1_xboole_0),k1_xboole_0)|~v1_xboole_0(esk3_0)|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_5107, negated_conjecture, (v4_relat_1(esk22_1(k1_xboole_0),k1_xboole_0)|~v1_xboole_0(esk3_0)|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_507, plain, (v1_xboole_0(k2_relset_1(X1,X2))|~v1_xboole_0(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_508, plain, (v1_funct_1(k2_relset_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5149, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(X2),k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~r2_hidden(esk12_1(X2),u1_struct_0(k1_yellow_2(X1,X2,X3)))|~v1_yellow_0(X2))).#trainneg
cnf(i_0_5146, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X1)|~m1_subset_1(esk12_1(X2),u1_struct_0(k1_yellow_2(X1,X2,X3)))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~r2_hidden(esk12_1(X2),u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_5160, plain, (v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|~r2_hidden(esk12_1(X1),u1_struct_0(k1_yellow_2(X3,X1,X2)))|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_5164, plain, (v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k5_yellow_0(X1,X2))|~r2_hidden(esk12_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_5165, plain, (v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k5_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(k5_yellow_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,X2))))).#trainneg
cnf(i_0_5182, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|~r2_hidden(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_849, plain, (k2_zfmisc_1(X1,X2)=k10_xtuple_0(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5179, plain, (v5_relat_1(k10_xtuple_0(k1_xboole_0),X1))).#trainneg
cnf(i_0_511, plain, (m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_512, plain, (v3_ordinal1(k1_relset_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_513, plain, (v1_xboole_0(X1)|~v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v1_xboole_0(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5280, plain, (v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k1_yellow_2(X2,X1,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_1(X1),k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_5285, plain, (v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|~r2_hidden(esk12_1(X1),k2_relset_1(u1_struct_0(X1),X2))|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_5284, plain, (v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|~r2_hidden(esk12_1(X1),k10_xtuple_0(X2))|~m1_subset_1(k10_xtuple_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_5289, plain, (v1_xboole_0(u1_struct_0(k1_yellow_2(X3,X1,X2)))|v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(esk12_1(X1),u1_struct_0(k1_yellow_2(X3,X1,X2)))|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_5286, plain, (v1_yellow_0(k1_yellow_2(X2,X1,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_1(X1),u1_struct_0(k1_yellow_2(X2,X1,X3)))|~m1_subset_1(u1_struct_0(k1_yellow_2(X2,X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X2))).#trainneg
cnf(i_0_5293, plain, (v1_xboole_0(X2)|v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k5_yellow_0(X1,X2))|~m1_subset_1(esk12_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_5463, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k10_xtuple_0(X4)))|~v1_xboole_0(X4)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1251, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4980, negated_conjecture, (v3_ordinal1(esk14_1(k10_xtuple_0(esk3_0)))|v1_xboole_0(k10_xtuple_0(esk3_0))|~v3_ordinal1(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_5110, negated_conjecture, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,esk22_1(k1_xboole_0))!=g1_orders_2(X1,X2)|~v1_xboole_0(esk3_0)|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_514, plain, (l1_orders_2(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5162, negated_conjecture, (v1_yellow_0(k5_yellow_0(esk2_0,k1_xboole_0))|~r2_hidden(esk12_1(esk2_0),k1_xboole_0)|~v1_yellow_0(esk2_0)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_5163, negated_conjecture, (v1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~r2_hidden(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_515, plain, (v1_orders_2(g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2)))|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_516, plain, (k1_relset_1(k2_zfmisc_1(X1,X1),X2)=X3|g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X2))!=g1_orders_2(X4,X3)|~v4_relat_1(X2,k2_zfmisc_1(X1,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5556, plain, (u1_struct_0(X1)=k10_xtuple_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|X1!=k1_yellow_2(X3,X4,X2)|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~v1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X2)|~l1_orders_2(X4)|~l1_orders_2(X3)|~v5_relat_1(X2,u1_struct_0(X4))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5607, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2)|~m1_subset_1(esk12_1(X2),u1_struct_0(k1_yellow_2(X1,X2,X3))))).#trainneg
cnf(i_0_5582, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|~r2_hidden(esk12_1(X2),u1_struct_0(k1_yellow_2(X4,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~l1_orders_2(X4))).#trainneg
cnf(i_0_1252, plain, (v1_partfun1(X1,X2)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_517, plain, (X1=X2|g1_orders_2(X1,k1_relset_1(k2_zfmisc_1(X1,X1),X3))!=g1_orders_2(X2,X4)|~v4_relat_1(X3,k2_zfmisc_1(X1,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_518, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_519, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5610, plain, (v1_xboole_0(k10_xtuple_0(X1))|v1_yellow_0(k1_yellow_2(X3,X2,X1))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(k10_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk12_1(X2),k10_xtuple_0(X1))|~l1_orders_2(X2)|~v1_yellow_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_orders_2(X3))).#trainneg
cnf(i_0_5611, plain, (v1_xboole_0(k2_relset_1(u1_struct_0(X1),X2))|v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk12_1(X1),k2_relset_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~v1_yellow_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_5572, plain, (v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k1_yellow_2(X2,X1,X3))|v2_struct_0(X2)|v1_xboole_0(k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_yellow_0(X1)|~m1_subset_1(esk12_1(X1),k10_xtuple_0(X3)))).#trainneg
cnf(i_0_5173, plain, (v5_relat_1(k10_xtuple_0(esk28_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_4520, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))=esk30_1(k1_xboole_0)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_520, plain, (k1_relset_1(k2_zfmisc_1(X1,X2),X3)=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_829, plain, (X1=k1_xboole_0|v1_funct_2(k1_xboole_0,X2,X1)|k1_relset_1(X2,k1_xboole_0)!=X2)).#trainneg
cnf(i_0_4001, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_521, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_522, plain, (v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_xboole_0(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5684, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(X2),u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_5747, plain, (v1_xboole_0(u1_struct_0(k1_yellow_2(X4,X1,X2)))|v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~m1_subset_1(esk12_1(X1),u1_struct_0(k1_yellow_2(X4,X1,X2)))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~v1_yellow_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~l1_orders_2(X4))).#trainneg
cnf(i_0_1753, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_5797, negated_conjecture, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,esk30_1(k1_xboole_0))!=g1_orders_2(X1,X2)|~v1_xboole_0(esk3_0)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_523, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_1109, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5782, negated_conjecture, (v1_relat_1(esk30_1(k1_xboole_0))|~l1_orders_2(k5_yellow_0(esk2_0,k1_xboole_0))|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_5792, negated_conjecture, (v5_relat_1(esk30_1(k1_xboole_0),k1_xboole_0)|~v1_xboole_0(esk3_0)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_4083, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_524, plain, (v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|X2!=k1_xboole_0|k1_relset_1(k2_zfmisc_1(X1,X2),X3)!=k1_xboole_0|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_525, plain, (v1_xboole_0(k1_relset_1(X1,X2))|~v1_xboole_0(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_850, plain, (k2_zfmisc_1(X1,X2)=k10_xtuple_0(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5899, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_526, plain, (v1_funct_1(k1_relset_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5809, plain, (v1_funct_2(k1_xboole_0,X1,X2)|k9_xtuple_0(k1_xboole_0)!=X1)).#trainneg
cnf(i_0_5869, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))),k5_numbers))).#trainneg
cnf(i_0_529, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_530, plain, (v1_partfun1(esk29_2(X1,X2),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5864, plain, (v1_xboole_0(u1_struct_0(k5_yellow_0(X2,k10_xtuple_0(X3))))|v1_yellow_0(k1_yellow_2(X4,X2,X3))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(X2),u1_struct_0(k5_yellow_0(X2,k10_xtuple_0(X3))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~l1_orders_2(X1)|~v1_yellow_0(X2)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_5861, plain, (v1_xboole_0(u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))))|v1_yellow_0(k1_yellow_2(X4,X2,X3))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(X2),u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~l1_orders_2(X1)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_5687, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(X2),u1_struct_0(k5_yellow_0(X2,k10_xtuple_0(X3))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(X2)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_5667, plain, (u1_struct_0(X1)=k10_xtuple_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|X1!=k5_yellow_0(X3,k2_relset_1(u1_struct_0(X3),X2))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~v1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X2)|~l1_orders_2(X3)|~l1_orders_2(X4))).#trainneg
cnf(i_0_6179, plain, (v1_funct_2(k1_xboole_0,k9_xtuple_0(k1_xboole_0),X1))).#trainneg
cnf(i_0_6180, plain, (v1_funct_2(k1_xboole_0,X1,X2)|k1_xboole_0!=X1)).#trainneg
cnf(i_0_534, plain, (v1_partfun1(esk17_1(k2_zfmisc_1(X1,X2)),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5859, negated_conjecture, (v1_xboole_0(u1_struct_0(k1_yellow_2(X3,esk2_0,esk3_0)))|v1_yellow_0(k1_yellow_2(X2,esk2_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X3,esk2_0,esk3_0)))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(esk2_0))|~l1_orders_2(X3))).#trainneg
cnf(i_0_5678, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|v1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_5672, negated_conjecture, (v1_xboole_0(u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0)))|v1_yellow_0(k1_yellow_2(X2,esk2_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0)))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~v1_yellow_0(esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_6130, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k10_xtuple_0(X4)))|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_900, plain, (k2_zfmisc_1(X1,X2)=k9_xtuple_0(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5794, negated_conjecture, (v4_relat_1(esk30_1(k1_xboole_0),k1_xboole_0)|~v1_xboole_0(esk3_0)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_6339, plain, (v1_funct_2(k1_xboole_0,k1_xboole_0,X1))).#trainneg
cnf(i_0_5795, negated_conjecture, (v1_funct_1(esk30_1(k1_xboole_0))|~v1_xboole_0(esk3_0)|~v1_funct_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_5919, negated_conjecture, (v5_relat_1(esk20_1(k1_xboole_0),k1_xboole_0)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_5927, plain, (v5_relat_1(k2_zfmisc_1(X4,X5),X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_537, plain, (v1_partfun1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_540, plain, (v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v1_xboole_0(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5876, negated_conjecture, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,esk20_1(k1_xboole_0))!=g1_orders_2(X1,X2)|~v1_card_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_541, plain, (v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v1_xboole_0(X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_5994, plain, (v5_relat_1(k2_zfmisc_1(X2,X3),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_542, plain, (v1_partfun1(u1_orders_2(X1),u1_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_543, plain, (v3_orders_2(k5_yellow_0(X1,X2))|~m1_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6020, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(X2),k2_zfmisc_1(X4,X5))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2)|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_816, plain, (u1_struct_0(X1)=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X4)|v2_struct_0(X2)|X1!=k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~v1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_6714, plain, (u1_struct_0(X1)=k2_relset_1(u1_struct_0(X2),u1_orders_2(X2))|v2_struct_0(X2)|X1!=k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),u1_orders_2(X2)))|~m1_yellow_0(X1,X2)|~v4_yellow_0(X1,X2)|~v1_orders_2(X1)|~v1_funct_2(u1_orders_2(X2),u1_struct_0(X2),u1_struct_0(X2))|~v1_funct_1(u1_orders_2(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_6522, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(X4)))|~v1_xboole_0(X4)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_544, plain, (v4_orders_2(k5_yellow_0(X1,X2))|~m1_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6340, plain, (k1_relset_1(X1,k1_xboole_0)=X1|X2=k1_xboole_0|k1_xboole_0!=X1)).#trainneg
cnf(i_0_3657, plain, (v1_partfun1(k2_zfmisc_1(X4,X5),X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_545, plain, (v5_orders_2(k5_yellow_0(X1,X2))|~m1_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_549, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v3_card_1(X1,k1_xboole_0)|~v1_card_1(k1_xboole_0)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5147, plain, (r1_orders_2(k1_yellow_2(X1,X2,X3),esk12_1(X2),X4)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X1)|~r2_hidden(X4,X5)|~m1_subset_1(X4,u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_subset_1(esk12_1(X2),u1_struct_0(k1_yellow_2(X1,X2,X3)))|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~r2_hidden(esk12_1(X2),u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_1326, plain, (r1_orders_2(X1,esk4_3(X2,u1_struct_0(X3),X4),esk4_3(X2,u1_struct_0(X3),X4))|r1_lattice3(X2,u1_struct_0(X3),X4)|v2_struct_0(X2)|~r2_hidden(esk4_3(X2,u1_struct_0(X3),X4),u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~m1_yellow_0(X3,X2)|~v4_yellow_0(X1,X3)|~v4_yellow_0(X3,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_5603, plain, (v1_yellow_0(k5_yellow_0(X2,k10_xtuple_0(X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_1(X2),u1_struct_0(k5_yellow_0(X2,k10_xtuple_0(X3))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_6308, negated_conjecture, (u1_struct_0(X1)=k10_xtuple_0(X2)|v2_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|v2_struct_0(X3)|X1!=k5_yellow_0(X3,k2_relset_1(u1_struct_0(X3),X2))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~v1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),u1_struct_0(X3))))|~v1_funct_2(X2,k10_xtuple_0(esk3_0),u1_struct_0(X3))|~v1_funct_1(X2)|~l1_orders_2(X3))).#trainneg
cnf(i_0_4403, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_901, plain, (k2_zfmisc_1(X1,X2)=k9_xtuple_0(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6874, negated_conjecture, (~v1_xboole_0(k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_550, plain, (X1=k1_xboole_0|v1_funct_2(esk29_2(X2,X1),X2,X1)|k1_relset_1(X2,esk29_2(X2,X1))!=X2)).#trainneg
cnf(i_0_6571, plain, (k1_relset_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|X1=k1_xboole_0)).#trainneg
cnf(i_0_7062, negated_conjecture, (~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_551, plain, (X1=k1_xboole_0|v1_funct_2(esk28_1(X1),k5_numbers,X1)|k1_relset_1(k5_numbers,esk28_1(X1))!=k5_numbers)).#trainneg
cnf(i_0_552, plain, (X1=k1_xboole_0|v1_funct_2(esk26_1(k2_zfmisc_1(X2,X1)),X2,X1)|v1_xboole_0(k2_zfmisc_1(X2,X1))|k1_relset_1(X2,esk26_1(k2_zfmisc_1(X2,X1)))!=X2)).#trainneg
cnf(i_0_6716, plain, (u1_struct_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))))=k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))|v2_struct_0(X1)|~m1_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))),X1)|~v4_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))),X1)|~v1_orders_2(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))))|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_6832, negated_conjecture, (u1_struct_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)))=k10_xtuple_0(X2)|v2_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|v2_struct_0(X1)|~m1_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)),X1)|~v4_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)),X1)|~v1_orders_2(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),u1_struct_0(X1))))|~v1_funct_2(X2,k10_xtuple_0(esk3_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_7072, plain, (v1_xboole_0(X1)|k1_relset_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|~v3_card_1(X1,X2))).#trainneg
cnf(i_0_7119, plain, (v1_xboole_0(esk30_1(X1))|k1_relset_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|~v1_card_1(X1))).#trainneg
cnf(i_0_7955, plain, (v1_xboole_0(u1_orders_2(X1))|k1_relset_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|~l1_orders_2(X1))).#trainneg
cnf(i_0_7401, plain, (k1_funct_1(X1,k1_xboole_0)=X3|k1_relset_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7015, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(X4)))|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1077, plain, (esk22_1(k1_xboole_0)=k2_zfmisc_1(X1,X2)|~v1_xboole_0(X2)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_7067, plain, (k1_relset_1(k1_xboole_0,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_4894, negated_conjecture, (m1_subset_1(esk14_1(k10_xtuple_0(esk3_0)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_553, plain, (X1=k1_xboole_0|v1_funct_2(esk19_1(k2_zfmisc_1(X2,X1)),X2,X1)|v1_xboole_0(k2_zfmisc_1(X2,X1))|k1_relset_1(X2,esk19_1(k2_zfmisc_1(X2,X1)))!=X2)).#trainneg
cnf(i_0_6795, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~v3_card_1(esk3_0,k1_xboole_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_6794, plain, (k5_numbers=k1_xboole_0|~v3_card_1(esk28_1(k5_numbers),k1_xboole_0)|~v1_card_1(k1_xboole_0)|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_554, plain, (X1=k1_xboole_0|v1_funct_2(esk17_1(k2_zfmisc_1(X2,X1)),X2,X1)|k1_relset_1(X2,esk17_1(k2_zfmisc_1(X2,X1)))!=X2)).#trainneg
cnf(i_0_555, plain, (X1=k1_xboole_0|v1_funct_2(esk16_1(k2_zfmisc_1(X2,X1)),X2,X1)|v1_xboole_0(k2_zfmisc_1(X2,X1))|k1_relset_1(X2,esk16_1(k2_zfmisc_1(X2,X1)))!=X2)).#trainneg
cnf(i_0_8352, negated_conjecture, (u1_struct_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)))=k10_xtuple_0(X2)|v2_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|v2_struct_0(X1)|~m1_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)),X1)|~v1_orders_2(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),u1_struct_0(X1))))|~v1_funct_2(X2,k10_xtuple_0(esk3_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1722, plain, (r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),X3),esk12_1(k1_yellow_2(X4,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X4)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X4,X2,X3))|~v4_yellow_0(X1,k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk4_3(X1,k2_relset_1(u1_struct_0(X2),X3),esk12_1(k1_yellow_2(X4,X2,X3))),u1_struct_0(X1))|~m1_subset_1(k2_relset_1(u1_struct_0(X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X4,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_8593, plain, (k9_xtuple_0(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_8600, negated_conjecture, (v3_ordinal1(esk14_1(k10_xtuple_0(esk3_0)))|~v3_ordinal1(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_557, plain, (X1=k1_xboole_0|v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X2,X1))),X2,X1)|k1_relset_1(X2,esk14_1(k1_zfmisc_1(k2_zfmisc_1(X2,X1))))!=X2)).#trainneg
cnf(i_0_560, plain, (X1=k1_xboole_0|v1_funct_2(k1_relset_1(k2_zfmisc_1(X2,X1),X3),X2,X1)|k1_relset_1(X2,k1_relset_1(k2_zfmisc_1(X2,X1),X3))!=X2|~v4_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_561, plain, (X1=k1_xboole_0|v1_funct_2(k2_relset_1(k2_zfmisc_1(X2,X1),X3),X2,X1)|k1_relset_1(X2,k2_relset_1(k2_zfmisc_1(X2,X1),X3))!=X2|~v5_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3076, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),u1_orders_2(X1)),u1_struct_0(X2))|v2_struct_0(X1)|~m1_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),u1_orders_2(X1))),X2)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_5464, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k10_xtuple_0(X4)))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1325, plain, (r1_orders_2(X1,esk4_3(X2,X3,X4),esk4_3(X2,X3,X4))|r1_lattice3(X2,X3,X4)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X5))|~r2_hidden(esk4_3(X2,X3,X4),u1_struct_0(X1))|~m1_yellow_0(X1,X5)|~m1_yellow_0(X5,X2)|~v4_yellow_0(X1,X5)|~v4_yellow_0(X5,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk4_3(X2,X3,X4),u1_struct_0(X5)))).#trainneg
cnf(i_0_1287, plain, (k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))=k5_yellow_0(X2,k10_xtuple_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_8810, negated_conjecture, (r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k2_relset_1(u1_struct_0(X2),X3),esk12_1(k1_yellow_2(X4,X2,X3)))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),k10_xtuple_0(esk3_0))|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk4_3(k1_yellow_2(X1,esk2_0,esk3_0),k2_relset_1(u1_struct_0(X2),X3),esk12_1(k1_yellow_2(X4,X2,X3))),k10_xtuple_0(esk3_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_563, plain, (k1_relset_1(X1,esk29_2(X1,X2))=X1|X2=k1_xboole_0|~v1_funct_2(esk29_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_564, plain, (k1_relset_1(k5_numbers,esk28_1(X1))=k5_numbers|X1=k1_xboole_0|~v1_funct_2(esk28_1(X1),k5_numbers,X1))).#trainneg
cnf(i_0_565, plain, (k1_relset_1(X1,esk26_1(k2_zfmisc_1(X1,X2)))=X1|X2=k1_xboole_0|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(esk26_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_8831, plain, (r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),X3),esk12_1(k1_yellow_2(X4,X2,X3)))|v2_struct_0(X4)|v2_struct_0(X2)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X4,X2,X3))|~v4_yellow_0(X1,k1_yellow_2(X4,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_8923, negated_conjecture, (r1_orders_2(X2,esk4_3(X3,X4,X5),esk4_3(X3,X4,X5))|r1_lattice3(X3,X4,X5)|v2_struct_0(X3)|v2_struct_0(X1)|~r2_hidden(esk4_3(X3,X4,X5),u1_struct_0(X2))|~m1_yellow_0(X2,k1_yellow_2(X1,esk2_0,esk3_0))|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),X3)|~v4_yellow_0(X2,k1_yellow_2(X1,esk2_0,esk3_0))|~v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),X3)|~m1_subset_1(esk4_3(X3,X4,X5),k10_xtuple_0(esk3_0))|~m1_subset_1(X5,u1_struct_0(X3))|~l1_orders_2(X3)|~v3_orders_2(X3)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4846, plain, (v1_xboole_0(k10_xtuple_0(X3))|m1_subset_1(X4,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1078, plain, (esk20_1(k1_xboole_0)=k2_zfmisc_1(X1,X2)|~v1_xboole_0(X2)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_4093, negated_conjecture, (k10_xtuple_0(esk3_0)=X1|g1_orders_2(k10_xtuple_0(esk3_0),u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))!=g1_orders_2(X1,X2))).#trainneg
cnf(i_0_566, plain, (k1_relset_1(X1,esk19_1(k2_zfmisc_1(X1,X2)))=X1|X2=k1_xboole_0|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(esk19_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_567, plain, (k1_relset_1(X1,esk17_1(k2_zfmisc_1(X1,X2)))=X1|X2=k1_xboole_0|~v1_funct_2(esk17_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_568, plain, (k1_relset_1(X1,esk16_1(k2_zfmisc_1(X1,X2)))=X1|X2=k1_xboole_0|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(esk16_1(k2_zfmisc_1(X1,X2)),X1,X2))).#trainneg
cnf(i_0_9026, plain, (r1_orders_2(X1,esk12_1(k1_yellow_2(X2,X3,X4)),X5)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X5,k2_relset_1(u1_struct_0(X3),X4))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(esk12_1(k1_yellow_2(X2,X3,X4)),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk12_1(k1_yellow_2(X2,X3,X4)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X2,X3,X4))|~v4_yellow_0(X1,k1_yellow_2(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_orders_2(X3)|~l1_orders_2(X2)|~v1_yellow_0(k1_yellow_2(X2,X3,X4)))).#trainneg
cnf(i_0_8992, negated_conjecture, (r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(X3),esk12_1(k1_yellow_2(X4,X2,X3)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),k10_xtuple_0(esk3_0))|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk4_3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(X3),esk12_1(k1_yellow_2(X4,X2,X3))),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~l1_orders_2(X1)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~v5_relat_1(X3,u1_struct_0(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_9042, plain, (r1_lattice3(X1,k10_xtuple_0(X3),esk12_1(k1_yellow_2(X4,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X4)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X4,X2,X3))|~v4_yellow_0(X1,k1_yellow_2(X4,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X4)|~l1_orders_2(X1)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~v5_relat_1(X3,u1_struct_0(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_9025, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),k2_relset_1(u1_struct_0(X2),X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3)))).#trainneg
cnf(i_0_4707, plain, (m1_subset_1(k2_relset_1(u1_struct_0(X1),u1_orders_2(X1)),k1_zfmisc_1(k10_xtuple_0(u1_orders_2(X1))))|v2_struct_0(X1)|~v1_funct_2(u1_orders_2(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_9371, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X5,X3)),k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~l1_orders_2(X2)|~l1_orders_2(X1)|~r2_hidden(esk12_1(k1_yellow_2(X4,X5,X3)),u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X5,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X5,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X5))|~l1_orders_2(X5)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X5,X3)))).#trainneg
cnf(i_0_9392, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),k10_xtuple_0(X3))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),k2_relset_1(u1_struct_0(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3)))).#trainneg
cnf(i_0_9372, plain, (r1_orders_2(X1,esk12_1(k1_yellow_2(X2,X3,X4)),X5)|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X5,k10_xtuple_0(X4))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(esk12_1(k1_yellow_2(X2,X3,X4)),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk12_1(k1_yellow_2(X2,X3,X4)),u1_struct_0(X1))|~m1_yellow_0(X1,k1_yellow_2(X2,X3,X4))|~v4_yellow_0(X1,k1_yellow_2(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_orders_2(X3)|~l1_orders_2(X2)|~v1_yellow_0(k1_yellow_2(X2,X3,X4)))).#trainneg
cnf(i_0_1086, plain, (esk30_1(k1_xboole_0)=k2_zfmisc_1(X1,X2)|~v1_xboole_0(X2)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_4663, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_570, plain, (k1_relset_1(X1,esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=X1|X2=k1_xboole_0|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2))).#trainneg
cnf(i_0_573, plain, (k1_relset_1(X1,k1_relset_1(k2_zfmisc_1(X1,X2),X3))=X1|X2=k1_xboole_0|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_574, plain, (k1_relset_1(X1,k2_relset_1(k2_zfmisc_1(X1,X2),X3))=X1|X2=k1_xboole_0|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_9439, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(k1_yellow_2(X5,X4,X3)),k10_xtuple_0(X3))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X5,X4,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X5,X4,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~m1_subset_1(esk12_1(k1_yellow_2(X5,X4,X3)),k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~l1_orders_2(X5)|~v1_yellow_0(k1_yellow_2(X5,X4,X3)))).#trainneg
cnf(i_0_9467, plain, (r1_lattice3(X1,X2,esk12_1(k1_yellow_2(X3,X4,X5)))|v2_struct_0(X4)|v2_struct_0(X3)|~m1_subset_1(esk12_1(k1_yellow_2(X3,X4,X5)),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk12_1(k1_yellow_2(X3,X4,X5)),u1_struct_0(X1))|~r2_hidden(esk4_3(X1,X2,esk12_1(k1_yellow_2(X3,X4,X5))),k10_xtuple_0(X5))|~m1_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~v4_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(esk4_3(X1,X2,esk12_1(k1_yellow_2(X3,X4,X5))),u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_orders_2(X4)|~l1_orders_2(X3)|~v1_yellow_0(k1_yellow_2(X3,X4,X5)))).#trainneg
cnf(i_0_9456, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),k10_xtuple_0(X3))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~v5_relat_1(X3,u1_struct_0(X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_9438, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(k1_yellow_2(X5,X4,X3)),k2_relset_1(u1_struct_0(X2),X3))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X5,X4,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X5,X4,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~m1_subset_1(esk12_1(k1_yellow_2(X5,X4,X3)),k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~l1_orders_2(X5)|~v1_yellow_0(k1_yellow_2(X5,X4,X3)))).#trainneg
cnf(i_0_9447, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X5,X4,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X5,X4,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~m1_subset_1(esk12_1(k1_yellow_2(X5,X4,X3)),k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~l1_orders_2(X5)|~v1_yellow_0(k1_yellow_2(X5,X4,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X5,X4,X3)),u1_struct_0(k1_yellow_2(X1,X2,X3))))).#trainneg
cnf(i_0_9800, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~r2_hidden(esk12_1(k1_yellow_2(X4,X5,X3)),u1_struct_0(k1_yellow_2(X6,X2,X3)))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X5,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X5,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X5,X3)),k10_xtuple_0(X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X5))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X5)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X5,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X6),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X6),u1_struct_0(X2))|~l1_orders_2(X6))).#trainneg
cnf(i_0_1268, plain, (esk30_1(k1_xboole_0)=k2_zfmisc_1(X1,X2)|~v1_xboole_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_582, plain, (m1_yellow_0(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(u1_orders_2(X2))))).#trainneg
cnf(i_0_586, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9818, plain, (v1_xboole_0(k10_xtuple_0(X3))|v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X5,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X5,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X5,X3)),k10_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X5))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X5)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X5,X3)))).#trainneg
cnf(i_0_9468, plain, (r1_orders_2(X1,esk12_1(k1_yellow_2(X2,X3,X4)),X5)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(esk12_1(k1_yellow_2(X2,X3,X4)),u1_struct_0(X1))|~m1_yellow_0(X1,X6)|~v4_yellow_0(X1,X6)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X6))|~m1_subset_1(esk12_1(k1_yellow_2(X2,X3,X4)),u1_struct_0(X6))|~l1_orders_2(X6)|~r2_hidden(esk12_1(k1_yellow_2(X2,X3,X4)),u1_struct_0(X6))|~r2_hidden(X5,k10_xtuple_0(X4))|~m1_yellow_0(X6,k1_yellow_2(X2,X3,X4))|~v4_yellow_0(X6,k1_yellow_2(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_orders_2(X3)|~l1_orders_2(X2)|~v1_yellow_0(k1_yellow_2(X2,X3,X4)))).#trainneg
cnf(i_0_9789, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(k10_xtuple_0(X3))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),k10_xtuple_0(X3)))).#trainneg
cnf(i_0_9337, plain, (r1_lattice3(X1,X2,esk12_1(k1_yellow_2(X3,X4,X5)))|v2_struct_0(X3)|v2_struct_0(X4)|~m1_subset_1(esk12_1(k1_yellow_2(X3,X4,X5)),u1_struct_0(X1))|~l1_orders_2(X1)|~r2_hidden(esk12_1(k1_yellow_2(X3,X4,X5)),u1_struct_0(X1))|~r2_hidden(esk4_3(X1,X2,esk12_1(k1_yellow_2(X3,X4,X5))),k2_relset_1(u1_struct_0(X4),X5))|~m1_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~v4_yellow_0(X1,k1_yellow_2(X3,X4,X5))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(esk4_3(X1,X2,esk12_1(k1_yellow_2(X3,X4,X5))),u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_orders_2(X4)|~l1_orders_2(X3)|~v1_yellow_0(k1_yellow_2(X3,X4,X5)))).#trainneg
cnf(i_0_9391, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_1(k1_yellow_2(X4,X2,X3)),k2_relset_1(u1_struct_0(X2),X3))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),k2_relset_1(u1_struct_0(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3)))).#trainneg
cnf(i_0_8530, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk22_1(k1_xboole_0)))|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1269, plain, (esk22_1(k1_xboole_0)=k2_zfmisc_1(X1,X2)|~v1_xboole_0(X1)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_4965, plain, (r1_tarski(u1_struct_0(esk27_1(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))),k5_numbers))).#trainneg
cnf(i_0_587, plain, (X1=k10_xtuple_0(esk29_2(X2,X3))|r2_hidden(esk7_2(esk29_2(X2,X3),X1),k9_xtuple_0(esk29_2(X2,X3)))|r2_hidden(esk6_2(esk29_2(X2,X3),X1),X1))).#trainneg
cnf(i_0_3959, negated_conjecture, (v3_ordinal1(u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0)))))).#trainneg
cnf(i_0_588, plain, (X1=k10_xtuple_0(esk28_1(X2))|r2_hidden(esk7_2(esk28_1(X2),X1),k9_xtuple_0(esk28_1(X2)))|r2_hidden(esk6_2(esk28_1(X2),X1),X1))).#trainneg
cnf(i_0_589, plain, (X1=k10_xtuple_0(esk25_2(X2,X3))|r2_hidden(esk7_2(esk25_2(X2,X3),X1),k9_xtuple_0(esk25_2(X2,X3)))|r2_hidden(esk6_2(esk25_2(X2,X3),X1),X1)|v2_struct_0(X3)|~v1_funct_1(esk25_2(X2,X3))|~l1_orders_2(X3)|~l1_orders_2(X2)|~v3_orders_2(X3))).#trainneg
cnf(i_0_10071, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_9400, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(k1_yellow_2(X1,X2,X3)))|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),k2_relset_1(u1_struct_0(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),u1_struct_0(k1_yellow_2(X1,X2,X3))))).#trainneg
cnf(i_0_1270, plain, (esk20_1(k1_xboole_0)=k2_zfmisc_1(X1,X2)|~v1_xboole_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_590, plain, (X1=k10_xtuple_0(esk22_1(X2))|r2_hidden(esk7_2(esk22_1(X2),X1),k9_xtuple_0(esk22_1(X2)))|r2_hidden(esk6_2(esk22_1(X2),X1),X1)|~v1_funct_1(esk22_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1563, plain, (m1_subset_1(esk12_1(g1_orders_2(k5_numbers,esk28_1(k5_numbers))),k5_numbers)|~v1_yellow_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))).#trainneg
cnf(i_0_591, plain, (X1=k10_xtuple_0(esk20_1(X2))|r2_hidden(esk7_2(esk20_1(X2),X1),k9_xtuple_0(esk20_1(X2)))|r2_hidden(esk6_2(esk20_1(X2),X1),X1)|~v1_funct_1(esk20_1(X2))|~v1_card_1(X2))).#trainneg
cnf(i_0_593, plain, (r2_hidden(X1,X2)|r1_lattice3(X4,k9_xtuple_0(X3),X5)|X1!=k1_funct_1(X3,esk4_3(X4,k9_xtuple_0(X3),X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(X4))|~l1_orders_2(X4))).#trainneg
cnf(i_0_10490, negated_conjecture, (r2_hidden(X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(X3,k1_funct_1(esk3_0,X4))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|k9_xtuple_0(X3)!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10501, negated_conjecture, (m1_subset_1(k1_funct_1(esk3_0,X1),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10502, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(k1_yellow_2(X2,X1,X3),esk12_1(X1),k1_funct_1(esk3_0,X4))|v2_struct_0(X2)|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(esk12_1(X1),u1_struct_0(k1_yellow_2(X2,X1,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(k1_funct_1(esk3_0,X4),u1_struct_0(k1_yellow_2(X2,X1,X3)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(k1_yellow_2(X2,X1,X3))|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_yellow_0(X1)|X5!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9466, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|v2_struct_0(X5)|v2_struct_0(X4)|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,esk12_1(k1_yellow_2(X4,X5,X6)))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X5,X6)),u1_struct_0(esk1_0))|~r2_hidden(esk12_1(k1_yellow_2(X4,X5,X6)),u1_struct_0(esk1_0))|~r2_hidden(X3,k10_xtuple_0(X6))|~m1_yellow_0(esk1_0,k1_yellow_2(X4,X5,X6))|~v4_yellow_0(esk1_0,k1_yellow_2(X4,X5,X6))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X5))|~v1_funct_1(X6)|~l1_orders_2(X5)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X5,X6)))).#trainneg
cnf(i_0_10521, plain, (v1_xboole_0(k2_relset_1(u1_struct_0(X2),X3))|v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~v4_yellow_0(k1_yellow_2(X1,X2,X3),k1_yellow_2(X4,X2,X3))|~m1_subset_1(esk12_1(k1_yellow_2(X4,X2,X3)),k2_relset_1(u1_struct_0(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v1_yellow_0(k1_yellow_2(X4,X2,X3)))).#trainneg
cnf(i_0_10499, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,X3))|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,X3),u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v1_yellow_0(X1)|u1_struct_0(X1)!=k10_xtuple_0(esk3_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9218, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk20_1(k1_xboole_0)))|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_10725, plain, (v3_ordinal1(esk12_1(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))|~v1_yellow_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))).#trainneg
cnf(i_0_594, plain, (m1_subset_1(esk4_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_598, plain, (r2_hidden(X1,X2)|X1!=k1_funct_1(X3,esk5_3(X3,X4,X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_10737, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,X2),k1_funct_1(esk3_0,X2))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(k1_funct_1(esk3_0,X2),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10848, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(X1)),k1_funct_1(esk3_0,k1_funct_1(esk3_0,X2)))|u1_struct_0(X1)!=k10_xtuple_0(esk3_0)|~r2_hidden(esk12_1(X1),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X1)|~v4_yellow_0(esk1_0,X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X1)),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_10829, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|v2_struct_0(X3)|v2_struct_0(X4)|X1!=k1_funct_1(esk3_0,esk12_1(k5_yellow_0(X4,k10_xtuple_0(X5))))|X2!=k1_funct_1(esk3_0,X6)|~r2_hidden(esk12_1(k5_yellow_0(X4,k10_xtuple_0(X5))),u1_struct_0(esk1_0))|~r2_hidden(X6,k10_xtuple_0(X5))|~m1_yellow_0(esk1_0,k5_yellow_0(X4,k10_xtuple_0(X5)))|~v4_yellow_0(esk1_0,k5_yellow_0(X4,k10_xtuple_0(X5)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_orders_2(X4)|~l1_orders_2(X3)|~v1_yellow_0(k5_yellow_0(X4,k10_xtuple_0(X5)))|~m1_subset_1(k10_xtuple_0(X5),k1_zfmisc_1(u1_struct_0(X4))))).#trainneg
cnf(i_0_10953, plain, (r2_hidden(k1_funct_1(X1,esk5_3(X1,X2,X3)),X4)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10960, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,X3,X3)|~r2_hidden(X3,u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk5_3(esk3_0,X2,X3),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_599, plain, (v3_ordinal1(esk4_3(X1,X2,X3))|r1_lattice3(X1,X2,X3)|~v3_ordinal1(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_601, plain, (k1_funct_1(esk29_2(X1,X2),esk7_2(esk29_2(X1,X2),X3))=esk6_2(esk29_2(X1,X2),X3)|X3=k10_xtuple_0(esk29_2(X1,X2))|r2_hidden(esk6_2(esk29_2(X1,X2),X3),X3))).#trainneg
cnf(i_0_602, plain, (k1_funct_1(esk28_1(X1),esk7_2(esk28_1(X1),X2))=esk6_2(esk28_1(X1),X2)|X2=k10_xtuple_0(esk28_1(X1))|r2_hidden(esk6_2(esk28_1(X1),X2),X2))).#trainneg
cnf(i_0_10934, plain, (v1_yellow_0(k1_yellow_2(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|r1_lattice3(X4,u1_struct_0(k1_yellow_2(X1,X2,X3)),X5)|~v1_xboole_0(k2_relset_1(u1_struct_0(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk4_3(X4,u1_struct_0(k1_yellow_2(X1,X2,X3)),X5),k2_relset_1(u1_struct_0(X2),X3))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(X5,u1_struct_0(X4))|~l1_orders_2(X4))).#trainneg
cnf(i_0_6249, negated_conjecture, (v1_xboole_0(u1_struct_0(k5_yellow_0(X1,k10_xtuple_0(X2))))|v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(esk3_0),u1_struct_0(X1))))|~m1_subset_1(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,k10_xtuple_0(X2))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_2(X2,k10_xtuple_0(esk3_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_9598, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk30_1(k1_xboole_0)))|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_603, plain, (k1_funct_1(esk25_2(X1,X2),esk7_2(esk25_2(X1,X2),X3))=esk6_2(esk25_2(X1,X2),X3)|X3=k10_xtuple_0(esk25_2(X1,X2))|r2_hidden(esk6_2(esk25_2(X1,X2),X3),X3)|v2_struct_0(X2)|~v1_funct_1(esk25_2(X1,X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v3_orders_2(X2))).#trainneg
cnf(i_0_604, plain, (k1_funct_1(esk22_1(X1),esk7_2(esk22_1(X1),X2))=esk6_2(esk22_1(X1),X2)|X2=k10_xtuple_0(esk22_1(X1))|r2_hidden(esk6_2(esk22_1(X1),X2),X2)|~v1_funct_1(esk22_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_605, plain, (k1_funct_1(esk20_1(X1),esk7_2(esk20_1(X1),X2))=esk6_2(esk20_1(X1),X2)|X2=k10_xtuple_0(esk20_1(X1))|r2_hidden(esk6_2(esk20_1(X1),X2),X2)|~v1_funct_1(esk20_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_11099, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,X2,X2)|X3!=k10_xtuple_0(esk3_0)|~r2_hidden(X2,u1_struct_0(X1))|~r2_hidden(X2,X3)|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0))).#trainneg
cnf(i_0_11096, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X2)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_11095, negated_conjecture, (r1_orders_2(esk2_0,X1,X1)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X1,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_11044, plain, (~v1_xboole_0(X1)|X1!=k10_xtuple_0(X2)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_11093, negated_conjecture, (r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,esk5_3(esk3_0,X2,X3))|~r2_hidden(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,esk1_0)|~v4_yellow_0(esk1_0,esk1_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(esk5_3(esk3_0,X2,X3),u1_struct_0(esk1_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_3848, negated_conjecture, (r1_orders_2(X1,esk12_1(esk2_0),X2)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(esk12_1(esk2_0),u1_struct_0(X3))|~l1_orders_2(X3)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(X3))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(X3,esk2_0)|~v4_yellow_0(X3,esk2_0)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_607, plain, (v3_ordinal1(esk25_2(X1,X2))|v2_struct_0(X2)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v3_orders_2(X2))).#trainneg
cnf(i_0_608, plain, (v1_xboole_0(esk25_2(X1,X2))|v2_struct_0(X2)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v3_orders_2(X2))).#trainneg
cnf(i_0_3293, plain, (v1_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)))|v1_xboole_0(u1_struct_0(k5_yellow_0(X1,u1_struct_0(X2))))|~m1_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)),X2)|~v4_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)),X2)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_yellow_0(X2)|~m1_subset_1(esk12_1(X2),u1_struct_0(k5_yellow_0(X1,u1_struct_0(X2)))))).#trainneg
cnf(i_0_774, plain, (k1_relset_1(X2,u1_orders_2(k5_yellow_0(X1,X2)))=X2|X2=k1_xboole_0|~v1_funct_2(u1_orders_2(k5_yellow_0(X1,X2)),X2,X2)|~l1_orders_2(k5_yellow_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_6662, negated_conjecture, (v3_orders_2(k1_yellow_2(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_10940, negated_conjecture, (r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),esk12_1(k1_yellow_2(X2,X3,esk3_0)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~r2_hidden(esk12_1(k1_yellow_2(X2,X3,esk3_0)),k10_xtuple_0(esk3_0))|~m1_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),k1_yellow_2(X2,X3,esk3_0))|~v4_yellow_0(k1_yellow_2(X1,esk2_0,esk3_0),k1_yellow_2(X2,X3,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(X3))|~l1_orders_2(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(k1_yellow_2(X2,X3,esk3_0))|~m1_subset_1(esk12_1(k1_yellow_2(X2,X3,esk3_0)),u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0)))).#trainneg
cnf(i_0_629, plain, (k5_yellow_0(X1,X2)=k5_yellow_0(X1,X3)|u1_struct_0(k5_yellow_0(X1,X2))!=X3|~m1_yellow_0(k5_yellow_0(X1,X2),X1)|~v1_orders_2(k5_yellow_0(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1294, plain, (u1_struct_0(k5_yellow_0(X2,k10_xtuple_0(X3)))=k2_relset_1(u1_struct_0(X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~m1_subset_1(k10_xtuple_0(X3),k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_9970, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk30_1(k1_xboole_0)))|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5292, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk3_0))|v1_yellow_0(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_610, plain, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))|v2_struct_0(X3)|~r2_hidden(X1,esk25_2(X2,X3))|~l1_orders_2(X3)|~l1_orders_2(X2)|~v3_orders_2(X3))).#trainneg
cnf(i_0_3158, negated_conjecture, (v1_orders_2(g1_orders_2(k10_xtuple_0(esk3_0),u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))))).#trainneg
cnf(i_0_3235, plain, (m1_subset_1(u1_struct_0(esk27_1(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_6766, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_613, plain, (k1_relset_1(u1_struct_0(X1),esk25_2(X1,X2))=u1_struct_0(X1)|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X2)|~v1_funct_2(esk25_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v3_orders_2(X2))).#trainneg
cnf(i_0_11173, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,X2,X2)|~r2_hidden(X2,u1_struct_0(X1))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0))).#trainneg
cnf(i_0_11239, plain, (k1_relset_1(k10_xtuple_0(X1),u1_orders_2(k1_yellow_2(X3,X2,X1)))=k10_xtuple_0(X1)|k10_xtuple_0(X1)=k1_xboole_0|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(k10_xtuple_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(u1_orders_2(k1_yellow_2(X3,X2,X1)),k10_xtuple_0(X1),k10_xtuple_0(X1))|~l1_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_orders_2(X3))).#trainneg
cnf(i_0_11207, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(esk12_1(esk1_0),u1_struct_0(esk1_0))|~r2_hidden(X1,X2)|~m1_yellow_0(esk1_0,esk1_0)|~v4_yellow_0(esk1_0,esk1_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_11206, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),k1_funct_1(esk3_0,esk5_3(esk3_0,X1,X2)))|X1!=k10_xtuple_0(esk3_0)|~r2_hidden(esk12_1(esk1_0),u1_struct_0(esk1_0))|~r2_hidden(X2,X1)|~m1_yellow_0(esk1_0,esk1_0)|~v4_yellow_0(esk1_0,esk1_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk5_3(esk3_0,X1,X2)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_11177, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|~r2_hidden(X1,k10_xtuple_0(esk3_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_10730, negated_conjecture, (r2_hidden(X1,X2)|r1_lattice3(X3,u1_struct_0(esk1_0),X4)|u1_struct_0(esk2_0)=k1_xboole_0|X1!=k1_funct_1(esk3_0,esk4_3(X3,u1_struct_0(esk1_0),X4))|X2!=k10_xtuple_0(esk3_0)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3))).#trainneg
cnf(i_0_10315, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk22_1(k1_xboole_0)))|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4144, plain, (m1_subset_1(k10_xtuple_0(esk28_1(k5_numbers)),k1_zfmisc_1(k5_numbers))|v2_struct_0(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_4063, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))=X1|g1_orders_2(k10_xtuple_0(esk3_0),u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))!=g1_orders_2(X2,X1))).#trainneg
cnf(i_0_615, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),esk25_2(X1,X1)))|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11461, plain, (m1_subset_1(k10_xtuple_0(esk28_1(k5_numbers)),k1_zfmisc_1(k5_numbers))|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_616, plain, (esk25_2(X1,X1)=X2|v2_struct_0(X1)|g1_orders_2(u1_struct_0(X1),esk25_2(X1,X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_617, plain, (u1_struct_0(X1)=X2|v2_struct_0(X1)|g1_orders_2(u1_struct_0(X1),esk25_2(X1,X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11418, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,X2,X2)|~r2_hidden(X2,u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(X2,k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_11441, negated_conjecture, (r1_lattice3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~r2_hidden(esk4_3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11436, negated_conjecture, (r1_orders_2(esk2_0,X1,X2)|u1_struct_0(esk2_0)=k1_xboole_0|X2!=k1_funct_1(esk3_0,X3)|X1!=k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk12_1(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk1_0))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk1_0))|~r2_hidden(X3,k10_xtuple_0(esk3_0))|~m1_yellow_0(esk1_0,esk2_0)|~v4_yellow_0(esk1_0,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11417, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,k1_funct_1(X2,esk5_3(X2,X3,X4)),k1_funct_1(X2,esk5_3(X2,X3,X4)))|~r2_hidden(k1_funct_1(X2,esk5_3(X2,X3,X4)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|k10_xtuple_0(esk3_0)!=k10_xtuple_0(X2)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_11432, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(esk2_0,k1_funct_1(esk3_0,esk12_1(esk1_0)),X1)|~r2_hidden(esk12_1(esk1_0),u1_struct_0(esk1_0))|~r2_hidden(X1,k10_xtuple_0(esk3_0))|~m1_yellow_0(esk1_0,esk1_0)|~v4_yellow_0(esk1_0,esk1_0)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_11463, negated_conjecture, (u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)))=u1_orders_2(X1)|g1_orders_2(k10_xtuple_0(esk3_0),u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))!=X1|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_3729, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X5,X6)|~v1_xboole_0(X4)|~v1_xboole_0(X1)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_11411, plain, (k1_relset_1(k10_xtuple_0(X3),esk25_2(k1_yellow_2(X1,X2,X3),X4))=k10_xtuple_0(X3)|u1_struct_0(X4)=k1_xboole_0|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_orders_2(X4)|~l1_orders_2(k1_yellow_2(X1,X2,X3))|~v3_orders_2(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_11438, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,X3)|~v4_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X3))|~l1_orders_2(X3)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X3))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(X3,esk2_0)|~v4_yellow_0(X3,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11408, plain, (k1_relset_1(X2,esk25_2(k5_yellow_0(X1,X2),X3))=X2|u1_struct_0(X3)=k1_xboole_0|v2_struct_0(X3)|~l1_orders_2(X3)|~l1_orders_2(k5_yellow_0(X1,X2))|~v3_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_11526, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(esk1_0)))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(esk4_3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_11533, negated_conjecture, (r1_lattice3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(esk1_0)))|u1_struct_0(esk2_0)=k1_xboole_0|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~r2_hidden(esk12_1(esk1_0),u1_struct_0(esk1_0))|~r2_hidden(esk4_3(esk2_0,X1,k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_yellow_0(esk1_0,esk1_0)|~v4_yellow_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_11521, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0)))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(X1))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~l1_orders_2(X1)|~m1_subset_1(esk4_3(X1,X2,k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0)))).#trainneg
cnf(i_0_11527, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(esk2_0,k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_11528, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(esk2_0,k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk12_1(esk1_0))),X1)|X1!=k1_funct_1(esk3_0,X2)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk1_0))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(esk1_0,esk2_0)|~v4_yellow_0(esk1_0,esk2_0)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,k1_funct_1(esk3_0,esk12_1(esk1_0))),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_621, plain, (esk25_2(X1,X2)=k1_xboole_0|v2_struct_0(X2)|u1_struct_0(X2)!=k1_xboole_0|~v1_funct_2(esk25_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v3_orders_2(X2))).#trainneg
cnf(i_0_11481, plain, (v3_ordinal1(k10_xtuple_0(esk28_1(k5_numbers)))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_11483, plain, (v1_funct_1(k10_xtuple_0(esk28_1(k5_numbers)))|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v1_funct_2(esk28_1(k5_numbers),k5_numbers,k5_numbers))).#trainneg
cnf(i_0_623, plain, (v1_xboole_0(esk25_2(X1,X2))|v2_struct_0(X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X2)|~l1_orders_2(X1)|~v3_orders_2(X2))).#trainneg
cnf(i_0_11534, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|~r2_hidden(esk12_1(esk1_0),u1_struct_0(esk1_0))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(esk1_0,esk1_0)|~v4_yellow_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_11599, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),X2)|v1_xboole_0(u1_struct_0(X3))|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~r2_hidden(X2,k10_xtuple_0(esk3_0))|~m1_yellow_0(X3,esk2_0)|~m1_yellow_0(X1,X3)|~v4_yellow_0(X3,esk2_0)|~v4_yellow_0(X1,X3)|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X3))|~l1_orders_2(X3)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X3)))).#trainneg
cnf(i_0_11439, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(esk1_0)),k1_funct_1(esk3_0,esk5_3(esk3_0,X2,X3)))|u1_struct_0(esk2_0)=k1_xboole_0|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,esk5_3(esk3_0,X2,X3)),u1_struct_0(X1))|~m1_subset_1(k1_funct_1(esk3_0,esk5_3(esk3_0,X2,X3)),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(esk2_0))|X2!=k10_xtuple_0(esk3_0)|~r2_hidden(esk12_1(esk1_0),u1_struct_0(esk1_0))|~r2_hidden(X3,X2)|~m1_yellow_0(esk1_0,esk1_0)|~v4_yellow_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_11406, negated_conjecture, (k1_relset_1(k10_xtuple_0(esk3_0),esk25_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),X1))=k10_xtuple_0(esk3_0)|u1_struct_0(X1)=k1_xboole_0|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11622, negated_conjecture, (k1_relset_1(k10_xtuple_0(esk3_0),esk25_2(k1_yellow_2(X2,esk2_0,esk3_0),X1))=k10_xtuple_0(esk3_0)|u1_struct_0(X1)=k1_xboole_0|v2_struct_0(X1)|v2_struct_0(X2)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_10658, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(esk20_1(k1_xboole_0)))|~v1_card_1(k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_11376, plain, (v3_ordinal1(u1_struct_0(esk27_1(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))))|~v3_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_11378, plain, (v1_funct_1(u1_struct_0(esk27_1(g1_orders_2(k5_numbers,esk28_1(k5_numbers)))))|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers))).#trainneg
cnf(i_0_6253, plain, (v1_xboole_0(u1_struct_0(k5_yellow_0(X1,k10_xtuple_0(X2))))|v1_yellow_0(k1_yellow_2(X3,X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k1_yellow_2(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(X6),u1_struct_0(X1))))|~m1_subset_1(esk12_1(X1),u1_struct_0(k5_yellow_0(X1,k10_xtuple_0(X2))))|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_2(X2,k10_xtuple_0(X6),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~l1_orders_2(k1_yellow_2(X4,X5,X6))|~v1_yellow_0(X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X5))|~v1_funct_1(X6)|~l1_orders_2(X5)|~l1_orders_2(X4))).#trainneg
cnf(i_0_3813, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X5)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_11624, plain, (k1_relset_1(k2_relset_1(u1_struct_0(X1),X2),esk25_2(k1_yellow_2(X4,X1,X2),X3))=k2_relset_1(u1_struct_0(X1),X2)|u1_struct_0(X3)=k1_xboole_0|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X1)|~m1_subset_1(k2_relset_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X3)|~l1_orders_2(X1)|~v3_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_orders_2(X4))).#trainneg
cnf(i_0_3435, negated_conjecture, (r1_orders_2(X1,esk4_3(X2,k10_xtuple_0(esk3_0),X3),esk4_3(X2,k10_xtuple_0(esk3_0),X3))|r1_lattice3(X2,k10_xtuple_0(esk3_0),X3)|~r2_hidden(esk4_3(X2,k10_xtuple_0(esk3_0),X3),u1_struct_0(X1))|~m1_yellow_0(X1,X4)|~v4_yellow_0(X1,X4)|~m1_subset_1(esk4_3(X2,k10_xtuple_0(esk3_0),X3),u1_struct_0(X1))|~m1_subset_1(esk4_3(X2,k10_xtuple_0(esk3_0),X3),u1_struct_0(X4))|~l1_orders_2(X4)|~r2_hidden(esk4_3(X2,k10_xtuple_0(esk3_0),X3),u1_struct_0(X4))|~m1_yellow_0(X4,esk2_0)|~v4_yellow_0(X4,esk2_0)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_627, plain, (X1=k10_xtuple_0(X2)|v1_xboole_0(X1)|esk6_2(X2,X1)!=k1_funct_1(X2,X3)|~v1_relat_1(X2)|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X2)|~m1_subset_1(esk6_2(X2,X1),X1))).#trainneg
cnf(i_0_2735, negated_conjecture, (v5_relat_1(u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))),k10_xtuple_0(esk3_0))|~l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_628, plain, (esk27_1(X1)=k5_yellow_0(X1,X2)|u1_struct_0(esk27_1(X1))!=X2|~m1_yellow_0(esk27_1(X1),X1)|~v1_orders_2(esk27_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_632, plain, (r1_orders_2(X1,X2,X3)|X3!=k1_funct_1(esk29_2(u1_struct_0(X4),u1_struct_0(X1)),X5)|X2!=k1_funct_1(esk29_2(u1_struct_0(X4),u1_struct_0(X1)),X6)|~r1_orders_2(X4,X6,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|~v5_orders_3(esk29_2(u1_struct_0(X4),u1_struct_0(X1)),X4,X1)|~v1_funct_2(esk29_2(u1_struct_0(X4),u1_struct_0(X1)),u1_struct_0(X4),u1_struct_0(X1))|~l1_orders_2(X1)|~l1_orders_2(X4))).#trainneg
cnf(i_0_4587, negated_conjecture, (v1_yellow_0(k1_yellow_2(X1,X2,esk3_0))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_hidden(esk12_1(esk2_0),u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),esk3_0))))|~m1_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~v4_yellow_0(k1_yellow_2(X1,X2,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk12_1(esk2_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(esk3_0,u1_struct_0(X3),u1_struct_0(X2))|~l1_orders_2(k1_yellow_2(X1,X2,esk3_0))|~l1_orders_2(X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~v1_yellow_0(esk2_0))).#trainneg
cnf(i_0_5674, plain, (v1_xboole_0(u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))))|v1_yellow_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(esk12_1(X2),u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_11820, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(esk1_0)),u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0)))).#trainneg
cnf(i_0_1275, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_647, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_partfun1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X1)|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2737, negated_conjecture, (v4_relat_1(u1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))),k10_xtuple_0(esk3_0))|~l1_orders_2(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0))))).#trainneg
cnf(i_0_633, plain, (r1_orders_2(X1,X2,X3)|v1_xboole_0(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))|X3!=k1_funct_1(esk26_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X5)|X2!=k1_funct_1(esk26_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X6)|~r1_orders_2(X4,X6,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|~v5_orders_3(esk26_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X4,X1)|~v1_funct_2(esk26_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(esk26_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~l1_orders_2(X1)|~l1_orders_2(X4))).#trainneg
cnf(i_0_634, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|X3!=k1_funct_1(esk25_2(X4,X1),X5)|X2!=k1_funct_1(esk25_2(X4,X1),X6)|~r1_orders_2(X4,X6,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|~v5_orders_3(esk25_2(X4,X1),X4,X1)|~v1_funct_2(esk25_2(X4,X1),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(esk25_2(X4,X1))|~l1_orders_2(X1)|~l1_orders_2(X4)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11813, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|r1_lattice3(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0)))|v2_struct_0(X1)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(esk1_0)),k10_xtuple_0(esk3_0))|~m1_subset_1(esk4_3(k5_yellow_0(esk2_0,k10_xtuple_0(esk3_0)),k10_xtuple_0(esk3_0),k1_funct_1(esk3_0,esk12_1(esk1_0))),k10_xtuple_0(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1))).#trainneg
cnf(i_0_11894, negated_conjecture, (r1_orders_2(X1,esk4_3(X2,k10_xtuple_0(esk3_0),X3),esk4_3(X2,k10_xtuple_0(esk3_0),X3))|r1_lattice3(X2,k10_xtuple_0(esk3_0),X3)|v1_xboole_0(u1_struct_0(X4))|~r2_hidden(esk4_3(X2,k10_xtuple_0(esk3_0),X3),u1_struct_0(X1))|~m1_yellow_0(X4,esk2_0)|~m1_yellow_0(X1,X4)|~v4_yellow_0(X4,esk2_0)|~v4_yellow_0(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X4)|~l1_orders_2(X2)|~m1_subset_1(esk4_3(X2,k10_xtuple_0(esk3_0),X3),u1_struct_0(X4)))).#trainneg
cnf(i_0_5548, plain, (r1_tarski(k2_relset_1(u1_struct_0(X1),k10_xtuple_0(X2)),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X4)|~m1_yellow_0(k5_yellow_0(X1,k2_relset_1(u1_struct_0(X1),k10_xtuple_0(X2))),X3)|~v1_funct_2(k10_xtuple_0(X2),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(k10_xtuple_0(X2))|~l1_orders_2(X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~v5_relat_1(X2,k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_5039, negated_conjecture, (v1_xboole_0(u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)))|m1_subset_1(esk4_3(k1_yellow_2(X1,esk2_0,esk3_0),X2,X3),u1_struct_0(esk2_0))|v2_struct_0(X1)|r1_lattice3(k1_yellow_2(X1,esk2_0,esk3_0),X2,X3)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(k1_yellow_2(X1,esk2_0,esk3_0)))|~l1_orders_2(k1_yellow_2(X1,esk2_0,esk3_0)))).#trainneg
cnf(i_0_648, plain, (k1_relset_1(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5))=X1|X3=k1_xboole_0|v1_xboole_0(X2)|~v1_funct_2(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_635, plain, (r1_orders_2(X1,X2,X3)|v1_xboole_0(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))|X3!=k1_funct_1(esk19_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X5)|X2!=k1_funct_1(esk19_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X6)|~r1_orders_2(X4,X6,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|~v5_orders_3(esk19_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X4,X1)|~v1_funct_2(esk19_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(esk19_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~l1_orders_2(X1)|~l1_orders_2(X4))).#trainneg
cnf(i_0_636, plain, (r1_orders_2(X1,X2,X3)|X3!=k1_funct_1(esk17_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X5)|X2!=k1_funct_1(esk17_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X6)|~r1_orders_2(X4,X6,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|~v5_orders_3(esk17_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X4,X1)|~v1_funct_2(esk17_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(esk17_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~l1_orders_2(X1)|~l1_orders_2(X4))).#trainneg
cnf(i_0_12471, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|r1_orders_2(X2,esk4_3(X1,k10_xtuple_0(esk3_0),X3),esk4_3(X1,k10_xtuple_0(esk3_0),X3))|r1_lattice3(X1,k10_xtuple_0(esk3_0),X3)|~r2_hidden(esk4_3(X1,k10_xtuple_0(esk3_0),X3),u1_struct_0(X2))|~m1_yellow_0(X1,esk2_0)|~m1_yellow_0(X2,X1)|~v4_yellow_0(X1,esk2_0)|~v4_yellow_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_12491, negated_conjecture, (r1_orders_2(X1,esk4_3(k1_yellow_2(X2,esk2_0,esk3_0),X3,X4),esk4_3(k1_yellow_2(X2,esk2_0,esk3_0),X3,X4))|v1_xboole_0(u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0)))|r1_lattice3(k1_yellow_2(X2,esk2_0,esk3_0),X3,X4)|v2_struct_0(X2)|~r2_hidden(esk4_3(k1_yellow_2(X2,esk2_0,esk3_0),X3,X4),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~m1_subset_1(X4,u1_struct_0(k1_yellow_2(X2,esk2_0,esk3_0)))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_11424, negated_conjecture, (k1_relset_1(k10_xtuple_0(esk3_0),u1_orders_2(k1_yellow_2(X2,esk2_0,esk3_0)))=k10_xtuple_0(esk3_0)|k10_xtuple_0(esk3_0)=k1_xboole_0|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(u1_orders_2(k1_yellow_2(X2,esk2_0,esk3_0)),k10_xtuple_0(esk3_0),k10_xtuple_0(esk3_0))|~v1_funct_2(esk3_0,u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_orders_2(X1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~l1_orders_2(X2))).#trainneg
cnf(i_0_6019, plain, (v1_xboole_0(u1_struct_0(X1))|v1_yellow_0(k1_yellow_2(X2,X1,X3))|v2_struct_0(X2)|~r2_hidden(esk12_1(X1),k2_zfmisc_1(X4,X5))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_yellow_0(X1)|~v1_xboole_0(X3)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_5600, plain, (v1_yellow_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(esk12_1(X2),u1_struct_0(k5_yellow_0(X2,k2_relset_1(u1_struct_0(X2),X3))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_orders_2(X2)|~l1_orders_2(X1)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_12108, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_11379, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,u1_struct_0(esk27_1(g1_orders_2(k5_numbers,esk28_1(k5_numbers))))))).#trainneg
cnf(i_0_8538, plain, (esk17_1(esk22_1(k1_xboole_0))=X2|g1_orders_2(X1,esk17_1(esk22_1(k1_xboole_0)))!=g1_orders_2(X3,X2)|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_649, plain, (l1_orders_2(g1_orders_2(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X3,X4)))|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~m1_subset_1(X4,X2)|~v1_funct_2(X3,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_637, plain, (r1_orders_2(X1,X2,X3)|v1_xboole_0(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))|X3!=k1_funct_1(esk16_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X5)|X2!=k1_funct_1(esk16_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X6)|~r1_orders_2(X4,X6,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|~v5_orders_3(esk16_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),X4,X1)|~v1_funct_2(esk16_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(esk16_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~l1_orders_2(X1)|~l1_orders_2(X4))).#trainneg
cnf(i_0_639, plain, (r1_orders_2(X1,X2,X3)|X3!=k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))),X5)|X2!=k1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))),X6)|~r1_orders_2(X4,X6,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|~v5_orders_3(esk14_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))),X4,X1)|~v1_funct_2(esk14_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(esk14_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))))|~l1_orders_2(X1)|~l1_orders_2(X4))).#trainneg
cnf(i_0_3292, plain, (v1_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)))|~r2_hidden(esk12_1(X2),u1_struct_0(X2))|~m1_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)),X2)|~v4_yellow_0(k5_yellow_0(X1,u1_struct_0(X2)),X2)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_1739, negated_conjecture, (r1_orders_2(X1,k1_funct_1(esk3_0,esk12_1(X2)),k1_funct_1(esk3_0,esk4_3(X2,X3,X4)))|v1_xboole_0(u1_struct_0(X2))|r1_lattice3(X2,X3,X4)|~r2_hidden(k1_funct_1(esk3_0,esk12_1(X2)),u1_struct_0(X1))|~m1_yellow_0(X1,esk2_0)|~v4_yellow_0(X1,esk2_0)|~m1_subset_1(k1_funct_1(esk3_0,esk4_3(X2,X3,X4)),u1_struct_0(X1))|~m1_subset_1(k1_funct_1(esk3_0,esk4_3(X2,X3,X4)),u1_struct_0(esk2_0))|~m1_subset_1(k1_funct_1(esk3_0,esk12_1(X2)),u1_struct_0(esk2_0))|~r2_hidden(esk12_1(X2),u1_struct_0(esk1_0))|~m1_yellow_0(esk1_0,X2)|~v4_yellow_0(esk1_0,X2)|~m1_subset_1(esk4_3(X2,X3,X4),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_4344, plain, (r1_tarski(k10_xtuple_0(X1),k2_relset_1(u1_struct_0(X2),X1))|v2_struct_0(k1_yellow_2(X3,X4,X5))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k10_xtuple_0(X5),u1_struct_0(X2))))|~v1_funct_2(X1,k10_xtuple_0(X5),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_orders_2(X2)|~l1_orders_2(k1_yellow_2(X3,X4,X5))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_orders_2(X4)|~l1_orders_2(X3))).#trainneg
cnf(i_0_12716, plain, (esk17_1(esk22_1(k1_xboole_0))=u1_orders_2(X1)|g1_orders_2(X2,esk17_1(esk22_1(k1_xboole_0)))!=X1|~v7_ordinal1(k1_xboole_0)|~v1_xboole_0(X2)|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_650, plain, (v1_orders_2(g1_orders_2(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X3,X4)))|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))|~m1_subset_1(X4,X2)|~v1_funct_2(X3,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_642, plain, (r1_orders_2(X1,X2,X3)|X3!=k1_funct_1(k1_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),X6)|X2!=k1_funct_1(k1_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),X7)|~r1_orders_2(X4,X7,X6)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X7,u1_struct_0(X4))|~v5_orders_3(k1_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(k1_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5))|~l1_orders_2(X1)|~l1_orders_2(X4)|~v4_relat_1(X5,k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))|~v1_relat_1(X5))).#trainneg
cnf(i_0_643, plain, (r1_orders_2(X1,X2,X3)|X3!=k1_funct_1(k2_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),X6)|X2!=k1_funct_1(k2_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),X7)|~r1_orders_2(X4,X7,X6)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X7,u1_struct_0(X4))|~v5_orders_3(k2_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(k2_relset_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)),X5))|~l1_orders_2(X1)|~l1_orders_2(X4)|~v5_relat_1(X5,k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1)))|~v1_relat_1(X5))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 132
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 249
# Removed in clause preprocessing      : 27
# Initial clauses in saturation        : 222
# Processed clauses                    : 2041
# ...of these trivial                  : 34
# ...subsumed                          : 926
# ...remaining for further processing  : 1081
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 63
# Backward-rewritten                   : 167
# Generated clauses                    : 12917
# ...of the previous two non-trivial   : 11883
# Contextual simplify-reflections      : 240
# Paramodulations                      : 12813
# Factorizations                       : 3
# Equation resolutions                 : 101
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 849
#    Positive orientable unit clauses  : 95
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 16
#    Non-unit-clauses                  : 738
# Current number of unprocessed clauses: 9684
# ...number of literals in the above   : 83893
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1312
# Proof object given clauses           : 117
# Proof search given clauses           : 1081
# Clause-clause subsumption calls (NU) : 618705
# Rec. Clause-clause subsumption calls : 26975
# Non-unit clause-clause subsumptions  : 1126
# Unit Clause-clause subsumption calls : 5635
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 49
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 397954

# -------------------------------------------------
# User time                : 31.147 s
# System time              : 0.697 s
# Total time               : 31.844 s
# Maximum resident set size: 389452 pages
31.14user 0.77system 0:32.84elapsed 97%CPU (0avgtext+0avgdata 389452maxresident)k
0inputs+632outputs (0major+223629minor)pagefaults 0swaps
