2020-05-31 16:07:19.942104: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55
2020-05-31 16:07:20.049734: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 16:07:20.360320: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 16:07:21.138522: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 16:07:23.359565: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3417467 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.713 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(cc13_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>(((((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v1_yellow_0(X1))&v24_waybel_0(X1))=>((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&v3_lattice3(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', cc13_waybel_0)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', cc1_lattice3)).
fof(cc4_waybel_3, axiom, ![X1]:(l1_orders_2(X1)=>(((~(v2_struct_0(X1))&v3_orders_2(X1))&v3_waybel_3(X1))=>(((~(v2_struct_0(X1))&v3_orders_2(X1))&v24_waybel_0(X1))&v2_waybel_3(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', cc4_waybel_3)).
fof(t6_waybel31, conjecture, ![X1]:((((~(v2_struct_0(X1))&v2_pre_topc(X1))&v6_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:(((((((v3_orders_2(X2)&v4_orders_2(X2))&v5_orders_2(X2))&v1_yellow_0(X2))&v1_lattice3(X2))&v3_waybel_3(X2))&l1_orders_2(X2))=>(k2_yellow_1(u1_pre_topc(X1))=X2=>(v8_struct_0(X1)|k2_waybel23(X1)=k1_waybel31(X2))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t6_waybel31)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t1_yellow_1)).
fof(t5_waybel31, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:((((((((v3_orders_2(X2)&v4_orders_2(X2))&v5_orders_2(X2))&v1_yellow_0(X2))&v1_lattice3(X2))&v2_lattice3(X2))&v3_waybel_3(X2))&l1_orders_2(X2))=>(k2_yellow_1(u1_pre_topc(X1))=X2=>![X3]:(((v1_tops_2(X3,X1)&v1_cantor_1(X3,X1))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>(v6_waybel23(k11_waybel_0(X2,X4),X2)&m1_waybel23(k11_waybel_0(X2,X4),X2)))))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t5_waybel31)).
fof(cc1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>((~(v2_struct_0(X1))&v3_lattice3(X1))=>((~(v2_struct_0(X1))&v1_lattice3(X1))&v2_lattice3(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', cc1_yellow_0)).
fof(t4_waybel31, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:((((((v3_orders_2(X2)&v4_orders_2(X2))&v5_orders_2(X2))&v1_lattice3(X2))&v3_waybel_3(X2))&l1_orders_2(X2))=>(k2_yellow_1(u1_pre_topc(X1))=X2=>![X3]:((v6_waybel23(X3,X2)&m1_waybel23(X3,X2))=>((v1_tops_2(X3,X1)&v1_cantor_1(X3,X1))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t4_waybel31)).
fof(fraenkel_a_1_0_waybel23, axiom, ![X1, X2]:(l1_pre_topc(X2)=>(r2_hidden(X1,a_1_0_waybel23(X2))<=>?[X3]:(((v1_cantor_1(X3,X2)&v1_tops_2(X3,X2))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))&X1=k1_card_1(X3)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fraenkel_a_1_0_waybel23)).
fof(fraenkel_a_1_2_waybel31, axiom, ![X1, X2]:((((~(v2_struct_0(X2))&v2_pre_topc(X2))&v6_pre_topc(X2))&l1_pre_topc(X2))=>(r2_hidden(X1,a_1_2_waybel31(X2))<=>?[X3]:(((v1_tops_2(X3,X2)&v1_cantor_1(X3,X2))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))&X1=k1_card_1(X3)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fraenkel_a_1_2_waybel31)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', dt_k2_yellow_1)).
fof(fc5_yellow_1, axiom, ![X1]:(((v1_orders_2(k2_yellow_1(X1))&v3_orders_2(k2_yellow_1(X1)))&v4_orders_2(k2_yellow_1(X1)))&v5_orders_2(k2_yellow_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fc5_yellow_1)).
fof(t64_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>r1_tarski(X2,u1_pre_topc(X1))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t64_tops_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t3_subset)).
fof(fc9_yellow_1, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_pre_topc(X1))&l1_pre_topc(X1))=>((~(v7_struct_0(k2_yellow_1(u1_pre_topc(X1))))&v1_orders_2(k2_yellow_1(u1_pre_topc(X1))))&v3_lattice3(k2_yellow_1(u1_pre_topc(X1))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fc9_yellow_1)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', dt_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', cc1_subset_1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', redefinition_m2_subset_1)).
fof(t27_yellow15, axiom, ![X1]:((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_yellow_0(X1))&v1_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(~(v1_finset_1(X2))=>k1_card_1(X2)=k1_card_1(k11_waybel_0(X1,X2))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t27_yellow15)).
fof(dt_m1_waybel23, axiom, ![X1]:((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v3_waybel_3(X1))&l1_orders_2(X1))=>![X2]:(m1_waybel23(X2,X1)=>m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', dt_m1_waybel23)).
fof(rc4_waybel23, axiom, ![X1]:((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v3_waybel_3(X1))&l1_orders_2(X1))=>?[X2]:(m1_waybel23(X2,X1)&v6_waybel23(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', rc4_waybel23)).
fof(cc5_waybel23, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v6_waybel23(X2,X1)=>~(v1_xboole_0(X2))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', cc5_waybel23)).
fof(dt_k11_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', dt_k11_waybel_0)).
fof(t30_yellow15, axiom, ![X1]:(((v2_pre_topc(X1)&v6_pre_topc(X1))&l1_pre_topc(X1))=>(~(v8_struct_0(X1))=>![X2]:(((v1_tops_2(X2,X1)&v1_cantor_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>~(v1_finset_1(X2))))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t30_yellow15)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t1_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', d3_tarski)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t4_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t2_subset)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fc1_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', dt_l1_orders_2)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', t7_boole)).
fof(fraenkel_a_1_3_waybel31, axiom, ![X1, X2]:(((((((v3_orders_2(X2)&v4_orders_2(X2))&v5_orders_2(X2))&v1_yellow_0(X2))&v1_lattice3(X2))&v3_waybel_3(X2))&l1_orders_2(X2))=>(r2_hidden(X1,a_1_3_waybel31(X2))<=>?[X3]:((v6_waybel23(X3,X2)&m1_waybel23(X3,X2))&X1=k1_card_1(X3)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fraenkel_a_1_3_waybel31)).
fof(fc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>~(v1_xboole_0(u1_pre_topc(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fc1_pre_topc)).
fof(fc8_card_1, axiom, ![X1]:(~(v1_finset_1(X1))=>(~(v1_finset_1(k1_card_1(X1)))&v1_card_1(k1_card_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fc8_card_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', d10_xboole_0)).
fof(fc6_card_1, axiom, ![X1]:(v1_finset_1(X1)=>(v1_finset_1(k1_card_1(X1))&v1_card_1(k1_card_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fc6_card_1)).
fof(fraenkel_a_1_0_waybel31, axiom, ![X1, X2]:((((((v3_orders_2(X2)&v4_orders_2(X2))&v5_orders_2(X2))&v1_lattice3(X2))&v3_waybel_3(X2))&l1_orders_2(X2))=>(r2_hidden(X1,a_1_0_waybel31(X2))<=>?[X3]:((v6_waybel23(X3,X2)&m1_waybel23(X3,X2))&X1=k1_card_1(X3)))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', fraenkel_a_1_0_waybel31)).
fof(d1_waybel31, axiom, ![X1]:((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v1_lattice3(X1))&v3_waybel_3(X1))&l1_orders_2(X1))=>k1_waybel31(X1)=k1_setfam_1(a_1_0_waybel31(X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', d1_waybel31)).
fof(d5_waybel23, axiom, ![X1]:(l1_pre_topc(X1)=>k2_waybel23(X1)=k1_setfam_1(a_1_0_waybel23(X1))), file('/scratch/mptp/allpred/preds__0.005/t6_waybel31', d5_waybel23)).
fof(c_0_39, plain, ![X69]:(((((~v2_struct_0(X69)|(v2_struct_0(X69)|~v3_orders_2(X69)|~v4_orders_2(X69)|~v5_orders_2(X69)|~v1_lattice3(X69)|~v1_yellow_0(X69)|~v24_waybel_0(X69))|~l1_orders_2(X69))&(v3_orders_2(X69)|(v2_struct_0(X69)|~v3_orders_2(X69)|~v4_orders_2(X69)|~v5_orders_2(X69)|~v1_lattice3(X69)|~v1_yellow_0(X69)|~v24_waybel_0(X69))|~l1_orders_2(X69)))&(v4_orders_2(X69)|(v2_struct_0(X69)|~v3_orders_2(X69)|~v4_orders_2(X69)|~v5_orders_2(X69)|~v1_lattice3(X69)|~v1_yellow_0(X69)|~v24_waybel_0(X69))|~l1_orders_2(X69)))&(v5_orders_2(X69)|(v2_struct_0(X69)|~v3_orders_2(X69)|~v4_orders_2(X69)|~v5_orders_2(X69)|~v1_lattice3(X69)|~v1_yellow_0(X69)|~v24_waybel_0(X69))|~l1_orders_2(X69)))&(v3_lattice3(X69)|(v2_struct_0(X69)|~v3_orders_2(X69)|~v4_orders_2(X69)|~v5_orders_2(X69)|~v1_lattice3(X69)|~v1_yellow_0(X69)|~v24_waybel_0(X69))|~l1_orders_2(X69))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc13_waybel_0])])])])).
fof(c_0_40, plain, ![X47]:(~l1_orders_2(X47)|(~v1_lattice3(X47)|~v2_struct_0(X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_41, plain, (v3_lattice3(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v24_waybel_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_42, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_43, plain, ![X65]:((((~v2_struct_0(X65)|(v2_struct_0(X65)|~v3_orders_2(X65)|~v3_waybel_3(X65))|~l1_orders_2(X65))&(v3_orders_2(X65)|(v2_struct_0(X65)|~v3_orders_2(X65)|~v3_waybel_3(X65))|~l1_orders_2(X65)))&(v24_waybel_0(X65)|(v2_struct_0(X65)|~v3_orders_2(X65)|~v3_waybel_3(X65))|~l1_orders_2(X65)))&(v2_waybel_3(X65)|(v2_struct_0(X65)|~v3_orders_2(X65)|~v3_waybel_3(X65))|~l1_orders_2(X65))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc4_waybel_3])])])])).
fof(c_0_44, negated_conjecture, ~(![X1]:((((~(v2_struct_0(X1))&v2_pre_topc(X1))&v6_pre_topc(X1))&l1_pre_topc(X1))=>![X2]:(((((((v3_orders_2(X2)&v4_orders_2(X2))&v5_orders_2(X2))&v1_yellow_0(X2))&v1_lattice3(X2))&v3_waybel_3(X2))&l1_orders_2(X2))=>(k2_yellow_1(u1_pre_topc(X1))=X2=>(v8_struct_0(X1)|k2_waybel23(X1)=k1_waybel31(X2)))))), inference(assume_negation,[status(cth)],[t6_waybel31])).
cnf(c_0_45, plain, (v2_struct_0(X1)|v3_lattice3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v24_waybel_0(X1)), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (~v2_struct_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_47, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_48, plain, ![X14]:(u1_struct_0(k2_yellow_1(X14))=X14&u1_orders_2(k2_yellow_1(X14))=k1_yellow_1(X14)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
fof(c_0_49, negated_conjecture, ((((~v2_struct_0(esk1_0)&v2_pre_topc(esk1_0))&v6_pre_topc(esk1_0))&l1_pre_topc(esk1_0))&(((((((v3_orders_2(esk2_0)&v4_orders_2(esk2_0))&v5_orders_2(esk2_0))&v1_yellow_0(esk2_0))&v1_lattice3(esk2_0))&v3_waybel_3(esk2_0))&l1_orders_2(esk2_0))&(k2_yellow_1(u1_pre_topc(esk1_0))=esk2_0&(~v8_struct_0(esk1_0)&k2_waybel23(esk1_0)!=k1_waybel31(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_44])])])])).
fof(c_0_50, plain, ![X7, X8, X9, X10]:((v6_waybel23(k11_waybel_0(X8,X10),X8)|X9!=X10|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X8)))|(~v1_tops_2(X9,X7)|~v1_cantor_1(X9,X7)|~m1_subset_1(X9,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X7)))))|k2_yellow_1(u1_pre_topc(X7))!=X8|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_yellow_0(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v3_waybel_3(X8)|~l1_orders_2(X8))|(v2_struct_0(X7)|~v2_pre_topc(X7)|~l1_pre_topc(X7)))&(m1_waybel23(k11_waybel_0(X8,X10),X8)|X9!=X10|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X8)))|(~v1_tops_2(X9,X7)|~v1_cantor_1(X9,X7)|~m1_subset_1(X9,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X7)))))|k2_yellow_1(u1_pre_topc(X7))!=X8|(~v3_orders_2(X8)|~v4_orders_2(X8)|~v5_orders_2(X8)|~v1_yellow_0(X8)|~v1_lattice3(X8)|~v2_lattice3(X8)|~v3_waybel_3(X8)|~l1_orders_2(X8))|(v2_struct_0(X7)|~v2_pre_topc(X7)|~l1_pre_topc(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_waybel31])])])])])).
fof(c_0_51, plain, ![X50]:(((~v2_struct_0(X50)|(v2_struct_0(X50)|~v3_lattice3(X50))|~l1_orders_2(X50))&(v1_lattice3(X50)|(v2_struct_0(X50)|~v3_lattice3(X50))|~l1_orders_2(X50)))&(v2_lattice3(X50)|(v2_struct_0(X50)|~v3_lattice3(X50))|~l1_orders_2(X50))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_yellow_0])])])])).
cnf(c_0_52, plain, (v3_lattice3(X1)|~v24_waybel_0(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_47])).
fof(c_0_54, plain, ![X11, X12, X13]:(((v1_tops_2(X13,X11)|(~v6_waybel23(X13,X12)|~m1_waybel23(X13,X12))|k2_yellow_1(u1_pre_topc(X11))!=X12|(~v3_orders_2(X12)|~v4_orders_2(X12)|~v5_orders_2(X12)|~v1_lattice3(X12)|~v3_waybel_3(X12)|~l1_orders_2(X12))|(v2_struct_0(X11)|~v2_pre_topc(X11)|~l1_pre_topc(X11)))&(v1_cantor_1(X13,X11)|(~v6_waybel23(X13,X12)|~m1_waybel23(X13,X12))|k2_yellow_1(u1_pre_topc(X11))!=X12|(~v3_orders_2(X12)|~v4_orders_2(X12)|~v5_orders_2(X12)|~v1_lattice3(X12)|~v3_waybel_3(X12)|~l1_orders_2(X12))|(v2_struct_0(X11)|~v2_pre_topc(X11)|~l1_pre_topc(X11))))&(m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X11))))|(~v6_waybel23(X13,X12)|~m1_waybel23(X13,X12))|k2_yellow_1(u1_pre_topc(X11))!=X12|(~v3_orders_2(X12)|~v4_orders_2(X12)|~v5_orders_2(X12)|~v1_lattice3(X12)|~v3_waybel_3(X12)|~l1_orders_2(X12))|(v2_struct_0(X11)|~v2_pre_topc(X11)|~l1_pre_topc(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_waybel31])])])])])).
cnf(c_0_55, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, negated_conjecture, (k2_yellow_1(u1_pre_topc(esk1_0))=esk2_0), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, plain, (v6_waybel23(k11_waybel_0(X1,X2),X1)|v2_struct_0(X4)|X3!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(X3,X4)|~v1_cantor_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|k2_yellow_1(u1_pre_topc(X4))!=X1|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, plain, (v3_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v3_waybel_3(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53])])).
cnf(c_0_60, plain, (v1_cantor_1(X1,X2)|v2_struct_0(X2)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3)|k2_yellow_1(u1_pre_topc(X2))!=X3|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_62, negated_conjecture, (k2_yellow_1(u1_pre_topc(esk1_0))=esk2_0), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_63, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_64, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_65, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_66, plain, (v2_struct_0(X4)|v6_waybel23(k11_waybel_0(X1,X2),X1)|X2!=X3|k2_yellow_1(u1_pre_topc(X4))!=X1|~v2_pre_topc(X4)|~l1_pre_topc(X4)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_tops_2(X3,X4)|~v1_cantor_1(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_67, plain, (v2_struct_0(X1)|v2_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (v3_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_59, c_0_46])).
cnf(c_0_69, plain, (m1_waybel23(k11_waybel_0(X1,X2),X1)|v2_struct_0(X4)|X3!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(X3,X4)|~v1_cantor_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|k2_yellow_1(u1_pre_topc(X4))!=X1|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_70, plain, (v2_struct_0(X2)|v1_cantor_1(X1,X2)|k2_yellow_1(u1_pre_topc(X2))!=X3|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3)), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_71, negated_conjecture, (u1_pre_topc(esk1_0)=u1_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_72, negated_conjecture, (l1_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_73, negated_conjecture, (v2_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_74, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_75, plain, (v6_waybel23(k11_waybel_0(X1,X2),X1)|v2_struct_0(X3)|k2_yellow_1(u1_pre_topc(X3))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,X3)|~v1_tops_2(X2,X3)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_66])])).
cnf(c_0_76, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_77, plain, (v2_struct_0(X4)|m1_waybel23(k11_waybel_0(X1,X2),X1)|X2!=X3|k2_yellow_1(u1_pre_topc(X4))!=X1|~v2_pre_topc(X4)|~l1_pre_topc(X4)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_tops_2(X3,X4)|~v1_cantor_1(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_78, plain, (v1_tops_2(X1,X2)|v2_struct_0(X2)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3)|k2_yellow_1(u1_pre_topc(X2))!=X3|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_79, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3)|k2_yellow_1(u1_pre_topc(X2))!=X3|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_80, negated_conjecture, (v1_cantor_1(X1,esk1_0)|k2_yellow_1(u1_struct_0(esk2_0))!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_81, negated_conjecture, (k2_yellow_1(u1_struct_0(esk2_0))=esk2_0), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_71])])).
cnf(c_0_82, negated_conjecture, (v6_waybel23(k11_waybel_0(X1,X2),X1)|k2_yellow_1(u1_struct_0(esk2_0))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_83, plain, (v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_76, c_0_46])).
cnf(c_0_84, plain, (m1_waybel23(k11_waybel_0(X1,X2),X1)|v2_struct_0(X3)|k2_yellow_1(u1_pre_topc(X3))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,X3)|~v1_tops_2(X2,X3)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_77])])).
cnf(c_0_85, plain, (v2_struct_0(X2)|v1_tops_2(X1,X2)|k2_yellow_1(u1_pre_topc(X2))!=X3|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_86, plain, (v2_struct_0(X2)|m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|k2_yellow_1(u1_pre_topc(X2))!=X3|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3)), inference(evalgc,[status(thm)],[c_0_79])).
fof(c_0_87, plain, ![X52, X53, X55]:(((((v1_cantor_1(esk5_2(X52,X53),X53)|~r2_hidden(X52,a_1_0_waybel23(X53))|~l1_pre_topc(X53))&(v1_tops_2(esk5_2(X52,X53),X53)|~r2_hidden(X52,a_1_0_waybel23(X53))|~l1_pre_topc(X53)))&(m1_subset_1(esk5_2(X52,X53),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X53))))|~r2_hidden(X52,a_1_0_waybel23(X53))|~l1_pre_topc(X53)))&(X52=k1_card_1(esk5_2(X52,X53))|~r2_hidden(X52,a_1_0_waybel23(X53))|~l1_pre_topc(X53)))&(~v1_cantor_1(X55,X53)|~v1_tops_2(X55,X53)|~m1_subset_1(X55,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X53))))|X52!=k1_card_1(X55)|r2_hidden(X52,a_1_0_waybel23(X53))|~l1_pre_topc(X53))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_1_0_waybel23])])])])])).
fof(c_0_88, plain, ![X56, X57, X59]:(((((v1_tops_2(esk6_2(X56,X57),X57)|~r2_hidden(X56,a_1_2_waybel31(X57))|(v2_struct_0(X57)|~v2_pre_topc(X57)|~v6_pre_topc(X57)|~l1_pre_topc(X57)))&(v1_cantor_1(esk6_2(X56,X57),X57)|~r2_hidden(X56,a_1_2_waybel31(X57))|(v2_struct_0(X57)|~v2_pre_topc(X57)|~v6_pre_topc(X57)|~l1_pre_topc(X57))))&(m1_subset_1(esk6_2(X56,X57),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X57))))|~r2_hidden(X56,a_1_2_waybel31(X57))|(v2_struct_0(X57)|~v2_pre_topc(X57)|~v6_pre_topc(X57)|~l1_pre_topc(X57))))&(X56=k1_card_1(esk6_2(X56,X57))|~r2_hidden(X56,a_1_2_waybel31(X57))|(v2_struct_0(X57)|~v2_pre_topc(X57)|~v6_pre_topc(X57)|~l1_pre_topc(X57))))&(~v1_tops_2(X59,X57)|~v1_cantor_1(X59,X57)|~m1_subset_1(X59,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X57))))|X56!=k1_card_1(X59)|r2_hidden(X56,a_1_2_waybel31(X57))|(v2_struct_0(X57)|~v2_pre_topc(X57)|~v6_pre_topc(X57)|~l1_pre_topc(X57)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_1_2_waybel31])])])])])])).
cnf(c_0_89, negated_conjecture, (v1_cantor_1(X1,esk1_0)|esk2_0!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_90, negated_conjecture, (v6_waybel23(k11_waybel_0(X1,X2),X1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_81]), c_0_83])).
cnf(c_0_91, negated_conjecture, (m1_waybel23(k11_waybel_0(X1,X2),X1)|k2_yellow_1(u1_struct_0(esk2_0))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_92, negated_conjecture, (v1_tops_2(X1,esk1_0)|k2_yellow_1(u1_struct_0(esk2_0))!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_93, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|k2_yellow_1(u1_struct_0(esk2_0))!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_94, plain, (r2_hidden(X3,a_1_0_waybel23(X2))|~v1_cantor_1(X1,X2)|~v1_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|X3!=k1_card_1(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_95, plain, (v1_cantor_1(esk6_2(X1,X2),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_96, plain, (r2_hidden(X3,a_1_2_waybel31(X2))|v2_struct_0(X2)|~v1_tops_2(X1,X2)|~v1_cantor_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|X3!=k1_card_1(X1)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_97, negated_conjecture, (v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_waybel23(k11_waybel_0(X1,X2),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_98, negated_conjecture, (m1_waybel23(k11_waybel_0(X1,X2),X1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_81]), c_0_83])).
cnf(c_0_99, negated_conjecture, (v6_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_100, negated_conjecture, (v1_tops_2(X1,esk1_0)|esk2_0!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(rw,[status(thm)],[c_0_92, c_0_81])).
cnf(c_0_101, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(rw,[status(thm)],[c_0_93, c_0_81])).
cnf(c_0_102, plain, (r2_hidden(X3,a_1_0_waybel23(X2))|X3!=k1_card_1(X1)|~l1_pre_topc(X2)|~v1_tops_2(X1,X2)|~v1_cantor_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_103, plain, (v2_struct_0(X2)|v1_cantor_1(esk6_2(X1,X2),X2)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_104, plain, (v1_tops_2(esk6_2(X1,X2),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_105, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_106, plain, (v2_struct_0(X2)|r2_hidden(X3,a_1_2_waybel31(X2))|X3!=k1_card_1(X1)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~v1_tops_2(X1,X2)|~v1_cantor_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_107, negated_conjecture, (v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_108, negated_conjecture, (v6_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_109, negated_conjecture, (v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_waybel23(k11_waybel_0(X1,X2),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_90])])).
cnf(c_0_110, negated_conjecture, (m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X1|~m1_waybel23(k11_waybel_0(X1,X2),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_90])])).
fof(c_0_111, plain, ![X25]:(v1_orders_2(k2_yellow_1(X25))&l1_orders_2(k2_yellow_1(X25))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_112, plain, ![X36]:(((v1_orders_2(k2_yellow_1(X36))&v3_orders_2(k2_yellow_1(X36)))&v4_orders_2(k2_yellow_1(X36)))&v5_orders_2(k2_yellow_1(X36))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_113, plain, ![X23, X24]:((~v1_tops_2(X24,X23)|r1_tarski(X24,u1_pre_topc(X23))|~m1_subset_1(X24,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X23))))|~l1_pre_topc(X23))&(~r1_tarski(X24,u1_pre_topc(X23))|v1_tops_2(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X23))))|~l1_pre_topc(X23))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_tops_2])])])])).
fof(c_0_114, plain, ![X37, X38]:((~m1_subset_1(X37,k1_zfmisc_1(X38))|r1_tarski(X37,X38))&(~r1_tarski(X37,X38)|m1_subset_1(X37,k1_zfmisc_1(X38)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_115, plain, ![X49]:(((~v7_struct_0(k2_yellow_1(u1_pre_topc(X49)))|(v2_struct_0(X49)|~v2_pre_topc(X49)|~l1_pre_topc(X49)))&(v1_orders_2(k2_yellow_1(u1_pre_topc(X49)))|(v2_struct_0(X49)|~v2_pre_topc(X49)|~l1_pre_topc(X49))))&(v3_lattice3(k2_yellow_1(u1_pre_topc(X49)))|(v2_struct_0(X49)|~v2_pre_topc(X49)|~l1_pre_topc(X49)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc9_yellow_1])])])])).
cnf(c_0_116, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk6_2(X3,X2))|~m1_subset_1(esk6_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk6_2(X3,X2),X2)|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_117, plain, (v2_struct_0(X2)|v1_tops_2(esk6_2(X1,X2),X2)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_118, plain, (v2_struct_0(X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_119, plain, (X1=k1_card_1(esk6_2(X1,X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_120, negated_conjecture, (r2_hidden(X1,a_1_2_waybel31(esk1_0))|X1!=k1_card_1(k11_waybel_0(X2,X3))|~m1_subset_1(k11_waybel_0(X2,X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(k11_waybel_0(X2,X3),esk1_0)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_121, negated_conjecture, (v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_109, c_0_98])).
cnf(c_0_122, negated_conjecture, (m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_110, c_0_98])).
cnf(c_0_123, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_124, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_125, plain, (v4_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_126, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_127, plain, (v1_tops_2(X1,X2)|~r1_tarski(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_128, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_129, plain, (v3_lattice3(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_115])).
fof(c_0_130, plain, ![X152, X153, X154]:(v1_xboole_0(X152)|v1_xboole_0(X153)|~m1_subset_1(X153,k1_zfmisc_1(X152))|(~m2_subset_1(X154,X152,X153)|m1_subset_1(X154,X152))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])).
fof(c_0_131, plain, ![X95, X96]:(~v1_xboole_0(X95)|(~m1_subset_1(X96,k1_zfmisc_1(X95))|v1_xboole_0(X96))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_132, plain, ![X110, X111, X112]:((~m2_subset_1(X112,X110,X111)|m1_subset_1(X112,X111)|(v1_xboole_0(X110)|v1_xboole_0(X111)|~m1_subset_1(X111,k1_zfmisc_1(X110))))&(~m1_subset_1(X112,X111)|m2_subset_1(X112,X110,X111)|(v1_xboole_0(X110)|v1_xboole_0(X111)|~m1_subset_1(X111,k1_zfmisc_1(X110))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_133, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk6_2(X3,X2))|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_116, c_0_117]), c_0_118])).
cnf(c_0_134, plain, (k1_card_1(esk6_2(X1,X2))=X1|v2_struct_0(X2)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_135, negated_conjecture, (r2_hidden(X1,a_1_2_waybel31(esk1_0))|X1!=k1_card_1(k11_waybel_0(X2,X3))|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_120, c_0_121]), c_0_122])).
fof(c_0_136, plain, ![X26, X27]:(~v3_orders_2(X26)|~v4_orders_2(X26)|~v5_orders_2(X26)|~v1_yellow_0(X26)|~v1_lattice3(X26)|~l1_orders_2(X26)|(~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|(v1_finset_1(X27)|k1_card_1(X27)=k1_card_1(k11_waybel_0(X26,X27))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_yellow15])])])])).
cnf(c_0_137, plain, (l1_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_123])).
cnf(c_0_138, plain, (v5_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_139, plain, (v4_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_140, plain, (v3_orders_2(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_141, plain, (v1_tops_2(X1,X2)|~l1_pre_topc(X2)|~r1_tarski(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_142, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_143, plain, (v2_struct_0(X1)|v3_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_144, negated_conjecture, (v1_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_145, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_146, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_147, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_148, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_132])).
fof(c_0_149, plain, ![X127, X128]:(~v3_orders_2(X127)|~v4_orders_2(X127)|~v5_orders_2(X127)|~v1_lattice3(X127)|~v3_waybel_3(X127)|~l1_orders_2(X127)|(~m1_waybel23(X128,X127)|m1_subset_1(X128,k1_zfmisc_1(u1_struct_0(X127))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_waybel23])])])).
fof(c_0_150, plain, ![X163]:((m1_waybel23(esk11_1(X163),X163)|(~v3_orders_2(X163)|~v4_orders_2(X163)|~v5_orders_2(X163)|~v1_lattice3(X163)|~v3_waybel_3(X163)|~l1_orders_2(X163)))&(v6_waybel23(esk11_1(X163),X163)|(~v3_orders_2(X163)|~v4_orders_2(X163)|~v5_orders_2(X163)|~v1_lattice3(X163)|~v3_waybel_3(X163)|~l1_orders_2(X163)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rc4_waybel23])])])])).
fof(c_0_151, plain, ![X91, X92]:(v2_struct_0(X91)|~l1_orders_2(X91)|(~m1_subset_1(X92,k1_zfmisc_1(u1_struct_0(X91)))|(~v6_waybel23(X92,X91)|~v1_xboole_0(X92)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_waybel23])])])])).
cnf(c_0_152, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134])])])).
cnf(c_0_153, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_2_waybel31(esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_135])])).
cnf(c_0_154, plain, (v1_finset_1(X2)|k1_card_1(X2)=k1_card_1(k11_waybel_0(X1,X2))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_136])).
fof(c_0_155, plain, ![X61, X62]:(v2_struct_0(X61)|~l1_orders_2(X61)|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|m1_subset_1(k11_waybel_0(X61,X62),k1_zfmisc_1(u1_struct_0(X61)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k11_waybel_0])])])).
cnf(c_0_156, plain, (v6_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v1_tops_2(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_75]), c_0_61]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_157, plain, (v1_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_pre_topc(X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_158, negated_conjecture, (v3_waybel_3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_159, negated_conjecture, (v1_yellow_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_160, plain, (v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_143]), c_0_137])])])).
cnf(c_0_161, negated_conjecture, (v1_lattice3(esk2_0)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_162, negated_conjecture, (l1_orders_2(esk2_0)), inference(evalgc,[status(thm)],[c_0_145])).
fof(c_0_163, plain, ![X33, X34]:(~v2_pre_topc(X33)|~v6_pre_topc(X33)|~l1_pre_topc(X33)|(v8_struct_0(X33)|(~v1_tops_2(X34,X33)|~v1_cantor_1(X34,X33)|~m1_subset_1(X34,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X33))))|~v1_finset_1(X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_yellow15])])])])).
cnf(c_0_164, plain, (v1_cantor_1(X1,X2)|v2_struct_0(X2)|~m1_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v6_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_70]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_165, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_166, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_167, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_168, plain, (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~m1_waybel23(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_169, plain, (m1_waybel23(esk11_1(X1),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_170, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v6_waybel23(X2,X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_171, plain, (v6_waybel23(esk11_1(X1),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_172, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_173, plain, (k1_card_1(k11_waybel_0(X1,X2))=k1_card_1(X2)|v1_finset_1(X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_174, plain, (v2_struct_0(X1)|m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_175, plain, (v6_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_176, negated_conjecture, (v3_waybel_3(esk2_0)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_177, negated_conjecture, (v1_yellow_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_178, negated_conjecture, (v2_lattice3(k2_yellow_1(u1_struct_0(esk2_0)))|v2_struct_0(k2_yellow_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_179, negated_conjecture, (~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_161]), c_0_162])])])).
cnf(c_0_180, plain, (v8_struct_0(X1)|~v2_pre_topc(X1)|~v6_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_181, negated_conjecture, (v1_cantor_1(X1,esk1_0)|~m1_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v6_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v3_waybel_3(k2_yellow_1(u1_struct_0(esk2_0)))|~v1_lattice3(k2_yellow_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_182, negated_conjecture, (~v8_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_183, plain, (v1_tops_2(X1,X2)|v2_struct_0(X2)|~m1_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v6_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_85]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_184, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~m1_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v6_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_86]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_185, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_186, plain, (m2_subset_1(X1,X2,X3)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[c_0_167, c_0_166])).
cnf(c_0_187, plain, (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~m1_waybel23(X2,X1)), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_188, plain, (m1_waybel23(esk11_1(X1),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_189, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~v1_xboole_0(X2)|~v6_waybel23(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_190, plain, (v6_waybel23(esk11_1(X1),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_171])).
fof(c_0_191, plain, ![X71, X72]:(~r2_hidden(X71,X72)|m1_subset_1(X71,X72)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_192, plain, ![X17, X18, X19, X20, X21]:((~r1_tarski(X17,X18)|(~r2_hidden(X19,X17)|r2_hidden(X19,X18)))&((r2_hidden(esk3_2(X20,X21),X20)|r1_tarski(X20,X21))&(~r2_hidden(esk3_2(X20,X21),X21)|r1_tarski(X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_193, plain, ![X97, X98, X99]:(~r2_hidden(X97,X98)|~m1_subset_1(X98,k1_zfmisc_1(X99))|m1_subset_1(X97,X99)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_194, negated_conjecture, (r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173])])).
cnf(c_0_195, plain, (v2_struct_0(X1)|m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_196, negated_conjecture, (v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v2_lattice3(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_71]), c_0_81]), c_0_81]), c_0_81]), c_0_81]), c_0_176]), c_0_81]), c_0_161]), c_0_81]), c_0_177]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_197, negated_conjecture, (v2_lattice3(esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_178, c_0_81]), c_0_81]), c_0_179])).
cnf(c_0_198, negated_conjecture, (v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_199, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_200, negated_conjecture, (v3_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_201, plain, (m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v1_tops_2(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_84]), c_0_61]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_202, plain, (v8_struct_0(X1)|~v2_pre_topc(X1)|~v6_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_203, negated_conjecture, (v1_cantor_1(X1,esk1_0)|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_181, c_0_81]), c_0_81]), c_0_81]), c_0_176]), c_0_81]), c_0_161])])).
cnf(c_0_204, negated_conjecture, (~v8_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_205, negated_conjecture, (v1_tops_2(X1,esk1_0)|~m1_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v6_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v3_waybel_3(k2_yellow_1(u1_struct_0(esk2_0)))|~v1_lattice3(k2_yellow_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_206, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v6_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v3_waybel_3(k2_yellow_1(u1_struct_0(esk2_0)))|~v1_lattice3(k2_yellow_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_71]), c_0_72]), c_0_73])]), c_0_74])])).
fof(c_0_207, plain, ![X101, X102]:(~m1_subset_1(X101,X102)|(v1_xboole_0(X102)|r2_hidden(X101,X102))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_208, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186])])).
cnf(c_0_209, plain, (m1_subset_1(esk11_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_210, plain, (v2_struct_0(X1)|~v1_xboole_0(esk11_1(X1))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190])])).
cnf(c_0_211, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_191])).
cnf(c_0_212, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_213, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_214, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel23(esk1_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_195])])).
cnf(c_0_215, negated_conjecture, (v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_196, c_0_197])])).
cnf(c_0_216, negated_conjecture, (v5_orders_2(esk2_0)), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_217, negated_conjecture, (v4_orders_2(esk2_0)), inference(evalgc,[status(thm)],[c_0_199])).
cnf(c_0_218, negated_conjecture, (v3_orders_2(esk2_0)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_219, plain, (m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_201, c_0_157])).
cnf(c_0_220, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(X1,esk1_0)|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203]), c_0_72]), c_0_108]), c_0_73])]), c_0_204])])).
cnf(c_0_221, negated_conjecture, (v1_tops_2(X1,esk1_0)|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_205, c_0_81]), c_0_81]), c_0_81]), c_0_176]), c_0_81]), c_0_161])])).
cnf(c_0_222, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_206, c_0_81]), c_0_81]), c_0_81]), c_0_176]), c_0_81]), c_0_161])])).
cnf(c_0_223, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_224, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_225, plain, (v1_xboole_0(esk11_1(X1))|m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,esk11_1(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_226, plain, (~v1_xboole_0(esk11_1(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_210, c_0_209]), c_0_46])).
cnf(c_0_227, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_228, plain, (r1_tarski(X1,X2)|r2_hidden(esk3_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_229, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_230, plain, (r1_tarski(X1,u1_pre_topc(X2))|~v1_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_231, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_213])).
fof(c_0_232, plain, ![X94]:(~v2_struct_0(X94)|~l1_struct_0(X94)|v1_xboole_0(u1_struct_0(X94))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
fof(c_0_233, plain, ![X51]:(~l1_orders_2(X51)|l1_struct_0(X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_234, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_214, c_0_46])).
cnf(c_0_235, negated_conjecture, (v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_215]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_236, negated_conjecture, (m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_71]), c_0_81]), c_0_81]), c_0_81]), c_0_197]), c_0_81]), c_0_176]), c_0_81]), c_0_161]), c_0_81]), c_0_177]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_237, negated_conjecture, (~v1_finset_1(X1)|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_220, c_0_221]), c_0_222])).
cnf(c_0_238, negated_conjecture, (v1_tops_2(X1,esk1_0)|~r1_tarski(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_71]), c_0_72])])])).
cnf(c_0_239, plain, (m1_subset_1(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v2_struct_0(X1)|~m1_waybel23(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_190]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_240, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_241, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_242, plain, (m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X1,esk11_1(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_225, c_0_226])).
cnf(c_0_243, plain, (m1_subset_1(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228])])).
cnf(c_0_244, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_245, plain, (r1_tarski(X1,u1_pre_topc(X2))|~l1_pre_topc(X2)|~v1_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_246, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k11_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_195])])).
cnf(c_0_247, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_248, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_233])).
cnf(c_0_249, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_173])])).
cnf(c_0_250, negated_conjecture, (v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(csr,[status(thm)],[c_0_235, c_0_236])).
cnf(c_0_251, negated_conjecture, (v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_215]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_252, negated_conjecture, (m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_215]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_253, negated_conjecture, (~v1_finset_1(k11_waybel_0(esk2_0,X1))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_90]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_254, negated_conjecture, (v1_tops_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_142])])).
fof(c_0_255, plain, ![X75, X76]:(~r2_hidden(X75,X76)|~v1_xboole_0(X76)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_256, plain, (m1_subset_1(esk11_1(k2_yellow_1(X1)),k1_zfmisc_1(X1))|~v3_waybel_3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_61]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_257, negated_conjecture, (m1_subset_1(esk11_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(esk11_1(esk2_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_71]), c_0_81]), c_0_81]), c_0_81]), c_0_81]), c_0_176]), c_0_81]), c_0_161]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_258, negated_conjecture, (~v1_finset_1(esk11_1(esk2_0))|~m1_waybel23(esk11_1(esk2_0),esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_190]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_259, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk3_2(X1,X2),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241])])).
cnf(c_0_260, plain, (m1_subset_1(esk3_2(esk11_1(X1),X2),u1_struct_0(X1))|r1_tarski(esk11_1(X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243])])).
cnf(c_0_261, plain, (v1_xboole_0(esk11_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_209])])).
cnf(c_0_262, plain, (r2_hidden(X1,u1_pre_topc(X2))|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(X3,X2)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245])])).
fof(c_0_263, plain, ![X42, X43, X45]:((((v6_waybel23(esk4_2(X42,X43),X43)|~r2_hidden(X42,a_1_3_waybel31(X43))|(~v3_orders_2(X43)|~v4_orders_2(X43)|~v5_orders_2(X43)|~v1_yellow_0(X43)|~v1_lattice3(X43)|~v3_waybel_3(X43)|~l1_orders_2(X43)))&(m1_waybel23(esk4_2(X42,X43),X43)|~r2_hidden(X42,a_1_3_waybel31(X43))|(~v3_orders_2(X43)|~v4_orders_2(X43)|~v5_orders_2(X43)|~v1_yellow_0(X43)|~v1_lattice3(X43)|~v3_waybel_3(X43)|~l1_orders_2(X43))))&(X42=k1_card_1(esk4_2(X42,X43))|~r2_hidden(X42,a_1_3_waybel31(X43))|(~v3_orders_2(X43)|~v4_orders_2(X43)|~v5_orders_2(X43)|~v1_yellow_0(X43)|~v1_lattice3(X43)|~v3_waybel_3(X43)|~l1_orders_2(X43))))&(~v6_waybel23(X45,X43)|~m1_waybel23(X45,X43)|X42!=k1_card_1(X45)|r2_hidden(X42,a_1_3_waybel31(X43))|(~v3_orders_2(X43)|~v4_orders_2(X43)|~v5_orders_2(X43)|~v1_yellow_0(X43)|~v1_lattice3(X43)|~v3_waybel_3(X43)|~l1_orders_2(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_1_3_waybel31])])])])])).
cnf(c_0_264, plain, (m1_subset_1(esk3_2(k11_waybel_0(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_228])])).
cnf(c_0_265, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_266, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_267, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(X1),a_1_0_waybel23(esk1_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_250]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_268, negated_conjecture, (v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(csr,[status(thm)],[c_0_251, c_0_236])).
cnf(c_0_269, negated_conjecture, (m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(csr,[status(thm)],[c_0_252, c_0_236])).
cnf(c_0_270, negated_conjecture, (~v1_finset_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_253, c_0_254]), c_0_236])).
cnf(c_0_271, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_255])).
cnf(c_0_272, negated_conjecture, (r2_hidden(k1_card_1(X2),a_1_2_waybel31(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_173])])).
cnf(c_0_273, negated_conjecture, (m1_subset_1(esk11_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_81]), c_0_176]), c_0_161])])])).
cnf(c_0_274, negated_conjecture, (m1_subset_1(esk11_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_188]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_275, negated_conjecture, (~v1_finset_1(esk11_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_188]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_276, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(esk11_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260])])).
cnf(c_0_277, plain, (~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_261, c_0_226])).
cnf(c_0_278, plain, (r2_hidden(X1,u1_pre_topc(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk6_2(X3,X2))|~m1_subset_1(esk6_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_117])])).
cnf(c_0_279, plain, (v1_cantor_1(esk5_2(X1,X2),X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_280, plain, (r2_hidden(X3,a_1_3_waybel31(X2))|~v6_waybel23(X1,X2)|~m1_waybel23(X1,X2)|X3!=k1_card_1(X1)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_281, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_282, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(k11_waybel_0(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_264])])).
cnf(c_0_283, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_266])])).
fof(c_0_284, plain, ![X109]:(~v2_pre_topc(X109)|~l1_pre_topc(X109)|~v1_xboole_0(u1_pre_topc(X109))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_pre_topc])])])).
cnf(c_0_285, negated_conjecture, (v1_finset_1(X1)|r2_hidden(k1_card_1(X1),a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_267, c_0_268]), c_0_269]), c_0_270])).
cnf(c_0_286, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_271])).
cnf(c_0_287, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_273]), c_0_274]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])]), c_0_275])])).
cnf(c_0_288, plain, (r1_tarski(esk11_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_276, c_0_277])).
cnf(c_0_289, plain, (r2_hidden(X1,u1_pre_topc(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk6_2(X3,X2))|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[c_0_278, c_0_118])).
cnf(c_0_290, plain, (v1_cantor_1(esk5_2(X1,X2),X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_291, plain, (v1_tops_2(esk5_2(X1,X2),X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_292, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_293, plain, (r2_hidden(X3,a_1_3_waybel31(X2))|X3!=k1_card_1(X1)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~v6_waybel23(X1,X2)|~m1_waybel23(X1,X2)), inference(evalgc,[status(thm)],[c_0_280])).
cnf(c_0_294, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_295, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(k11_waybel_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_282, c_0_283])).
cnf(c_0_296, plain, (~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(u1_pre_topc(X1))), inference(split_conjunct,[status(thm)],[c_0_284])).
cnf(c_0_297, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(esk6_2(X2,X1))|~m1_subset_1(esk6_2(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(esk6_2(X2,X1),X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_103])])).
cnf(c_0_298, negated_conjecture, (v1_finset_1(esk6_2(X1,X2))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|v2_struct_0(X2)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk6_2(X1,X2),esk1_0)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_134])])).
cnf(c_0_299, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287])])).
cnf(c_0_300, negated_conjecture, (v1_tops_2(esk11_1(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_288]), c_0_274]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_301, plain, (v1_cantor_1(esk11_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(X1)|~m1_waybel23(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_190]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_302, plain, (r2_hidden(esk3_2(esk6_2(X1,X2),X3),u1_pre_topc(X2))|v2_struct_0(X2)|r1_tarski(esk6_2(X1,X2),X3)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_228])])).
cnf(c_0_303, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk5_2(X3,X2))|~m1_subset_1(esk5_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk5_2(X3,X2),X2)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X3,a_1_0_waybel23(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_290])])).
cnf(c_0_304, plain, (v1_tops_2(esk5_2(X1,X2),X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_305, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))), inference(evalgc,[status(thm)],[c_0_292])).
cnf(c_0_306, plain, (X1=k1_card_1(esk5_2(X1,X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_307, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(k11_waybel_0(X2,X3))|~m1_waybel23(k11_waybel_0(X2,X3),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_90])])).
cnf(c_0_308, plain, (m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_295])])).
cnf(c_0_309, plain, (~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(u1_pre_topc(X1))), inference(evalgc,[status(thm)],[c_0_296])).
fof(c_0_310, plain, ![X85]:((~v1_finset_1(k1_card_1(X85))|v1_finset_1(X85))&(v1_card_1(k1_card_1(X85))|v1_finset_1(X85))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc8_card_1])])])])).
cnf(c_0_311, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(esk6_2(X2,X1))|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_297, c_0_117]), c_0_118])).
cnf(c_0_312, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_103]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_313, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_299, c_0_300])])])).
cnf(c_0_314, negated_conjecture, (v1_cantor_1(esk11_1(esk2_0),esk1_0)|~m1_waybel23(esk11_1(esk2_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_71]), c_0_81]), c_0_81]), c_0_81]), c_0_81]), c_0_176]), c_0_81]), c_0_161]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_315, plain, (r1_tarski(esk6_2(X1,X2),u1_pre_topc(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_302])])).
cnf(c_0_316, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk5_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_303, c_0_304]), c_0_305])).
cnf(c_0_317, plain, (k1_card_1(esk5_2(X1,X2))=X1|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))), inference(evalgc,[status(thm)],[c_0_306])).
cnf(c_0_318, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(k11_waybel_0(X2,X3))|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_307, c_0_98])).
cnf(c_0_319, plain, (v1_xboole_0(X1)|m1_subset_1(k11_waybel_0(k2_yellow_1(X1),X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_61]), c_0_137])])])).
cnf(c_0_320, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_71]), c_0_72]), c_0_73])])])).
cnf(c_0_321, plain, (v1_finset_1(X1)|~v1_finset_1(k1_card_1(X1))), inference(split_conjunct,[status(thm)],[c_0_310])).
cnf(c_0_322, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_312]), c_0_72]), c_0_108]), c_0_73])]), c_0_204]), c_0_74])])).
cnf(c_0_323, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0))|~m1_waybel23(esk11_1(esk2_0),esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_314])])).
cnf(c_0_324, negated_conjecture, (r1_tarski(esk6_2(X1,esk1_0),u1_struct_0(esk2_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_71]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
fof(c_0_325, plain, ![X15, X16]:(((r1_tarski(X15,X16)|X15!=X16)&(r1_tarski(X16,X15)|X15!=X16))&(~r1_tarski(X15,X16)|~r1_tarski(X16,X15)|X15=X16)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_326, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_317])])])).
cnf(c_0_327, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(X2))|v1_finset_1(X3)|X1!=k1_card_1(X3)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_173])])).
cnf(c_0_328, negated_conjecture, (m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_81]), c_0_320])])).
cnf(c_0_329, plain, (v8_struct_0(X1)|~v1_finset_1(esk5_2(X2,X1))|~m1_subset_1(esk5_2(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(esk5_2(X2,X1),X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_290])])).
cnf(c_0_330, plain, (v1_finset_1(X1)|~v1_finset_1(k1_card_1(X1))), inference(evalgc,[status(thm)],[c_0_321])).
cnf(c_0_331, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_322, c_0_118]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_332, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_188]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_333, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_324])])).
cnf(c_0_334, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_325])).
cnf(c_0_335, plain, (r1_tarski(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(esk3_2(X1,a_1_2_waybel31(X2)),a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_326])])).
cnf(c_0_336, plain, (m1_subset_1(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_326])])).
cnf(c_0_337, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_328]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_338, negated_conjecture, (v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_295]), c_0_162])]), c_0_320])])).
cnf(c_0_339, plain, (v8_struct_0(X1)|~v1_finset_1(esk5_2(X2,X1))|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_329, c_0_304]), c_0_305])).
cnf(c_0_340, plain, (v1_finset_1(esk5_2(X1,X2))|~v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_317])])).
cnf(c_0_341, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_241]), c_0_332])])).
cnf(c_0_342, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_241]), c_0_332])])).
cnf(c_0_343, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_344, plain, (r1_tarski(a_1_0_waybel23(X1),a_1_2_waybel31(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_228])])).
cnf(c_0_345, plain, (r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_0_waybel23(X1))|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_228])])).
cnf(c_0_346, plain, (m1_subset_1(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|v1_xboole_0(a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_241])])).
cnf(c_0_347, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)), inference(csr,[status(thm)],[c_0_337, c_0_338])).
cnf(c_0_348, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_340])])).
cnf(c_0_349, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_342])])).
cnf(c_0_350, plain, (a_1_2_waybel31(X1)=a_1_0_waybel23(X1)|v2_struct_0(X1)|~r1_tarski(a_1_2_waybel31(X1),a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_344])])).
cnf(c_0_351, plain, (r1_tarski(a_1_2_waybel31(X1),a_1_0_waybel23(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_345])])).
cnf(c_0_352, plain, (v1_xboole_0(a_1_0_waybel23(X1))|m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),a_1_2_waybel31(X1))|v2_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_243])])).
cnf(c_0_353, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_228])])).
cnf(c_0_354, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_107]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_355, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_349]), c_0_72]), c_0_108]), c_0_73])]), c_0_204])])).
cnf(c_0_356, plain, (a_1_2_waybel31(X1)=a_1_0_waybel23(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_351])])).
cnf(c_0_357, plain, (r2_hidden(k1_card_1(esk6_2(X1,X2)),a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_133])])).
cnf(c_0_358, negated_conjecture, (r1_tarski(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk3_2(X1,a_1_2_waybel31(esk1_0)),a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_349]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_359, plain, (r1_tarski(a_1_0_waybel23(X1),X2)|m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),a_1_2_waybel31(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_352, c_0_353])).
cnf(c_0_360, negated_conjecture, (r1_tarski(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel23(esk1_0)),a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_349])])).
cnf(c_0_361, plain, (v6_waybel23(esk4_2(X1,X2),X2)|~r2_hidden(X1,a_1_3_waybel31(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_362, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(k11_waybel_0(esk2_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X2,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_354, c_0_254]), c_0_269]), c_0_270])).
cnf(c_0_363, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_356]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_364, plain, (m1_subset_1(k1_card_1(esk6_2(X1,X2)),a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_357])])).
cnf(c_0_365, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_359]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_366, negated_conjecture, (r1_tarski(a_1_2_waybel31(esk1_0),a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_243])])).
fof(c_0_367, plain, ![X73]:((v1_finset_1(k1_card_1(X73))|~v1_finset_1(X73))&(v1_card_1(k1_card_1(X73))|~v1_finset_1(X73))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_card_1])])])).
cnf(c_0_368, plain, (v6_waybel23(esk4_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_361])).
cnf(c_0_369, plain, (m1_waybel23(esk4_2(X1,X2),X2)|~r2_hidden(X1,a_1_3_waybel31(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_370, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|X1!=k1_card_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X2,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_362, c_0_173]), c_0_162]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_371, negated_conjecture, (~v1_finset_1(k1_card_1(esk6_2(X1,esk1_0)))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_363, c_0_364]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_372, negated_conjecture, (a_1_2_waybel31(esk1_0)=a_1_0_waybel23(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_365]), c_0_366])])])).
cnf(c_0_373, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_367])).
cnf(c_0_374, plain, (v1_cantor_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~m1_waybel23(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_368]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_375, plain, (m1_waybel23(esk4_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_369])).
cnf(c_0_376, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk6_2(X1,esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_370, c_0_103]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_377, negated_conjecture, (~v1_finset_1(k1_card_1(esk6_2(X1,esk1_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))), inference(rw,[status(thm)],[c_0_371, c_0_372])).
cnf(c_0_378, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1)), inference(evalgc,[status(thm)],[c_0_373])).
cnf(c_0_379, plain, (v1_cantor_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_375]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_380, plain, (v1_tops_2(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~m1_waybel23(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_368]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_381, plain, (m1_subset_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~m1_waybel23(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_368]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_382, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk6_2(X1,esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_376, c_0_372])])).
cnf(c_0_383, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_326]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_384, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~v1_finset_1(esk6_2(X1,esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_378])])).
cnf(c_0_385, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))))|~m1_subset_1(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))),X2)|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_379])])).
cnf(c_0_386, plain, (v1_tops_2(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_375]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_387, plain, (m1_subset_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_375]), c_0_137]), c_0_138]), c_0_139]), c_0_140])])])).
cnf(c_0_388, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk11_1(esk2_0))|~m1_waybel23(esk11_1(esk2_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_370, c_0_314]), c_0_274]), c_0_273])]), c_0_275])])).
cnf(c_0_389, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk6_2(X2,esk1_0))|~r2_hidden(X2,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X2,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_382, c_0_383]), c_0_384])).
fof(c_0_390, plain, ![X79, X80, X82]:((((v6_waybel23(esk7_2(X79,X80),X80)|~r2_hidden(X79,a_1_0_waybel31(X80))|(~v3_orders_2(X80)|~v4_orders_2(X80)|~v5_orders_2(X80)|~v1_lattice3(X80)|~v3_waybel_3(X80)|~l1_orders_2(X80)))&(m1_waybel23(esk7_2(X79,X80),X80)|~r2_hidden(X79,a_1_0_waybel31(X80))|(~v3_orders_2(X80)|~v4_orders_2(X80)|~v5_orders_2(X80)|~v1_lattice3(X80)|~v3_waybel_3(X80)|~l1_orders_2(X80))))&(X79=k1_card_1(esk7_2(X79,X80))|~r2_hidden(X79,a_1_0_waybel31(X80))|(~v3_orders_2(X80)|~v4_orders_2(X80)|~v5_orders_2(X80)|~v1_lattice3(X80)|~v3_waybel_3(X80)|~l1_orders_2(X80))))&(~v6_waybel23(X82,X80)|~m1_waybel23(X82,X80)|X79!=k1_card_1(X82)|r2_hidden(X79,a_1_0_waybel31(X80))|(~v3_orders_2(X80)|~v4_orders_2(X80)|~v5_orders_2(X80)|~v1_lattice3(X80)|~v3_waybel_3(X80)|~l1_orders_2(X80)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_1_0_waybel31])])])])])).
cnf(c_0_391, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))))|~r2_hidden(X3,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_385, c_0_386]), c_0_387])).
cnf(c_0_392, plain, (X1=k1_card_1(esk4_2(X1,X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_263])).
cnf(c_0_393, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk11_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_188]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_394, negated_conjecture, (r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_389])])).
cnf(c_0_395, plain, (r2_hidden(X3,a_1_0_waybel31(X2))|~v6_waybel23(X1,X2)|~m1_waybel23(X1,X2)|X3!=k1_card_1(X1)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_396, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(esk4_2(X2,esk2_0))|~r2_hidden(X2,a_1_3_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_71]), c_0_81]), c_0_81]), c_0_81]), c_0_176]), c_0_81]), c_0_161]), c_0_81]), c_0_177]), c_0_72]), c_0_73])]), c_0_74])])).
cnf(c_0_397, plain, (k1_card_1(esk4_2(X1,X2))=X1|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_392])).
cnf(c_0_398, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_3_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_393])])).
cnf(c_0_399, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_134]), c_0_372]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_400, plain, (r2_hidden(X3,a_1_0_waybel31(X2))|X3!=k1_card_1(X1)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~v6_waybel23(X1,X2)|~m1_waybel23(X1,X2)), inference(evalgc,[status(thm)],[c_0_395])).
cnf(c_0_401, plain, (v6_waybel23(esk7_2(X1,X2),X2)|~r2_hidden(X1,a_1_0_waybel31(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_402, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_3_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_397]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])])).
cnf(c_0_403, negated_conjecture, (~v1_xboole_0(a_1_3_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_398])])).
cnf(c_0_404, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_118]), c_0_372]), c_0_72]), c_0_108]), c_0_73])]), c_0_74])])).
cnf(c_0_405, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|X1!=k1_card_1(esk4_2(X3,X2))|~m1_waybel23(esk4_2(X3,X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X3,a_1_3_waybel31(X2))|~v1_yellow_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_368])])).
cnf(c_0_406, plain, (v6_waybel23(esk7_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_401])).
cnf(c_0_407, plain, (m1_waybel23(esk7_2(X1,X2),X2)|~r2_hidden(X1,a_1_0_waybel31(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_408, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,a_1_3_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_241]), c_0_403])])).
cnf(c_0_409, negated_conjecture, (r1_tarski(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(esk3_2(X1,a_1_3_waybel31(esk2_0)),a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_404])])).
cnf(c_0_410, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|X1!=k1_card_1(esk4_2(X3,X2))|~r2_hidden(X3,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_405, c_0_375])).
cnf(c_0_411, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(esk7_2(X3,X2))|~m1_waybel23(esk7_2(X3,X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X3,a_1_0_waybel31(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_406])])).
cnf(c_0_412, plain, (m1_waybel23(esk7_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_407])).
cnf(c_0_413, plain, (X1=k1_card_1(esk7_2(X1,X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_414, negated_conjecture, (r1_tarski(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel23(esk1_0)),a_1_3_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_408])])).
cnf(c_0_415, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_228])])).
cnf(c_0_416, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_397])])])).
cnf(c_0_417, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(esk7_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_411, c_0_412])).
cnf(c_0_418, plain, (k1_card_1(esk7_2(X1,X2))=X1|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2))), inference(evalgc,[status(thm)],[c_0_413])).
cnf(c_0_419, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_142])])).
cnf(c_0_420, negated_conjecture, (r1_tarski(a_1_3_waybel31(esk2_0),a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_243])])).
cnf(c_0_421, negated_conjecture, (m1_subset_1(a_1_0_waybel23(esk1_0),k1_zfmisc_1(a_1_3_waybel31(esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_415])])).
cnf(c_0_422, negated_conjecture, (r2_hidden(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_404]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_423, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_418])])])).
cnf(c_0_424, negated_conjecture, (a_1_3_waybel31(esk2_0)=a_1_0_waybel23(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_420]), c_0_421])])])).
cnf(c_0_425, negated_conjecture, (r1_tarski(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(esk3_2(X1,a_1_0_waybel31(esk2_0)),a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_422])])).
cnf(c_0_426, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_424]), c_0_162]), c_0_176]), c_0_161]), c_0_177]), c_0_216]), c_0_217]), c_0_218])])])).
fof(c_0_427, plain, ![X60]:(~v3_orders_2(X60)|~v4_orders_2(X60)|~v5_orders_2(X60)|~v1_lattice3(X60)|~v3_waybel_3(X60)|~l1_orders_2(X60)|k1_waybel31(X60)=k1_setfam_1(a_1_0_waybel31(X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel31])])).
cnf(c_0_428, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_0_waybel31(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_425, c_0_228])])).
cnf(c_0_429, negated_conjecture, (r2_hidden(esk3_2(a_1_0_waybel31(esk2_0),X1),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_0_waybel31(esk2_0),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_228])])).
fof(c_0_430, plain, ![X32]:(~l1_pre_topc(X32)|k2_waybel23(X32)=k1_setfam_1(a_1_0_waybel23(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_waybel23])])).
cnf(c_0_431, plain, (k1_waybel31(X1)=k1_setfam_1(a_1_0_waybel31(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_427])).
cnf(c_0_432, negated_conjecture, (a_1_0_waybel31(esk2_0)=a_1_0_waybel23(esk1_0)|~r1_tarski(a_1_0_waybel31(esk2_0),a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_428])])).
cnf(c_0_433, negated_conjecture, (r1_tarski(a_1_0_waybel31(esk2_0),a_1_0_waybel23(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_429])])).
cnf(c_0_434, plain, (k2_waybel23(X1)=k1_setfam_1(a_1_0_waybel23(X1))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_430])).
cnf(c_0_435, plain, (k1_setfam_1(a_1_0_waybel31(X1))=k1_waybel31(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_431])).
cnf(c_0_436, negated_conjecture, (a_1_0_waybel31(esk2_0)=a_1_0_waybel23(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_432, c_0_433])])])).
cnf(c_0_437, negated_conjecture, (k2_waybel23(esk1_0)!=k1_waybel31(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_438, plain, (k1_setfam_1(a_1_0_waybel23(X1))=k2_waybel23(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_434])).
cnf(c_0_439, negated_conjecture, (k1_setfam_1(a_1_0_waybel23(esk1_0))=k1_waybel31(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_436]), c_0_162]), c_0_176]), c_0_161]), c_0_216]), c_0_217]), c_0_218])])])).
cnf(c_0_440, negated_conjecture, (k1_waybel31(esk2_0)!=k2_waybel23(esk1_0)), inference(evalgc,[status(thm)],[c_0_437])).
cnf(c_0_441, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_439]), c_0_72])]), c_0_440]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 442
# Proof object clause steps            : 363
# Proof object formula steps           : 79
# Proof object conjectures             : 156
# Proof object clause conjectures      : 153
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 76
# Proof object initial formulas used   : 39
# Proof object generating inferences   : 158
# Proof object simplifying inferences  : 525
# Training examples: 239 positive, 2061 negative
# Training: Positive examples begin
cnf(c_0_73, negated_conjecture, (v2_pre_topc(esk1_0))).# trainpos
cnf(c_0_108, negated_conjecture, (v6_pre_topc(esk1_0))).# trainpos
cnf(c_0_72, negated_conjecture, (l1_pre_topc(esk1_0))).# trainpos
cnf(c_0_218, negated_conjecture, (v3_orders_2(esk2_0))).# trainpos
cnf(c_0_217, negated_conjecture, (v4_orders_2(esk2_0))).# trainpos
cnf(c_0_216, negated_conjecture, (v5_orders_2(esk2_0))).# trainpos
cnf(c_0_177, negated_conjecture, (v1_yellow_0(esk2_0))).# trainpos
cnf(c_0_161, negated_conjecture, (v1_lattice3(esk2_0))).# trainpos
cnf(c_0_176, negated_conjecture, (v3_waybel_3(esk2_0))).# trainpos
cnf(c_0_162, negated_conjecture, (l1_orders_2(esk2_0))).# trainpos
cnf(c_0_74, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_204, negated_conjecture, (~v8_struct_0(esk1_0))).# trainpos
cnf(c_0_440, negated_conjecture, (k1_waybel31(esk2_0)!=k2_waybel23(esk1_0))).# trainpos
cnf(c_0_140, plain, (v3_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_139, plain, (v4_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_62, negated_conjecture, (k2_yellow_1(u1_pre_topc(esk1_0))=esk2_0)).# trainpos
cnf(c_0_138, plain, (v5_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_137, plain, (l1_orders_2(k2_yellow_1(X1)))).# trainpos
cnf(c_0_61, plain, (u1_struct_0(k2_yellow_1(X1))=X1)).# trainpos
cnf(c_0_46, plain, (~v2_struct_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_286, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_266, plain, (l1_struct_0(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_309, plain, (~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(u1_pre_topc(X1)))).# trainpos
cnf(c_0_330, plain, (v1_finset_1(X1)|~v1_finset_1(k1_card_1(X1)))).# trainpos
cnf(c_0_378, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).# trainpos
cnf(c_0_438, plain, (k1_setfam_1(a_1_0_waybel23(X1))=k2_waybel23(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_53, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_67, plain, (v2_struct_0(X1)|v2_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).# trainpos
cnf(c_0_227, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_294, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_142, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_265, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_343, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_228, plain, (r1_tarski(X1,X2)|r2_hidden(esk3_2(X1,X2),X1))).# trainpos
cnf(c_0_240, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2))).# trainpos
cnf(c_0_166, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_241, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_244, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1))).# trainpos
cnf(c_0_189, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~v1_xboole_0(X2)|~v6_waybel23(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_143, plain, (v2_struct_0(X1)|v3_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_231, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_317, plain, (k1_card_1(esk5_2(X1,X2))=X1|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).# trainpos
cnf(c_0_304, plain, (v1_tops_2(esk5_2(X1,X2),X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).# trainpos
cnf(c_0_290, plain, (v1_cantor_1(esk5_2(X1,X2),X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).# trainpos
cnf(c_0_435, plain, (k1_setfam_1(a_1_0_waybel31(X1))=k1_waybel31(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_202, plain, (v8_struct_0(X1)|~v2_pre_topc(X1)|~v6_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v1_tops_2(X2,X1)|~v1_cantor_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_45, plain, (v2_struct_0(X1)|v3_lattice3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v24_waybel_0(X1))).# trainpos
cnf(c_0_190, plain, (v6_waybel23(esk11_1(X1),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_188, plain, (m1_waybel23(esk11_1(X1),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_305, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).# trainpos
cnf(c_0_245, plain, (r1_tarski(X1,u1_pre_topc(X2))|~l1_pre_topc(X2)|~v1_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))).# trainpos
cnf(c_0_141, plain, (v1_tops_2(X1,X2)|~l1_pre_topc(X2)|~r1_tarski(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))).# trainpos
cnf(c_0_167, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_195, plain, (v2_struct_0(X1)|m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_165, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_134, plain, (k1_card_1(esk6_2(X1,X2))=X1|v2_struct_0(X2)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).# trainpos
cnf(c_0_117, plain, (v2_struct_0(X2)|v1_tops_2(esk6_2(X1,X2),X2)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).# trainpos
cnf(c_0_103, plain, (v2_struct_0(X2)|v1_cantor_1(esk6_2(X1,X2),X2)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).# trainpos
cnf(c_0_118, plain, (v2_struct_0(X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).# trainpos
cnf(c_0_187, plain, (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~m1_waybel23(X2,X1))).# trainpos
cnf(c_0_406, plain, (v6_waybel23(esk7_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).# trainpos
cnf(c_0_102, plain, (r2_hidden(X3,a_1_0_waybel23(X2))|X3!=k1_card_1(X1)|~l1_pre_topc(X2)|~v1_tops_2(X1,X2)|~v1_cantor_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))).# trainpos
cnf(c_0_412, plain, (m1_waybel23(esk7_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).# trainpos
cnf(c_0_85, plain, (v2_struct_0(X2)|v1_tops_2(X1,X2)|k2_yellow_1(u1_pre_topc(X2))!=X3|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3))).# trainpos
cnf(c_0_70, plain, (v2_struct_0(X2)|v1_cantor_1(X1,X2)|k2_yellow_1(u1_pre_topc(X2))!=X3|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3))).# trainpos
cnf(c_0_418, plain, (k1_card_1(esk7_2(X1,X2))=X1|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).# trainpos
cnf(c_0_397, plain, (k1_card_1(esk4_2(X1,X2))=X1|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2)))).# trainpos
cnf(c_0_368, plain, (v6_waybel23(esk4_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2)))).# trainpos
cnf(c_0_375, plain, (m1_waybel23(esk4_2(X1,X2),X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2)))).# trainpos
cnf(c_0_173, plain, (k1_card_1(k11_waybel_0(X1,X2))=k1_card_1(X2)|v1_finset_1(X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_400, plain, (r2_hidden(X3,a_1_0_waybel31(X2))|X3!=k1_card_1(X1)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~v6_waybel23(X1,X2)|~m1_waybel23(X1,X2))).# trainpos
cnf(c_0_106, plain, (v2_struct_0(X2)|r2_hidden(X3,a_1_2_waybel31(X2))|X3!=k1_card_1(X1)|~v2_pre_topc(X2)|~v6_pre_topc(X2)|~l1_pre_topc(X2)|~v1_tops_2(X1,X2)|~v1_cantor_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))).# trainpos
cnf(c_0_293, plain, (r2_hidden(X3,a_1_3_waybel31(X2))|X3!=k1_card_1(X1)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_yellow_0(X2)|~v1_lattice3(X2)|~v3_waybel_3(X2)|~l1_orders_2(X2)|~v6_waybel23(X1,X2)|~m1_waybel23(X1,X2))).# trainpos
cnf(c_0_86, plain, (v2_struct_0(X2)|m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|k2_yellow_1(u1_pre_topc(X2))!=X3|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~v1_lattice3(X3)|~v3_waybel_3(X3)|~l1_orders_2(X3)|~v6_waybel23(X1,X3)|~m1_waybel23(X1,X3))).# trainpos
cnf(c_0_66, plain, (v2_struct_0(X4)|v6_waybel23(k11_waybel_0(X1,X2),X1)|X2!=X3|k2_yellow_1(u1_pre_topc(X4))!=X1|~v2_pre_topc(X4)|~l1_pre_topc(X4)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_tops_2(X3,X4)|~v1_cantor_1(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))).# trainpos
cnf(c_0_77, plain, (v2_struct_0(X4)|m1_waybel23(k11_waybel_0(X1,X2),X1)|X2!=X3|k2_yellow_1(u1_pre_topc(X4))!=X1|~v2_pre_topc(X4)|~l1_pre_topc(X4)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_tops_2(X3,X4)|~v1_cantor_1(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))).# trainpos
cnf(c_0_71, negated_conjecture, (u1_pre_topc(esk1_0)=u1_struct_0(esk2_0))).# trainpos
cnf(c_0_75, plain, (v6_waybel23(k11_waybel_0(X1,X2),X1)|v2_struct_0(X3)|k2_yellow_1(u1_pre_topc(X3))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,X3)|~v1_tops_2(X2,X3)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).# trainpos
cnf(c_0_179, negated_conjecture, (~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_160, plain, (v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_59, plain, (v3_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v3_waybel_3(X1))).# trainpos
cnf(c_0_157, plain, (v1_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_pre_topc(X2))))).# trainpos
cnf(c_0_84, plain, (m1_waybel23(k11_waybel_0(X1,X2),X1)|v2_struct_0(X3)|k2_yellow_1(u1_pre_topc(X3))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,X3)|~v1_tops_2(X2,X3)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).# trainpos
cnf(c_0_164, plain, (v1_cantor_1(X1,X2)|v2_struct_0(X2)|~m1_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v6_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_183, plain, (v1_tops_2(X1,X2)|v2_struct_0(X2)|~m1_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v6_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_419, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_262, plain, (r2_hidden(X1,u1_pre_topc(X2))|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(X3,X2)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_184, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~m1_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v6_waybel23(X1,k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_329, plain, (v8_struct_0(X1)|~v1_finset_1(esk5_2(X2,X1))|~m1_subset_1(esk5_2(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(esk5_2(X2,X1),X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).# trainpos
cnf(c_0_259, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk3_2(X1,X2),X2))).# trainpos
cnf(c_0_353, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_283, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_243, plain, (m1_subset_1(esk3_2(X1,X2),X1)|r1_tarski(X1,X2))).# trainpos
cnf(c_0_340, plain, (v1_finset_1(esk5_2(X1,X2))|~v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).# trainpos
cnf(c_0_209, plain, (m1_subset_1(esk11_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_297, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(esk6_2(X2,X1))|~m1_subset_1(esk6_2(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(esk6_2(X2,X1),X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).# trainpos
cnf(c_0_210, plain, (v2_struct_0(X1)|~v1_xboole_0(esk11_1(X1))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_208, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_246, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k11_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2))).# trainpos
cnf(c_0_116, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk6_2(X3,X2))|~m1_subset_1(esk6_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk6_2(X3,X2),X2)|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_303, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk5_2(X3,X2))|~m1_subset_1(esk5_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk5_2(X3,X2),X2)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X3,a_1_0_waybel23(X2)))).# trainpos
cnf(c_0_411, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(esk7_2(X3,X2))|~m1_waybel23(esk7_2(X3,X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X3,a_1_0_waybel31(X2)))).# trainpos
cnf(c_0_405, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|X1!=k1_card_1(esk4_2(X3,X2))|~m1_waybel23(esk4_2(X3,X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X3,a_1_3_waybel31(X2))|~v1_yellow_0(X2))).# trainpos
cnf(c_0_81, negated_conjecture, (k2_yellow_1(u1_struct_0(esk2_0))=esk2_0)).# trainpos
cnf(c_0_76, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_238, negated_conjecture, (v1_tops_2(X1,esk1_0)|~r1_tarski(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_320, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_156, plain, (v6_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v1_tops_2(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_91, negated_conjecture, (m1_waybel23(k11_waybel_0(X1,X2),X1)|k2_yellow_1(u1_struct_0(esk2_0))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_348, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(X2))).# trainpos
cnf(c_0_93, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|k2_yellow_1(u1_struct_0(esk2_0))!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).# trainpos
cnf(c_0_178, negated_conjecture, (v2_lattice3(k2_yellow_1(u1_struct_0(esk2_0)))|v2_struct_0(k2_yellow_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_92, negated_conjecture, (v1_tops_2(X1,esk1_0)|k2_yellow_1(u1_struct_0(esk2_0))!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).# trainpos
cnf(c_0_80, negated_conjecture, (v1_cantor_1(X1,esk1_0)|k2_yellow_1(u1_struct_0(esk2_0))!=X2|~m1_waybel23(X1,X2)|~v6_waybel23(X1,X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).# trainpos
cnf(c_0_82, negated_conjecture, (v6_waybel23(k11_waybel_0(X1,X2),X1)|k2_yellow_1(u1_struct_0(esk2_0))!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_206, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v6_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v3_waybel_3(k2_yellow_1(u1_struct_0(esk2_0)))|~v1_lattice3(k2_yellow_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_181, negated_conjecture, (v1_cantor_1(X1,esk1_0)|~m1_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v6_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v3_waybel_3(k2_yellow_1(u1_struct_0(esk2_0)))|~v1_lattice3(k2_yellow_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_264, plain, (m1_subset_1(esk3_2(k11_waybel_0(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).# trainpos
cnf(c_0_261, plain, (v1_xboole_0(esk11_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_205, negated_conjecture, (v1_tops_2(X1,esk1_0)|~m1_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v6_waybel23(X1,k2_yellow_1(u1_struct_0(esk2_0)))|~v3_waybel_3(k2_yellow_1(u1_struct_0(esk2_0)))|~v1_lattice3(k2_yellow_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_201, plain, (m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v1_tops_2(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_301, plain, (v1_cantor_1(esk11_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(X1)|~m1_waybel23(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_278, plain, (r2_hidden(X1,u1_pre_topc(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk6_2(X3,X2))|~m1_subset_1(esk6_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_374, plain, (v1_cantor_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~m1_waybel23(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2))))).# trainpos
cnf(c_0_357, plain, (r2_hidden(k1_card_1(esk6_2(X1,X2)),a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_152, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_220, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(X1,esk1_0)|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0))).# trainpos
cnf(c_0_109, negated_conjecture, (v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_waybel23(k11_waybel_0(X1,X2),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1))).# trainpos
cnf(c_0_236, negated_conjecture, (m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).# trainpos
cnf(c_0_380, plain, (v1_tops_2(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~m1_waybel23(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2))))).# trainpos
cnf(c_0_239, plain, (m1_subset_1(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v2_struct_0(X1)|~m1_waybel23(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_110, negated_conjecture, (m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X1|~m1_waybel23(k11_waybel_0(X1,X2),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1))).# trainpos
cnf(c_0_97, negated_conjecture, (v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_waybel23(k11_waybel_0(X1,X2),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1))).# trainpos
cnf(c_0_381, plain, (m1_subset_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~m1_waybel23(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2))))).# trainpos
cnf(c_0_307, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(k11_waybel_0(X2,X3))|~m1_waybel23(k11_waybel_0(X2,X3),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0))).# trainpos
cnf(c_0_196, negated_conjecture, (v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v2_lattice3(esk2_0))).# trainpos
cnf(c_0_254, negated_conjecture, (v1_tops_2(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_282, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(k11_waybel_0(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).# trainpos
cnf(c_0_327, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(X2))|v1_finset_1(X3)|X1!=k1_card_1(X3)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).# trainpos
cnf(c_0_338, negated_conjecture, (v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_308, plain, (m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).# trainpos
cnf(c_0_253, negated_conjecture, (~v1_finset_1(k11_waybel_0(esk2_0,X1))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0))).# trainpos
cnf(c_0_251, negated_conjecture, (v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).# trainpos
cnf(c_0_258, negated_conjecture, (~v1_finset_1(esk11_1(esk2_0))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).# trainpos
cnf(c_0_252, negated_conjecture, (m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).# trainpos
cnf(c_0_235, negated_conjecture, (v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).# trainpos
cnf(c_0_257, negated_conjecture, (m1_subset_1(esk11_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).# trainpos
cnf(c_0_120, negated_conjecture, (r2_hidden(X1,a_1_2_waybel31(esk1_0))|X1!=k1_card_1(k11_waybel_0(X2,X3))|~m1_subset_1(k11_waybel_0(X2,X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(k11_waybel_0(X2,X3),esk1_0)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).# trainpos
cnf(c_0_275, negated_conjecture, (~v1_finset_1(esk11_1(esk2_0)))).# trainpos
cnf(c_0_274, negated_conjecture, (m1_subset_1(esk11_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_256, plain, (m1_subset_1(esk11_1(k2_yellow_1(X1)),k1_zfmisc_1(X1))|~v3_waybel_3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1)))).# trainpos
cnf(c_0_153, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_2_waybel31(esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_314, negated_conjecture, (v1_cantor_1(esk11_1(esk2_0),esk1_0)|~m1_waybel23(esk11_1(esk2_0),esk2_0))).# trainpos
cnf(c_0_319, plain, (v1_xboole_0(X1)|m1_subset_1(k11_waybel_0(k2_yellow_1(X1),X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_273, negated_conjecture, (m1_subset_1(esk11_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_328, negated_conjecture, (m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_172, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_194, negated_conjecture, (r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_272, negated_conjecture, (r2_hidden(k1_card_1(X2),a_1_2_waybel31(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_225, plain, (v1_xboole_0(esk11_1(X1))|m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,esk11_1(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_214, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel23(esk1_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_260, plain, (m1_subset_1(esk3_2(esk11_1(X1),X2),u1_struct_0(X1))|r1_tarski(esk11_1(X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_287, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0))).# trainpos
cnf(c_0_299, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0))).# trainpos
cnf(c_0_276, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(esk11_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_249, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_337, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_300, negated_conjecture, (v1_tops_2(esk11_1(esk2_0),esk1_0))).# trainpos
cnf(c_0_267, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(X1),a_1_0_waybel23(esk1_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0))).# trainpos
cnf(c_0_354, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0))).# trainpos
cnf(c_0_313, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).# trainpos
cnf(c_0_370, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|X1!=k1_card_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X2,esk1_0))).# trainpos
cnf(c_0_298, negated_conjecture, (v1_finset_1(esk6_2(X1,X2))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|v2_struct_0(X2)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk6_2(X1,X2),esk1_0)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_323, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).# trainpos
cnf(c_0_332, negated_conjecture, (~v1_xboole_0(a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_326, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_312, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_388, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk11_1(esk2_0))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).# trainpos
cnf(c_0_335, plain, (r1_tarski(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(esk3_2(X1,a_1_2_waybel31(X2)),a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_393, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk11_1(esk2_0)))).# trainpos
cnf(c_0_322, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_423, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).# trainpos
cnf(c_0_336, plain, (m1_subset_1(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_416, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).# trainpos
cnf(c_0_344, plain, (r1_tarski(a_1_0_waybel23(X1),a_1_2_waybel31(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_398, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_3_waybel31(esk2_0)))).# trainpos
cnf(c_0_331, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_403, negated_conjecture, (~v1_xboole_0(a_1_3_waybel31(esk2_0)))).# trainpos
cnf(c_0_350, plain, (a_1_2_waybel31(X1)=a_1_0_waybel23(X1)|v2_struct_0(X1)|~r1_tarski(a_1_2_waybel31(X1),a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_341, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_376, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk6_2(X1,esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_346, plain, (m1_subset_1(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|v1_xboole_0(a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).# trainpos
cnf(c_0_352, plain, (v1_xboole_0(a_1_0_waybel23(X1))|m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),a_1_2_waybel31(X1))|v2_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_302, plain, (r2_hidden(esk3_2(esk6_2(X1,X2),X3),u1_pre_topc(X2))|v2_struct_0(X2)|r1_tarski(esk6_2(X1,X2),X3)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_315, plain, (r1_tarski(esk6_2(X1,X2),u1_pre_topc(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_379, plain, (v1_cantor_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_324, negated_conjecture, (r1_tarski(esk6_2(X1,esk1_0),u1_struct_0(esk2_0))|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_333, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_364, plain, (m1_subset_1(k1_card_1(esk6_2(X1,X2)),a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_342, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_345, plain, (r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_0_waybel23(X1))|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_351, plain, (r1_tarski(a_1_2_waybel31(X1),a_1_0_waybel23(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_349, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_383, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_356, plain, (a_1_2_waybel31(X1)=a_1_0_waybel23(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_355, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_358, negated_conjecture, (r1_tarski(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk3_2(X1,a_1_2_waybel31(esk1_0)),a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_360, negated_conjecture, (r1_tarski(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel23(esk1_0)),a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_363, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_366, negated_conjecture, (r1_tarski(a_1_2_waybel31(esk1_0),a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_365, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_372, negated_conjecture, (a_1_2_waybel31(esk1_0)=a_1_0_waybel23(esk1_0))).# trainpos
cnf(c_0_371, negated_conjecture, (~v1_finset_1(k1_card_1(esk6_2(X1,esk1_0)))|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).# trainpos
cnf(c_0_384, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~v1_finset_1(esk6_2(X1,esk1_0)))).# trainpos
cnf(c_0_386, plain, (v1_tops_2(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_382, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk6_2(X1,esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_387, plain, (m1_subset_1(esk4_2(X1,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_394, negated_conjecture, (r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_399, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_404, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_385, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))))|~m1_subset_1(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))),X2)|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~v2_pre_topc(X2))).# trainpos
cnf(c_0_409, negated_conjecture, (r1_tarski(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(esk3_2(X1,a_1_3_waybel31(esk2_0)),a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_422, negated_conjecture, (r2_hidden(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_396, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(esk4_2(X2,esk2_0))|~r2_hidden(X2,a_1_3_waybel31(esk2_0)))).# trainpos
cnf(c_0_402, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).# trainpos
cnf(c_0_425, negated_conjecture, (r1_tarski(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(esk3_2(X1,a_1_0_waybel31(esk2_0)),a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_415, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)))).# trainpos
cnf(c_0_428, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_0_waybel31(esk2_0)))).# trainpos
cnf(c_0_408, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).# trainpos
cnf(c_0_414, negated_conjecture, (r1_tarski(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel23(esk1_0)),a_1_3_waybel31(esk2_0)))).# trainpos
cnf(c_0_421, negated_conjecture, (m1_subset_1(a_1_0_waybel23(esk1_0),k1_zfmisc_1(a_1_3_waybel31(esk2_0))))).# trainpos
cnf(c_0_432, negated_conjecture, (a_1_0_waybel31(esk2_0)=a_1_0_waybel23(esk1_0)|~r1_tarski(a_1_0_waybel31(esk2_0),a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_420, negated_conjecture, (r1_tarski(a_1_3_waybel31(esk2_0),a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_424, negated_conjecture, (a_1_3_waybel31(esk2_0)=a_1_0_waybel23(esk1_0))).# trainpos
cnf(c_0_426, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel31(esk2_0)))).# trainpos
cnf(c_0_429, negated_conjecture, (r2_hidden(esk3_2(a_1_0_waybel31(esk2_0),X1),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_0_waybel31(esk2_0),X1))).# trainpos
cnf(c_0_433, negated_conjecture, (r1_tarski(a_1_0_waybel31(esk2_0),a_1_0_waybel23(esk1_0)))).# trainpos
cnf(c_0_436, negated_conjecture, (a_1_0_waybel31(esk2_0)=a_1_0_waybel23(esk1_0))).# trainpos
cnf(c_0_439, negated_conjecture, (k1_setfam_1(a_1_0_waybel23(esk1_0))=k1_waybel31(esk2_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_188, plain, (k5_numbers=k4_ordinal1)).#trainneg
cnf(i_0_85, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_165, plain, (v1_xboole_0(esk9_0))).#trainneg
cnf(i_0_232, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_197, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_206, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_233, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_198, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_31, plain, (v1_orders_2(k2_yellow_1(X1)))).#trainneg
cnf(i_0_126, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_180, plain, (v1_xboole_0(esk10_1(X1)))).#trainneg
cnf(i_0_196, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_195, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_84, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_178, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_87, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_108, plain, (v6_pre_topc(X1)|~v2_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_190, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_128, plain, (v1_yellow_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v8_struct_0(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_199, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_236, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_194, plain, (~v2_struct_0(X1)|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_79, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_93, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_257, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_109, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_158, plain, (m1_subset_1(esk8_1(X1),X1))).#trainneg
cnf(i_0_49, plain, (v1_relat_2(u1_orders_2(k2_yellow_1(X1))))).#trainneg
cnf(i_0_48, plain, (v4_relat_2(u1_orders_2(k2_yellow_1(X1))))).#trainneg
cnf(i_0_47, plain, (v8_relat_2(u1_orders_2(k2_yellow_1(X1))))).#trainneg
cnf(i_0_168, plain, (v1_relat_1(u1_orders_2(k2_yellow_1(X1))))).#trainneg
cnf(i_0_240, plain, (v1_xboole_0(esk13_2(X1,X2)))).#trainneg
cnf(i_0_239, plain, (v1_relat_1(esk13_2(X1,X2)))).#trainneg
cnf(i_0_23, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_24, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_181, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_46, plain, (v1_partfun1(u1_orders_2(k2_yellow_1(X1)),X1))).#trainneg
cnf(i_0_148, plain, (v1_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_187, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~v2_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))).#trainneg
cnf(i_0_212, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_179, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_153, plain, (v8_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_112, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_238, plain, (v4_relat_1(esk13_2(X1,X2),X1))).#trainneg
cnf(i_0_237, plain, (v5_relat_1(esk13_2(X1,X2),X2))).#trainneg
cnf(i_0_57, plain, (v1_card_1(k2_waybel23(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_114, plain, (v1_xboole_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_78, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_177, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_218, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_62, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_252, plain, (v2_struct_0(X1)|v1_yellow_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v25_waybel_0(X1))).#trainneg
cnf(i_0_147, plain, (v2_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_226, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_250, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_136, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_55, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_166, plain, (v2_struct_0(X1)|v3_yellow_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_97, plain, (v1_finset_1(u1_struct_0(X1))|~v8_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_45, plain, (m1_subset_1(u1_orders_2(k2_yellow_1(X1)),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_241, plain, (m1_subset_1(esk13_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_255, plain, (v1_yellow_3(X1)|~l1_orders_2(X1)|~v1_xboole_0(u1_orders_2(X1)))).#trainneg
cnf(i_0_156, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_182, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_60, plain, (v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v7_struct_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_80, plain, (v2_struct_0(X1)|v2_waybel_3(X1)|~v3_orders_2(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_159, plain, (v3_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_waybel_8(X1))).#trainneg
cnf(i_0_138, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_101, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v2_waybel_8(X1))).#trainneg
cnf(i_0_228, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_100, plain, (v2_struct_0(X1)|v1_waybel_8(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v2_waybel_8(X1))).#trainneg
cnf(i_0_137, plain, (v2_struct_0(X1)|v25_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_258, plain, (v2_struct_0(X1)|v6_waybel23(esk14_1(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_191, plain, (v2_yellow_0(X1)|~v3_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v24_waybel_0(X1))).#trainneg
cnf(i_0_256, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_145, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_144, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_50, plain, (g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=X1|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_150, plain, (l1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_151, plain, (v1_orders_2(g1_orders_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_259, plain, (v2_struct_0(X1)|m1_subset_1(esk14_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_186, plain, (v2_struct_0(X1)|v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_234, plain, (v4_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v4_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_220, plain, (v3_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v24_waybel_0(X1)|~v2_waybel_3(X1))).#trainneg
cnf(i_0_44, plain, (v1_card_1(k1_waybel31(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_260, plain, (v2_struct_0(X1)|v1_waybel25(X1)|~v2_pre_topc(X1)|~v6_pre_topc(X1)|~l1_pre_topc(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_213, plain, (v2_struct_0(X1)|v1_yellow11(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_waybel_1(X1))).#trainneg
cnf(i_0_183, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_184, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_133, plain, (v2_struct_0(X1)|v2_waybel_1(k2_yellow_1(u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_117, plain, (v3_lattice3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v8_struct_0(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_200, plain, (v2_waybel_8(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v3_waybel_8(X1))).#trainneg
cnf(i_0_207, plain, (v5_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v5_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_248, plain, (v2_lattice3(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_235, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v4_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_208, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v5_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_249, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_37, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_33, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_34, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_115, plain, (v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)|~v1_xboole_0(k11_waybel_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_229, plain, (v2_struct_0(X1)|v3_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_242, plain, (v3_waybel_8(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_169, plain, (v2_waybel_2(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_211, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_135, plain, (v2_waybel23(X2,X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~m1_waybel23(X2,X1))).#trainneg
cnf(i_0_189, plain, (v7_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1)|~v1_xboole_0(X2)|~m1_waybel23(X2,X1))).#trainneg
cnf(i_0_143, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_231, plain, (m1_waybel23(esk12_1(X1),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v3_waybel_3(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_149, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(k11_waybel_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_146, plain, (v1_waybel_0(k11_waybel_0(X1,X2),X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_251, plain, (v1_waybel_0(X2,X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~v2_waybel23(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_466, plain, (v1_tops_2(u1_pre_topc(X1),X1)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_311, plain, (v1_yellow_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_271, plain, (esk9_0=k1_xboole_0)).#trainneg
cnf(i_0_270, plain, (esk10_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_325, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_395, plain, (v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_281, plain, (esk13_2(X1,X2)=k1_xboole_0)).#trainneg
cnf(i_0_285, plain, (v2_struct_0(k2_yellow_1(X1))|~v2_struct_0(g1_orders_2(X1,u1_orders_2(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_309, negated_conjecture, (~l1_struct_0(esk1_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_464, plain, (u1_pre_topc(X1)=X2|~r1_tarski(u1_pre_topc(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_521, plain, (v1_finset_1(k1_card_1(X2))|v1_finset_1(X2)|~v1_finset_1(k11_waybel_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_463, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_pre_topc(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(X1,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_470, plain, (v1_finset_1(k11_waybel_0(X1,X2))|v2_struct_0(X1)|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_308, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_307, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_434, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_340, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_447, plain, (X1=X2|g1_orders_2(X1,u1_orders_2(k2_yellow_1(X1)))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_290, plain, (v1_finset_1(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_345, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_274, plain, (~m1_subset_1(k4_ordinal1,k4_ordinal1))).#trainneg
cnf(i_0_452, plain, (v2_waybel_8(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_291, plain, (v1_finset_1(X1)|~m1_subset_1(k1_card_1(X1),k4_ordinal1))).#trainneg
cnf(i_0_312, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_314, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_320, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_352, plain, (v1_finset_1(u1_pre_topc(X1))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_456, plain, (v7_struct_0(X1)|~v1_xboole_0(esk11_1(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_310, plain, (v1_yellow_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_struct_0(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_301, plain, (k1_card_1(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_442, plain, (u1_orders_2(k2_yellow_1(X1))=X2|g1_orders_2(X1,u1_orders_2(k2_yellow_1(X1)))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_300, plain, (X1=k1_card_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_416, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_287, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk8_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_323, plain, (v3_ordinal1(esk8_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_390, plain, (v1_xboole_0(u1_orders_2(k2_yellow_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_327, plain, (v3_ordinal1(u1_orders_2(k2_yellow_1(X1)))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_367, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_448, plain, (u1_struct_0(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_421, plain, (k1_card_1(X1)=X1|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_377, plain, (v1_finset_1(esk14_1(X1))|v2_struct_0(X1)|~v1_finset_1(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_483, plain, (v1_finset_1(esk6_2(X1,X2))|v2_struct_0(X2)|~v1_finset_1(X1)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_355, plain, (m1_subset_1(u1_pre_topc(k2_yellow_1(X1)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_460, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,esk5_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_474, plain, (m1_subset_1(k11_waybel_0(k2_yellow_1(X1),X2),k1_zfmisc_1(X1))|v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_386, plain, (v2_struct_0(X1)|~v1_xboole_0(esk14_1(X1))|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_359, plain, (v4_relat_1(u1_orders_2(k2_yellow_1(X1)),X1))).#trainneg
cnf(i_0_339, plain, (v1_finset_1(esk8_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_443, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_334, plain, (v1_relat_1(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_422, plain, (v1_card_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_316, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X1))|~r2_hidden(X2,u1_orders_2(k2_yellow_1(X1))))).#trainneg
cnf(i_0_363, plain, (v5_relat_1(u1_orders_2(k2_yellow_1(X1)),X1))).#trainneg
cnf(i_0_341, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~v1_finset_1(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_344, plain, (v1_xboole_0(esk8_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_520, plain, (v1_xboole_0(k1_card_1(X2))|v1_finset_1(X2)|~v1_xboole_0(k11_waybel_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_490, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X1,esk6_2(X3,X2))|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_420, plain, (v1_finset_1(X1)|~v1_finset_1(esk5_2(X1,X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_481, plain, (k1_card_1(X1)=X1|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_480, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~v1_finset_1(esk6_2(X1,X2))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_365, plain, (g1_orders_2(X1,u1_orders_2(k2_yellow_1(X1)))=k2_yellow_1(X1))).#trainneg
cnf(i_0_380, plain, (m1_subset_1(esk14_1(k2_yellow_1(X1)),k1_zfmisc_1(X1))|v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_426, plain, (v1_finset_1(u1_orders_2(X1))|~v1_finset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_485, plain, (v7_struct_0(X1)|~v1_xboole_0(esk12_1(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_432, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_403, plain, (~l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v2_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v2_lattice3(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_346, plain, (v1_xboole_0(u1_orders_2(k2_yellow_1(X1)))|~v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_429, plain, (~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~r2_hidden(X2,u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_351, plain, (v3_ordinal1(u1_pre_topc(X1))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_493, plain, (v1_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(X1,X2)|~v3_waybel_3(X2))).#trainneg
cnf(i_0_509, plain, (v1_finset_1(esk7_2(X1,X2))|~v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_358, plain, (v4_relat_1(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_362, plain, (v5_relat_1(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_533, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(esk11_1(X2))|~m1_waybel23(esk11_1(X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_526, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk14_1(X2))|~m1_waybel23(esk14_1(X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_532, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk14_1(X2))|~m1_waybel23(esk14_1(X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_527, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|X1!=k1_card_1(esk11_1(X2))|~m1_waybel23(esk11_1(X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_500, plain, (m1_subset_1(esk7_2(X1,X2),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_507, plain, (k1_card_1(X1)=X1|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_482, plain, (v1_card_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_488, plain, (v1_finset_1(esk6_2(X1,X2))|v2_struct_0(X2)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_498, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|X1!=k1_card_1(esk5_2(X3,X2))|~m1_subset_1(esk5_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk5_2(X3,X2),X2)|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_415, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk8_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_376, plain, (v3_ordinal1(esk14_1(X1))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_446, plain, (X1=X2|g1_orders_2(X1,esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_370, plain, (l1_orders_2(g1_orders_2(X1,esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_419, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk5_2(X1,X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_427, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_515, plain, (v1_finset_1(esk4_2(X1,X2))|~v1_finset_1(X1)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_374, plain, (v1_orders_2(g1_orders_2(X1,esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))).#trainneg
cnf(i_0_506, plain, (v1_finset_1(X1)|~v1_finset_1(esk7_2(X1,X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_530, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk6_2(X3,X2))|~m1_subset_1(esk6_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk6_2(X3,X2),X2)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X3,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_458, plain, (v1_finset_1(esk5_2(X1,X2))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_519, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2))|~v1_yellow_0(X2))).#trainneg
cnf(i_0_513, plain, (k1_card_1(X1)=X1|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_512, plain, (v1_finset_1(X1)|~v1_finset_1(esk4_2(X1,X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_535, plain, (r2_hidden(X1,a_1_3_waybel31(X2))|X1!=k1_card_1(esk4_2(X3,X2))|~m1_waybel23(esk4_2(X3,X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X3,a_1_3_waybel31(X2)))).#trainneg
cnf(i_0_492, plain, (m1_subset_1(esk6_2(X1,k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|v2_struct_0(k2_yellow_1(X2))|~r2_hidden(X1,a_1_2_waybel31(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2))|~v6_pre_topc(k2_yellow_1(X2))|~v2_pre_topc(k2_yellow_1(X2)))).#trainneg
cnf(i_0_494, plain, (m1_subset_1(esk12_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_496, plain, (v2_struct_0(X1)|~v1_xboole_0(esk7_2(X2,X1))|~m1_subset_1(esk7_2(X2,X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X2,a_1_0_waybel31(X1))|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_433, plain, (v5_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_441, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=X2|g1_orders_2(X1,esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_424, plain, (v1_xboole_0(esk5_2(X1,X2))|~v1_xboole_0(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_425, plain, (v3_ordinal1(u1_orders_2(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1))).#trainneg
cnf(i_0_436, plain, (v4_relat_1(u1_orders_2(X1),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_389, plain, (v1_xboole_0(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_393, plain, (v1_xboole_0(esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_462, plain, (m1_subset_1(esk5_2(X1,k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~r2_hidden(X1,a_1_0_waybel23(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2)))).#trainneg
cnf(i_0_528, plain, (r2_hidden(X1,a_1_0_waybel31(X2))|X1!=k1_card_1(esk7_2(X3,X2))|~m1_waybel23(esk7_2(X3,X2),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X3,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_479, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~v1_xboole_0(esk6_2(X1,X2))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_412, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk14_1(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_484, plain, (v1_xboole_0(esk6_2(X1,X2))|v2_struct_0(X2)|~v1_xboole_0(X1)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_508, plain, (v1_card_1(X1)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_514, plain, (v1_card_1(X1)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_517, plain, (v2_struct_0(X1)|~v1_xboole_0(esk4_2(X2,X1))|~m1_subset_1(esk4_2(X2,X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~r2_hidden(X2,a_1_3_waybel31(X1))|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_430, plain, (l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_431, plain, (v1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_417, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,u1_orders_2(k2_yellow_1(X2))))).#trainneg
cnf(i_0_457, plain, (v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_510, plain, (v1_xboole_0(esk7_2(X1,X2))|~v1_xboole_0(X1)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_516, plain, (v1_xboole_0(esk4_2(X1,X2))|~v1_xboole_0(X1)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_428, plain, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~r2_hidden(X1,u1_orders_2(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_469, plain, (v3_ordinal1(k11_waybel_0(X1,X2))|v2_struct_0(X1)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_487, plain, (v3_ordinal1(esk6_2(X1,X2))|v2_struct_0(X2)|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_559, negated_conjecture, (~v1_finset_1(u1_struct_0(esk1_0))|~l1_orders_2(esk1_0))).#trainneg
cnf(i_0_565, plain, (u1_pre_topc(X1)=u1_pre_topc(X2)|~m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(u1_pre_topc(X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(u1_pre_topc(X1),X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_550, negated_conjecture, (v1_tops_2(u1_struct_0(esk2_0),esk1_0))).#trainneg
cnf(i_0_547, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(X1,esk1_0))).#trainneg
cnf(i_0_544, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_239, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_447, plain, (X1=X2|k2_yellow_1(X1)!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_663, plain, (v1_card_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_560, negated_conjecture, (~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_551, plain, (v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_541, negated_conjecture, (v3_lattice3(k2_yellow_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_568, plain, (v1_finset_1(k1_card_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(k11_waybel_0(X2,X1),k4_ordinal1))).#trainneg
cnf(i_0_574, plain, (m1_subset_1(esk3_2(u1_pre_topc(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(u1_pre_topc(X1),X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_570, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_pre_topc(X2)))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_584, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_566, plain, (u1_pre_topc(X1)=X2|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_598, plain, (v3_waybel_3(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_237, plain, (v5_relat_1(k1_xboole_0,X2))).#trainneg
cnf(i_0_542, negated_conjecture, (v2_waybel_1(k2_yellow_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_442, plain, (u1_orders_2(k2_yellow_1(X1))=X2|k2_yellow_1(X1)!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_181, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_238, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_587, plain, (~m1_subset_1(k1_numbers,k4_ordinal1))).#trainneg
cnf(i_0_543, negated_conjecture, (~v7_struct_0(k2_yellow_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_585, plain, (v2_struct_0(k2_yellow_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_564, negated_conjecture, (u1_struct_0(esk2_0)=X1|~r1_tarski(u1_struct_0(esk2_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(X1,esk1_0))).#trainneg
cnf(i_0_573, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_798, plain, (~v2_lattice3(k2_yellow_1(X1))|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_578, plain, (v1_finset_1(k1_card_1(X1))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~v1_finset_1(u1_struct_0(X2)))).#trainneg
cnf(i_0_729, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(X2))).#trainneg
cnf(i_0_602, plain, (v1_finset_1(k1_card_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(k1_card_1(k11_waybel_0(X2,X1)),k4_ordinal1))).#trainneg
cnf(i_0_647, plain, (v1_yellow_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_603, plain, (v1_finset_1(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_710, plain, (u1_struct_0(X1)=X2|X1!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_567, plain, (v1_finset_1(k1_card_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~v1_xboole_0(k11_waybel_0(X2,X1)))).#trainneg
cnf(i_0_599, plain, (~m1_subset_1(k1_card_1(k1_numbers),k4_ordinal1))).#trainneg
cnf(i_0_600, plain, (~m1_subset_1(k1_card_1(k4_ordinal1),k4_ordinal1))).#trainneg
cnf(i_0_646, plain, (v1_yellow_0(k2_yellow_1(X1))|~v1_finset_1(X1)|~v2_lattice3(k2_yellow_1(X1)))).#trainneg
cnf(i_0_597, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v7_struct_0(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1))).#trainneg
cnf(i_0_569, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_pre_topc(X2)))|v2_struct_0(X2)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_622, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk5_2(X2,X1),k4_ordinal1))).#trainneg
cnf(i_0_613, plain, (r2_hidden(X1,u1_pre_topc(X2))|~r2_hidden(X1,esk5_2(X3,X2))|~m1_subset_1(esk5_2(X3,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~r2_hidden(X3,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_623, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(k1_card_1(esk5_2(X2,X1)),k4_ordinal1))).#trainneg
cnf(i_0_727, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk6_2(X2,X1),k4_ordinal1))).#trainneg
cnf(i_0_714, plain, (v1_finset_1(esk11_1(X1))|~v1_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_764, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(esk5_2(X1,X2),k4_ordinal1))).#trainneg
cnf(i_0_611, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_648, plain, (v1_yellow_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(u1_struct_0(X1),k4_ordinal1))).#trainneg
cnf(i_0_604, plain, (v1_finset_1(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_639, negated_conjecture, (v1_finset_1(u1_struct_0(esk2_0))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_750, plain, (u1_orders_2(X1)=X2|X1!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v1_orders_2(X1))).#trainneg
cnf(i_0_575, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(u1_pre_topc(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,u1_pre_topc(X2)))).#trainneg
cnf(i_0_700, plain, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_802, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_662, plain, (k1_card_1(k1_xboole_0)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_661, plain, (v1_finset_1(k1_xboole_0)|~v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_676, plain, (X1=u1_orders_2(k2_yellow_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_580, plain, (v1_cantor_1(esk14_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|~m1_waybel23(esk14_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_803, plain, (r2_hidden(k1_card_1(esk11_1(X1)),a_1_3_waybel31(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_876, plain, (v1_finset_1(esk12_1(X1))|~v1_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_703, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k11_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2))).#trainneg
cnf(i_0_860, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_820, plain, (r2_hidden(k1_card_1(esk11_1(X1)),a_1_0_waybel31(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_728, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(k1_card_1(esk6_2(X2,X1)),k4_ordinal1))).#trainneg
cnf(i_0_621, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk5_2(X2,X1)))).#trainneg
cnf(i_0_866, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk4_2(X1,X2),k4_ordinal1))).#trainneg
cnf(i_0_649, plain, (v1_yellow_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_card_1(u1_struct_0(X1)),k4_ordinal1))).#trainneg
cnf(i_0_631, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_672, plain, (m1_subset_1(esk3_2(k4_ordinal1,X1),k1_numbers)|r1_tarski(k4_ordinal1,X1))).#trainneg
cnf(i_0_582, plain, (v1_cantor_1(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~m1_waybel23(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(u1_pre_topc(X2)))))).#trainneg
cnf(i_0_624, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_650, plain, (X1=u1_struct_0(X2)|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_723, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,u1_pre_topc(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2)))).#trainneg
cnf(i_0_765, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(k1_card_1(esk5_2(X1,X2)),k4_ordinal1))).#trainneg
cnf(i_0_745, plain, (r2_hidden(X1,X2)|v2_struct_0(k2_yellow_1(X2))|~r2_hidden(X1,k11_waybel_0(k2_yellow_1(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_742, plain, (v1_finset_1(k11_waybel_0(k2_yellow_1(X1),X2))|v2_struct_0(k2_yellow_1(X1))|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_823, plain, (v1_finset_1(esk7_2(X1,X2))|~v1_finset_1(u1_struct_0(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_838, plain, (r2_hidden(k1_card_1(esk5_2(X1,X2)),a_1_0_waybel23(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_837, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_704, plain, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_730, plain, (m1_subset_1(esk3_2(esk5_2(X1,X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|r1_tarski(esk5_2(X1,X2),X3)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_726, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk6_2(X2,X1)))).#trainneg
cnf(i_0_630, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_591, plain, (v1_finset_1(esk6_2(X1,X2))|v2_struct_0(X2)|~v1_xboole_0(X1)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_724, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,u1_pre_topc(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2)))).#trainneg
cnf(i_0_817, plain, (v1_xboole_0(u1_pre_topc(k2_yellow_1(X1)))|m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X2,u1_pre_topc(k2_yellow_1(X1)))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_592, plain, (v1_finset_1(esk5_2(X1,X2))|~v1_xboole_0(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_825, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk7_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_869, plain, (r2_hidden(k1_card_1(esk4_2(X1,X2)),a_1_3_waybel31(X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_852, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk7_2(X1,X2),k4_ordinal1))).#trainneg
cnf(i_0_826, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,esk7_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_768, plain, (k1_card_1(X1)=X1|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_770, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk6_2(X1,X2),k4_ordinal1))).#trainneg
cnf(i_0_907, plain, (r2_hidden(k1_card_1(esk5_2(X1,X2)),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_640, plain, (u1_pre_topc(X1)=X2|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(X2,X1)|~l1_pre_topc(X1)|~v1_xboole_0(u1_pre_topc(X1)))).#trainneg
cnf(i_0_785, plain, (v1_finset_1(esk14_1(k2_yellow_1(X1)))|v2_struct_0(k2_yellow_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_855, plain, (r2_hidden(k1_card_1(esk6_2(X1,X2)),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_911, plain, (r2_hidden(k1_card_1(esk7_2(X1,X2)),a_1_0_waybel31(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_945, plain, (r2_hidden(k1_card_1(esk4_2(X1,X2)),a_1_0_waybel31(X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_628, plain, (X1=X2|g1_orders_2(X1,k2_zfmisc_1(X1,X1))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_633, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_673, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_606, plain, (v1_tops_2(esk14_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|~m1_waybel23(esk14_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_651, plain, (u1_struct_0(X1)=k1_xboole_0|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_678, plain, (u1_orders_2(k2_yellow_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_758, plain, (X1=esk8_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_949, negated_conjecture, (v1_orders_2(esk2_0))).#trainneg
cnf(i_0_607, plain, (v1_tops_2(esk11_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(X1)|~m1_waybel23(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_608, plain, (v1_tops_2(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~m1_waybel23(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(u1_pre_topc(X2)))))).#trainneg
cnf(i_0_1142, plain, (u1_struct_0(X1)=X2|X1!=k2_yellow_1(X2)|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1230, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_1024, negated_conjecture, (m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_1002, plain, (v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_994, plain, (X1=X2|k2_yellow_1(X1)!=k2_yellow_1(X2))).#trainneg
cnf(i_0_1139, plain, (v1_finset_1(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_617, plain, (m1_subset_1(esk14_1(k2_yellow_1(u1_pre_topc(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v2_struct_0(X1)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|~m1_waybel23(esk14_1(k2_yellow_1(u1_pre_topc(X1))),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_956, negated_conjecture, (v1_relat_2(u1_orders_2(esk2_0)))).#trainneg
cnf(i_0_1222, plain, (k1_card_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_619, plain, (m1_subset_1(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~m1_waybel23(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(u1_pre_topc(X2)))))).#trainneg
cnf(i_0_1516, plain, (v1_relat_2(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_625, plain, (l1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1171, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~m1_subset_1(X2,k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1301, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(X1),X2),X3),X1)|v2_struct_0(k2_yellow_1(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(X1),X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1397, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(k1_card_1(esk5_2(X2,X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1314, plain, (v8_struct_0(X1)|~v1_finset_1(k1_card_1(esk5_2(X2,X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1270, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~m1_subset_1(X2,k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1081, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1189, plain, (X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1093, negated_conjecture, (r2_hidden(X1,a_1_0_waybel31(X2))|X1!=k1_card_1(k11_waybel_0(X2,X3))|~m1_waybel23(k11_waybel_0(X2,X3),X2)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~v1_yellow_0(X2))).#trainneg
cnf(i_0_957, negated_conjecture, (v4_relat_2(u1_orders_2(esk2_0)))).#trainneg
cnf(i_0_958, negated_conjecture, (v8_relat_2(u1_orders_2(esk2_0)))).#trainneg
cnf(i_0_626, plain, (v1_orders_2(g1_orders_2(X1,k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1517, plain, (v4_relat_2(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1009, negated_conjecture, (~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_1010, negated_conjecture, (~m1_subset_1(k1_card_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_1518, plain, (v8_relat_2(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_627, plain, (k2_zfmisc_1(X1,X1)=X2|g1_orders_2(X1,k2_zfmisc_1(X1,X1))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_629, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1023, negated_conjecture, (v7_struct_0(X1)|~v1_xboole_0(k11_waybel_0(X1,X2))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_1460, plain, (k1_xboole_0=X1|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1149, plain, (v3_lattice3(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v7_struct_0(X1)|~v2_lattice3(X1))).#trainneg
cnf(i_0_1363, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_959, negated_conjecture, (v1_relat_1(u1_orders_2(esk2_0)))).#trainneg
cnf(i_0_1279, plain, (r1_tarski(k4_ordinal1,k1_numbers))).#trainneg
cnf(i_0_632, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_655, plain, (k1_xboole_0=X1|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2)))).#trainneg
cnf(i_0_1044, plain, (v1_finset_1(esk5_2(X1,X2))|~v1_finset_1(u1_pre_topc(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_981, negated_conjecture, (u1_pre_topc(X1)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(u1_struct_0(esk2_0),X1)|~l1_pre_topc(X1)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(u1_pre_topc(X1),esk1_0))).#trainneg
cnf(i_0_1126, plain, (r1_tarski(k11_waybel_0(k2_yellow_1(X1),X2),X3)|m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(X1),X2),X3),X1)|v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1215, plain, (r1_tarski(X1,k1_numbers)|~r2_hidden(esk3_2(X1,k1_numbers),k4_ordinal1))).#trainneg
cnf(i_0_982, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(X1,esk1_0))).#trainneg
cnf(i_0_965, negated_conjecture, (g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))=esk2_0)).#trainneg
cnf(i_0_1258, plain, (r2_hidden(esk3_2(k11_waybel_0(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1193, plain, (k11_waybel_0(X1,X2)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k11_waybel_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1069, plain, (l1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1071, plain, (k1_xboole_0=X1|g1_orders_2(X2,k1_xboole_0)!=g1_orders_2(X3,X1))).#trainneg
cnf(i_0_656, plain, (k1_xboole_0=X1|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2)))).#trainneg
cnf(i_0_1188, plain, (k1_numbers=k4_ordinal1|~m1_subset_1(k1_numbers,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_1358, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|v1_xboole_0(u1_pre_topc(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2))|~m1_subset_1(X1,u1_pre_topc(k2_yellow_1(X2))))).#trainneg
cnf(i_0_1283, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_983, negated_conjecture, (X1=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(X1,esk1_0))).#trainneg
cnf(i_0_1344, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(k1_card_1(esk6_2(X2,X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_1122, plain, (v1_finset_1(k1_card_1(k11_waybel_0(X1,X2)))|v2_struct_0(X1)|~v1_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1161, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(X2,esk2_0)|~v6_waybel23(X2,esk2_0))).#trainneg
cnf(i_0_1053, negated_conjecture, (v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~m1_waybel23(esk11_1(X1),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1254, plain, (m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1355, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1070, plain, (v1_orders_2(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_960, negated_conjecture, (v1_finset_1(u1_orders_2(esk2_0))|~v1_finset_1(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_670, plain, (k1_card_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1405, negated_conjecture, (v1_finset_1(esk14_1(esk2_0))|~v1_finset_1(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_721, plain, (v1_finset_1(u1_pre_topc(k2_yellow_1(X1)))|~v1_finset_1(k1_zfmisc_1(X1))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_671, plain, (u1_struct_0(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_677, plain, (u1_orders_2(k2_yellow_1(X1))=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1604, negated_conjecture, (~v1_finset_1(k11_waybel_0(X1,X2))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1718, negated_conjecture, (r2_hidden(X1,a_1_0_waybel31(X2))|v1_finset_1(X3)|X1!=k1_card_1(X3)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1561, negated_conjecture, (m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1565, negated_conjecture, (m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_1713, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_3_waybel31(X1))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1587, plain, (v1_yellow_0(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v2_struct_0(k2_yellow_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1618, plain, (v1_relat_2(k1_xboole_0))).#trainneg
cnf(i_0_1730, plain, (v4_relat_2(k1_xboole_0))).#trainneg
cnf(i_0_1870, negated_conjecture, (u1_orders_2(esk2_0)=k1_xboole_0|g1_orders_2(X1,k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_679, plain, (v3_ordinal1(u1_orders_2(k2_yellow_1(X1)))|~v1_card_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1741, plain, (v8_relat_2(k1_xboole_0))).#trainneg
cnf(i_0_1578, plain, (~v2_struct_0(k2_yellow_1(k1_numbers)))).#trainneg
cnf(i_0_1579, plain, (~v2_struct_0(k2_yellow_1(k4_ordinal1)))).#trainneg
cnf(i_0_2097, plain, (v1_finset_1(k1_card_1(X2))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_680, plain, (l1_orders_2(g1_orders_2(X1,esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2)))|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2))).#trainneg
cnf(i_0_681, plain, (v1_orders_2(g1_orders_2(X1,esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2)))|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2))).#trainneg
cnf(i_0_1716, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel31(X1))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1859, negated_conjecture, (u1_struct_0(esk2_0)=X1|g1_orders_2(X1,X2)!=esk2_0)).#trainneg
cnf(i_0_1864, plain, (r1_tarski(k11_waybel_0(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1842, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0))).#trainneg
cnf(i_0_1797, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_pre_topc(X1)))).#trainneg
cnf(i_0_1483, plain, (v5_relat_1(u1_orders_2(X1),k1_xboole_0)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1567, negated_conjecture, (X1=u1_struct_0(esk2_0)|k2_yellow_1(X1)!=esk2_0)).#trainneg
cnf(i_0_682, plain, (esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2)=X3|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2)|g1_orders_2(X1,esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X2))!=g1_orders_2(X4,X3))).#trainneg
cnf(i_0_1514, plain, (v1_partfun1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1586, negated_conjecture, (~v2_struct_0(k2_yellow_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_683, plain, (X1=X2|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X3)|g1_orders_2(X1,esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X1)),X3))!=g1_orders_2(X2,X4))).#trainneg
cnf(i_0_684, plain, (v1_xboole_0(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1536, negated_conjecture, (v1_tops_2(esk11_1(esk2_0),esk1_0)|~m1_waybel23(esk11_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_1705, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~m1_subset_1(X2,k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_1686, negated_conjecture, (r2_hidden(esk3_2(k11_waybel_0(esk2_0,X1),X2),u1_struct_0(esk2_0))|r1_tarski(k11_waybel_0(esk2_0,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1583, plain, (v1_finset_1(k1_card_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~v2_struct_0(k2_yellow_1(k11_waybel_0(X2,X1))))).#trainneg
cnf(i_0_1560, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k11_waybel_0(X2,X3))|~m1_subset_1(k11_waybel_0(X2,X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1629, plain, (v8_struct_0(X1)|v1_xboole_0(a_1_0_waybel23(X1))|~m1_subset_1(X2,k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1484, plain, (v4_relat_1(u1_orders_2(X1),k1_xboole_0)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1710, plain, (esk8_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_1858, negated_conjecture, (u1_orders_2(esk2_0)=X1|g1_orders_2(X2,X1)!=esk2_0)).#trainneg
cnf(i_0_685, plain, (v5_relat_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X2)|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))).#trainneg
cnf(i_0_1080, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1542, negated_conjecture, (~v1_finset_1(esk14_1(esk2_0))|~m1_waybel23(esk14_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_686, plain, (v1_relat_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))).#trainneg
cnf(i_0_687, plain, (v1_xboole_0(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1563, negated_conjecture, (k11_waybel_0(X1,X2)=u1_struct_0(esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k11_waybel_0(X1,X2)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1550, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|~v1_finset_1(u1_struct_0(X1))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1553, negated_conjecture, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k11_waybel_0(X2,X3))|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1481, plain, (l1_orders_2(g1_orders_2(k1_xboole_0,u1_orders_2(X1)))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1709, plain, (esk14_1(k2_yellow_1(k1_xboole_0))=k1_xboole_0|v2_struct_0(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_688, plain, (v4_relat_1(esk3_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3),X1)|r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3))).#trainneg
cnf(i_0_1068, plain, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1776, plain, (k1_numbers=k4_ordinal1|~r1_tarski(k1_numbers,k4_ordinal1))).#trainneg
cnf(i_0_2079, plain, (v1_card_1(u1_struct_0(X2))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_689, plain, (v1_finset_1(esk3_2(k1_zfmisc_1(X1),X2))|r1_tarski(k1_zfmisc_1(X1),X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_690, plain, (v1_xboole_0(esk3_2(k1_zfmisc_1(X1),X2))|r1_tarski(k1_zfmisc_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1900, plain, (v1_finset_1(k11_waybel_0(X1,X2))|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1804, plain, (k11_waybel_0(X1,X2)=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r1_tarski(u1_struct_0(X1),k11_waybel_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1185, plain, (esk5_2(X2,X1)=u1_pre_topc(X1)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(esk5_2(X2,X1)))|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1832, negated_conjecture, (r1_tarski(k11_waybel_0(esk2_0,X1),X2)|m1_subset_1(esk3_2(k11_waybel_0(esk2_0,X1),X2),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1799, negated_conjecture, (u1_pre_topc(X1)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(u1_struct_0(esk2_0),X1)|~l1_pre_topc(X1)|~m1_waybel23(u1_pre_topc(X1),esk2_0)|~v6_waybel23(u1_pre_topc(X1),esk2_0))).#trainneg
cnf(i_0_1026, plain, (v8_struct_0(X1)|v1_xboole_0(a_1_0_waybel23(X1))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1910, negated_conjecture, (m1_subset_1(esk11_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1465, plain, (v4_orders_2(g1_orders_2(k1_xboole_0,u1_orders_2(X1)))|~l1_orders_2(X1)|~v4_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_691, plain, (m1_subset_1(X1,X2)|r1_tarski(k1_zfmisc_1(X2),X3)|~r2_hidden(X1,esk3_2(k1_zfmisc_1(X2),X3)))).#trainneg
cnf(i_0_1749, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1778, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_693, plain, (v3_ordinal1(esk3_2(X1,X2))|r1_tarski(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_695, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk14_1(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_1552, negated_conjecture, (k11_waybel_0(X1,X2)=u1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k11_waybel_0(X1,X2)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1395, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(esk6_2(k1_card_1(esk5_2(X2,X1)),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1655, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_1151, plain, (v1_finset_1(esk6_2(X1,X2))|v2_struct_0(X2)|~v1_finset_1(u1_pre_topc(X2))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1576, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(esk5_2(X2,X1))))).#trainneg
cnf(i_0_1466, plain, (v5_orders_2(g1_orders_2(k1_xboole_0,u1_orders_2(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1206, negated_conjecture, (v1_finset_1(u1_struct_0(esk2_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_696, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk8_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1218, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v1_card_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1975, plain, (v1_finset_1(k1_card_1(X2))|~v1_finset_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_697, plain, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X1,esk6_2(X3,X2))|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_698, plain, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,esk5_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_985, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_1261, plain, (v8_struct_0(X1)|v1_xboole_0(a_1_0_waybel23(X1))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1038, negated_conjecture, (m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X1|~m1_waybel23(esk11_1(X1),X1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1891, negated_conjecture, (X1=u1_struct_0(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(X1))|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0))).#trainneg
cnf(i_0_1187, plain, (esk3_2(u1_pre_topc(X1),X2)=u1_struct_0(X1)|r1_tarski(u1_pre_topc(X1),X2)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(esk3_2(u1_pre_topc(X1),X2)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1025, plain, (v8_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1803, plain, (k11_waybel_0(X1,X2)=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k11_waybel_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1305, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~v1_finset_1(u1_struct_0(X2)))).#trainneg
cnf(i_0_1909, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk11_1(X2))|~m1_subset_1(esk11_1(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|esk2_0!=X2|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1653, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~m1_subset_1(esk6_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_1303, plain, (v1_finset_1(k1_card_1(X1))|v2_struct_0(k2_yellow_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_lattice3(k2_yellow_1(X2))|~v1_yellow_0(k2_yellow_1(X2))|~v1_finset_1(X2))).#trainneg
cnf(i_0_2463, plain, (v2_struct_0(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_2481, plain, (v1_card_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_699, plain, (r2_hidden(X1,X2)|r1_tarski(k1_zfmisc_1(X2),X3)|~r2_hidden(X1,esk3_2(k1_zfmisc_1(X2),X3)))).#trainneg
cnf(i_0_701, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,u1_orders_2(k2_yellow_1(X2))))).#trainneg
cnf(i_0_702, plain, (r2_hidden(X1,k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~r2_hidden(X1,u1_orders_2(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_2622, negated_conjecture, (m1_subset_1(esk11_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2212, negated_conjecture, (esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v2_struct_0(k2_yellow_1(k11_waybel_0(X1,X2))))).#trainneg
cnf(i_0_2438, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2439, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2437, negated_conjecture, (v2_struct_0(X1)|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_954, negated_conjecture, (v1_partfun1(u1_orders_2(esk2_0),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_2189, plain, (v1_relat_1(k1_card_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2195, plain, (v1_relat_2(k1_card_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_711, plain, (k1_card_1(esk3_2(a_1_0_waybel23(X1),X2))=esk3_2(a_1_0_waybel23(X1),X2)|r1_tarski(a_1_0_waybel23(X1),X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_712, plain, (k1_card_1(X1)=X1|v1_xboole_0(a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_2453, negated_conjecture, (~v1_finset_1(u1_struct_0(esk2_0))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2649, negated_conjecture, (r2_hidden(esk3_2(esk11_1(X1),X2),u1_struct_0(esk2_0))|r1_tarski(esk11_1(X1),X2)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2361, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_2408, negated_conjecture, (r2_hidden(esk3_2(k11_waybel_0(X1,X2),X3),u1_struct_0(esk2_0))|r1_tarski(k11_waybel_0(X1,X2),X3)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2442, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(X1)),X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(X1)),X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2452, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2289, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v1_finset_1(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_2456, negated_conjecture, (m1_subset_1(esk3_2(k11_waybel_0(X1,X2),X3),u1_struct_0(X1))|r1_tarski(k11_waybel_0(X1,X2),X3)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_963, negated_conjecture, (v4_relat_1(u1_orders_2(esk2_0),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1584, negated_conjecture, (v1_finset_1(u1_struct_0(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1486, plain, (k1_xboole_0=X2|g1_orders_2(k1_xboole_0,u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_713, plain, (v3_ordinal1(esk11_1(X1))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2472, plain, (v3_ordinal1(k1_xboole_0)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1207, negated_conjecture, (v1_finset_1(u1_struct_0(esk2_0))|~m1_subset_1(k1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k4_ordinal1))).#trainneg
cnf(i_0_1884, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v1_card_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))))).#trainneg
cnf(i_0_2196, plain, (v4_relat_2(k1_card_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2197, plain, (v8_relat_2(k1_card_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_716, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk11_1(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_717, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,esk11_1(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_2552, plain, (r2_hidden(esk3_2(esk14_1(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(esk14_1(X1),X2)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2210, negated_conjecture, (esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k11_waybel_0(X1,X2),k4_ordinal1))).#trainneg
cnf(i_0_2436, negated_conjecture, (v2_struct_0(k2_yellow_1(u1_struct_0(X1)))|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk14_1(k2_yellow_1(u1_struct_0(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(k2_yellow_1(u1_struct_0(X1))),esk1_0)|~v1_tops_2(esk14_1(k2_yellow_1(u1_struct_0(X1))),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2329, negated_conjecture, (~v1_xboole_0(k1_card_1(k11_waybel_0(X1,X2)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_2454, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(u1_struct_0(X1),esk1_0)|~v1_tops_2(u1_struct_0(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_964, negated_conjecture, (v5_relat_1(u1_orders_2(esk2_0),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1072, plain, (X1=X2|g1_orders_2(X1,k1_xboole_0)!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_1243, negated_conjecture, (v1_cantor_1(esk14_1(esk2_0),esk1_0)|~m1_waybel23(esk14_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_720, plain, (v3_ordinal1(u1_pre_topc(k2_yellow_1(X1)))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1)))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_731, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(esk5_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,esk5_2(X3,X2)))).#trainneg
cnf(i_0_2293, negated_conjecture, (k11_waybel_0(esk2_0,X1)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k11_waybel_0(esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_2354, negated_conjecture, (m1_subset_1(esk11_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_2363, plain, (v8_struct_0(X1)|v1_xboole_0(a_1_0_waybel23(X1))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_2211, negated_conjecture, (esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_card_1(k11_waybel_0(X1,X2)),k4_ordinal1))).#trainneg
cnf(i_0_1482, plain, (v1_orders_2(g1_orders_2(k1_xboole_0,u1_orders_2(X1)))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_732, plain, (v3_ordinal1(k11_waybel_0(k2_yellow_1(X1),X2))|v2_struct_0(k2_yellow_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_733, plain, (l1_orders_2(g1_orders_2(X1,k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X1)),X2)))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_734, plain, (v1_orders_2(g1_orders_2(X1,k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X1)),X2)))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_2308, negated_conjecture, (m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_2638, plain, (k11_waybel_0(k2_yellow_1(X1),X2)=X1|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k11_waybel_0(k2_yellow_1(X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2449, negated_conjecture, (v2_struct_0(k2_yellow_1(u1_struct_0(X1)))|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~v1_tops_2(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2235, negated_conjecture, (m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_2594, plain, (r2_hidden(esk3_2(esk6_2(X1,X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|r1_tarski(esk6_2(X1,X2),X3)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_2562, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~m1_subset_1(esk6_2(X2,X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2567, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X2)))).#trainneg
cnf(i_0_1109, plain, (~v2_struct_0(k2_yellow_1(X1))|~v3_waybel_3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1))|~v1_yellow_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_735, plain, (k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X1)),X2)=X3|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1)))|g1_orders_2(X1,k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X1)),X2))!=g1_orders_2(X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1444, negated_conjecture, (v1_tops_2(esk14_1(esk2_0),esk1_0)|~m1_waybel23(esk14_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_736, plain, (X1=X2|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1)))|g1_orders_2(X1,k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X1)),X3))!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_737, plain, (v1_xboole_0(k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X2)),X3))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2643, plain, (v1_finset_1(k1_card_1(esk11_1(X1)))|~v1_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2625, negated_conjecture, (esk11_1(esk2_0)=u1_struct_0(esk2_0)|~m1_waybel23(esk11_1(esk2_0),esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk11_1(esk2_0))))).#trainneg
cnf(i_0_2404, negated_conjecture, (r1_tarski(k11_waybel_0(esk2_0,X1),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_738, plain, (v5_relat_1(k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X2)),X3),X2)|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_739, plain, (v1_relat_1(k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X2)),X3))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_740, plain, (v1_xboole_0(k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X2)),X3))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2272, plain, (v1_finset_1(k11_waybel_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2446, negated_conjecture, (r1_tarski(u1_pre_topc(X1),X2)|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk3_2(u1_pre_topc(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(u1_pre_topc(X1),X2),esk1_0)|~v1_tops_2(esk3_2(u1_pre_topc(X1),X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2288, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v2_struct_0(k2_yellow_1(k11_waybel_0(esk2_0,X1))))).#trainneg
cnf(i_0_2237, negated_conjecture, (k11_waybel_0(esk2_0,X1)=u1_struct_0(esk2_0)|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k11_waybel_0(esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_2586, plain, (v1_finset_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_966, negated_conjecture, (m1_subset_1(esk14_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1485, plain, (u1_orders_2(X1)=X2|g1_orders_2(k1_xboole_0,u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_741, plain, (v4_relat_1(k11_waybel_0(k2_yellow_1(k2_zfmisc_1(X1,X2)),X3),X1)|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1498, plain, (g1_orders_2(X1,k1_xboole_0)=k2_yellow_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2172, plain, (v1_card_1(u1_orders_2(k2_yellow_1(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_744, plain, (m1_subset_1(X1,X2)|v2_struct_0(k2_yellow_1(X2))|~r2_hidden(X1,k11_waybel_0(k2_yellow_1(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2447, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~v1_cantor_1(k1_xboole_0,esk1_0)|~v1_tops_2(k1_xboole_0,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1253, plain, (~v1_xboole_0(esk4_2(k1_card_1(esk11_1(X1)),X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1168, plain, (r2_hidden(esk3_2(esk5_2(X1,X2),X3),u1_pre_topc(X2))|r1_tarski(esk5_2(X1,X2),X3)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2695, negated_conjecture, (~v1_finset_1(u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_2889, plain, (r1_tarski(esk14_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2965, negated_conjecture, (~v1_finset_1(k1_card_1(k11_waybel_0(X1,X2)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3181, negated_conjecture, (u1_orders_2(esk2_0)=k1_xboole_0|k2_yellow_1(X1)!=esk2_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2691, negated_conjecture, (~v1_xboole_0(esk11_1(esk2_0)))).#trainneg
cnf(i_0_2692, negated_conjecture, (~m1_subset_1(esk11_1(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_2685, plain, (v1_card_1(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_751, plain, (v1_card_1(esk3_2(a_1_0_waybel23(X1),X2))|r1_tarski(a_1_0_waybel23(X1),X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_752, plain, (v1_card_1(X1)|v1_xboole_0(a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_2817, negated_conjecture, (r1_tarski(k11_waybel_0(X1,X2),u1_struct_0(esk2_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2814, negated_conjecture, (m1_subset_1(esk3_2(k11_waybel_0(X1,X2),X3),u1_struct_0(esk2_0))|r1_tarski(k11_waybel_0(X1,X2),X3)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_755, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_2694, negated_conjecture, (~v2_struct_0(k2_yellow_1(esk11_1(esk2_0))))).#trainneg
cnf(i_0_2726, negated_conjecture, (~v1_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2752, plain, (v1_relat_1(k1_card_1(X1))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_2765, plain, (v1_relat_2(k1_card_1(X1))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_756, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~m1_subset_1(k2_zfmisc_1(X1,X1),k4_ordinal1))).#trainneg
cnf(i_0_757, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~m1_subset_1(k1_card_1(k2_zfmisc_1(X1,X1)),k4_ordinal1))).#trainneg
cnf(i_0_2890, plain, (r2_hidden(esk3_2(esk14_1(k2_yellow_1(X1)),X2),X1)|r1_tarski(esk14_1(k2_yellow_1(X1)),X2)|v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_2746, plain, (v1_relat_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_759, plain, (esk8_1(k1_zfmisc_1(X1))=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2693, negated_conjecture, (~m1_subset_1(k1_card_1(esk11_1(esk2_0)),k4_ordinal1))).#trainneg
cnf(i_0_2866, plain, (v4_relat_2(k1_card_1(X1))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_2879, plain, (v8_relat_2(k1_card_1(X1))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_760, plain, (esk8_1(k1_zfmisc_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_761, plain, (m1_subset_1(esk3_2(esk6_2(X1,X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|r1_tarski(esk6_2(X1,X2),X3)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3029, plain, (v1_finset_1(k11_waybel_0(k2_yellow_1(X1),X2))|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2737, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(esk11_1(X1),esk2_0)|~v6_waybel23(esk11_1(X1),esk2_0))).#trainneg
cnf(i_0_2824, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(esk4_2(k1_card_1(esk11_1(X1)),X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(esk4_2(k1_card_1(esk11_1(X1)),X1),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(esk11_1(X1)),X1),esk2_0))).#trainneg
cnf(i_0_2759, plain, (v1_relat_2(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3155, negated_conjecture, (v3_ordinal1(esk14_1(esk2_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_762, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|v1_xboole_0(esk6_2(X3,X2))|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,esk6_2(X3,X2)))).#trainneg
cnf(i_0_763, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2)))).#trainneg
cnf(i_0_767, plain, (k1_card_1(esk3_2(a_1_2_waybel31(X1),X2))=esk3_2(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2807, negated_conjecture, (r1_tarski(esk11_1(X1),u1_struct_0(esk2_0))|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2950, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(X1,X2)),a_1_2_waybel31(esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2816, negated_conjecture, (m1_subset_1(esk3_2(k11_waybel_0(X1,X2),X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(k11_waybel_0(X1,X2),X3)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2813, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_2915, plain, (esk11_1(k2_yellow_1(k1_xboole_0))=k1_xboole_0|~v3_waybel_3(k2_yellow_1(k1_xboole_0))|~v1_lattice3(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_2960, negated_conjecture, (~m1_subset_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),k4_ordinal1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2860, plain, (v4_relat_2(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2725, negated_conjecture, (v3_ordinal1(esk11_1(esk2_0))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_769, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2)))).#trainneg
cnf(i_0_771, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(k1_card_1(esk6_2(X1,X2)),k4_ordinal1))).#trainneg
cnf(i_0_775, plain, (v3_ordinal1(esk14_1(k2_yellow_1(X1)))|v2_struct_0(k2_yellow_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3065, plain, (v1_finset_1(X1)|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_2738, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(esk14_1(X1),esk2_0)|~v6_waybel23(esk14_1(X1),esk2_0))).#trainneg
cnf(i_0_2804, negated_conjecture, (m1_subset_1(esk3_2(esk11_1(X1),X2),u1_struct_0(esk2_0))|r1_tarski(esk11_1(X1),X2)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3044, negated_conjecture, (v2_struct_0(k2_yellow_1(u1_struct_0(X1)))|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk2_0)|~v6_waybel23(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk2_0))).#trainneg
cnf(i_0_2722, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk11_1(X1))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1574, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(esk6_2(X2,X1))))).#trainneg
cnf(i_0_2873, plain, (v8_relat_2(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3186, plain, (u1_orders_2(X1)=k1_xboole_0|X1!=k2_yellow_1(X2)|~v1_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_776, plain, (l1_orders_2(g1_orders_2(X1,esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X1)))))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_3214, plain, (v1_card_1(u1_orders_2(k2_yellow_1(X1)))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_777, plain, (v1_orders_2(g1_orders_2(X1,esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X1)))))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_778, plain, (esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X1)))=X2|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1)))|g1_orders_2(X1,esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X3,X2))).#trainneg
cnf(i_0_2886, plain, (m1_subset_1(esk3_2(esk14_1(X1),X2),u1_struct_0(X1))|r1_tarski(esk14_1(X1),X2)|v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_3025, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~v1_tops_2(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2972, plain, (v1_xboole_0(a_1_0_waybel23(X1))|v8_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2736, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(esk12_1(X1),esk2_0)|~v6_waybel23(esk12_1(X1),esk2_0))).#trainneg
cnf(i_0_3062, plain, (v1_finset_1(k1_card_1(esk6_2(X1,X2)))|v2_struct_0(X2)|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_2825, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_1696, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(k1_card_1(esk5_2(X2,X1))))).#trainneg
cnf(i_0_3183, plain, (u1_orders_2(k2_yellow_1(X1))=k1_xboole_0|k2_yellow_1(X1)!=k2_yellow_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_779, plain, (X1=X2|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1)))|g1_orders_2(X1,esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X1))))!=g1_orders_2(X2,X3))).#trainneg
cnf(i_0_780, plain, (v1_xboole_0(esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X2))))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_781, plain, (v5_relat_1(esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X2))),X2)|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3779, plain, (v1_xboole_0(a_1_0_waybel23(X1))|v8_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_1499, plain, (v4_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_782, plain, (v1_relat_1(esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X2))))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_783, plain, (v1_xboole_0(esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X2))))|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_784, plain, (v4_relat_1(esk14_1(k2_yellow_1(k2_zfmisc_1(X1,X2))),X1)|v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3246, negated_conjecture, (v1_tops_2(esk14_1(esk2_0),esk1_0)|~m1_subset_1(esk14_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1500, plain, (v5_orders_2(g1_orders_2(X1,k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3235, negated_conjecture, (~m1_subset_1(u1_struct_0(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_787, plain, (m1_subset_1(X1,X2)|v2_struct_0(k2_yellow_1(X2))|~r2_hidden(X1,esk14_1(k2_yellow_1(X2))))).#trainneg
cnf(i_0_788, plain, (r2_hidden(X1,X2)|v2_struct_0(k2_yellow_1(X2))|~r2_hidden(X1,esk14_1(k2_yellow_1(X2))))).#trainneg
cnf(i_0_3665, plain, (~v3_waybel_3(k2_yellow_1(k1_xboole_0))|~v1_lattice3(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_3657, plain, (v1_xboole_0(a_1_2_waybel31(X1))|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3341, negated_conjecture, (~m1_subset_1(esk5_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),k4_ordinal1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3250, negated_conjecture, (v3_ordinal1(esk11_1(esk2_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_3182, plain, (u1_orders_2(X1)=k1_xboole_0|g1_orders_2(k1_xboole_0,u1_orders_2(X1))!=k2_yellow_1(X2)|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3236, negated_conjecture, (~m1_subset_1(k1_card_1(u1_struct_0(esk2_0)),k4_ordinal1))).#trainneg
cnf(i_0_3240, negated_conjecture, (~v2_struct_0(k2_yellow_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_792, plain, (v1_finset_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))).#trainneg
cnf(i_0_3645, negated_conjecture, (esk3_2(k11_waybel_0(X1,X2),X3)=u1_struct_0(esk1_0)|r1_tarski(k11_waybel_0(X1,X2),X3)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk3_2(k11_waybel_0(X1,X2),X3)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3327, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3303, negated_conjecture, (k11_waybel_0(X1,X2)=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_793, plain, (v1_finset_1(u1_orders_2(X1))|~l1_orders_2(X1)|~m1_subset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k4_ordinal1))).#trainneg
cnf(i_0_3241, negated_conjecture, (~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_3152, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_794, plain, (v1_finset_1(u1_orders_2(X1))|~l1_orders_2(X1)|~m1_subset_1(k1_card_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))),k4_ordinal1))).#trainneg
cnf(i_0_795, plain, (X1=u1_orders_2(X2)|~v1_xboole_0(X1)|~v1_xboole_0(u1_struct_0(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_3328, negated_conjecture, (~v1_xboole_0(a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3391, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),X2)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_finset_1(esk3_2(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),X2))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3409, plain, (v8_struct_0(X1)|r1_tarski(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),X2)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_finset_1(esk3_2(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),X2))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3247, plain, (r1_tarski(esk14_1(k2_yellow_1(X1)),X1)|v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_3570, plain, (r1_tarski(esk6_2(X2,X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1468, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~l1_pre_topc(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_3429, plain, (r1_tarski(esk14_1(k2_yellow_1(k4_ordinal1)),k1_numbers))).#trainneg
cnf(i_0_796, plain, (u1_orders_2(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3316, negated_conjecture, (~m1_subset_1(k1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k4_ordinal1))).#trainneg
cnf(i_0_797, plain, (u1_orders_2(X1)=k1_xboole_0|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3550, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0))).#trainneg
cnf(i_0_3027, negated_conjecture, (v1_finset_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2))|r2_hidden(X3,a_1_3_waybel31(X1))|v1_xboole_0(u1_struct_0(X1))|X3!=k1_card_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~v1_tops_2(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3340, negated_conjecture, (~v1_xboole_0(esk5_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3268, negated_conjecture, (esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(k11_waybel_0(X1,X2)))).#trainneg
cnf(i_0_1699, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(k1_card_1(esk5_2(X2,X1)))))).#trainneg
cnf(i_0_1262, plain, (m1_subset_1(k1_card_1(esk11_1(X1)),a_1_0_waybel31(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_805, plain, (v1_xboole_0(esk14_1(k2_yellow_1(X1)))|m1_subset_1(X2,X1)|v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X2,esk14_1(k2_yellow_1(X1))))).#trainneg
cnf(i_0_806, plain, (v1_xboole_0(esk14_1(X1))|m1_subset_1(X2,u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,esk14_1(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2806, negated_conjecture, (m1_subset_1(esk3_2(esk11_1(X1),X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk11_1(X1),X2)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3012, negated_conjecture, (m1_subset_1(esk4_2(k1_card_1(X2),X1),k1_zfmisc_1(u1_struct_0(X1)))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3405, plain, (v8_struct_0(X1)|r1_tarski(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),X2)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3389, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),X2)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4231, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_2_waybel31(esk1_0))|esk2_0!=X1|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4274, plain, (v8_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(X2)))).#trainneg
cnf(i_0_808, plain, (v1_xboole_0(esk8_1(k1_zfmisc_1(X1)))|m1_subset_1(X2,X1)|~m1_subset_1(X2,esk8_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_4246, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_2_waybel31(esk1_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4230, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3861, plain, (v1_xboole_0(a_1_2_waybel31(X1))|v8_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_4197, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4196, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4250, negated_conjecture, (v1_finset_1(u1_struct_0(X1))|r2_hidden(k1_card_1(u1_struct_0(X1)),a_1_2_waybel31(esk1_0))|esk2_0!=X1|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(u1_struct_0(X1),esk1_0)|~v1_tops_2(u1_struct_0(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_811, plain, (v1_xboole_0(esk3_2(k1_zfmisc_1(X1),X2))|m1_subset_1(X3,X1)|r1_tarski(k1_zfmisc_1(X1),X2)|~m1_subset_1(X3,esk3_2(k1_zfmisc_1(X1),X2)))).#trainneg
cnf(i_0_813, plain, (v1_xboole_0(u1_orders_2(k2_yellow_1(X1)))|m1_subset_1(X2,k2_zfmisc_1(X1,X1))|~m1_subset_1(X2,u1_orders_2(k2_yellow_1(X1))))).#trainneg
cnf(i_0_814, plain, (v1_xboole_0(u1_orders_2(X1))|m1_subset_1(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~m1_subset_1(X2,u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4226, negated_conjecture, (v1_finset_1(esk14_1(X1))|r2_hidden(k1_card_1(esk14_1(X1)),a_1_2_waybel31(esk1_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3975, negated_conjecture, (r1_tarski(esk14_1(esk2_0),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_4216, negated_conjecture, (v1_finset_1(u1_struct_0(X1))|r2_hidden(k1_card_1(u1_struct_0(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(u1_struct_0(X1),esk1_0)|~v1_tops_2(u1_struct_0(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3821, plain, (a_1_0_waybel23(X1)=k1_xboole_0|v8_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_4288, plain, (m1_subset_1(esk8_1(esk14_1(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_815, plain, (v1_xboole_0(k11_waybel_0(k2_yellow_1(X1),X2))|m1_subset_1(X3,X1)|v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X3,k11_waybel_0(k2_yellow_1(X1),X2))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3929, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_816, plain, (v1_xboole_0(k11_waybel_0(X1,X2))|m1_subset_1(X3,u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,k11_waybel_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4353, negated_conjecture, (v1_tops_2(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),esk1_0)|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~m1_subset_1(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3970, plain, (v1_tops_2(esk14_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|~m1_subset_1(esk14_1(k2_yellow_1(u1_pre_topc(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4347, negated_conjecture, (v1_tops_2(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),esk1_0)|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~m1_subset_1(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4192, negated_conjecture, (v1_finset_1(esk14_1(X1))|r2_hidden(k1_card_1(esk14_1(X1)),a_1_0_waybel23(esk1_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3973, plain, (esk14_1(k2_yellow_1(X1))=X1|v2_struct_0(k2_yellow_1(X1))|~r1_tarski(X1,esk14_1(k2_yellow_1(X1))))).#trainneg
cnf(i_0_4354, plain, (v1_tops_2(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),X2)|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~m1_subset_1(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1248, plain, (m1_subset_1(k1_card_1(esk11_1(X1)),a_1_3_waybel31(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_822, plain, (v3_ordinal1(esk7_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_828, plain, (v1_xboole_0(esk7_2(X1,X2))|m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X3,esk7_2(X1,X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_4351, plain, (X1=esk14_1(k2_yellow_1(a_1_0_waybel23(X2)))|v8_struct_0(X2)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X2)))|~r1_tarski(X1,esk14_1(k2_yellow_1(a_1_0_waybel23(X2))))|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3841, negated_conjecture, (esk14_1(esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(esk14_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk14_1(esk2_0))))).#trainneg
cnf(i_0_4348, plain, (v1_tops_2(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),X2)|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~m1_subset_1(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4350, plain, (esk14_1(k2_yellow_1(a_1_0_waybel23(X1)))=X2|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~m1_subset_1(X2,k1_zfmisc_1(esk14_1(k2_yellow_1(a_1_0_waybel23(X1)))))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4355, plain, (m1_subset_1(esk14_1(k2_yellow_1(a_1_2_waybel31(X1))),k1_zfmisc_1(X2))|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3231, plain, (r1_tarski(esk5_2(X1,X2),u1_pre_topc(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1582, plain, (v1_finset_1(X1)|~v2_struct_0(k2_yellow_1(k1_card_1(X1))))).#trainneg
cnf(i_0_3296, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))))).#trainneg
cnf(i_0_829, plain, (m1_subset_1(esk7_2(X1,k2_yellow_1(X2)),k1_zfmisc_1(X2))|~r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(X2)))|~v3_waybel_3(k2_yellow_1(X2))|~v1_lattice3(k2_yellow_1(X2)))).#trainneg
cnf(i_0_3295, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))))).#trainneg
cnf(i_0_830, plain, (k1_card_1(esk3_2(a_1_0_waybel31(X1),X2))=esk3_2(a_1_0_waybel31(X1),X2)|r1_tarski(a_1_0_waybel31(X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_831, plain, (k1_card_1(X1)=X1|v1_xboole_0(a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_3427, plain, (m1_subset_1(esk3_2(esk14_1(k2_yellow_1(X1)),X2),X1)|r1_tarski(esk14_1(k2_yellow_1(X1)),X2)|v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_4730, negated_conjecture, (~v1_finset_1(k1_card_1(esk11_1(esk2_0)))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_4812, plain, (v1_finset_1(esk6_2(X1,X2))|v2_struct_0(X2)|~v2_struct_0(k2_yellow_1(X1))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_4727, negated_conjecture, (~v1_xboole_0(k1_card_1(esk11_1(esk2_0)))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_834, plain, (v1_card_1(esk3_2(a_1_2_waybel31(X1),X2))|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_835, plain, (v1_card_1(X1)|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_4387, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(k1_card_1(esk6_2(X2,X1))))|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_4389, negated_conjecture, (~v2_struct_0(k2_yellow_1(k1_card_1(k11_waybel_0(X1,X2))))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4807, plain, (esk5_2(X2,X1)=u1_pre_topc(X1)|~r1_tarski(u1_pre_topc(X1),esk5_2(X2,X1))|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3297, plain, (v3_ordinal1(k1_xboole_0)|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_3199, plain, (k1_xboole_0=X1|k2_yellow_1(X2)!=g1_orders_2(X3,X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_836, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_839, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|X1!=k1_xboole_0|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X3,X2)))).#trainneg
cnf(i_0_842, plain, (m1_subset_1(esk3_2(esk8_1(k1_zfmisc_1(X1)),X2),X1)|r1_tarski(esk8_1(k1_zfmisc_1(X1)),X2))).#trainneg
cnf(i_0_4757, plain, (esk14_1(k2_yellow_1(a_1_0_waybel23(X1)))=esk14_1(k2_yellow_1(a_1_0_waybel23(X2)))|v8_struct_0(X2)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X2)))|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4756, plain, (esk14_1(k2_yellow_1(a_1_2_waybel31(X1)))=esk14_1(k2_yellow_1(a_1_0_waybel23(X2)))|v8_struct_0(X2)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X2)))|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4552, negated_conjecture, (~v1_finset_1(k1_card_1(u1_struct_0(X1)))|esk2_0!=X1|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(u1_struct_0(X1),esk1_0)|~v1_tops_2(u1_struct_0(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4386, plain, (v8_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(esk3_2(a_1_0_waybel23(X1),X2))))).#trainneg
cnf(i_0_1771, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(k1_card_1(esk6_2(X2,X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_3206, plain, (v1_card_1(u1_orders_2(k2_yellow_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4724, negated_conjecture, (~v1_xboole_0(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_4725, negated_conjecture, (~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k4_ordinal1)|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_4721, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_0_waybel23(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0)|~v1_tops_2(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_4811, plain, (v1_finset_1(esk7_2(X1,X2))|~v2_struct_0(k2_yellow_1(X1))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_3310, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3993, plain, (m1_subset_1(esk14_1(k2_yellow_1(k4_ordinal1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_847, plain, (v1_xboole_0(esk3_2(a_1_0_waybel23(X1),X2))|r1_tarski(a_1_0_waybel23(X1),X2)|~v1_xboole_0(esk5_2(esk3_2(a_1_0_waybel23(X1),X2),X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3441, plain, (v1_relat_1(u1_struct_0(X2))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_3594, plain, (v1_relat_2(u1_struct_0(X2))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_848, plain, (v1_xboole_0(X1)|v1_xboole_0(a_1_0_waybel23(X2))|~v1_xboole_0(esk5_2(X1,X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_4357, plain, (X1=esk14_1(k2_yellow_1(a_1_2_waybel31(X2)))|v8_struct_0(X2)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X2)))|v2_struct_0(X2)|~r1_tarski(X1,esk14_1(k2_yellow_1(a_1_2_waybel31(X2))))|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_4422, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_2_waybel31(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4810, negated_conjecture, (r1_tarski(esk5_2(X1,esk1_0),u1_struct_0(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_3640, negated_conjecture, (m1_subset_1(k1_card_1(X2),a_1_2_waybel31(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4356, plain, (esk14_1(k2_yellow_1(a_1_2_waybel31(X1)))=X2|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(esk14_1(k2_yellow_1(a_1_2_waybel31(X1)))))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3995, plain, (esk14_1(k2_yellow_1(k4_ordinal1))=k1_numbers|~r1_tarski(k1_numbers,esk14_1(k2_yellow_1(k4_ordinal1))))).#trainneg
cnf(i_0_4609, negated_conjecture, (esk14_1(esk2_0)=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),esk14_1(esk2_0)))).#trainneg
cnf(i_0_3704, plain, (v4_relat_2(u1_struct_0(X2))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_3751, plain, (v8_relat_2(u1_struct_0(X2))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_853, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(k1_card_1(esk7_2(X1,X2)),k4_ordinal1))).#trainneg
cnf(i_0_5124, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_4931, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(X2)))).#trainneg
cnf(i_0_5338, negated_conjecture, (v1_tops_2(esk5_2(X1,esk1_0),esk1_0)|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_5292, plain, (v1_relat_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_856, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_xboole_0|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X3,X2)))).#trainneg
cnf(i_0_859, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_861, plain, (m1_subset_1(esk4_2(esk3_2(a_1_3_waybel31(X1),X2),X1),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(a_1_3_waybel31(X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5334, negated_conjecture, (m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_4967, negated_conjecture, (esk5_2(X1,esk1_0)=u1_struct_0(esk2_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r1_tarski(u1_struct_0(esk2_0),esk5_2(X1,esk1_0)))).#trainneg
cnf(i_0_2715, negated_conjecture, (esk11_1(X1)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk11_1(X1)))|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5311, plain, (v1_relat_2(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_862, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,a_1_3_waybel31(X2)))).#trainneg
cnf(i_0_5275, plain, (v3_ordinal1(esk14_1(k2_yellow_1(k4_ordinal1)))|~v3_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5410, plain, (v4_relat_2(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_5429, plain, (v8_relat_2(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_863, plain, (k1_card_1(esk3_2(a_1_3_waybel31(X1),X2))=esk3_2(a_1_3_waybel31(X1),X2)|r1_tarski(a_1_3_waybel31(X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_864, plain, (k1_card_1(X1)=X1|v1_xboole_0(a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,a_1_3_waybel31(X2)))).#trainneg
cnf(i_0_4951, plain, (esk11_1(X1)=u1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(esk11_1(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4952, plain, (esk11_1(X1)=u1_struct_0(X1)|~r1_tarski(u1_struct_0(X1),esk11_1(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3244, plain, (esk14_1(X1)=u1_struct_0(X1)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),esk14_1(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_4243, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_2_waybel31(esk1_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_4813, plain, (v1_finset_1(esk5_2(X1,X2))|~v2_struct_0(k2_yellow_1(X1))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_4349, plain, (m1_subset_1(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),k1_zfmisc_1(X2))|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5400, plain, (v4_relat_2(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_3105, negated_conjecture, (esk11_1(esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk11_1(esk2_0))))).#trainneg
cnf(i_0_867, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(k1_card_1(esk4_2(X1,X2)),k4_ordinal1))).#trainneg
cnf(i_0_3243, plain, (esk14_1(X1)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(esk14_1(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_3773, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk2_0)|~v6_waybel23(k11_waybel_0(k2_yellow_1(u1_struct_0(X1)),X2),esk2_0))).#trainneg
cnf(i_0_5370, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(k1_card_1(k11_waybel_0(X1,X2)),a_1_2_waybel31(esk1_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2959, negated_conjecture, (~v1_xboole_0(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3063, plain, (v1_finset_1(k1_card_1(esk5_2(X1,X2)))|v2_struct_0(X2)|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_5419, plain, (v8_relat_2(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1100, negated_conjecture, (m1_subset_1(esk14_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_waybel23(esk14_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_873, plain, (m1_subset_1(esk6_2(esk3_2(a_1_2_waybel31(k2_yellow_1(X1)),X2),k2_yellow_1(X1)),k1_zfmisc_1(k1_zfmisc_1(X1)))|v2_struct_0(k2_yellow_1(X1))|r1_tarski(a_1_2_waybel31(k2_yellow_1(X1)),X2)|~l1_pre_topc(k2_yellow_1(X1))|~v6_pre_topc(k2_yellow_1(X1))|~v2_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_3255, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk11_1(esk2_0)))).#trainneg
cnf(i_0_3159, negated_conjecture, (esk14_1(esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk14_1(esk2_0))))).#trainneg
cnf(i_0_874, plain, (m1_subset_1(esk6_2(X1,k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|v2_struct_0(k2_yellow_1(X2))|v1_xboole_0(a_1_2_waybel31(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2))|~v6_pre_topc(k2_yellow_1(X2))|~v2_pre_topc(k2_yellow_1(X2))|~m1_subset_1(X1,a_1_2_waybel31(k2_yellow_1(X2))))).#trainneg
cnf(i_0_875, plain, (v3_ordinal1(esk12_1(X1))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4458, plain, (a_1_2_waybel31(X1)=u1_struct_0(X2)|v8_struct_0(X1)|v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_3002, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1))|r2_hidden(X3,a_1_3_waybel31(X1))|X3!=k1_card_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_2897, negated_conjecture, (v2_struct_0(k2_yellow_1(u1_struct_0(X1)))|esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk14_1(k2_yellow_1(u1_struct_0(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(esk14_1(k2_yellow_1(u1_struct_0(X1))),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(esk14_1(k2_yellow_1(u1_struct_0(X1))),esk2_0)|~v6_waybel23(esk14_1(k2_yellow_1(u1_struct_0(X1))),esk2_0))).#trainneg
cnf(i_0_5384, plain, (esk14_1(k2_yellow_1(a_1_2_waybel31(X1)))=k1_xboole_0|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4814, plain, (v1_finset_1(esk4_2(X1,X2))|~v2_struct_0(k2_yellow_1(X1))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_3817, plain, (v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_878, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk12_1(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_3256, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk11_1(esk2_0)))).#trainneg
cnf(i_0_3994, plain, (esk14_1(k2_yellow_1(k4_ordinal1))=k1_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk14_1(k2_yellow_1(k4_ordinal1)))))).#trainneg
cnf(i_0_2084, plain, (k1_card_1(X2)=X1|~v1_xboole_0(X2)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_879, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,esk12_1(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_881, plain, (v1_xboole_0(esk12_1(X1))|m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,esk12_1(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3972, plain, (esk14_1(k2_yellow_1(X1))=X1|v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(esk14_1(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_3431, negated_conjecture, (r2_hidden(esk3_2(esk14_1(esk2_0),X1),u1_struct_0(esk2_0))|r1_tarski(esk14_1(esk2_0),X1))).#trainneg
cnf(i_0_4339, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(esk4_2(k1_card_1(X2),X3),k1_zfmisc_1(u1_struct_0(X3)))|v1_finset_1(X2)|esk2_0!=X3|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5323, plain, (esk14_1(k2_yellow_1(a_1_2_waybel31(X1)))=esk14_1(k2_yellow_1(a_1_2_waybel31(X2)))|v8_struct_0(X2)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X2)))|v2_struct_0(X2)|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5539, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5465, plain, (v1_relat_1(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_882, plain, (m1_subset_1(esk12_1(k2_yellow_1(X1)),k1_zfmisc_1(X1))|~v3_waybel_3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1)))).#trainneg
cnf(i_0_5578, plain, (v1_relat_2(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_4730, negated_conjecture, (~v1_finset_1(k1_card_1(esk11_1(esk2_0)))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_4727, negated_conjecture, (~v1_xboole_0(k1_card_1(esk11_1(esk2_0)))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_5714, plain, (v4_relat_2(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_5768, plain, (v8_relat_2(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_884, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_xboole_0|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X3,X2)))).#trainneg
cnf(i_0_5749, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(k1_card_1(X2),a_1_2_waybel31(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5897, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(esk4_2(k1_card_1(X2),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~v3_waybel_3(X1))).#trainneg
cnf(i_0_5894, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(esk4_2(k1_card_1(X2),X1),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5457, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(X2))|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_5678, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(k1_card_1(X1),a_1_2_waybel31(esk1_0))|v1_finset_1(X1)|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_5806, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X2),X1))|r2_hidden(X3,a_1_3_waybel31(X1))|v1_finset_1(X2)|X3!=k1_card_1(esk4_2(k1_card_1(X2),X1))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(X2),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(esk4_2(k1_card_1(X2),X1),esk1_0)|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X2),X1),esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5813, plain, (v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4724, negated_conjecture, (~v1_xboole_0(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_4725, negated_conjecture, (~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k4_ordinal1)|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_893, plain, (v1_xboole_0(esk5_2(esk3_2(a_1_0_waybel23(X1),X2),X1))|r1_tarski(a_1_0_waybel23(X1),X2)|~v1_xboole_0(esk3_2(a_1_0_waybel23(X1),X2))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_894, plain, (v1_xboole_0(esk5_2(X1,X2))|v1_xboole_0(a_1_0_waybel23(X2))|~v1_xboole_0(X1)|~l1_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_5538, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_0_waybel23(esk1_0))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k1_card_1(X1),esk1_0))).#trainneg
cnf(i_0_5767, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~v1_finset_1(u1_pre_topc(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_5532, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk5_2(X1,X2),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_5449, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(esk3_2(a_1_2_waybel31(X1),X2))))).#trainneg
cnf(i_0_5807, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1))|r2_hidden(X3,a_1_3_waybel31(X1))|X3!=k1_card_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),X1),esk2_0))).#trainneg
cnf(i_0_3837, plain, (v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4721, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_0_waybel23(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_897, plain, (v3_ordinal1(u1_orders_2(X1))|~l1_orders_2(X1)|~v1_card_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))).#trainneg
cnf(i_0_5558, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(k1_card_1(esk5_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk5_2(X1,esk1_0),esk1_0)|~v1_tops_2(esk5_2(X1,esk1_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_5560, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(k1_card_1(esk5_2(X1,esk1_0)),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk5_2(X1,esk1_0),esk1_0)|~v1_tops_2(esk5_2(X1,esk1_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_5535, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_2_waybel31(X1),X2),esk1_0)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5458, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk5_2(X2,esk1_0))|~m1_subset_1(esk5_2(X2,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(X2,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1291, plain, (u1_struct_0(X1)=u1_struct_0(X2)|~l1_orders_2(X2)|~v2_struct_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_4228, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_899, plain, (X1=esk8_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_900, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_901, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6088, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(esk4_2(k1_card_1(X2),X1),k1_zfmisc_1(u1_struct_0(X1)))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_5982, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_6165, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_3_waybel31(esk2_0))|X3!=k1_card_1(k11_waybel_0(X1,X2))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6111, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X2),X3))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X4,a_1_3_waybel31(X3))|v1_finset_1(X2)|X4!=k1_card_1(esk4_2(k1_card_1(X2),X3))|esk2_0!=X3|~m1_subset_1(esk4_2(k1_card_1(X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(esk4_2(k1_card_1(X2),X3),esk1_0)|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X2),X3),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6084, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_6060, negated_conjecture, (v1_finset_1(esk11_1(X1))|m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6016, negated_conjecture, (m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_2845, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,k1_card_1(X4))!=g1_orders_2(X1,X3)|~v2_struct_0(k2_yellow_1(X2))|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_902, plain, (X1=esk8_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_903, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_904, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6083, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6057, negated_conjecture, (v1_finset_1(esk12_1(X1))|m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_905, plain, (m1_subset_1(esk5_2(esk3_2(a_1_0_waybel23(k2_yellow_1(X1)),X2),k2_yellow_1(X1)),k1_zfmisc_1(k1_zfmisc_1(X1)))|r1_tarski(a_1_0_waybel23(k2_yellow_1(X1)),X2)|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_6234, negated_conjecture, (~v1_xboole_0(a_1_0_waybel23(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_906, plain, (m1_subset_1(esk5_2(X1,k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|v1_xboole_0(a_1_0_waybel23(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2))|~m1_subset_1(X1,a_1_0_waybel23(k2_yellow_1(X2))))).#trainneg
cnf(i_0_908, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_xboole_0|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk5_2(X3,X2)))).#trainneg
cnf(i_0_6348, negated_conjecture, (r1_tarski(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0))|v1_finset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)))|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)),esk1_0)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6188, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_6218, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel23(X2))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X2)|~m1_waybel23(esk5_2(X1,X2),esk2_0)|~v6_waybel23(esk5_2(X1,X2),esk2_0))).#trainneg
cnf(i_0_6055, negated_conjecture, (v1_finset_1(esk14_1(X1))|m1_subset_1(esk4_2(k1_card_1(esk14_1(X1)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6353, negated_conjecture, (r2_hidden(esk3_2(esk5_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk5_2(X1,esk1_0),X2)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6098, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(X1),a_1_2_waybel31(esk1_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0))).#trainneg
cnf(i_0_6045, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(k1_card_1(X2),a_1_2_waybel31(esk1_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_6961, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(X3,a_1_3_waybel31(esk2_0))|X3!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_cantor_1(k11_waybel_0(X2,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|esk2_0!=X2|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0)|~v3_waybel_3(X2))).#trainneg
cnf(i_0_915, plain, (v1_xboole_0(esk3_2(a_1_2_waybel31(X1),X2))|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~v1_xboole_0(esk6_2(esk3_2(a_1_2_waybel31(X1),X2),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6940, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6166, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0))).#trainneg
cnf(i_0_6058, negated_conjecture, (m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_6228, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(X2))|v1_finset_1(X3)|X1!=k1_card_1(esk4_2(k1_card_1(X3),X2))|esk2_0!=X2|~m1_waybel23(esk4_2(k1_card_1(X3),X2),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(X3),X2),esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_6189, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_6996, negated_conjecture, (esk4_2(k1_card_1(esk11_1(X1)),esk2_0)=u1_struct_0(esk2_0)|v1_finset_1(esk11_1(X1))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0)))|esk2_0!=X1|~v1_cantor_1(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7021, negated_conjecture, (v3_ordinal1(esk12_1(esk2_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_916, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~v1_xboole_0(esk6_2(X1,X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_6032, negated_conjecture, (~v1_xboole_0(k1_card_1(esk11_1(esk2_0)))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_6707, negated_conjecture, (~v2_struct_0(k2_yellow_1(k1_card_1(esk11_1(esk2_0))))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_6099, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(X1),a_1_2_waybel31(esk1_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0))).#trainneg
cnf(i_0_919, plain, (v1_xboole_0(esk6_2(esk3_2(a_1_2_waybel31(X1),X2),X1))|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~v1_xboole_0(esk3_2(a_1_2_waybel31(X1),X2))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6867, plain, (v4_relat_1(k1_card_1(X3),X1)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_920, plain, (v1_xboole_0(esk6_2(X1,X2))|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~v1_xboole_0(X1)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_921, plain, (v1_card_1(esk3_2(a_1_0_waybel31(X1),X2))|r1_tarski(a_1_0_waybel31(X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6928, negated_conjecture, (v1_finset_1(k1_card_1(X1))|m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(u1_struct_0(X2)))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_6110, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k1_card_1(X1))|r2_hidden(X3,a_1_3_waybel31(X2))|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_cantor_1(k1_card_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(k1_card_1(X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_6349, negated_conjecture, (v1_finset_1(esk6_2(X1,X2))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_waybel23(esk6_2(X1,X2),esk2_0)|~v6_waybel23(esk6_2(X1,X2),esk2_0))).#trainneg
cnf(i_0_6958, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(X3,a_1_3_waybel31(X2))|v2_struct_0(X2)|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_cantor_1(k11_waybel_0(X2,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_7037, negated_conjecture, (v1_finset_1(esk12_1(esk2_0))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk12_1(esk2_0))|~m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(esk2_0),esk1_0)|~v1_tops_2(esk12_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_922, plain, (v1_card_1(X1)|v1_xboole_0(a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_7217, plain, (v5_relat_1(k1_card_1(X3),X2)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7597, negated_conjecture, (~v1_xboole_0(k1_card_1(esk11_1(esk2_0))))).#trainneg
cnf(i_0_7673, plain, (v4_relat_1(k1_card_1(X1),X2)|~v1_xboole_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_7780, plain, (v5_relat_1(k1_card_1(X1),X2)|~v1_xboole_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_923, plain, (v1_card_1(esk3_2(a_1_3_waybel31(X1),X2))|r1_tarski(a_1_3_waybel31(X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_924, plain, (v1_card_1(X1)|v1_xboole_0(a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,a_1_3_waybel31(X2)))).#trainneg
cnf(i_0_7412, negated_conjecture, (r2_hidden(k1_card_1(X1),a_1_0_waybel23(esk1_0))|v1_finset_1(X1)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_7398, negated_conjecture, (m1_subset_1(esk3_2(esk5_2(X1,esk1_0),X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk5_2(X1,esk1_0),X2)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7322, negated_conjecture, (~v1_xboole_0(a_1_0_waybel23(esk1_0))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_931, plain, (m1_subset_1(esk3_2(u1_orders_2(k2_yellow_1(X1)),X2),k2_zfmisc_1(X1,X1))|r1_tarski(u1_orders_2(k2_yellow_1(X1)),X2))).#trainneg
cnf(i_0_937, plain, (v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|r1_tarski(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~r2_hidden(esk3_2(X2,k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))),u1_orders_2(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_7534, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0)|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_7511, negated_conjecture, (r2_hidden(k1_card_1(X2),a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7396, negated_conjecture, (m1_subset_1(esk3_2(esk5_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk5_2(X1,esk1_0),X2)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7600, negated_conjecture, (m1_subset_1(k1_card_1(X1),a_1_2_waybel31(esk1_0))|v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_7487, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk6_2(X2,X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_6233, negated_conjecture, (m1_subset_1(k1_card_1(esk11_1(esk2_0)),a_1_0_waybel23(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_938, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)))|~r2_hidden(X1,u1_orders_2(X2))|~l1_orders_2(X2))).#trainneg
cnf(i_0_7633, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_2_waybel31(esk1_0))|v1_finset_1(X2)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7435, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_2_waybel31(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_7295, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|v1_finset_1(esk12_1(X1))|X2!=k1_card_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6693, negated_conjecture, (m1_subset_1(k1_card_1(esk11_1(esk2_0)),a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_941, plain, (r2_hidden(k1_card_1(esk7_2(X1,X2)),a_1_3_waybel31(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_7403, negated_conjecture, (m1_subset_1(k1_card_1(X1),a_1_2_waybel31(esk1_0))|v1_finset_1(X1)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_7632, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_2_waybel31(esk1_0))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k1_card_1(X1),esk1_0))).#trainneg
cnf(i_0_7711, negated_conjecture, (v1_finset_1(k1_card_1(X1))|m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k1_card_1(X1),esk1_0))).#trainneg
cnf(i_0_7599, negated_conjecture, (~v2_struct_0(k2_yellow_1(k1_card_1(esk11_1(esk2_0))))|~m1_waybel23(esk11_1(esk2_0),esk2_0))).#trainneg
cnf(i_0_955, negated_conjecture, (m1_subset_1(u1_orders_2(esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_6089, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_6949, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|v1_finset_1(X3)|X1!=k1_card_1(X3)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1885, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_970, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)))|~r2_hidden(X1,u1_orders_2(esk2_0)))).#trainneg
cnf(i_0_7843, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|~m1_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7842, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_0_waybel23(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8040, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|r2_hidden(k1_card_1(X2),a_1_2_waybel31(esk1_0))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_8237, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7962, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_3_waybel31(esk2_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7929, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_3_waybel31(esk2_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_7932, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7931, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_3_waybel31(esk2_0))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7883, negated_conjecture, (~v1_xboole_0(a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_990, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_992, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_7928, negated_conjecture, (v1_finset_1(esk12_1(X1))|r2_hidden(k1_card_1(esk12_1(X1)),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8044, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|v2_struct_0(X2)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_8045, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_7835, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8271, negated_conjecture, (esk4_2(k1_card_1(X1),esk2_0)=u1_struct_0(esk2_0)|v1_finset_1(X1)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk4_2(k1_card_1(X1),esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_8015, negated_conjecture, (v3_ordinal1(k1_card_1(esk11_1(esk2_0)))|~v3_ordinal1(a_1_0_waybel23(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_7056, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,k1_card_1(X2))!=g1_orders_2(X1,X3)|~v1_xboole_0(k1_card_1(X2))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1027, plain, (v3_ordinal1(esk3_2(u1_pre_topc(X1),X2))|r1_tarski(u1_pre_topc(X1),X2)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1012, plain, (v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~v1_xboole_0(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2))|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k1_zfmisc_1(u1_pre_topc(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1013, plain, (r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(u1_pre_topc(X2))))|v2_struct_0(X2)|X1!=k1_card_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X2)),X3))|~m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X2)),X3),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_pre_topc(X2)))|~v1_cantor_1(X3,X2)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_8162, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_8163, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8116, plain, (~v1_xboole_0(esk4_2(X1,X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_7857, negated_conjecture, (esk3_2(esk5_2(X1,esk1_0),X2)=u1_struct_0(esk1_0)|r1_tarski(esk5_2(X1,esk1_0),X2)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk3_2(esk5_2(X1,esk1_0),X2)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_8079, negated_conjecture, (v3_ordinal1(k1_card_1(esk11_1(esk2_0)))|~v3_ordinal1(a_1_2_waybel31(esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_7059, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,u1_struct_0(X5))!=g1_orders_2(X1,X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4)|~l1_orders_2(X5)|~v2_struct_0(X5))).#trainneg
cnf(i_0_1028, plain, (v1_finset_1(esk3_2(u1_pre_topc(X1),X2))|r1_tarski(u1_pre_topc(X1),X2)|~v1_finset_1(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6251, negated_conjecture, (~v1_xboole_0(esk5_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_1014, plain, (r2_hidden(X1,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|v2_struct_0(X2)|X1!=k1_card_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X2)),X3))|~m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X2)),X3),k2_yellow_1(u1_pre_topc(X2)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_pre_topc(X2)))|~v1_cantor_1(X3,X2)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1015, plain, (v1_tops_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X1)|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_7925, negated_conjecture, (v1_finset_1(esk14_1(X1))|r2_hidden(k1_card_1(esk14_1(X1)),a_1_3_waybel31(esk2_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6194, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r2_hidden(esk3_2(a_1_0_waybel23(X1),X2),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),X2),esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1029, plain, (v1_xboole_0(esk3_2(u1_pre_topc(X1),X2))|r1_tarski(u1_pre_topc(X1),X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1016, plain, (v1_cantor_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X1)|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_1017, plain, (m1_subset_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k2_yellow_1(u1_pre_topc(X1)))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_8046, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_7349, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)))|r1_tarski(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0))|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_waybel23(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)),esk2_0)|~v6_waybel23(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel23(esk1_0)),esk2_0))).#trainneg
cnf(i_0_7753, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(X3,a_1_3_waybel31(X2))|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(k11_waybel_0(X2,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(esk4_2(k1_card_1(X1),X2),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(X1),X2),esk2_0))).#trainneg
cnf(i_0_1030, plain, (m1_subset_1(X1,u1_struct_0(X2))|r1_tarski(u1_pre_topc(X2),X3)|~r2_hidden(X1,esk3_2(u1_pre_topc(X2),X3))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1031, plain, (r2_hidden(X1,u1_struct_0(X2))|r1_tarski(u1_pre_topc(X2),X3)|~r2_hidden(X1,esk3_2(u1_pre_topc(X2),X3))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1033, plain, (v1_xboole_0(esk3_2(u1_pre_topc(X1),X2))|m1_subset_1(X3,u1_struct_0(X1))|r1_tarski(u1_pre_topc(X1),X2)|~m1_subset_1(X3,esk3_2(u1_pre_topc(X1),X2))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6163, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk5_2(X1,esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_7665, plain, (v4_relat_1(k1_card_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1035, plain, (r1_tarski(u1_pre_topc(k2_yellow_1(X1)),X2)|m1_subset_1(esk3_2(u1_pre_topc(k2_yellow_1(X1)),X2),k1_zfmisc_1(X1))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_8897, plain, (v4_relat_1(u1_orders_2(k2_yellow_1(X3)),X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1036, negated_conjecture, (r1_tarski(u1_struct_0(esk2_0),X1)|m1_subset_1(esk3_2(u1_struct_0(esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1043, plain, (v3_ordinal1(esk5_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(u1_pre_topc(X2)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8361, plain, (m1_subset_1(a_1_0_waybel23(X1),k1_zfmisc_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8436, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8788, negated_conjecture, (r1_tarski(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0))|v1_finset_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0)))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0)),esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8365, plain, (v3_ordinal1(X1)|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_7772, plain, (v5_relat_1(k1_card_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1045, plain, (v1_xboole_0(esk5_2(X1,X2))|~v1_xboole_0(u1_pre_topc(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8527, negated_conjecture, (~v1_xboole_0(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0)))).#trainneg
cnf(i_0_8895, plain, (v4_relat_1(u1_struct_0(X3),X2)|~v1_xboole_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_6424, plain, (u1_struct_0(X2)=X1|~l1_orders_2(X2)|~v2_struct_0(k2_yellow_1(X1))|~v2_struct_0(X2))).#trainneg
cnf(i_0_1046, plain, (m1_subset_1(X1,u1_pre_topc(X2))|~r2_hidden(X1,esk5_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8426, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel31(esk2_0))|v1_finset_1(k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8362, plain, (a_1_2_waybel31(X1)=a_1_0_waybel23(X1)|v2_struct_0(X1)|~m1_subset_1(a_1_2_waybel31(X1),k1_zfmisc_1(a_1_0_waybel23(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8530, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_8417, negated_conjecture, (r2_hidden(k1_card_1(X2),a_1_0_waybel31(esk2_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1315, plain, (v8_struct_0(X1)|~v1_xboole_0(esk5_2(k1_card_1(esk5_2(X2,X1)),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_8522, negated_conjecture, (m1_subset_1(k1_card_1(esk11_1(esk2_0)),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_1049, plain, (v1_xboole_0(esk5_2(X1,X2))|m1_subset_1(X3,u1_pre_topc(X2))|~m1_subset_1(X3,esk5_2(X1,X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_7026, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk12_1(esk2_0)))).#trainneg
cnf(i_0_6252, negated_conjecture, (~m1_subset_1(esk5_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k4_ordinal1)|~v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_3818, plain, (X1=a_1_0_waybel23(X2)|v8_struct_0(X2)|~v1_xboole_0(X1)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v2_struct_0(k2_yellow_1(u1_pre_topc(X2))))).#trainneg
cnf(i_0_1057, plain, (esk6_2(X2,X1)=u1_pre_topc(X1)|v2_struct_0(X1)|~m1_subset_1(esk6_2(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(esk6_2(X2,X1)))|~l1_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8871, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(k1_card_1(esk5_2(X1,esk1_0)),a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_8490, negated_conjecture, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk3_2(X1,k1_zfmisc_1(u1_struct_0(esk1_0))),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_8381, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8377, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|v2_struct_0(X2)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_8383, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r2_hidden(esk3_2(a_1_0_waybel23(X1),X2),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_0_waybel23(X1),X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_0_waybel23(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8789, negated_conjecture, (r1_tarski(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0))|v1_finset_1(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)),esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8620, negated_conjecture, (r2_hidden(X1,a_1_0_waybel31(esk2_0))|X1!=k1_card_1(k11_waybel_0(esk2_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X2,esk1_0))).#trainneg
cnf(i_0_3438, plain, (v1_relat_1(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_7027, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk12_1(esk2_0)))).#trainneg
cnf(i_0_4448, plain, (X1=a_1_2_waybel31(X2)|v8_struct_0(X2)|v2_struct_0(X2)|~v1_xboole_0(X1)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v2_struct_0(k2_yellow_1(u1_pre_topc(X2))))).#trainneg
cnf(i_0_1124, plain, (v3_ordinal1(esk3_2(k11_waybel_0(X1,X2),X3))|r1_tarski(k11_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~v3_ordinal1(u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1127, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8673, negated_conjecture, (r2_hidden(k1_card_1(esk5_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(esk5_2(X1,esk1_0)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_8881, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3591, plain, (v1_relat_2(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_1128, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_6120, plain, (a_1_2_waybel31(X1)=k1_xboole_0|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4452, plain, (a_1_2_waybel31(X1)=k1_xboole_0|v8_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_8379, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_8378, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8672, negated_conjecture, (r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(esk6_2(X1,esk1_0)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_3701, plain, (v4_relat_2(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_1133, plain, (v1_yellow_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v1_xboole_0(u1_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))))|~l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v5_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v4_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v3_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v2_lattice3(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_3504, plain, (v1_card_1(esk8_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1140, plain, (u1_struct_0(g1_orders_2(X1,X2))=X1|~v1_orders_2(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_1150, plain, (v3_ordinal1(esk6_2(X1,X2))|v2_struct_0(X2)|~v3_ordinal1(k1_zfmisc_1(u1_pre_topc(X2)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_6220, plain, (r1_tarski(a_1_2_waybel31(X1),X2)|v8_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk3_2(a_1_2_waybel31(X1),X2)))).#trainneg
cnf(i_0_7006, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|r2_hidden(k1_card_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0)),a_1_2_waybel31(esk1_0))|v1_finset_1(esk11_1(X1))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|esk2_0!=X1|~v1_cantor_1(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6103, negated_conjecture, (v1_finset_1(esk4_2(X1,X3))|v1_finset_1(esk5_2(X1,X2))|r2_hidden(X4,a_1_3_waybel31(X3))|X4!=k1_card_1(esk4_2(X1,X3))|esk2_0!=X3|~m1_subset_1(esk4_2(X1,X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(esk4_2(X1,X3),esk1_0)|~v1_cantor_1(esk5_2(X1,X2),esk1_0)|~v1_tops_2(esk4_2(X1,X3),esk1_0)|~v1_tops_2(esk5_2(X1,X2),esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8992, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|r2_hidden(k1_card_1(X2),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_9061, plain, (v4_relat_1(u1_struct_0(X1),X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_9100, negated_conjecture, (~v1_xboole_0(a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_9391, plain, (v1_yellow_0(k2_yellow_1(X1))|~v1_xboole_0(X1)|~v2_lattice3(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1153, plain, (m1_subset_1(X1,u1_pre_topc(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk6_2(X3,X2))|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_9310, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_9098, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_0_waybel31(esk2_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1156, plain, (v1_xboole_0(esk6_2(X1,X2))|m1_subset_1(X3,u1_pre_topc(X2))|v2_struct_0(X2)|~m1_subset_1(X3,esk6_2(X1,X2))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_9104, negated_conjecture, (~v1_xboole_0(esk7_2(k1_card_1(esk11_1(esk2_0)),esk2_0)))).#trainneg
cnf(i_0_8955, plain, (v4_relat_1(u1_orders_2(k2_yellow_1(X1)),X2)|~v1_xboole_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_1158, plain, (v8_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_subset_1(esk5_2(esk3_2(a_1_0_waybel23(X1),X2),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1159, plain, (v8_struct_0(X1)|v1_xboole_0(a_1_0_waybel23(X1))|~m1_subset_1(esk5_2(X2,X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_9311, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_9143, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|r2_hidden(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_0_waybel31(esk2_0))|v1_finset_1(X1)|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_9266, negated_conjecture, (r2_hidden(X1,a_1_0_waybel31(esk2_0))|v1_finset_1(X2)|X1!=k1_card_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X2,esk1_0))).#trainneg
cnf(i_0_9255, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)))|r1_tarski(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~m1_waybel23(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)),esk2_0)|~v6_waybel23(esk3_2(a_1_0_waybel23(X1),a_1_2_waybel31(esk1_0)),esk2_0))).#trainneg
cnf(i_0_8984, plain, (v1_finset_1(a_1_0_waybel23(X1))|v2_struct_0(X1)|~v1_finset_1(a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_9130, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_0_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8993, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_9238, negated_conjecture, (r1_tarski(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1164, plain, (v7_struct_0(k2_yellow_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~v1_xboole_0(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1165, plain, (m1_subset_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k1_zfmisc_1(u1_pre_topc(X1)))|v2_struct_0(X1)|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_cantor_1(X2,X1)|~v2_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1169, plain, (r2_hidden(X1,u1_pre_topc(X2))|v1_xboole_0(esk5_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,esk5_2(X3,X2)))).#trainneg
cnf(i_0_9129, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk11_1(X1)),a_1_0_waybel31(esk2_0))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9209, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|~r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(X1,esk1_0))))).#trainneg
cnf(i_0_9256, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_0_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_8960, negated_conjecture, (esk3_2(u1_struct_0(esk2_0),X1)=u1_struct_0(esk1_0)|r1_tarski(u1_struct_0(esk2_0),X1)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk3_2(u1_struct_0(esk2_0),X1))))).#trainneg
cnf(i_0_9409, plain, (v1_card_1(esk8_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9181, negated_conjecture, (v3_ordinal1(k1_card_1(esk11_1(esk2_0)))|~v3_ordinal1(a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_1172, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~m1_subset_1(esk6_2(esk3_2(a_1_2_waybel31(X1),X2),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8242, negated_conjecture, (v1_finset_1(u1_orders_2(esk2_0))|~v1_finset_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1173, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~m1_subset_1(esk6_2(X2,X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_1174, plain, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_subset_1(esk5_2(esk3_2(a_1_0_waybel23(X1),X2),X1),k4_ordinal1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6966, negated_conjecture, (esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6910, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|r2_hidden(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_2_waybel31(esk1_0))|v1_finset_1(X1)|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_7809, negated_conjecture, (r1_tarski(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0))|v1_finset_1(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)))|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)),esk1_0)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8164, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_3748, plain, (v8_relat_2(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_1175, plain, (v1_finset_1(X1)|v1_xboole_0(a_1_0_waybel23(X2))|~m1_subset_1(esk5_2(X1,X2),k4_ordinal1)|~l1_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_8243, negated_conjecture, (v1_xboole_0(u1_orders_2(esk2_0))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8503, negated_conjecture, (u1_struct_0(esk1_0)=X1|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(X1))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_7556, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_3_waybel31(X4))|v1_finset_1(X2)|X3!=k1_card_1(esk4_2(k1_card_1(X2),X4))|esk2_0!=X4|~m1_waybel23(esk4_2(k1_card_1(X2),X4),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(X2),X4),esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X4)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X4)|~v3_waybel_3(X4)|~v1_lattice3(X4)|~v1_yellow_0(X4)|~v5_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7908, plain, (r1_tarski(u1_orders_2(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_7428, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_7806, plain, (r1_tarski(esk14_1(k2_yellow_1(a_1_0_waybel23(X1))),a_1_2_waybel31(X1))|v2_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6190, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_8505, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,X2)|~m1_subset_1(X2,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_9002, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))))))).#trainneg
cnf(i_0_1178, plain, (esk12_1(X1)=u1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(esk12_1(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2519, plain, (v4_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v2_struct_0(k2_yellow_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8241, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_9481, plain, (v4_relat_1(X1,X2)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1180, plain, (esk8_1(k1_zfmisc_1(X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(esk8_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_9497, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|r2_hidden(k1_card_1(X2),a_1_0_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_9599, negated_conjecture, (r2_hidden(k1_card_1(X1),a_1_0_waybel31(esk2_0))|v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_9568, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X2,a_1_0_waybel31(esk2_0))|X2!=k1_card_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0))).#trainneg
cnf(i_0_10008, plain, (v4_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_8708, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,X2)!=g1_orders_2(X1,X3)|~v1_xboole_0(X4)|~v1_xboole_0(X5)|~v2_struct_0(k2_yellow_1(X2)))).#trainneg
cnf(i_0_1181, plain, (esk7_2(X2,X1)=u1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(esk7_2(X2,X1)))|~r2_hidden(X2,a_1_0_waybel31(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1182, plain, (esk6_2(X2,X1)=k1_zfmisc_1(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(X1)),k1_zfmisc_1(esk6_2(X2,X1)))|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_9933, negated_conjecture, (r2_hidden(esk3_2(X1,X2),u1_struct_0(esk1_0))|r1_tarski(X1,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_9567, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_0_waybel31(esk2_0))|X3!=k1_card_1(k11_waybel_0(X1,X2))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9803, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)))|r1_tarski(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0))|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_waybel23(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)),esk2_0)|~v6_waybel23(esk3_2(a_1_2_waybel31(X1),a_1_2_waybel31(esk1_0)),esk2_0))).#trainneg
cnf(i_0_9681, negated_conjecture, (esk6_2(k1_card_1(X1),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(X1)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(X1),esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_9678, negated_conjecture, (esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9004, plain, (v3_ordinal1(X1)|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_1184, plain, (esk5_2(X2,X1)=k1_zfmisc_1(u1_struct_0(X1))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(X1)),k1_zfmisc_1(esk5_2(X2,X1)))|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_9730, negated_conjecture, (v1_finset_1(u1_orders_2(esk2_0))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_9019, plain, (v5_relat_1(u1_orders_2(k2_yellow_1(X3)),X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1186, plain, (esk3_2(k1_zfmisc_1(X1),X2)=X1|r1_tarski(k1_zfmisc_1(X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(esk3_2(k1_zfmisc_1(X1),X2))))).#trainneg
cnf(i_0_1190, plain, (k2_zfmisc_1(X1,X1)=u1_orders_2(k2_yellow_1(X1))|~m1_subset_1(k2_zfmisc_1(X1,X1),k1_zfmisc_1(u1_orders_2(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_9564, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(X2,a_1_0_waybel31(esk2_0))|X2!=k1_card_1(esk5_2(X1,esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_9888, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X2)))|v2_struct_0(X2)|~r2_hidden(X1,esk14_1(k2_yellow_1(a_1_0_waybel23(X2))))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_6920, negated_conjecture, (v1_finset_1(esk6_2(X1,X2))|m1_subset_1(esk4_2(X1,X3),k1_zfmisc_1(u1_struct_0(X3)))|v2_struct_0(X2)|esk2_0!=X3|~m1_waybel23(k11_waybel_0(X3,esk6_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(X3,esk6_2(X1,X2)),esk2_0)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_7035, negated_conjecture, (v1_finset_1(esk12_1(esk2_0))|r2_hidden(k1_card_1(esk12_1(esk2_0)),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(esk2_0),esk1_0)|~v1_tops_2(esk12_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_6912, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|r2_hidden(X3,a_1_3_waybel31(X2))|v1_finset_1(X1)|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_1345, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk5_2(k1_card_1(esk6_2(X2,X1)),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_7060, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,u1_orders_2(X5))!=g1_orders_2(X1,X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4)|~v1_xboole_0(u1_struct_0(X5))|~l1_orders_2(X5))).#trainneg
cnf(i_0_1191, plain, (k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))=u1_orders_2(X1)|~m1_subset_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),k1_zfmisc_1(u1_orders_2(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_9731, negated_conjecture, (v1_finset_1(u1_orders_2(esk2_0))|~m1_subset_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)),k4_ordinal1))).#trainneg
cnf(i_0_9733, negated_conjecture, (v1_finset_1(u1_orders_2(esk2_0))|~v2_struct_0(k2_yellow_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1192, plain, (k11_waybel_0(k2_yellow_1(X1),X2)=X1|v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k11_waybel_0(k2_yellow_1(X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1194, negated_conjecture, (k1_zfmisc_1(u1_struct_0(esk1_0))=u1_struct_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8376, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_7008, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|v1_finset_1(esk11_1(X1))|X2!=k1_card_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|esk2_0!=X1|~v1_cantor_1(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8320, negated_conjecture, (v1_finset_1(esk12_1(X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk12_1(X1))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8204, negated_conjecture, (v1_finset_1(X1)|m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_9218, negated_conjecture, (r2_hidden(k1_card_1(esk5_2(X1,esk1_0)),a_1_0_waybel31(esk2_0))|v1_finset_1(esk5_2(X1,esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_8117, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_1394, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk6_2(k1_card_1(esk5_2(X2,X1)),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1195, plain, (u1_pre_topc(k2_yellow_1(X1))=k1_zfmisc_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(u1_pre_topc(k2_yellow_1(X1))))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_2571, plain, (v5_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0))|~v2_struct_0(k2_yellow_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9017, plain, (v5_relat_1(u1_struct_0(X3),X2)|~v1_xboole_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_1196, plain, (u1_pre_topc(X1)=k1_zfmisc_1(u1_struct_0(X1))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(X1)),k1_zfmisc_1(u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1198, plain, (v1_yellow_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~m1_subset_1(u1_struct_0(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))),k4_ordinal1)|~l1_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v5_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v4_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v3_orders_2(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1)))|~v2_lattice3(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_7625, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk5_2(X1,X2),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8011, negated_conjecture, (v1_finset_1(X1)|~v1_xboole_0(esk6_2(k1_card_1(X1),esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_8668, negated_conjecture, (r2_hidden(esk3_2(a_1_2_waybel31(esk1_0),X1),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_2_waybel31(esk1_0),X1)|~m1_subset_1(esk6_2(esk3_2(a_1_2_waybel31(esk1_0),X1),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_7959, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_3_waybel31(esk2_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_1415, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk6_2(k1_card_1(esk6_2(X2,X1)),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_9099, negated_conjecture, (m1_subset_1(k1_card_1(esk11_1(esk2_0)),a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_1278, plain, (v3_ordinal1(esk3_2(k4_ordinal1,X1))|r1_tarski(k4_ordinal1,X1)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_10023, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v1_card_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1208, plain, (u1_orders_2(g1_orders_2(X1,X2))=X2|~v1_orders_2(g1_orders_2(X1,X2))|~l1_orders_2(g1_orders_2(X1,X2)))).#trainneg
cnf(i_0_1213, plain, (v1_xboole_0(u1_pre_topc(X1))|m1_subset_1(esk8_1(u1_pre_topc(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10093, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_10090, negated_conjecture, (m1_subset_1(esk3_2(X1,X2),u1_struct_0(esk1_0))|r1_tarski(X1,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_10036, negated_conjecture, (m1_subset_1(k1_card_1(X1),a_1_0_waybel31(esk2_0))|v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_10435, plain, (v5_orders_2(g1_orders_2(k1_xboole_0,k1_xboole_0)))).#trainneg
cnf(i_0_1235, plain, (esk8_1(k1_zfmisc_1(X1))=u1_orders_2(k2_yellow_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10183, plain, (v5_relat_1(u1_orders_2(k2_yellow_1(X1)),X2)|~v1_xboole_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_10463, plain, (v5_relat_1(u1_struct_0(X1),X2)|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_1236, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=u1_orders_2(k2_yellow_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1237, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=u1_orders_2(k2_yellow_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10316, negated_conjecture, (k11_waybel_0(esk2_0,X1)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(k11_waybel_0(esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_10211, negated_conjecture, (m1_subset_1(k1_card_1(esk12_1(esk2_0)),a_1_2_waybel31(esk1_0))|v1_finset_1(esk12_1(esk2_0))|~m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(esk2_0),esk1_0)|~v1_tops_2(esk12_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_10528, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0))|v1_finset_1(X1)|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_10140, negated_conjecture, (esk6_2(k1_card_1(X2),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(X2),esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10275, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk5_2(X2,X1))|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_10387, plain, (v1_finset_1(esk4_2(X1,X2))|~v1_finset_1(u1_struct_0(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_10455, plain, (v5_relat_1(u1_struct_0(X1),X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_10540, negated_conjecture, (v3_ordinal1(k1_card_1(esk11_1(esk2_0)))|~v3_ordinal1(a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_1240, plain, (u1_orders_2(k2_yellow_1(X1))=u1_orders_2(k2_yellow_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7515, negated_conjecture, (r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk4_2(X2,esk2_0))|~m1_waybel23(esk4_2(X2,esk2_0),esk2_0)|~r2_hidden(X2,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_7742, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k11_waybel_0(X2,k1_card_1(X1)))|v1_finset_1(k1_card_1(X1))|r2_hidden(X3,a_1_3_waybel31(X2))|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,k1_card_1(X1)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_cantor_1(k11_waybel_0(X2,k1_card_1(X1)),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(k11_waybel_0(X2,k1_card_1(X1)),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_10964, plain, (v5_relat_1(X1,X2)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1241, plain, (u1_orders_2(X1)=u1_orders_2(k2_yellow_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1242, plain, (u1_struct_0(X1)=u1_orders_2(k2_yellow_1(X2))|~v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1247, plain, (v1_cantor_1(esk11_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(X1)|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_10916, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0))).#trainneg
cnf(i_0_10944, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~v1_xboole_0(esk5_2(esk3_2(a_1_2_waybel31(X1),X2),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8356, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_3_waybel31(esk2_0))|v2_struct_0(X1)|X3!=k1_card_1(k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7293, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|r2_hidden(k1_card_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0)),a_1_2_waybel31(esk1_0))|v1_finset_1(esk12_1(X1))|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9677, negated_conjecture, (esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(k11_waybel_0(X1,X2))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0)))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1257, plain, (r2_hidden(X1,u1_pre_topc(X2))|v2_struct_0(X2)|v1_xboole_0(esk6_2(X3,X2))|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,esk6_2(X3,X2)))).#trainneg
cnf(i_0_7025, negated_conjecture, (esk12_1(esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk12_1(esk2_0))))).#trainneg
cnf(i_0_8205, negated_conjecture, (v1_finset_1(X1)|m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_9684, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(X1,esk1_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_9820, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_6955, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_8675, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_0_waybel23(esk1_0))|v1_finset_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(esk6_2(k1_card_1(k11_waybel_0(esk2_0,X1)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_3886, plain, (v1_finset_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_9236, negated_conjecture, (r1_tarski(esk14_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1259, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(k11_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~m1_subset_1(X1,k11_waybel_0(X2,X3)))).#trainneg
cnf(i_0_1260, plain, (v8_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7457, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k1_card_1(X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|esk2_0!=X3|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_cantor_1(k1_card_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(k1_card_1(X1),esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_11432, negated_conjecture, (v1_cantor_1(esk11_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_11444, negated_conjecture, (r2_hidden(k1_card_1(X1),a_1_0_waybel31(esk2_0))|v1_finset_1(X1)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_11455, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|~m1_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9732, negated_conjecture, (v1_finset_1(u1_orders_2(esk2_0))|~m1_subset_1(k1_card_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))),k4_ordinal1))).#trainneg
cnf(i_0_1271, plain, (v8_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~v1_xboole_0(esk5_2(esk3_2(a_1_0_waybel23(X1),X2),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1272, plain, (v8_struct_0(X1)|v1_xboole_0(a_1_0_waybel23(X1))|~v1_xboole_0(esk5_2(X2,X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_11454, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_11447, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_11631, negated_conjecture, (m1_subset_1(esk14_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_11539, negated_conjecture, (r2_hidden(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk6_2(X1,esk1_0),X2)|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_3984, plain, (v1_finset_1(u1_pre_topc(X1))|~v1_finset_1(k1_zfmisc_1(k1_xboole_0))|~l1_orders_2(X1)|~l1_pre_topc(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1273, plain, (v1_finset_1(esk3_2(a_1_3_waybel31(X1),X2))|r1_tarski(a_1_3_waybel31(X1),X2)|~m1_subset_1(esk4_2(esk3_2(a_1_3_waybel31(X1),X2),X1),k4_ordinal1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10025, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v1_card_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))))).#trainneg
cnf(i_0_10556, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))))).#trainneg
cnf(i_0_1274, plain, (v1_finset_1(k1_card_1(esk11_1(X1)))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),X1),k4_ordinal1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1275, plain, (v1_finset_1(X1)|v1_xboole_0(a_1_3_waybel31(X2))|~m1_subset_1(esk4_2(X1,X2),k4_ordinal1)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,a_1_3_waybel31(X2)))).#trainneg
cnf(i_0_11425, negated_conjecture, (r2_hidden(X1,a_1_0_waybel31(esk2_0))|X1!=k1_card_1(esk11_1(esk2_0)))).#trainneg
cnf(i_0_11527, negated_conjecture, (esk6_2(k1_card_1(X2),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(X2),esk1_0)))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10594, negated_conjecture, (u1_pre_topc(X1)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_tops_2(u1_struct_0(esk1_0),X1)|~l1_pre_topc(X1)|~m1_subset_1(u1_pre_topc(X1),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_10117, negated_conjecture, (esk6_2(k1_card_1(X1),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(k1_card_1(X1))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(X1),esk1_0)))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k1_card_1(X1),esk1_0))).#trainneg
cnf(i_0_7732, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k1_card_1(X1))|r2_hidden(X3,a_1_3_waybel31(X2))|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(k1_card_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(k1_card_1(X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(esk4_2(k1_card_1(X1),X2),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(X1),X2),esk2_0))).#trainneg
cnf(i_0_1276, plain, (v1_yellow_0(k2_yellow_1(X1))|~m1_subset_1(k1_card_1(X1),k4_ordinal1)|~v2_lattice3(k2_yellow_1(X1)))).#trainneg
cnf(i_0_10148, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))))))).#trainneg
cnf(i_0_10024, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))))))).#trainneg
cnf(i_0_1282, plain, (v1_cantor_1(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1284, plain, (esk8_1(k1_zfmisc_1(X1))=u1_struct_0(X2)|~l1_orders_2(X2)|~v2_struct_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5658, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_0_waybel23(esk1_0))|v1_finset_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_10620, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(k1_card_1(X2),a_1_0_waybel31(esk2_0))|v1_finset_1(X2)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1316, plain, (v8_struct_0(X1)|~m1_subset_1(esk5_2(k1_card_1(esk5_2(X2,X1)),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1472, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1285, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=u1_struct_0(X3)|~l1_orders_2(X3)|~v2_struct_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1286, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=u1_struct_0(X3)|~l1_orders_2(X3)|~v2_struct_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4721, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_12085, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(k1_card_1(X1),a_1_0_waybel23(esk1_0))|v1_finset_1(X1)|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_11707, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(X2),a_1_0_waybel31(esk2_0))|v1_finset_1(X2)|~m1_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11706, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_0_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_11672, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(esk11_1(esk2_0)))).#trainneg
cnf(i_0_6058, negated_conjecture, (m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_11699, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_12113, plain, (v1_xboole_0(u1_orders_2(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_3246, negated_conjecture, (v1_tops_2(esk14_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_6707, negated_conjecture, (~v2_struct_0(k2_yellow_1(k1_card_1(esk11_1(esk2_0)))))).#trainneg
cnf(i_0_1290, plain, (u1_orders_2(X1)=u1_struct_0(X2)|~l1_orders_2(X2)|~v2_struct_0(X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_11704, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r2_hidden(esk3_2(a_1_0_waybel23(X1),X2),a_1_0_waybel31(esk2_0))|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_0_waybel23(X1),X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_0_waybel23(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_11673, negated_conjecture, (r2_hidden(X1,a_1_2_waybel31(esk1_0))|X1!=k1_card_1(esk11_1(esk2_0)))).#trainneg
cnf(i_0_11702, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_0_waybel31(esk2_0))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_11717, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|r2_hidden(k1_card_1(X2),a_1_3_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0))).#trainneg
cnf(i_0_11816, negated_conjecture, (esk6_2(k1_card_1(X1),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|esk2_0!=X2|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(X1),esk1_0)))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0))).#trainneg
cnf(i_0_11758, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(X2))|~m1_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X2)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk5_2(X1,X2),esk1_0)|~v1_tops_2(esk5_2(X1,X2),esk1_0))).#trainneg
cnf(i_0_11757, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k1_card_1(X1),esk1_0))).#trainneg
cnf(i_0_4730, negated_conjecture, (~v1_finset_1(k1_card_1(esk11_1(esk2_0))))).#trainneg
cnf(i_0_12814, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(esk5_2(X1,X2))|esk2_0!=X3|~m1_waybel23(k11_waybel_0(X3,esk5_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(X3,esk5_2(X1,X2)),esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(X1,esk1_0)))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_12550, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(X2))|~m1_waybel23(k11_waybel_0(esk2_0,esk5_2(X1,X2)),esk2_0)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X2)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk5_2(X1,X2),esk1_0)|~v1_tops_2(esk5_2(X1,X2),esk1_0))).#trainneg
cnf(i_0_6251, negated_conjecture, (~v1_xboole_0(esk5_2(k1_card_1(esk11_1(esk2_0)),esk1_0)))).#trainneg
cnf(i_0_1302, plain, (r2_hidden(X1,X2)|v2_struct_0(k2_yellow_1(X2))|v1_xboole_0(k11_waybel_0(k2_yellow_1(X2),X3))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k11_waybel_0(k2_yellow_1(X2),X3)))).#trainneg
cnf(i_0_1306, plain, (m1_subset_1(k1_card_1(esk5_2(X1,X2)),a_1_0_waybel23(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1308, plain, (m1_subset_1(esk5_2(k1_card_1(esk5_2(X1,k2_yellow_1(X2))),k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|~l1_pre_topc(k2_yellow_1(X2))|~r2_hidden(X1,a_1_0_waybel23(k2_yellow_1(X2))))).#trainneg
cnf(i_0_12798, negated_conjecture, (r2_hidden(k1_card_1(esk11_1(esk2_0)),a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_11743, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r2_hidden(esk3_2(a_1_0_waybel23(X1),X2),a_1_3_waybel31(esk2_0))|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_0_waybel23(X1),X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_0_waybel23(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4724, negated_conjecture, (~v1_xboole_0(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)))).#trainneg
cnf(i_0_1309, plain, (v1_xboole_0(k1_card_1(esk5_2(X1,X2)))|~v1_xboole_0(esk5_2(k1_card_1(esk5_2(X1,X2)),X2))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_1312, plain, (v1_xboole_0(esk5_2(k1_card_1(esk5_2(X1,X2)),X2))|~v1_xboole_0(k1_card_1(esk5_2(X1,X2)))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_1313, plain, (v1_finset_1(k1_card_1(esk5_2(X1,X2)))|~m1_subset_1(esk5_2(k1_card_1(esk5_2(X1,X2)),X2),k4_ordinal1)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_11698, negated_conjecture, (v1_finset_1(esk6_2(X1,X2))|r2_hidden(X1,a_1_0_waybel31(esk2_0))|v2_struct_0(X2)|~m1_waybel23(k11_waybel_0(esk2_0,esk6_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk6_2(X1,X2)),esk2_0)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_11682, plain, (esk6_2(X2,X1)=u1_pre_topc(X1)|v2_struct_0(X1)|~r1_tarski(u1_pre_topc(X1),esk6_2(X2,X1))|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_11884, negated_conjecture, (v1_finset_1(esk7_2(X1,esk2_0))|r2_hidden(X2,a_1_0_waybel31(esk2_0))|X2!=k1_card_1(esk7_2(X1,esk2_0))|~m1_subset_1(esk7_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk7_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_12530, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|r2_hidden(k1_card_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0))).#trainneg
cnf(i_0_6252, negated_conjecture, (~m1_subset_1(esk5_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_1318, plain, (r2_hidden(k1_xboole_0,a_1_0_waybel23(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2)))).#trainneg
cnf(i_0_4725, negated_conjecture, (~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_1320, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1321, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_11741, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_3_waybel31(esk2_0))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_12821, negated_conjecture, (esk6_2(k1_card_1(X1),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(k1_card_1(X1))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(X1),esk1_0)))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_6233, negated_conjecture, (m1_subset_1(k1_card_1(esk11_1(esk2_0)),a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_12796, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r2_hidden(esk3_2(a_1_0_waybel23(X1),X2),a_1_0_waybel31(esk2_0))|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_0_waybel23(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),X2),esk1_0)|~v1_tops_2(esk3_2(a_1_0_waybel23(X1),X2),esk1_0))).#trainneg
cnf(i_0_11760, negated_conjecture, (v3_ordinal1(esk14_1(esk2_0))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_7534, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0))).#trainneg
cnf(i_0_1323, plain, (r2_hidden(esk3_2(u1_pre_topc(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(u1_pre_topc(X1),X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1324, plain, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(u1_pre_topc(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,u1_pre_topc(X2)))).#trainneg
cnf(i_0_6693, negated_conjecture, (m1_subset_1(k1_card_1(esk11_1(esk2_0)),a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_11733, plain, (r1_tarski(a_1_0_waybel23(X1),X2)|v8_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk3_2(a_1_0_waybel23(X1),X2)))).#trainneg
cnf(i_0_12806, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_0_waybel31(esk2_0))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(a_1_2_waybel31(X1),X2),esk1_0)|~v1_tops_2(esk3_2(a_1_2_waybel31(X1),X2),esk1_0))).#trainneg
cnf(i_0_11889, negated_conjecture, (v1_cantor_1(esk7_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_8079, negated_conjecture, (v3_ordinal1(k1_card_1(esk11_1(esk2_0)))|~v3_ordinal1(a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_8015, negated_conjecture, (v3_ordinal1(k1_card_1(esk11_1(esk2_0)))|~v3_ordinal1(a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1325, plain, (v3_ordinal1(esk3_2(esk5_2(X1,X2),X3))|r1_tarski(esk5_2(X1,X2),X3)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1326, plain, (v1_finset_1(esk3_2(esk5_2(X1,X2),X3))|r1_tarski(esk5_2(X1,X2),X3)|~v1_finset_1(u1_struct_0(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1327, plain, (v1_xboole_0(esk3_2(esk5_2(X1,X2),X3))|r1_tarski(esk5_2(X1,X2),X3)|~v1_xboole_0(u1_struct_0(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_11772, negated_conjecture, (m1_subset_1(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk6_2(X1,esk1_0),X2)|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_11738, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_3_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_12813, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(esk6_2(X1,X2))|v2_struct_0(X2)|esk2_0!=X3|~m1_waybel23(k11_waybel_0(X3,esk6_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(X3,esk6_2(X1,X2)),esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(X1,esk1_0)))|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_11850, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X3,a_1_3_waybel31(X4))|v2_struct_0(X2)|X3!=k1_card_1(esk4_2(X1,X4))|esk2_0!=X4|~m1_waybel23(esk4_2(X1,X4),esk2_0)|~v6_waybel23(esk4_2(X1,X4),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X4)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0)|~l1_orders_2(X4)|~v3_waybel_3(X4)|~v1_lattice3(X4)|~v1_yellow_0(X4)|~v5_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_11833, negated_conjecture, (esk6_2(esk3_2(a_1_0_waybel23(X1),X2),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(esk3_2(a_1_0_waybel23(X1),X2),esk1_0)))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),X2),esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_11849, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X3,a_1_3_waybel31(X4))|X3!=k1_card_1(esk4_2(X1,X4))|esk2_0!=X4|~m1_waybel23(esk4_2(X1,X4),esk2_0)|~v6_waybel23(esk4_2(X1,X4),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X4)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0)|~l1_orders_2(X4)|~v3_waybel_3(X4)|~v1_lattice3(X4)|~v1_yellow_0(X4)|~v5_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_7264, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0))|r2_hidden(X3,a_1_3_waybel31(esk2_0))|v1_finset_1(k11_waybel_0(X1,X2))|X3!=k1_card_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),esk1_0)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7810, negated_conjecture, (m1_subset_1(k1_card_1(X1),a_1_0_waybel23(esk1_0))|v1_finset_1(X1)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_12555, plain, (u1_orders_2(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1328, plain, (esk3_2(esk5_2(X2,X1),X3)=u1_struct_0(X1)|r1_tarski(esk5_2(X2,X1),X3)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(esk3_2(esk5_2(X2,X1),X3)))|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_13530, plain, (u1_orders_2(X2)=k1_xboole_0|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1329, plain, (m1_subset_1(X1,u1_struct_0(X2))|r1_tarski(esk5_2(X3,X2),X4)|~r2_hidden(X1,esk3_2(esk5_2(X3,X2),X4))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1330, plain, (r2_hidden(X1,u1_struct_0(X2))|r1_tarski(esk5_2(X3,X2),X4)|~r2_hidden(X1,esk3_2(esk5_2(X3,X2),X4))|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_13276, negated_conjecture, (r1_tarski(a_1_2_waybel31(X1),a_1_0_waybel31(esk2_0))|v1_finset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel31(esk2_0)))|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_2_waybel31(X1),a_1_0_waybel31(esk2_0)),esk1_0)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_13167, negated_conjecture, (v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_13220, negated_conjecture, (r1_tarski(a_1_0_waybel23(X1),a_1_0_waybel31(esk2_0))|v1_finset_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel31(esk2_0)))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel31(esk2_0)),esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_12896, negated_conjecture, (v1_cantor_1(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_13333, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X2),esk2_0))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|X3!=k1_card_1(esk4_2(k1_card_1(X2),esk2_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(X2),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(esk4_2(k1_card_1(X2),esk2_0),esk1_0)|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X2),esk2_0),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13232, plain, (r1_tarski(u1_pre_topc(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1332, plain, (v1_xboole_0(esk3_2(esk5_2(X1,X2),X3))|m1_subset_1(X4,u1_struct_0(X2))|r1_tarski(esk5_2(X1,X2),X3)|~m1_subset_1(X4,esk3_2(esk5_2(X1,X2),X3))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1335, plain, (r1_tarski(esk5_2(X1,k2_yellow_1(X2)),X3)|m1_subset_1(esk3_2(esk5_2(X1,k2_yellow_1(X2)),X3),k1_zfmisc_1(X2))|~r2_hidden(X1,a_1_0_waybel23(k2_yellow_1(X2)))|~l1_pre_topc(k2_yellow_1(X2)))).#trainneg
cnf(i_0_12900, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_13178, negated_conjecture, (esk6_2(esk3_2(a_1_2_waybel31(X1),X2),esk1_0)=u1_struct_0(esk2_0)|v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|esk2_0!=X3|~m1_waybel23(k11_waybel_0(X3,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~v6_waybel23(k11_waybel_0(X3,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(esk3_2(a_1_2_waybel31(X1),X2),esk1_0)))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_13171, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r2_hidden(esk3_2(a_1_2_waybel31(X1),X2),a_1_3_waybel31(esk2_0))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~m1_waybel23(k11_waybel_0(esk2_0,esk3_2(a_1_2_waybel31(X1),X2)),esk2_0)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(a_1_2_waybel31(X1),X2),esk1_0)|~v1_tops_2(esk3_2(a_1_2_waybel31(X1),X2),esk1_0))).#trainneg
cnf(i_0_13048, negated_conjecture, (v1_finset_1(esk6_2(X1,X2))|r2_hidden(X1,a_1_0_waybel31(esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~m1_waybel23(k11_waybel_0(esk2_0,esk6_2(X1,X2)),esk2_0)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(X1,X2),esk1_0))).#trainneg
cnf(i_0_12564, plain, (u1_orders_2(X1)=u1_struct_0(X2)|~l1_orders_2(X2)|~v2_struct_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1338, plain, (m1_subset_1(esk5_2(k1_card_1(esk6_2(X1,k2_yellow_1(X2))),k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|v2_struct_0(k2_yellow_1(X2))|~l1_pre_topc(k2_yellow_1(X2))|~r2_hidden(X1,a_1_2_waybel31(k2_yellow_1(X2)))|~v6_pre_topc(k2_yellow_1(X2))|~v2_pre_topc(k2_yellow_1(X2)))).#trainneg
cnf(i_0_1339, plain, (v1_xboole_0(k1_card_1(esk6_2(X1,X2)))|v2_struct_0(X2)|~v1_xboole_0(esk5_2(k1_card_1(esk6_2(X1,X2)),X2))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1342, plain, (v1_xboole_0(esk5_2(k1_card_1(esk6_2(X1,X2)),X2))|v2_struct_0(X2)|~v1_xboole_0(k1_card_1(esk6_2(X1,X2)))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_12899, negated_conjecture, (v1_finset_1(esk5_2(X1,X2))|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(X2))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X2)|~m1_waybel23(esk5_2(X1,X2),esk2_0)|~v6_waybel23(esk5_2(X1,X2),esk2_0))).#trainneg
cnf(i_0_12887, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))))|~m1_subset_1(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_tops_2(esk4_2(X3,k2_yellow_1(u1_pre_topc(X2))),X2)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X3,a_1_3_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~v1_yellow_0(k2_yellow_1(u1_pre_topc(X2))))).#trainneg
cnf(i_0_13244, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_12981, negated_conjecture, (esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))))).#trainneg
cnf(i_0_1343, plain, (v1_finset_1(k1_card_1(esk6_2(X1,X2)))|v2_struct_0(X2)|~m1_subset_1(esk5_2(k1_card_1(esk6_2(X1,X2)),X2),k4_ordinal1)|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1346, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(esk5_2(k1_card_1(esk6_2(X2,X1)),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_1348, plain, (r2_hidden(k1_xboole_0,a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2)))).#trainneg
cnf(i_0_13782, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(X3,a_1_3_waybel31(esk2_0))|X3!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(k11_waybel_0(X2,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_13872, negated_conjecture, (r1_tarski(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0))|v1_finset_1(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)))|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)),esk1_0)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_13743, negated_conjecture, (m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_13756, negated_conjecture, (m1_subset_1(esk6_2(k1_card_1(X1),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_13734, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_3_waybel31(esk2_0))|X3!=k1_card_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(k1_card_1(k11_waybel_0(X1,X2)),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_13736, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(k1_card_1(esk12_1(X1)),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_13724, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(X1)|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(X1,esk1_0)|~r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_14082, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,u1_struct_0(X4))!=g1_orders_2(X1,X3)|~l1_orders_2(X2)|~v2_struct_0(X2)|~l1_orders_2(X4)|~v2_struct_0(X4))).#trainneg
cnf(i_0_13889, plain, (v1_finset_1(u1_orders_2(X2))|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1351, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_13737, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0))).#trainneg
cnf(i_0_13710, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(X2,esk1_0))|~m1_subset_1(esk6_2(X2,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(X2,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_13732, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(X1)|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|esk2_0!=X3|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(X1,esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_13721, negated_conjecture, (v1_finset_1(esk4_2(X1,esk2_0))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_13785, plain, (u1_pre_topc(X1)=k1_zfmisc_1(u1_struct_0(X1))|~r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),u1_pre_topc(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_13752, negated_conjecture, (m1_subset_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|v1_finset_1(k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14180, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~m1_subset_1(esk5_2(X1,X2),k4_ordinal1)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_13665, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,k1_card_1(X4))!=g1_orders_2(X1,X3)|~l1_orders_2(X2)|~v2_struct_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1352, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~v1_xboole_0(esk6_2(esk3_2(a_1_2_waybel31(X1),X2),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_13899, plain, (v1_relat_2(u1_orders_2(X2))|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_13900, plain, (v4_relat_2(u1_orders_2(X2))|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1353, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~v1_xboole_0(esk6_2(X2,X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_13753, negated_conjecture, (m1_subset_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14052, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)|~l1_orders_2(X2)|~v2_struct_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_14255, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|r2_hidden(k1_card_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0))).#trainneg
cnf(i_0_14336, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(X1,esk1_0),esk1_0)|~v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14415, plain, (v1_finset_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_13901, plain, (v8_relat_2(u1_orders_2(X2))|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1357, plain, (r2_hidden(esk3_2(u1_pre_topc(k2_yellow_1(X1)),X2),k1_zfmisc_1(X1))|r1_tarski(u1_pre_topc(k2_yellow_1(X1)),X2)|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_14326, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(X2,esk1_0))|~m1_subset_1(X2,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14238, negated_conjecture, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)))|r1_tarski(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0))|v2_struct_0(X1)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_waybel23(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)),esk2_0)|~v6_waybel23(esk3_2(a_1_2_waybel31(X1),a_1_3_waybel31(esk2_0)),esk2_0))).#trainneg
cnf(i_0_14191, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_2_waybel31(X1))|~m1_subset_1(esk5_2(X2,X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_14639, plain, (v1_relat_2(u1_orders_2(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1360, plain, (v1_xboole_0(u1_pre_topc(k2_yellow_1(X1)))|m1_subset_1(esk8_1(u1_pre_topc(k2_yellow_1(X1))),k1_zfmisc_1(X1))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_13947, plain, (v1_card_1(u1_orders_2(X2))|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_14245, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)))).#trainneg
cnf(i_0_14332, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|m1_subset_1(k1_card_1(esk6_2(X1,esk1_0)),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(X1,esk1_0),esk1_0)|~v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14569, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0))|r2_hidden(k1_card_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0)),a_1_2_waybel31(esk1_0))|v1_finset_1(k11_waybel_0(X1,X2))|~m1_subset_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),esk1_0)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14644, plain, (v4_relat_2(u1_orders_2(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1364, plain, (m1_subset_1(esk3_2(esk7_2(X1,X2),X3),u1_struct_0(X2))|r1_tarski(esk7_2(X1,X2),X3)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1366, plain, (m1_subset_1(k1_card_1(esk4_2(X1,X2)),a_1_3_waybel31(X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_14251, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|m1_subset_1(k1_card_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)),a_1_2_waybel31(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0))).#trainneg
cnf(i_0_14325, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(X2,esk1_0))|~m1_subset_1(X2,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_13735, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|v1_finset_1(esk11_1(X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(k1_card_1(esk11_1(X1)),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_1372, plain, (m1_subset_1(esk4_2(k1_card_1(esk4_2(X1,X2)),X2),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2)))).#trainneg
cnf(i_0_1373, plain, (v1_finset_1(k1_card_1(esk4_2(X1,X2)))|~m1_subset_1(esk4_2(k1_card_1(esk4_2(X1,X2)),X2),k4_ordinal1)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2)))).#trainneg
cnf(i_0_13731, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_14324, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(X1,esk1_0)))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14588, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(esk6_2(k1_card_1(X2),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13855, plain, (m1_subset_1(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_13726, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(X1)|r2_hidden(X2,a_1_3_waybel31(esk2_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~r2_hidden(k1_card_1(X1),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_14880, negated_conjecture, (k1_card_1(X1)=X1|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14853, plain, (v8_relat_2(u1_orders_2(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1376, plain, (v1_finset_1(esk3_2(a_1_0_waybel31(X1),X2))|r1_tarski(a_1_0_waybel31(X1),X2)|~m1_subset_1(esk7_2(esk3_2(a_1_0_waybel31(X1),X2),X1),k4_ordinal1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1377, plain, (v1_finset_1(k1_card_1(esk11_1(X1)))|~m1_subset_1(esk7_2(k1_card_1(esk11_1(X1)),X1),k4_ordinal1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1378, plain, (v1_finset_1(X1)|v1_xboole_0(a_1_0_waybel31(X2))|~m1_subset_1(esk7_2(X1,X2),k4_ordinal1)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_14877, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_2_waybel31(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14986, negated_conjecture, (r2_hidden(esk3_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1),u1_struct_0(esk2_0))|r1_tarski(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1))).#trainneg
cnf(i_0_14891, negated_conjecture, (~m1_subset_1(X1,k4_ordinal1)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14790, plain, (m1_subset_1(a_1_2_waybel31(X1),k1_zfmisc_1(a_1_0_waybel23(X1)))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_14888, negated_conjecture, (~v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14865, negated_conjecture, (r2_hidden(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk6_2(X1,esk1_0),X2)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_14971, plain, (v1_card_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_13792, plain, (r1_tarski(u1_pre_topc(X1),k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1379, plain, (r2_hidden(esk3_2(esk7_2(X1,X2),X3),u1_struct_0(X2))|r1_tarski(esk7_2(X1,X2),X3)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1380, plain, (r2_hidden(X1,u1_struct_0(X2))|v1_xboole_0(esk7_2(X3,X2))|~r2_hidden(X3,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,esk7_2(X3,X2)))).#trainneg
cnf(i_0_1381, plain, (k1_card_1(esk8_1(a_1_2_waybel31(X1)))=esk8_1(a_1_2_waybel31(X1))|v1_xboole_0(a_1_2_waybel31(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_15133, negated_conjecture, (m1_subset_1(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk6_2(X1,esk1_0),X2)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_15050, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(X2),esk1_0))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(esk6_2(k1_card_1(X2),esk1_0)),a_1_2_waybel31(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(esk6_2(k1_card_1(X2),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(esk6_2(k1_card_1(X2),esk1_0),esk1_0)|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(X2),esk1_0),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1383, plain, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~m1_subset_1(esk6_2(esk3_2(a_1_2_waybel31(X1),X2),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1384, plain, (v1_finset_1(X1)|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~m1_subset_1(esk6_2(X1,X2),k4_ordinal1)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_15256, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(esk6_2(k1_card_1(X2),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_15277, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(X1,X2)),a_1_0_waybel23(esk1_0))|v1_finset_1(k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15259, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(esk6_2(k1_card_1(X2),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_1385, plain, (m1_subset_1(k1_card_1(esk5_2(X1,X2)),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1387, plain, (m1_subset_1(esk6_2(k1_card_1(esk5_2(X1,k2_yellow_1(X2))),k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|v2_struct_0(k2_yellow_1(X2))|~l1_pre_topc(k2_yellow_1(X2))|~v6_pre_topc(k2_yellow_1(X2))|~v2_pre_topc(k2_yellow_1(X2))|~r2_hidden(X1,a_1_0_waybel23(k2_yellow_1(X2))))).#trainneg
cnf(i_0_1388, plain, (v1_xboole_0(k1_card_1(esk5_2(X1,X2)))|v2_struct_0(X2)|~v1_xboole_0(esk6_2(k1_card_1(esk5_2(X1,X2)),X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_15381, negated_conjecture, (k1_card_1(X1)=X1|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_15534, negated_conjecture, (~m1_subset_1(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_15558, negated_conjecture, (~m1_subset_1(X1,a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1391, plain, (v1_xboole_0(esk6_2(k1_card_1(esk5_2(X1,X2)),X2))|v2_struct_0(X2)|~v1_xboole_0(k1_card_1(esk5_2(X1,X2)))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_15497, negated_conjecture, (~v1_finset_1(esk8_1(a_1_2_waybel31(esk1_0))))).#trainneg
cnf(i_0_15535, negated_conjecture, (~m1_subset_1(esk8_1(a_1_2_waybel31(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_1392, plain, (v1_finset_1(k1_card_1(esk5_2(X1,X2)))|v2_struct_0(X2)|~m1_subset_1(esk6_2(k1_card_1(esk5_2(X1,X2)),X2),k4_ordinal1)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_1398, plain, (r2_hidden(k1_xboole_0,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2)))).#trainneg
cnf(i_0_15309, negated_conjecture, (k1_card_1(esk8_1(a_1_2_waybel31(esk1_0)))=esk8_1(a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_15559, negated_conjecture, (~v2_struct_0(k2_yellow_1(esk8_1(a_1_2_waybel31(esk1_0)))))).#trainneg
cnf(i_0_1406, plain, (m1_subset_1(k1_card_1(esk6_2(X1,X2)),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1408, plain, (m1_subset_1(esk6_2(k1_card_1(esk6_2(X1,k2_yellow_1(X2))),k2_yellow_1(X2)),k1_zfmisc_1(k1_zfmisc_1(X2)))|v2_struct_0(k2_yellow_1(X2))|~l1_pre_topc(k2_yellow_1(X2))|~v6_pre_topc(k2_yellow_1(X2))|~v2_pre_topc(k2_yellow_1(X2))|~r2_hidden(X1,a_1_2_waybel31(k2_yellow_1(X2))))).#trainneg
cnf(i_0_15466, negated_conjecture, (v1_finset_1(X1)|m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~m1_subset_1(X1,a_1_2_waybel31(esk1_0)))).#trainneg
cnf(i_0_15521, negated_conjecture, (r1_tarski(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_15360, negated_conjecture, (v1_finset_1(X1)|m1_subset_1(k1_card_1(X1),a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_15738, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(k1_card_1(X2),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15517, negated_conjecture, (m1_subset_1(esk3_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1),u1_struct_0(esk2_0))|r1_tarski(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1))).#trainneg
cnf(i_0_1409, plain, (v1_xboole_0(k1_card_1(esk6_2(X1,X2)))|v2_struct_0(X2)|~v1_xboole_0(esk6_2(k1_card_1(esk6_2(X1,X2)),X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_11791, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)))))))).#trainneg
cnf(i_0_1412, plain, (v1_xboole_0(esk6_2(k1_card_1(esk6_2(X1,X2)),X2))|v2_struct_0(X2)|~v1_xboole_0(k1_card_1(esk6_2(X1,X2)))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_1413, plain, (v1_finset_1(k1_card_1(esk6_2(X1,X2)))|v2_struct_0(X2)|~m1_subset_1(esk6_2(k1_card_1(esk6_2(X1,X2)),X2),k4_ordinal1)|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_15358, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(X1)|r2_hidden(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_0_waybel23(esk1_0))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_15522, negated_conjecture, (r1_tarski(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_15520, negated_conjecture, (m1_subset_1(esk3_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1))).#trainneg
cnf(i_0_14404, plain, (k1_xboole_0=X1|g1_orders_2(k1_xboole_0,X2)!=g1_orders_2(X1,X3)|~l1_orders_2(X4)|~v2_struct_0(X4)|~v2_struct_0(k2_yellow_1(X2)))).#trainneg
cnf(i_0_12554, plain, (X1=u1_orders_2(X2)|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1416, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(esk6_2(k1_card_1(esk6_2(X2,X1)),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_1419, plain, (r2_hidden(k1_xboole_0,a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2)))).#trainneg
cnf(i_0_16333, negated_conjecture, (m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_16135, negated_conjecture, (v1_tops_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0)|~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_16412, plain, (u1_orders_2(X1)=u1_orders_2(X2)|~l1_orders_2(X2)|~v2_struct_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_15869, negated_conjecture, (k1_card_1(esk8_1(a_1_0_waybel23(esk1_0)))=esk8_1(a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1421, plain, (m1_subset_1(k1_card_1(esk7_2(X1,X2)),a_1_0_waybel31(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_15899, negated_conjecture, (~v1_finset_1(esk8_1(a_1_0_waybel23(esk1_0))))).#trainneg
cnf(i_0_15879, negated_conjecture, (~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_1423, plain, (v1_finset_1(k1_card_1(esk7_2(X1,X2)))|~m1_subset_1(esk7_2(k1_card_1(esk7_2(X1,X2)),X2),k4_ordinal1)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_0_waybel31(X2)))).#trainneg
cnf(i_0_16186, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(k1_card_1(X2),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_15903, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),X1)|~v1_finset_1(esk3_2(a_1_0_waybel23(esk1_0),X1)))).#trainneg
cnf(i_0_16275, negated_conjecture, (m1_subset_1(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_2_waybel31(esk1_0))|v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(X1)|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_16459, negated_conjecture, (k1_card_1(X1)=X1|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_15559, negated_conjecture, (~v2_struct_0(k2_yellow_1(esk8_1(a_1_0_waybel23(esk1_0)))))).#trainneg
cnf(i_0_1430, plain, (m1_subset_1(k1_card_1(esk4_2(X1,X2)),a_1_0_waybel31(X2))|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1432, plain, (v1_finset_1(k1_card_1(esk4_2(X1,X2)))|~m1_subset_1(esk7_2(k1_card_1(esk4_2(X1,X2)),X2),k4_ordinal1)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X1,a_1_3_waybel31(X2))|~v1_yellow_0(X2))).#trainneg
cnf(i_0_13758, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_16530, negated_conjecture, (~v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_14344, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_16513, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_16750, negated_conjecture, (v1_card_1(esk8_1(a_1_0_waybel23(esk1_0))))).#trainneg
cnf(i_0_16823, negated_conjecture, (~v1_xboole_0(esk8_1(a_1_0_waybel23(esk1_0))))).#trainneg
cnf(i_0_14865, negated_conjecture, (r2_hidden(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk6_2(X1,esk1_0),X2)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_15133, negated_conjecture, (r1_tarski(esk6_2(X1,esk1_0),X2)|m1_subset_1(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_13167, negated_conjecture, (v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_11539, negated_conjecture, (r2_hidden(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|r1_tarski(esk6_2(X1,esk1_0),X2)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1463, plain, (g1_orders_2(k1_xboole_0,u1_orders_2(X1))=X1|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_13698, plain, (u1_orders_2(X1)=k1_xboole_0|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1477, plain, (v3_ordinal1(u1_orders_2(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_11772, negated_conjecture, (r1_tarski(esk6_2(X1,esk1_0),X2)|m1_subset_1(esk3_2(esk6_2(X1,esk1_0),X2),u1_struct_0(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_15906, negated_conjecture, (~v1_finset_1(k1_card_1(esk5_2(X1,esk1_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_16756, negated_conjecture, (v1_finset_1(esk8_1(a_1_0_waybel23(esk1_0)))|r2_hidden(esk8_1(a_1_0_waybel23(esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_16471, negated_conjecture, (r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1490, plain, (r1_tarski(u1_pre_topc(X1),X2)|m1_subset_1(esk3_2(u1_pre_topc(X1),X2),k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_16478, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_16472, negated_conjecture, (r2_hidden(k1_card_1(esk5_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_15412, plain, (v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_9255, negated_conjecture, (r1_tarski(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0))|~m1_waybel23(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0)),esk2_0)|~v6_waybel23(esk3_2(a_1_0_waybel23(X1),a_1_0_waybel23(esk1_0)),esk2_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_14324, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk6_2(X1,esk1_0)))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1492, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1494, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_xboole_0,k1_xboole_0))|~r2_hidden(X1,u1_orders_2(X2))|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_16759, negated_conjecture, (v1_finset_1(esk8_1(a_1_0_waybel23(esk1_0)))|m1_subset_1(esk4_2(esk8_1(a_1_0_waybel23(esk1_0)),X1),k1_zfmisc_1(u1_struct_0(X1)))|esk2_0!=X1|~m1_waybel23(k11_waybel_0(X1,esk8_1(a_1_0_waybel23(esk1_0))),esk2_0)|~v6_waybel23(k11_waybel_0(X1,esk8_1(a_1_0_waybel23(esk1_0))),esk2_0)|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_16336, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)))).#trainneg
cnf(i_0_1603, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(k11_waybel_0(X2,X3))|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1496, plain, (r1_tarski(esk5_2(X1,X2),X3)|m1_subset_1(esk3_2(esk5_2(X1,X2),X3),k1_zfmisc_1(k1_xboole_0))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1501, plain, (v2_lattice3(g1_orders_2(X1,k1_xboole_0))|~v2_lattice3(k2_yellow_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_15760, negated_conjecture, (v1_finset_1(esk11_1(X1))|m1_subset_1(esk6_2(k1_card_1(esk11_1(X1)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_16995, negated_conjecture, (v1_finset_1(X1)|m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_17091, plain, (X1=k1_xboole_0|k2_yellow_1(X1)!=X2|~v1_orders_2(X2)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_14868, negated_conjecture, (r1_tarski(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel23(esk1_0)),a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_17074, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|X1!=esk2_0|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1528, plain, (esk8_1(k1_zfmisc_1(X1))=esk8_1(k1_zfmisc_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1529, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk8_1(k1_zfmisc_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1530, plain, (esk8_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=esk8_1(k1_zfmisc_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13906, plain, (v4_relat_1(u1_orders_2(X3),X2)|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_13908, plain, (v5_relat_1(u1_orders_2(X3),X2)|~l1_orders_2(X1)|~v2_struct_0(X1)|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_17153, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~v1_finset_1(esk5_2(X1,esk1_0)))).#trainneg
cnf(i_0_1534, plain, (u1_orders_2(X1)=esk8_1(k1_zfmisc_1(X2))|~v1_xboole_0(X2)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_17189, negated_conjecture, (m1_subset_1(k1_card_1(esk6_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_17068, negated_conjecture, (r1_tarski(esk6_2(X1,esk1_0),u1_struct_0(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_14326, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(X2,esk1_0))|~m1_subset_1(X2,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1538, plain, (v1_tops_2(esk11_1(k2_yellow_1(u1_pre_topc(X1))),X1)|v2_struct_0(X1)|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1541, plain, (v1_tops_2(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),X2)|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_17046, negated_conjecture, (m1_subset_1(esk3_2(esk6_2(X1,esk1_0),X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk6_2(X1,esk1_0),X2)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_17047, negated_conjecture, (r1_tarski(esk6_2(X1,esk1_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_17416, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|v1_finset_1(X3)|X1!=k1_card_1(X3)|esk2_0!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(X3,esk1_0)|~v1_tops_2(X3,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_17958, negated_conjecture, (v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_17723, plain, (v4_relat_1(u1_orders_2(X1),X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_14325, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(X2,esk1_0))|~m1_subset_1(X2,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1572, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~v2_struct_0(k2_yellow_1(esk7_2(X1,X2))))).#trainneg
cnf(i_0_8529, negated_conjecture, (v1_finset_1(k1_card_1(esk11_1(esk2_0)))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_1573, plain, (v1_finset_1(X1)|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v2_struct_0(k2_yellow_1(esk6_2(X1,X2))))).#trainneg
cnf(i_0_1575, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v2_struct_0(k2_yellow_1(esk5_2(X1,X2))))).#trainneg
cnf(i_0_18032, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_0_waybel23(esk1_0))|v2_struct_0(X1)|X3!=k1_card_1(k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_18005, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|r2_hidden(X3,a_1_0_waybel23(esk1_0))|v1_finset_1(X1)|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_17403, negated_conjecture, (r2_hidden(esk3_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1))).#trainneg
cnf(i_0_17729, plain, (v5_relat_1(u1_orders_2(X1),X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_13710, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(X2,esk1_0))|~r2_hidden(X2,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X2,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1577, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~v2_struct_0(k2_yellow_1(esk4_2(X1,X2))))).#trainneg
cnf(i_0_9103, negated_conjecture, (v1_finset_1(k1_card_1(esk11_1(esk2_0)))|~m1_subset_1(esk7_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_10145, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0)))))))))).#trainneg
cnf(i_0_1580, plain, (v1_finset_1(u1_orders_2(X1))|~l1_orders_2(X1)|~v2_struct_0(k2_yellow_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))).#trainneg
cnf(i_0_1581, plain, (v1_finset_1(u1_orders_2(k2_yellow_1(X1)))|~v2_struct_0(k2_yellow_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_17271, negated_conjecture, (m1_subset_1(k1_card_1(esk5_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_16189, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(k1_card_1(X2),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_10150, plain, (v3_ordinal1(X1)|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))))).#trainneg
cnf(i_0_14940, plain, (esk8_1(u1_pre_topc(k2_yellow_1(k1_xboole_0)))=k1_xboole_0|v1_xboole_0(u1_pre_topc(k2_yellow_1(k1_xboole_0)))|~l1_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_1592, plain, (m1_subset_1(esk11_1(k2_yellow_1(u1_pre_topc(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v2_struct_0(X1)|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X1)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1614, plain, (m1_subset_1(esk7_2(X1,k2_yellow_1(u1_pre_topc(X2))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel31(k2_yellow_1(u1_pre_topc(X2))))|~v3_waybel_3(k2_yellow_1(u1_pre_topc(X2)))|~v1_lattice3(k2_yellow_1(u1_pre_topc(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_15759, negated_conjecture, (v1_finset_1(esk11_1(X1))|m1_subset_1(esk6_2(k1_card_1(esk11_1(X1)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_17934, negated_conjecture, (esk3_2(esk6_2(X1,esk1_0),X2)=u1_struct_0(esk1_0)|r1_tarski(esk6_2(X1,esk1_0),X2)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk3_2(esk6_2(X1,esk1_0),X2)))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1626, plain, (v8_struct_0(X1)|r1_tarski(a_1_0_waybel23(X1),X2)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_12556, plain, (u1_orders_2(X1)=k1_xboole_0|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1630, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(esk14_1(X1)),X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(esk14_1(X1)),X2),X3)|v2_struct_0(k2_yellow_1(esk14_1(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(esk14_1(X1))))).#trainneg
cnf(i_0_1632, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(esk8_1(k1_zfmisc_1(X1))),X2),X3),X1)|r1_tarski(k11_waybel_0(k2_yellow_1(esk8_1(k1_zfmisc_1(X1))),X2),X3)|v2_struct_0(k2_yellow_1(esk8_1(k1_zfmisc_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(esk8_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_15756, negated_conjecture, (v1_finset_1(esk12_1(X1))|m1_subset_1(esk6_2(k1_card_1(esk12_1(X1)),esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_17067, negated_conjecture, (m1_subset_1(esk3_2(esk6_2(X1,esk1_0),X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk6_2(X1,esk1_0),X2)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_18232, negated_conjecture, (v1_finset_1(esk11_1(X1))|m1_subset_1(k1_card_1(esk11_1(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18231, negated_conjecture, (v1_finset_1(esk11_1(X1))|m1_subset_1(k1_card_1(esk11_1(X1)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18451, plain, (v4_orders_2(g1_orders_2(u1_struct_0(X1),k1_xboole_0))|~l1_orders_2(X1)|~v4_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1633, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(esk7_2(X1,X2)),X3),X4),u1_struct_0(X2))|r1_tarski(k11_waybel_0(k2_yellow_1(esk7_2(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(esk7_2(X1,X2)))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(esk7_2(X1,X2))))).#trainneg
cnf(i_0_18423, plain, (v1_partfun1(k1_xboole_0,X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1634, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(esk7_2(X1,X2)),X3),X4),u1_struct_0(X2))|r1_tarski(k11_waybel_0(k2_yellow_1(esk7_2(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(esk7_2(X1,X2)))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(esk7_2(X1,X2))))).#trainneg
cnf(i_0_1635, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(esk6_2(X1,X2)),X3),X4),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|r1_tarski(k11_waybel_0(k2_yellow_1(esk6_2(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(esk6_2(X1,X2)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(esk6_2(X1,X2))))).#trainneg
cnf(i_0_18374, negated_conjecture, (m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_18452, plain, (v5_orders_2(g1_orders_2(u1_struct_0(X1),k1_xboole_0))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_18296, plain, (m1_subset_1(k1_xboole_0,u1_pre_topc(k2_yellow_1(k1_xboole_0)))|v1_xboole_0(u1_pre_topc(k2_yellow_1(k1_xboole_0)))|~l1_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_1636, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(esk6_2(X1,X2)),X3),X4),u1_pre_topc(X2))|v2_struct_0(X2)|r1_tarski(k11_waybel_0(k2_yellow_1(esk6_2(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(esk6_2(X1,X2)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(esk6_2(X1,X2))))).#trainneg
cnf(i_0_18458, plain, (g1_orders_2(k1_xboole_0,k1_xboole_0)=X1|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_18292, negated_conjecture, (v3_ordinal1(u1_struct_0(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))))))))).#trainneg
cnf(i_0_1637, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(esk5_2(X1,X2)),X3),X4),k1_zfmisc_1(u1_struct_0(X2)))|r1_tarski(k11_waybel_0(k2_yellow_1(esk5_2(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(esk5_2(X1,X2)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(esk5_2(X1,X2))))).#trainneg
cnf(i_0_1638, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(esk5_2(X1,X2)),X3),X4),u1_pre_topc(X2))|r1_tarski(k11_waybel_0(k2_yellow_1(esk5_2(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(esk5_2(X1,X2)))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(esk5_2(X1,X2))))).#trainneg
cnf(i_0_18653, negated_conjecture, (esk3_2(esk6_2(X1,esk1_0),X2)=u1_struct_0(esk1_0)|r1_tarski(esk6_2(X1,esk1_0),X2)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk3_2(esk6_2(X1,esk1_0),X2)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_18228, negated_conjecture, (v1_finset_1(esk12_1(X1))|m1_subset_1(k1_card_1(esk12_1(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18450, plain, (g1_orders_2(u1_struct_0(X1),k1_xboole_0)=X1|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_15556, plain, (v1_xboole_0(a_1_2_waybel31(X1))|m1_subset_1(X2,a_1_0_waybel23(X1))|v2_struct_0(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_18434, plain, (g1_orders_2(X1,k1_xboole_0)=k2_yellow_1(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1639, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(k4_ordinal1),X1),X2),k1_numbers)|r1_tarski(k11_waybel_0(k2_yellow_1(k4_ordinal1),X1),X2)|v2_struct_0(k2_yellow_1(k4_ordinal1))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_1640, plain, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel31(X1)))|~m1_subset_1(esk7_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3),X1),k4_ordinal1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel31(X1))))).#trainneg
cnf(i_0_1642, plain, (k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3))=esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel31(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel31(X1))))).#trainneg
cnf(i_0_15555, plain, (v2_struct_0(X1)|~v1_xboole_0(a_1_0_waybel23(X1))|~r2_hidden(X2,a_1_2_waybel31(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_18225, negated_conjecture, (v1_finset_1(esk14_1(X1))|m1_subset_1(k1_card_1(esk14_1(X1)),a_1_0_waybel23(esk1_0))|v2_struct_0(X1)|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_16455, negated_conjecture, (v1_finset_1(X1)|m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_15401, plain, (r2_hidden(k1_card_1(esk6_2(X1,X2)),a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_10592, negated_conjecture, (m1_subset_1(esk8_1(u1_struct_0(esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1643, plain, (v1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel31(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel31(X1))))).#trainneg
cnf(i_0_1645, plain, (m1_subset_1(esk6_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(k2_yellow_1(X1))),X2),X3),k2_yellow_1(X1)),k1_zfmisc_1(k1_zfmisc_1(X1)))|v2_struct_0(k2_yellow_1(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(k2_yellow_1(X1))),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(k2_yellow_1(X1))))|~l1_pre_topc(k2_yellow_1(X1))|~v6_pre_topc(k2_yellow_1(X1))|~v2_pre_topc(k2_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_1646, plain, (v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_xboole_0(esk6_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_19197, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(k1_card_1(esk6_2(X2,X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_18952, plain, (v5_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1647, plain, (k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3))=esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_19212, negated_conjecture, (v3_ordinal1(esk8_1(u1_struct_0(esk2_0)))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1648, plain, (v1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_1649, plain, (v1_xboole_0(esk6_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_15873, negated_conjecture, (k1_card_1(esk3_2(a_1_0_waybel23(esk1_0),X1))=esk3_2(a_1_0_waybel23(esk1_0),X1)|r1_tarski(a_1_0_waybel23(esk1_0),X1))).#trainneg
cnf(i_0_16311, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X2),X3))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(esk4_2(k1_card_1(X2),X3)),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|esk2_0!=X3|~m1_waybel23(k11_waybel_0(X3,k11_waybel_0(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(X3,k11_waybel_0(X1,X2)),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(esk4_2(k1_card_1(X2),X3),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X2),X3),esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18078, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_0_waybel23(esk1_0))|X3!=k1_card_1(k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_18953, plain, (v4_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_18717, plain, (v4_orders_2(g1_orders_2(X1,k1_xboole_0))|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_1650, plain, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~m1_subset_1(esk6_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_1651, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),a_1_0_waybel23(X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_1652, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_xboole_0(esk6_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_15549, plain, (v1_finset_1(a_1_2_waybel31(X1))|v2_struct_0(X1)|~v1_finset_1(a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_19175, plain, (v1_finset_1(k1_card_1(esk6_2(X1,X2)))|v2_struct_0(X2)|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_17738, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(esk5_2(X1,esk1_0))))).#trainneg
cnf(i_0_18805, plain, (v5_orders_2(g1_orders_2(X1,k1_xboole_0))|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_18821, plain, (u1_pre_topc(k2_yellow_1(k1_xboole_0))=k1_xboole_0|m1_subset_1(k1_xboole_0,u1_pre_topc(k2_yellow_1(k1_xboole_0)))|~l1_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_9820, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel31(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1654, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_18934, negated_conjecture, (u1_orders_2(esk2_0)=k1_xboole_0|X1!=esk2_0|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_18289, negated_conjecture, (v3_ordinal1(u1_orders_2(esk2_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))))))))).#trainneg
cnf(i_0_1656, plain, (k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3))=esk3_2(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_3_waybel31(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_3_waybel31(X1))))).#trainneg
cnf(i_0_1657, plain, (v1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_3_waybel31(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_3_waybel31(X1))))).#trainneg
cnf(i_0_16310, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_0_waybel23(esk1_0))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_18036, negated_conjecture, (v1_finset_1(u1_struct_0(X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(u1_struct_0(X1))|esk2_0!=X1|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(u1_struct_0(X1),esk1_0)|~v1_tops_2(u1_struct_0(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15410, plain, (v1_xboole_0(a_1_0_waybel23(X1))|v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(X2)))).#trainneg
cnf(i_0_19194, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(k1_card_1(esk6_2(X2,X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_17492, negated_conjecture, (r1_tarski(esk8_1(k1_zfmisc_1(a_1_0_waybel23(esk1_0))),a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1660, plain, (m1_subset_1(esk4_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3),X1),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_3_waybel31(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_3_waybel31(X1))))).#trainneg
cnf(i_0_1661, plain, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_3_waybel31(X1)))|~m1_subset_1(esk4_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_3_waybel31(X1)),X2),X3),X1),k4_ordinal1)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_3_waybel31(X1))))).#trainneg
cnf(i_0_1662, plain, (m1_subset_1(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(k2_yellow_1(X1))),X2),X3),k2_yellow_1(X1)),k1_zfmisc_1(k1_zfmisc_1(X1)))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(k2_yellow_1(X1))),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(k2_yellow_1(X1))))|~l1_pre_topc(k2_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_19270, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(esk6_2(X2,X1)))).#trainneg
cnf(i_0_19463, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(k11_waybel_0(esk2_0,X2))|~m1_waybel23(k11_waybel_0(esk2_0,X2),esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_19475, plain, (v4_orders_2(X1)|~v2_struct_0(k2_yellow_1(u1_struct_0(X1)))|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_19543, plain, (v5_orders_2(X1)|~v2_struct_0(k2_yellow_1(u1_struct_0(X1)))|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1663, plain, (v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_xboole_0(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3),X1))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_1664, plain, (k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3))=esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_1665, plain, (v1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_19269, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(k1_card_1(esk6_2(X2,X1)))))).#trainneg
cnf(i_0_19547, plain, (v1_xboole_0(u1_pre_topc(k2_yellow_1(k1_xboole_0)))|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|u1_pre_topc(k2_yellow_1(k1_xboole_0))=k1_xboole_0|~l1_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_1666, plain, (v1_xboole_0(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3),X1))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_1667, plain, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~m1_subset_1(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3),X1),k4_ordinal1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_1668, plain, (v8_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_16352, negated_conjecture, (esk3_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1)=u1_struct_0(esk1_0)|r1_tarski(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(esk3_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),X1))))).#trainneg
cnf(i_0_19527, plain, (v1_finset_1(esk6_2(X1,X2))|v2_struct_0(X2)|~v1_finset_1(u1_pre_topc(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_18353, negated_conjecture, (m1_subset_1(esk7_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_19144, negated_conjecture, (esk4_2(X1,esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk4_2(X1,esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_1669, plain, (v8_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_1670, plain, (v8_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_xboole_0(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_1671, plain, (v8_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~m1_subset_1(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_15612, plain, (r1_tarski(esk7_2(X1,X2),u1_struct_0(X2))|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_16766, negated_conjecture, (v1_finset_1(esk8_1(a_1_0_waybel23(esk1_0)))|r2_hidden(esk8_1(a_1_0_waybel23(esk1_0)),a_1_0_waybel23(esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,esk8_1(a_1_0_waybel23(esk1_0))),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk8_1(a_1_0_waybel23(esk1_0))),esk2_0)|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_19151, negated_conjecture, (v1_finset_1(esk4_2(X1,esk2_0))|r2_hidden(k1_card_1(esk4_2(X1,esk2_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(X1,esk2_0),esk1_0)|~v1_tops_2(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_1672, plain, (v8_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_1674, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_orders_2(k2_yellow_1(X1))),X2),X3),k2_zfmisc_1(X1,X1))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_orders_2(k2_yellow_1(X1))),X2),X3)|v2_struct_0(k2_yellow_1(u1_orders_2(k2_yellow_1(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_orders_2(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_1675, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(k11_waybel_0(k2_yellow_1(X1),X2)),X3),X4),X1)|v2_struct_0(k2_yellow_1(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(k11_waybel_0(k2_yellow_1(X1),X2)),X3),X4)|v2_struct_0(k2_yellow_1(k11_waybel_0(k2_yellow_1(X1),X2)))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k11_waybel_0(k2_yellow_1(X1),X2))))).#trainneg
cnf(i_0_18989, plain, (k2_yellow_1(u1_struct_0(X1))=X1|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1)|~v2_struct_0(k2_yellow_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_19516, plain, (r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),a_1_0_waybel23(X1))|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1)))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_16949, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(X1,esk1_0),esk1_0)|~v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_19865, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|X1!=k1_card_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X2,esk1_0))).#trainneg
cnf(i_0_19994, plain, (r1_tarski(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_1676, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(k11_waybel_0(X1,X2)),X3),X4),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(k11_waybel_0(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(k11_waybel_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k11_waybel_0(X1,X2))))).#trainneg
cnf(i_0_19976, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(k2_yellow_1(k1_xboole_0))|~v2_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_1677, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(k11_waybel_0(X1,X2)),X3),X4),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(k11_waybel_0(X1,X2)),X3),X4)|v2_struct_0(k2_yellow_1(k11_waybel_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(k11_waybel_0(X1,X2))))).#trainneg
cnf(i_0_1679, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(k2_yellow_1(X1))),X2),X3),k1_zfmisc_1(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(k2_yellow_1(X1))),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(k2_yellow_1(X1))))|~l1_pre_topc(k2_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_19854, negated_conjecture, (r2_hidden(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_19848, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~r2_hidden(X2,a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(k2_yellow_1(esk6_2(X2,X1))))).#trainneg
cnf(i_0_20135, plain, (esk7_2(X2,X1)=u1_struct_0(X1)|~r1_tarski(u1_struct_0(X1),esk7_2(X2,X1))|~r2_hidden(X2,a_1_0_waybel31(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1680, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(k2_yellow_1(X1))),X2),X3),k1_zfmisc_1(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(k2_yellow_1(X1))),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(k2_yellow_1(X1))))|~l1_pre_topc(k2_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(k2_yellow_1(X1)))))).#trainneg
cnf(i_0_1681, plain, (r2_hidden(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_1682, plain, (m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_15262, plain, (m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),a_1_0_waybel23(X1))|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_15550, plain, (v1_xboole_0(a_1_2_waybel31(X1))|v2_struct_0(X1)|~v1_xboole_0(a_1_0_waybel23(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_19161, negated_conjecture, (v1_finset_1(esk4_2(X1,esk2_0))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(X1,esk2_0),esk1_0)|~v1_tops_2(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_16269, negated_conjecture, (r2_hidden(k1_card_1(esk4_2(k1_card_1(X1),X2)),u1_struct_0(esk1_0))|v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(X1)|~m1_subset_1(a_1_0_waybel23(esk1_0),u1_struct_0(esk2_0))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_20002, plain, (v1_tops_2(k1_xboole_0,k2_yellow_1(k1_xboole_0))|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_15050, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(X1),esk1_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(esk6_2(k1_card_1(X1),esk1_0)),a_1_0_waybel23(esk1_0))|esk2_0!=X2|~m1_subset_1(esk6_2(k1_card_1(X1),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_cantor_1(k11_waybel_0(X2,X1),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1685, plain, (r1_tarski(k11_waybel_0(k2_yellow_1(X1),X2),X1)|v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1701, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_1702, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~m1_subset_1(esk3_2(a_1_2_waybel31(X1),X2),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_14441, plain, (v8_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_19425, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(esk1_0),X1))|r2_hidden(esk3_2(a_1_0_waybel23(esk1_0),X1),a_1_3_waybel31(esk2_0))|r1_tarski(a_1_0_waybel23(esk1_0),X1)|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0))).#trainneg
cnf(i_0_13720, negated_conjecture, (~v1_finset_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_14257, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|r2_hidden(X1,a_1_3_waybel31(esk2_0))|X1!=k1_card_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0))).#trainneg
cnf(i_0_18105, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(X1)|r2_hidden(X3,a_1_0_waybel23(esk1_0))|X3!=k1_card_1(esk4_2(k1_card_1(X1),X2))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(esk4_2(k1_card_1(X1),X2),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(X1),X2),esk2_0))).#trainneg
cnf(i_0_19192, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k1_card_1(esk6_2(X2,X1)),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_18831, plain, (m1_subset_1(k1_xboole_0,u1_pre_topc(k2_yellow_1(k1_xboole_0)))|~l1_pre_topc(k2_yellow_1(k1_xboole_0))|~v2_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_10211, negated_conjecture, (v1_finset_1(esk12_1(esk2_0))|m1_subset_1(k1_card_1(esk12_1(esk2_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(esk2_0),esk1_0)|~v1_tops_2(esk12_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_17165, negated_conjecture, (~v1_xboole_0(esk5_2(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_1707, plain, (r2_hidden(esk3_2(k1_xboole_0,X1),X2)|r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_1711, plain, (esk3_2(k1_zfmisc_1(k1_xboole_0),X1)=k1_xboole_0|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1))).#trainneg
cnf(i_0_20620, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(X4,a_1_3_waybel31(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|X4!=k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_20448, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(X1,esk2_0)|~v6_waybel23(X1,esk2_0))).#trainneg
cnf(i_0_20727, plain, (v1_tops_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X1)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_20980, plain, (r1_tarski(k1_xboole_0,k1_numbers))).#trainneg
cnf(i_0_20983, negated_conjecture, (r1_tarski(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7035, negated_conjecture, (v1_finset_1(esk12_1(esk2_0))|r2_hidden(k1_card_1(esk12_1(esk2_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(esk2_0),esk1_0)|~v1_tops_2(esk12_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_20984, plain, (r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_20812, plain, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(u1_pre_topc(k2_yellow_1(k1_xboole_0)))|~l1_pre_topc(k2_yellow_1(k1_xboole_0))|~v2_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_20985, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),k1_numbers)|~r2_hidden(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_20621, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(X4,a_1_0_waybel31(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|X4!=k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_20724, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(X1),esk1_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(esk6_2(k1_card_1(X1),esk1_0)),a_1_0_waybel23(esk1_0))|esk2_0!=X2|~m1_subset_1(esk6_2(k1_card_1(X1),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(X1,esk1_0))).#trainneg
cnf(i_0_20758, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)),esk1_0))).#trainneg
cnf(i_0_20721, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(X1),esk1_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(esk6_2(k1_card_1(X1),esk1_0)),a_1_0_waybel23(esk1_0))|esk2_0!=X2|~m1_subset_1(esk6_2(k1_card_1(X1),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0))).#trainneg
cnf(i_0_20740, plain, (v1_tops_2(a_1_2_waybel31(X1),X2)|v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(a_1_2_waybel31(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_20622, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(X4,a_1_3_waybel31(esk2_0))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|X4!=k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_20739, negated_conjecture, (v1_tops_2(a_1_2_waybel31(X1),esk1_0)|v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(a_1_2_waybel31(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_20615, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)),a_1_3_waybel31(esk2_0))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_20671, negated_conjecture, (v1_finset_1(esk4_2(X1,esk2_0))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(X1,esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_20987, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),a_1_0_waybel23(esk1_0))|~m1_subset_1(k1_xboole_0,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_14255, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|r2_hidden(k1_card_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0))).#trainneg
cnf(i_0_1750, plain, (k2_zfmisc_1(X1,X2)=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_21076, negated_conjecture, (v1_tops_2(k1_xboole_0,esk1_0))).#trainneg
cnf(i_0_21083, plain, (v3_ordinal1(k1_xboole_0)|~l1_pre_topc(k2_yellow_1(k1_xboole_0))|~v2_pre_topc(k2_yellow_1(k1_xboole_0))|~v1_card_1(u1_pre_topc(k2_yellow_1(k1_xboole_0))))).#trainneg
cnf(i_0_1751, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1752, plain, (k2_zfmisc_1(X1,X2)=u1_orders_2(k2_yellow_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_21183, negated_conjecture, (r1_tarski(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(esk3_2(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)),esk2_0)|~v6_waybel23(esk3_2(a_1_0_waybel23(esk1_0),a_1_3_waybel31(esk2_0)),esk2_0))).#trainneg
cnf(i_0_21004, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(X4,a_1_3_waybel31(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|X4!=k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_pre_topc(X1)|~m1_waybel23(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk2_0)|~v6_waybel23(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk2_0))).#trainneg
cnf(i_0_21096, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(X4,a_1_0_waybel31(X1))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|X4!=k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_pre_topc(X1)|~m1_waybel23(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk2_0)|~v6_waybel23(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk2_0))).#trainneg
cnf(i_0_20613, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)),a_1_0_waybel23(esk1_0))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_21194, plain, (v1_tops_2(a_1_0_waybel23(X1),X2)|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~m1_subset_1(a_1_0_waybel23(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_21080, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_21077, plain, (v1_tops_2(k1_xboole_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_21090, plain, (m1_subset_1(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_14251, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|m1_subset_1(k1_card_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0))).#trainneg
cnf(i_0_1753, plain, (k2_zfmisc_1(X1,X2)=esk8_1(k1_zfmisc_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_21679, plain, (v2_struct_0(k2_yellow_1(k1_card_1(X3)))|~v1_xboole_0(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_22309, plain, (v2_struct_0(k2_yellow_1(k1_card_1(X1)))|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1757, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(X3)|~l1_orders_2(X3)|~v2_struct_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_21752, plain, (~r2_hidden(X2,u1_orders_2(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_21835, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_22092, plain, (esk14_1(k2_yellow_1(a_1_0_waybel23(X1)))=k1_xboole_0|v8_struct_0(X1)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2447, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~v1_cantor_1(k1_xboole_0,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_21728, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k1_xboole_0))|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1769, plain, (v8_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_21091, plain, (k1_zfmisc_1(k1_xboole_0)=k1_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~r2_hidden(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_1773, plain, (v8_struct_0(X1)|v1_xboole_0(a_1_0_waybel23(X1))|~v1_xboole_0(X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_20443, negated_conjecture, (v1_finset_1(esk4_2(X1,esk2_0))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_20546, negated_conjecture, (~v1_finset_1(k1_card_1(k11_waybel_0(esk2_0,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_20730, plain, (k11_waybel_0(k2_yellow_1(X1),X2)=X1|v2_struct_0(k2_yellow_1(X1))|~r1_tarski(X1,k11_waybel_0(k2_yellow_1(X1),X2))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_22301, plain, (v2_struct_0(k2_yellow_1(k1_card_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_14332, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|m1_subset_1(k1_card_1(esk6_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0))|~v1_cantor_1(esk6_2(X1,esk1_0),esk1_0)|~v1_tops_2(esk6_2(X1,esk1_0),esk1_0))).#trainneg
cnf(i_0_1779, plain, (k2_zfmisc_1(X1,X2)=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1780, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1781, plain, (k2_zfmisc_1(X1,X2)=u1_orders_2(k2_yellow_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_22643, plain, (v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|v8_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_22770, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v2_struct_0(k2_yellow_1(k1_card_1(k11_waybel_0(esk2_0,X1)))))).#trainneg
cnf(i_0_20759, negated_conjecture, (~v1_finset_1(esk4_2(X1,esk2_0))|~r2_hidden(X1,a_1_3_waybel31(esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_21206, negated_conjecture, (v1_tops_2(a_1_0_waybel23(X1),esk1_0)|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~m1_subset_1(a_1_0_waybel23(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1782, plain, (k2_zfmisc_1(X1,X2)=esk8_1(k1_zfmisc_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_21092, plain, (k1_zfmisc_1(k1_xboole_0)=k1_numbers|~r1_tarski(k1_numbers,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_22333, plain, (v1_finset_1(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_22342, plain, (v1_relat_2(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1786, plain, (k2_zfmisc_1(X1,X2)=u1_struct_0(X3)|~l1_orders_2(X3)|~v2_struct_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1791, plain, (esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)=k1_xboole_0|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_xboole_0(esk6_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_23636, negated_conjecture, (~r2_hidden(X1,a_1_3_waybel31(esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_23628, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v1_xboole_0(k1_card_1(k11_waybel_0(esk2_0,X1))))).#trainneg
cnf(i_0_20722, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(X1),esk1_0))|v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(esk6_2(k1_card_1(X1),esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(X1),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(k11_waybel_0(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_23611, plain, (a_1_0_waybel23(X1)=k1_xboole_0|v8_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1792, plain, (esk3_2(a_1_2_waybel31(X1),X2)=k1_xboole_0|v2_struct_0(X1)|r1_tarski(a_1_2_waybel31(X1),X2)|~v1_xboole_0(esk6_2(esk3_2(a_1_2_waybel31(X1),X2),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_22343, plain, (v4_relat_2(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_22344, plain, (v8_relat_2(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1793, plain, (k1_card_1(esk6_2(X1,X2))=k1_xboole_0|v2_struct_0(X2)|~v1_xboole_0(esk6_2(k1_card_1(esk6_2(X1,X2)),X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_1794, plain, (k1_card_1(esk5_2(X1,X2))=k1_xboole_0|v2_struct_0(X2)|~v1_xboole_0(esk6_2(k1_card_1(esk5_2(X1,X2)),X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_20442, negated_conjecture, (v1_finset_1(esk5_2(X1,esk1_0))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk5_2(X1,esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_17865, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(esk6_2(X1,esk1_0))))).#trainneg
cnf(i_0_17048, negated_conjecture, (r1_tarski(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_7293, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|r2_hidden(k1_card_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1795, plain, (k1_xboole_0=X1|v2_struct_0(X2)|v1_xboole_0(a_1_2_waybel31(X2))|~v1_xboole_0(esk6_2(X1,X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~m1_subset_1(X1,a_1_2_waybel31(X2)))).#trainneg
cnf(i_0_21837, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_22399, plain, (v1_card_1(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1796, plain, (v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_finset_1(u1_pre_topc(X2)))).#trainneg
cnf(i_0_1805, plain, (r2_hidden(X1,u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2))|~r2_hidden(X1,k11_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2))).#trainneg
cnf(i_0_17895, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),esk6_2(X1,esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_20441, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk6_2(X1,esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_7006, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|v1_finset_1(esk11_1(X1))|r2_hidden(k1_card_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1807, plain, (v1_xboole_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_23820, plain, (v1_finset_1(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_23829, plain, (v1_relat_2(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1810, plain, (esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k1_xboole_0)),X1),X2)=k1_xboole_0|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k1_xboole_0)),X1),X2)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_10474, plain, (v1_yellow_0(k2_yellow_1(X1))|~m1_subset_1(X1,k4_ordinal1)|~v2_lattice3(k2_yellow_1(X1)))).#trainneg
cnf(i_0_7830, negated_conjecture, (v1_finset_1(X1)|~v1_xboole_0(esk5_2(k1_card_1(X1),esk1_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_17956, negated_conjecture, (esk6_2(X1,esk1_0)=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),esk6_2(X1,esk1_0))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_6929, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(esk4_2(k1_card_1(X2),X3),k1_zfmisc_1(u1_struct_0(X3)))|v1_finset_1(X2)|esk2_0!=X3|~m1_waybel23(k11_waybel_0(X3,k11_waybel_0(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(X3,k11_waybel_0(X1,X2)),esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_22774, negated_conjecture, (k11_waybel_0(esk2_0,X1)=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),k11_waybel_0(esk2_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_14569, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),esk1_0)|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(k11_waybel_0(X1,X2)),esk1_0),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1811, plain, (l1_orders_2(g1_orders_2(X1,esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2),X3)))|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2),X3)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_24371, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_23830, plain, (v4_relat_2(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_23831, plain, (v8_relat_2(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1812, plain, (v1_orders_2(g1_orders_2(X1,esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2),X3)))|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2),X3)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_1813, plain, (esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2),X3)=X4|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2),X3)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))|g1_orders_2(X1,esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X2),X3))!=g1_orders_2(X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_24281, negated_conjecture, (m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_1814, plain, (X1=X2|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))|g1_orders_2(X1,esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1))),X3),X4))!=g1_orders_2(X2,X5)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_21093, plain, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_21727, plain, (v3_ordinal1(k1_xboole_0)|~l1_pre_topc(k2_yellow_1(k1_xboole_0))|~v2_pre_topc(k2_yellow_1(k1_xboole_0))|~v2_struct_0(k2_yellow_1(u1_pre_topc(k2_yellow_1(k1_xboole_0)))))).#trainneg
cnf(i_0_23886, plain, (v1_card_1(k2_zfmisc_1(X2,X3))|~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1815, plain, (v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4))|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1816, plain, (v5_relat_1(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4),X2)|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_18936, plain, (u1_orders_2(k2_yellow_1(X1))=k1_xboole_0|k2_yellow_1(X1)!=X2|~v1_orders_2(X2)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1817, plain, (v1_relat_1(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4))|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1818, plain, (v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4))|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1819, plain, (v4_relat_1(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4),X1)|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_24410, plain, (v1_finset_1(k1_card_1(esk5_2(X1,X2)))|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_24540, plain, (v1_tops_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X1)|v1_xboole_0(u1_pre_topc(X1))|~m1_subset_1(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1))))).#trainneg
cnf(i_0_24347, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|r2_hidden(k1_card_1(esk12_1(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(k1_card_1(esk12_1(X1)),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_19191, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk6_2(k1_card_1(esk6_2(X2,X1)),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1821, plain, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1822, plain, (v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1823, plain, (esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3)=X1|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(X1)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_24184, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(X2),esk1_0))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(k1_card_1(esk6_2(k1_card_1(X2),esk1_0)),a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|~m1_subset_1(esk6_2(k1_card_1(X2),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk6_2(k1_card_1(X2),esk1_0),esk1_0)|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(X2),esk1_0),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_24172, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|v1_finset_1(esk4_2(X1,X2))|r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk6_2(X1,esk1_0),esk1_0)|~v1_cantor_1(esk4_2(X1,X2),esk1_0)|~v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_19198, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk5_2(k1_card_1(esk6_2(X2,X1)),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1824, plain, (m1_subset_1(X1,X2)|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X2)),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(X2)))|~r2_hidden(X1,esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X2)),X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_21543, plain, (v1_card_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_21591, plain, (v1_finset_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1825, plain, (r2_hidden(X1,X2)|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X2)),X3),X4)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(X2)))|~r2_hidden(X1,esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X2)),X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1827, plain, (v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3))|m1_subset_1(X4,X1)|r1_tarski(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3)|v2_struct_0(k2_yellow_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,esk3_2(k11_waybel_0(k2_yellow_1(k1_zfmisc_1(X1)),X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_24396, plain, (v1_finset_1(esk3_2(a_1_2_waybel31(X1),X2))|r1_tarski(a_1_2_waybel31(X1),X2)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_22109, plain, (m1_subset_1(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_21593, plain, (v1_relat_2(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_21594, plain, (v4_relat_2(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1830, plain, (v3_ordinal1(esk3_2(k11_waybel_0(k2_yellow_1(X1),X2),X3))|r1_tarski(k11_waybel_0(k2_yellow_1(X1),X2),X3)|v2_struct_0(k2_yellow_1(X1))|~v3_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_24107, negated_conjecture, (~v1_finset_1(k1_card_1(esk4_2(X1,esk2_0)))|~m1_subset_1(esk4_2(k1_card_1(esk4_2(X1,esk2_0)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_24543, plain, (k11_waybel_0(k2_yellow_1(X1),X2)=X1|v1_xboole_0(X1)|~r1_tarski(X1,k11_waybel_0(k2_yellow_1(X1),X2))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_24425, plain, (v1_finset_1(X1)|v1_xboole_0(a_1_0_waybel23(X2))|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_19029, plain, (k1_xboole_0=X1|k2_yellow_1(X2)!=g1_orders_2(X3,X1)|~v2_struct_0(k2_yellow_1(X2)))).#trainneg
cnf(i_0_21595, plain, (v8_relat_2(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1834, plain, (r1_tarski(k11_waybel_0(k2_yellow_1(k4_ordinal1),X1),k1_numbers)|v2_struct_0(k2_yellow_1(k4_ordinal1))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_24890, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|r2_hidden(k1_card_1(esk12_1(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk2_0)|~v6_waybel23(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk2_0))).#trainneg
cnf(i_0_24889, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|r2_hidden(k1_card_1(esk12_1(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_25168, plain, (k1_xboole_0=X1|g1_orders_2(X2,X1)!=k2_yellow_1(k1_xboole_0))).#trainneg
cnf(i_0_24924, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(X1),esk1_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(k1_card_1(esk6_2(k1_card_1(X1),esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(k1_card_1(X1),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(X1),esk1_0),esk1_0)|~l1_orders_2(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0))).#trainneg
cnf(i_0_24592, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(esk4_2(k1_card_1(X2),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_waybel23(k11_waybel_0(esk2_0,k11_waybel_0(X1,X2)),esk2_0)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0))).#trainneg
cnf(i_0_23792, plain, (v1_finset_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_17090, plain, (X1=k1_xboole_0|g1_orders_2(X1,k1_xboole_0)!=X2|~v1_orders_2(X2)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1873, plain, (esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)=k1_xboole_0|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_0_waybel23(X1)))|~v1_xboole_0(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_0_waybel23(X1)),X2),X3),X1))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_0_waybel23(X1))))).#trainneg
cnf(i_0_21839, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_25034, plain, (v3_ordinal1(k1_zfmisc_1(k1_xboole_0))|~v3_ordinal1(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_1874, plain, (esk3_2(a_1_0_waybel23(X1),X2)=k1_xboole_0|r1_tarski(a_1_0_waybel23(X1),X2)|~v1_xboole_0(esk5_2(esk3_2(a_1_0_waybel23(X1),X2),X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1875, plain, (k1_card_1(esk6_2(X1,X2))=k1_xboole_0|v2_struct_0(X2)|~v1_xboole_0(esk5_2(k1_card_1(esk6_2(X1,X2)),X2))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_25133, negated_conjecture, (~v1_finset_1(k1_card_1(esk4_2(X1,esk2_0)))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_23799, plain, (v1_relat_2(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1876, plain, (k1_card_1(esk5_2(X1,X2))=k1_xboole_0|~v1_xboole_0(esk5_2(k1_card_1(esk5_2(X1,X2)),X2))|~l1_pre_topc(X2)|~r2_hidden(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_23404, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1877, plain, (k1_xboole_0=X1|v1_xboole_0(a_1_0_waybel23(X2))|~v1_xboole_0(esk5_2(X1,X2))|~l1_pre_topc(X2)|~m1_subset_1(X1,a_1_0_waybel23(X2)))).#trainneg
cnf(i_0_24177, negated_conjecture, (v1_finset_1(esk6_2(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0))|v1_finset_1(esk3_2(a_1_0_waybel23(esk1_0),X1))|r2_hidden(k1_card_1(esk6_2(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0)),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_0_waybel23(esk1_0),X1)|~m1_subset_1(esk6_2(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk6_2(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0),esk1_0)|~v1_cantor_1(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0)|~v1_tops_2(esk6_2(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0),esk1_0))).#trainneg
cnf(i_0_11856, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),X2))|r2_hidden(X3,a_1_3_waybel31(X4))|r1_tarski(a_1_0_waybel23(X1),X2)|X3!=k1_card_1(esk4_2(esk3_2(a_1_0_waybel23(X1),X2),X4))|esk2_0!=X4|~m1_waybel23(esk4_2(esk3_2(a_1_0_waybel23(X1),X2),X4),esk2_0)|~v6_waybel23(esk4_2(esk3_2(a_1_0_waybel23(X1),X2),X4),esk2_0)|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),X2),k1_zfmisc_1(u1_struct_0(X4)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),X2),esk1_0)|~v1_tops_2(esk3_2(a_1_0_waybel23(X1),X2),esk1_0)|~l1_orders_2(X4)|~v3_waybel_3(X4)|~v1_lattice3(X4)|~v1_yellow_0(X4)|~v5_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10943, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_xboole_0(esk5_2(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3),X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_17982, negated_conjecture, (v1_finset_1(esk11_1(X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk11_1(X1))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18850, plain, (g1_orders_2(k1_xboole_0,k1_xboole_0)=k2_yellow_1(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_24223, plain, (v4_relat_2(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1879, plain, (v1_xboole_0(u1_pre_topc(k2_yellow_1(X1)))|r2_hidden(esk8_1(u1_pre_topc(k2_yellow_1(X1))),k1_zfmisc_1(X1))|~l1_pre_topc(k2_yellow_1(X1)))).#trainneg
cnf(i_0_24800, plain, (u1_orders_2(k2_yellow_1(X1))=k1_xboole_0|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_12509, negated_conjecture, (v1_finset_1(k1_card_1(X1))|r2_hidden(k1_card_1(X1),a_1_0_waybel31(esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,k1_card_1(X1)),esk2_0)|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(k1_card_1(X1),esk1_0)|~v1_tops_2(k1_card_1(X1),esk1_0))).#trainneg
cnf(i_0_8736, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k11_waybel_0(esk2_0,X2))|~m1_subset_1(k11_waybel_0(esk2_0,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X2,esk1_0))).#trainneg
cnf(i_0_24230, plain, (v8_relat_2(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1882, plain, (v3_ordinal1(u1_orders_2(k2_yellow_1(X1)))|~v1_card_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(X1,X1)))))).#trainneg
cnf(i_0_22719, plain, (k1_zfmisc_1(k1_xboole_0)=k1_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_1883, plain, (v3_ordinal1(u1_orders_2(X1))|~l1_orders_2(X1)|~v1_card_1(k1_zfmisc_1(k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))).#trainneg
cnf(i_0_1887, negated_conjecture, (esk5_2(X1,esk1_0)=u1_struct_0(esk2_0)|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk5_2(X1,esk1_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_8403, negated_conjecture, (v1_finset_1(k1_card_1(k11_waybel_0(X1,X2)))|~m1_subset_1(esk7_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),k4_ordinal1)|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7834, negated_conjecture, (v1_finset_1(esk6_2(X1,X2))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|v2_struct_0(X2)|~m1_waybel23(k11_waybel_0(esk2_0,esk6_2(X1,X2)),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,esk6_2(X1,X2)),esk2_0)|~m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_20623, negated_conjecture, (v1_finset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|r2_hidden(X4,a_1_0_waybel23(esk1_0))|r1_tarski(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3)|v2_struct_0(k2_yellow_1(u1_pre_topc(X1)))|X4!=k1_card_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3))|esk2_0!=X1|~m1_subset_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~v1_tops_2(esk3_2(k11_waybel_0(k2_yellow_1(u1_pre_topc(X1)),X2),X3),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_15208, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk6_2(X2,esk1_0))|~r2_hidden(X2,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_24383, plain, (v1_card_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1915, plain, (v3_ordinal1(esk4_2(k1_card_1(esk11_1(X1)),X1))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14253, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|r2_hidden(k1_card_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)),a_1_3_waybel31(esk2_0))|~m1_subset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0)|~v1_tops_2(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0))).#trainneg
cnf(i_0_13856, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(esk5_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_17002, negated_conjecture, (v1_finset_1(X1)|~m1_subset_1(esk5_2(X1,esk1_0),k4_ordinal1)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_13243, negated_conjecture, (r2_hidden(esk3_2(u1_struct_0(esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(u1_struct_0(esk2_0),X1))).#trainneg
cnf(i_0_24171, negated_conjecture, (v1_finset_1(esk6_2(X1,esk1_0))|v1_finset_1(esk5_2(X1,X2))|r2_hidden(k1_card_1(esk6_2(X1,esk1_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk6_2(X1,esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk6_2(X1,esk1_0),esk1_0)|~v1_cantor_1(esk5_2(X1,X2),esk1_0)|~v1_tops_2(esk6_2(X1,esk1_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_25345, negated_conjecture, (~v1_finset_1(X1)|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_25275, plain, (X1=k1_xboole_0|~v2_struct_0(g1_orders_2(X1,k1_xboole_0)))).#trainneg
cnf(i_0_1916, plain, (v1_finset_1(esk4_2(k1_card_1(esk11_1(X1)),X1))|~v1_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_25470, plain, (g1_orders_2(k1_xboole_0,k1_xboole_0)=k2_yellow_1(k1_xboole_0))).#trainneg
cnf(i_0_1918, plain, (esk4_2(k1_card_1(esk11_1(X1)),X1)=u1_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(esk4_2(k1_card_1(esk11_1(X1)),X1)))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_25350, negated_conjecture, (~r2_hidden(X1,a_1_3_waybel31(esk2_0))|~v1_finset_1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_25245, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(esk4_2(k1_card_1(X2),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_25591, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_3_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_25590, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_25502, plain, (g1_orders_2(k1_xboole_0,k1_xboole_0)=k2_yellow_1(X1)|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_25240, negated_conjecture, (u1_orders_2(esk2_0)=k1_xboole_0|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_1919, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk4_2(k1_card_1(esk11_1(X2)),X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_17176, negated_conjecture, (~v1_xboole_0(esk6_2(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_1920, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,esk4_2(k1_card_1(esk11_1(X2)),X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1922, plain, (v1_xboole_0(esk4_2(k1_card_1(esk11_1(X1)),X1))|m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,esk4_2(k1_card_1(esk11_1(X1)),X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_25588, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_0_waybel31(X2))|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_25258, negated_conjecture, (m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_25596, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(esk1_0),X1))|r2_hidden(esk3_2(a_1_0_waybel23(esk1_0),X1),a_1_0_waybel31(esk2_0))|r1_tarski(a_1_0_waybel23(esk1_0),X1)|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(esk1_0),X1),esk1_0))).#trainneg
cnf(i_0_25241, plain, (u1_orders_2(X1)=k1_xboole_0|X1!=k2_yellow_1(k1_xboole_0)|~v1_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_25682, negated_conjecture, (esk5_2(X1,esk1_0)=u1_struct_0(esk2_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk5_2(X1,esk1_0))))).#trainneg
cnf(i_0_25250, negated_conjecture, (m1_subset_1(X1,a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_25589, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_25493, plain, (v1_xboole_0(u1_pre_topc(k2_yellow_1(k1_xboole_0)))|r2_hidden(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~l1_pre_topc(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_1924, plain, (m1_subset_1(esk4_2(k1_card_1(esk11_1(k2_yellow_1(X1))),k2_yellow_1(X1)),k1_zfmisc_1(X1))|~v3_waybel_3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1))|~v1_yellow_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_17171, negated_conjecture, (~m1_subset_1(esk5_2(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0),k4_ordinal1)|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_19196, plain, (v8_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(esk6_2(k1_card_1(esk6_2(X2,X1)),X1),k4_ordinal1)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~r2_hidden(X2,a_1_0_waybel23(X1)))).#trainneg
cnf(i_0_1927, plain, (v8_struct_0(X1)|v2_struct_0(X1)|r1_tarski(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3)|v2_struct_0(k2_yellow_1(a_1_2_waybel31(X1)))|~v1_xboole_0(esk3_2(k11_waybel_0(k2_yellow_1(a_1_2_waybel31(X1)),X2),X3))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(a_1_2_waybel31(X1))))).#trainneg
cnf(i_0_17181, negated_conjecture, (~m1_subset_1(esk6_2(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0),k4_ordinal1)|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_1931, plain, (v8_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(a_1_2_waybel31(X1))|~v1_xboole_0(X2)|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,a_1_2_waybel31(X1)))).#trainneg
cnf(i_0_8317, negated_conjecture, (v1_finset_1(esk14_1(X1))|r2_hidden(X2,a_1_3_waybel31(esk2_0))|v2_struct_0(X1)|X2!=k1_card_1(esk14_1(X1))|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_16778, negated_conjecture, (v1_finset_1(esk8_1(a_1_0_waybel23(esk1_0)))|r2_hidden(esk8_1(a_1_0_waybel23(esk1_0)),a_1_0_waybel31(esk2_0))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_14987, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)))).#trainneg
cnf(i_0_25833, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_26232, negated_conjecture, (m1_subset_1(k1_card_1(X1),a_1_3_waybel31(esk2_0))|v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_26102, plain, (r2_hidden(esk3_2(esk4_2(k1_card_1(esk11_1(X1)),X1),X2),u1_struct_0(X1))|r1_tarski(esk4_2(k1_card_1(esk11_1(X1)),X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26010, plain, (k2_yellow_1(k1_xboole_0)=k2_yellow_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_26220, negated_conjecture, (m1_subset_1(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_3_waybel31(esk2_0))|v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(X1)|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,X1),esk2_0)|~v6_waybel23(k11_waybel_0(X2,X1),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_26084, plain, (m1_subset_1(esk3_2(esk4_2(k1_card_1(esk11_1(X1)),X1),X2),u1_struct_0(X1))|r1_tarski(esk4_2(k1_card_1(esk11_1(X1)),X1),X2)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_25921, negated_conjecture, (r2_hidden(k1_card_1(esk4_2(X1,esk2_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_26230, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(X1,X2)),a_1_3_waybel31(esk2_0))|v1_finset_1(k11_waybel_0(X1,X2))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26229, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(X1,X2)),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18458, plain, (k2_yellow_1(k1_xboole_0)=X1|~v1_orders_2(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_25855, plain, (u1_orders_2(k2_yellow_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_26370, negated_conjecture, (~v1_finset_1(esk8_1(a_1_3_waybel31(esk2_0))))).#trainneg
cnf(i_0_26898, plain, (k2_yellow_1(X2)=k2_yellow_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_27132, plain, (~v2_lattice3(k2_yellow_1(k1_xboole_0)))).#trainneg
cnf(i_0_26417, plain, (r1_tarski(esk4_2(k1_card_1(esk11_1(X1)),X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_27066, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(k1_card_1(X2),a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_25916, negated_conjecture, (~r2_hidden(X1,a_1_3_waybel31(esk2_0))|~v2_struct_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_27157, plain, (v1_partfun1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_26064, negated_conjecture, (v1_partfun1(k1_xboole_0,u1_struct_0(esk2_0))|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_26819, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|k2_yellow_1(X1)!=esk2_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1941, plain, (k1_card_1(X3)=X1|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1942, plain, (r2_hidden(k1_card_1(X3),a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1943, plain, (r2_hidden(k1_card_1(X3),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_26218, negated_conjecture, (m1_subset_1(k1_card_1(esk5_2(X1,esk1_0)),a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_26235, negated_conjecture, (m1_subset_1(X1,a_1_3_waybel31(esk2_0))|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_27097, negated_conjecture, (r1_tarski(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel31(esk2_0)),a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_26011, negated_conjecture, (g1_orders_2(u1_struct_0(esk2_0),k1_xboole_0)=esk2_0|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_1944, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=k1_card_1(X4)|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X3,X2))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1945, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(X4)|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X3,X2))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_25834, negated_conjecture, (~v1_finset_1(X1)|~r2_hidden(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_26002, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_26697, negated_conjecture, (X1=u1_struct_0(esk2_0)|k2_yellow_1(k1_xboole_0)!=esk2_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_27186, negated_conjecture, (~v1_xboole_0(esk8_1(a_1_3_waybel31(esk2_0))))).#trainneg
cnf(i_0_27187, negated_conjecture, (~m1_subset_1(esk8_1(a_1_3_waybel31(esk2_0)),k4_ordinal1))).#trainneg
cnf(i_0_1948, plain, (k1_card_1(X3)=X1|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1949, plain, (r2_hidden(k1_card_1(X3),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_25807, negated_conjecture, (r1_tarski(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0))|v1_finset_1(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)),esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_26210, negated_conjecture, (m1_subset_1(k1_card_1(esk6_2(X1,esk1_0)),a_1_3_waybel31(esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_26234, negated_conjecture, (m1_subset_1(k1_card_1(X1),a_1_3_waybel31(esk2_0))|v1_finset_1(X1)|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_26012, negated_conjecture, (v4_orders_2(g1_orders_2(u1_struct_0(esk2_0),k1_xboole_0))|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_1950, plain, (r2_hidden(k1_card_1(X3),a_1_0_waybel23(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_27189, negated_conjecture, (~v2_struct_0(k2_yellow_1(esk8_1(a_1_3_waybel31(esk2_0)))))).#trainneg
cnf(i_0_1952, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|X1!=k1_card_1(X4)|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X3,X2))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_26013, negated_conjecture, (v5_orders_2(g1_orders_2(u1_struct_0(esk2_0),k1_xboole_0))|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_1953, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=k1_card_1(X4)|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk5_2(X3,X2))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_27188, negated_conjecture, (~m1_subset_1(k1_card_1(esk8_1(a_1_3_waybel31(esk2_0))),k4_ordinal1))).#trainneg
cnf(i_0_26231, negated_conjecture, (m1_subset_1(k1_card_1(u1_struct_0(X1)),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(u1_struct_0(X1),esk1_0)|~v1_tops_2(u1_struct_0(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_28404, negated_conjecture, (k1_card_1(X1)=X1|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28418, negated_conjecture, (~m1_subset_1(X1,k4_ordinal1)|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_27484, plain, (X1=X2|k2_yellow_1(X3)!=k2_yellow_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_28402, negated_conjecture, (v1_finset_1(X1)|m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28392, negated_conjecture, (m1_subset_1(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28199, negated_conjecture, (m1_subset_1(a_1_0_waybel23(esk1_0),k1_zfmisc_1(a_1_0_waybel31(esk2_0))))).#trainneg
cnf(i_0_28075, negated_conjecture, (~v1_finset_1(X1)|~m1_subset_1(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_27901, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(k1_card_1(X2),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_28413, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_27929, negated_conjecture, (~v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_24546, plain, (v1_finset_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_27875, negated_conjecture, (v1_tops_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0)|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_28389, negated_conjecture, (r1_tarski(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel31(esk2_0)),a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_27904, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(k1_card_1(X2),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(X2,esk1_0)|~v1_tops_2(X2,esk1_0))).#trainneg
cnf(i_0_26762, plain, (X1=X2|k2_yellow_1(X1)!=k2_yellow_1(k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_24553, plain, (v1_relat_2(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_25885, plain, (esk4_2(k1_card_1(esk11_1(k2_yellow_1(X1))),k2_yellow_1(X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(k1_card_1(esk11_1(k2_yellow_1(X1))),k2_yellow_1(X1))))|~v3_waybel_3(k2_yellow_1(X1))|~v1_lattice3(k2_yellow_1(X1))|~v1_yellow_0(k2_yellow_1(X1)))).#trainneg
cnf(i_0_26206, negated_conjecture, (m1_subset_1(k1_card_1(esk14_1(X1)),a_1_3_waybel31(esk2_0))|v1_finset_1(esk14_1(X1))|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26849, plain, (k1_xboole_0=X1|k2_yellow_1(X2)!=k2_yellow_1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_24670, plain, (v4_relat_2(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_26016, negated_conjecture, (~v2_struct_0(g1_orders_2(u1_struct_0(esk2_0),k1_xboole_0))|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_25798, negated_conjecture, (r1_tarski(X1,a_1_3_waybel31(esk2_0))|~m1_subset_1(esk3_2(X1,a_1_3_waybel31(esk2_0)),a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_26221, negated_conjecture, (m1_subset_1(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_3_waybel31(esk2_0))|v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k1_card_1(X1))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_24677, plain, (v8_relat_2(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_26005, plain, (k2_yellow_1(a_1_2_waybel31(X1))=k2_yellow_1(k1_xboole_0)|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_24759, plain, (v1_card_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_28684, negated_conjecture, (v1_finset_1(esk11_1(X1))|m1_subset_1(k1_card_1(esk11_1(X1)),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_29060, plain, (k2_yellow_1(a_1_0_waybel23(X1))=k2_yellow_1(k1_xboole_0)|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_28562, negated_conjecture, (~v1_finset_1(esk8_1(a_1_0_waybel31(esk2_0))))).#trainneg
cnf(i_0_28683, negated_conjecture, (v1_finset_1(esk11_1(X1))|m1_subset_1(k1_card_1(esk11_1(X1)),a_1_3_waybel31(esk2_0))|~m1_subset_1(esk11_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk11_1(X1),esk1_0)|esk2_0!=X1|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_28630, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_26224, negated_conjecture, (m1_subset_1(k1_card_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0)),a_1_3_waybel31(esk2_0))|v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26416, negated_conjecture, (m1_subset_1(esk3_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),X1))).#trainneg
cnf(i_0_28554, negated_conjecture, (v1_finset_1(a_1_0_waybel23(esk1_0))|~v1_finset_1(a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_26201, negated_conjecture, (m1_subset_1(esk3_2(a_1_0_waybel23(esk1_0),X1),a_1_3_waybel31(esk2_0))|r1_tarski(a_1_0_waybel23(esk1_0),X1))).#trainneg
cnf(i_0_26228, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_3_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_28434, negated_conjecture, (k1_card_1(esk8_1(a_1_3_waybel31(esk2_0)))=esk8_1(a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28574, negated_conjecture, (v1_finset_1(a_1_0_waybel23(esk1_0))|~v1_finset_1(a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28553, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v3_ordinal1(k1_zfmisc_1(a_1_0_waybel31(esk2_0))))).#trainneg
cnf(i_0_26006, plain, (k2_yellow_1(a_1_0_waybel23(X1))=k2_yellow_1(k1_xboole_0)|v8_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_28680, negated_conjecture, (v1_finset_1(esk12_1(X1))|m1_subset_1(k1_card_1(esk12_1(X1)),a_1_3_waybel31(esk2_0))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_28546, negated_conjecture, (m1_subset_1(k1_card_1(esk4_2(X1,esk2_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28661, negated_conjecture, (r1_tarski(a_1_3_waybel31(esk2_0),a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_28573, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v3_ordinal1(k1_zfmisc_1(a_1_3_waybel31(esk2_0))))).#trainneg
cnf(i_0_26866, plain, (u1_struct_0(X1)=k1_xboole_0|X1!=k2_yellow_1(X2)|~v1_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_26227, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_3_waybel31(esk2_0))|v1_finset_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_26400, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|m1_subset_1(k1_card_1(X2),a_1_3_waybel31(esk2_0))|v1_finset_1(X2)|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_28420, negated_conjecture, (~v1_xboole_0(X1)|~m1_subset_1(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28179, negated_conjecture, (v1_finset_1(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)))|r1_tarski(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~l1_pre_topc(X1)|~m1_waybel23(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)),esk2_0)|~v6_waybel23(esk3_2(a_1_0_waybel23(X1),a_1_3_waybel31(esk2_0)),esk2_0))).#trainneg
cnf(i_0_26155, negated_conjecture, (v1_finset_1(esk4_2(X1,esk2_0))|r2_hidden(k1_card_1(esk4_2(X1,esk2_0)),a_1_0_waybel23(esk1_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(X1,esk2_0),esk1_0)|~v1_tops_2(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_29335, negated_conjecture, (r1_tarski(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_29568, plain, (a_1_0_waybel23(X1)=X2|v8_struct_0(X1)|k2_yellow_1(k1_xboole_0)!=k2_yellow_1(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_28557, negated_conjecture, (a_1_0_waybel31(esk2_0)=a_1_0_waybel23(esk1_0)|~m1_subset_1(a_1_0_waybel31(esk2_0),k1_zfmisc_1(a_1_0_waybel23(esk1_0))))).#trainneg
cnf(i_0_29239, negated_conjecture, (~v1_xboole_0(esk8_1(a_1_0_waybel31(esk2_0))))).#trainneg
cnf(i_0_29240, negated_conjecture, (~m1_subset_1(esk8_1(a_1_0_waybel31(esk2_0)),k4_ordinal1))).#trainneg
cnf(i_0_26364, negated_conjecture, (r2_hidden(esk8_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_29524, plain, (g1_orders_2(a_1_0_waybel23(X1),k1_xboole_0)=k2_yellow_1(k1_xboole_0)|v8_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_26116, plain, (m1_subset_1(esk8_1(esk4_2(k1_card_1(esk11_1(X1)),X1)),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26166, negated_conjecture, (v1_finset_1(esk4_2(X1,esk2_0))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(X1,esk2_0),esk1_0)|~v1_tops_2(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_29276, negated_conjecture, (r2_hidden(esk3_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),X1),u1_struct_0(esk2_0))|r1_tarski(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),X1)|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_26148, negated_conjecture, (esk4_2(X1,esk2_0)=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(esk4_2(X1,esk2_0)))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_29350, negated_conjecture, (v1_finset_1(a_1_0_waybel23(esk1_0))|~m1_subset_1(a_1_0_waybel31(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_29242, negated_conjecture, (~v2_struct_0(k2_yellow_1(esk8_1(a_1_0_waybel31(esk2_0)))))).#trainneg
cnf(i_0_29241, negated_conjecture, (~m1_subset_1(k1_card_1(esk8_1(a_1_0_waybel31(esk2_0))),k4_ordinal1))).#trainneg
cnf(i_0_26279, negated_conjecture, (esk2_0!=X1|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(k2_yellow_1(u1_struct_0(X1)))),k2_yellow_1(u1_struct_0(X1))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk11_1(k2_yellow_1(u1_struct_0(X1)))),k2_yellow_1(u1_struct_0(X1))),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(k2_yellow_1(u1_struct_0(X1)))),k2_yellow_1(u1_struct_0(X1))),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v3_waybel_3(k2_yellow_1(u1_struct_0(X1)))|~v1_lattice3(k2_yellow_1(u1_struct_0(X1)))|~v1_yellow_0(k2_yellow_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_28950, plain, (g1_orders_2(a_1_2_waybel31(X1),k1_xboole_0)=k2_yellow_1(k1_xboole_0)|v8_struct_0(X1)|v2_struct_0(X1)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_26684, plain, (g1_orders_2(X1,u1_orders_2(k2_yellow_1(k1_xboole_0)))=k2_yellow_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_27873, plain, (esk4_2(k1_card_1(esk11_1(X1)),X1)=u1_struct_0(X1)|~r1_tarski(u1_struct_0(X1),esk4_2(k1_card_1(esk11_1(X1)),X1))|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_29352, negated_conjecture, (v1_finset_1(a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(a_1_0_waybel31(esk2_0))))).#trainneg
cnf(i_0_26704, plain, (X1=X2|g1_orders_2(X2,X3)!=k2_yellow_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_29479, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v1_card_1(k1_zfmisc_1(a_1_0_waybel31(esk2_0))))).#trainneg
cnf(i_0_29351, negated_conjecture, (v1_finset_1(a_1_0_waybel23(esk1_0))|~m1_subset_1(k1_card_1(a_1_0_waybel31(esk2_0)),k4_ordinal1))).#trainneg
cnf(i_0_29589, plain, (u1_struct_0(X1)=a_1_0_waybel23(X2)|v8_struct_0(X2)|X1!=k2_yellow_1(k1_xboole_0)|~v1_orders_2(X1)|~l1_orders_2(X1)|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_28571, negated_conjecture, (~v1_finset_1(k1_card_1(esk4_2(X1,esk2_0)))|~r2_hidden(X1,a_1_3_waybel31(esk2_0)))).#trainneg
cnf(i_0_28994, plain, (a_1_2_waybel31(X1)=X2|v8_struct_0(X1)|v2_struct_0(X1)|k2_yellow_1(k1_xboole_0)!=k2_yellow_1(X2)|~v1_finset_1(u1_pre_topc(X1))|~l1_pre_topc(X1)|~v6_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7928, negated_conjecture, (v1_finset_1(esk12_1(X1))|r2_hidden(k1_card_1(esk12_1(X1)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18374, negated_conjecture, (m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_12896, negated_conjecture, (v1_cantor_1(esk4_2(X1,esk2_0),esk1_0)|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_8320, negated_conjecture, (v1_finset_1(esk12_1(X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk12_1(X1))|esk2_0!=X1|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_25350, negated_conjecture, (~v1_finset_1(esk4_2(X1,esk2_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_29827, negated_conjecture, (m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_28023, negated_conjecture, (k1_xboole_0=X1|g1_orders_2(X2,X1)!=esk2_0|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_2042, plain, (u1_struct_0(X3)=X1|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2))|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_29483, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v1_card_1(k1_zfmisc_1(k1_zfmisc_1(a_1_0_waybel31(esk2_0)))))).#trainneg
cnf(i_0_2043, plain, (r2_hidden(u1_struct_0(X3),a_1_0_waybel23(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2))|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_2044, plain, (r2_hidden(u1_struct_0(X3),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X1,X2))|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_13782, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(X3,a_1_0_waybel23(esk1_0))|X3!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(k11_waybel_0(X2,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_29480, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(a_1_0_waybel31(esk2_0)))))))).#trainneg
cnf(i_0_2045, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|v2_struct_0(X2)|X1!=u1_struct_0(X4)|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X3,X2))|~l1_orders_2(X4)|~v2_struct_0(X4))).#trainneg
cnf(i_0_29481, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(a_1_0_waybel31(esk2_0))))))))).#trainneg
cnf(i_0_2046, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=u1_struct_0(X4)|~r2_hidden(X3,a_1_2_waybel31(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk6_2(X3,X2))|~l1_orders_2(X4)|~v2_struct_0(X4))).#trainneg
cnf(i_0_28420, negated_conjecture, (~v1_xboole_0(X1)|~m1_subset_1(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_29482, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(a_1_0_waybel31(esk2_0)))))))))).#trainneg
cnf(i_0_2049, plain, (u1_struct_0(X3)=X1|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2))|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_2050, plain, (r2_hidden(u1_struct_0(X3),a_1_2_waybel31(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2))|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_25921, negated_conjecture, (r2_hidden(k1_card_1(esk4_2(X1,esk2_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_13734, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_0_waybel23(esk1_0))|X3!=k1_card_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(k11_waybel_0(X1,X2)),esk2_0),esk1_0)|~v1_tops_2(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13736, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26014, negated_conjecture, (v2_lattice3(g1_orders_2(u1_struct_0(esk2_0),k1_xboole_0))|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_25684, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(esk4_2(X2,esk2_0))|~r2_hidden(X2,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_2051, plain, (r2_hidden(u1_struct_0(X3),a_1_0_waybel23(X2))|~r2_hidden(X1,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X1,X2))|~l1_orders_2(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_2053, plain, (r2_hidden(X1,a_1_0_waybel23(X2))|X1!=u1_struct_0(X4)|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v1_xboole_0(esk5_2(X3,X2))|~l1_orders_2(X4)|~v2_struct_0(X4))).#trainneg
cnf(i_0_30461, negated_conjecture, (r2_hidden(X1,a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_2054, plain, (r2_hidden(X1,a_1_2_waybel31(X2))|v2_struct_0(X2)|X1!=u1_struct_0(X4)|~r2_hidden(X3,a_1_0_waybel23(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2)|~v1_xboole_0(esk5_2(X3,X2))|~l1_orders_2(X4)|~v2_struct_0(X4))).#trainneg
cnf(i_0_30450, negated_conjecture, (r2_hidden(k1_card_1(esk7_2(X1,esk2_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_30365, negated_conjecture, (v3_ordinal1(a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(k1_zfmisc_1(k1_zfmisc_1(a_1_0_waybel31(esk2_0))))))).#trainneg
cnf(i_0_26224, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0))|v1_finset_1(esk12_1(X1))|m1_subset_1(k1_card_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0)),a_1_0_waybel23(esk1_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk12_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_cantor_1(esk12_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk12_1(X1)),esk2_0),esk1_0)|~v1_tops_2(esk12_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26400, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(k1_card_1(X2),a_1_0_waybel23(esk1_0))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26015, negated_conjecture, (v3_orders_2(g1_orders_2(u1_struct_0(esk2_0),k1_xboole_0))|k2_yellow_1(k1_xboole_0)!=esk2_0)).#trainneg
cnf(i_0_28630, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0)))).#trainneg
cnf(i_0_3290, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(k11_waybel_0(esk2_0,X1))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_13735, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|v1_finset_1(esk11_1(X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0))|esk2_0!=X1|~m1_subset_1(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk11_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(esk11_1(X1)),esk2_0),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2078, plain, (k1_card_1(u1_struct_0(X2))=u1_struct_0(X2)|~v1_xboole_0(X1)|~l1_orders_2(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_30416, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k11_waybel_0(X2,X1))|v1_finset_1(X1)|r2_hidden(X3,a_1_0_waybel23(esk1_0))|X3!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|esk2_0!=X2|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X2,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(k11_waybel_0(X2,X1),esk1_0)|~v1_tops_2(k11_waybel_0(X2,X1),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_13726, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(X1)|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~v6_waybel23(k11_waybel_0(esk2_0,X1),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0))).#trainneg
cnf(i_0_7037, negated_conjecture, (v1_finset_1(esk12_1(esk2_0))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(esk12_1(esk2_0))|~m1_subset_1(esk12_1(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk12_1(esk2_0),esk1_0)|~v1_tops_2(esk12_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_27096, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),a_1_0_waybel31(esk2_0))|~r2_hidden(k1_xboole_0,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_27993, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),a_1_0_waybel31(esk2_0))|~m1_subset_1(k1_xboole_0,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_2087, plain, (g1_orders_2(k1_card_1(X2),u1_orders_2(X1))=X1|~v1_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_8317, negated_conjecture, (v1_finset_1(esk14_1(X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk14_1(X1))|esk2_0!=X1|~m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_cantor_1(esk14_1(X1),esk1_0)|~v1_tops_2(esk14_1(X1),esk1_0)|~l1_orders_2(X1)|~v3_waybel_3(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13731, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k11_waybel_0(esk2_0,X1))|v1_finset_1(X1)|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0))).#trainneg
cnf(i_0_26221, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),X2))|v1_finset_1(k1_card_1(X1))|m1_subset_1(k1_card_1(esk4_2(k1_card_1(X1),X2)),a_1_0_waybel23(esk1_0))|esk2_0!=X2|~m1_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~v6_waybel23(k11_waybel_0(X2,k1_card_1(X1)),esk2_0)|~m1_subset_1(esk4_2(k1_card_1(X1),X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~v1_cantor_1(esk4_2(k1_card_1(X1),X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),X2),esk1_0)|~l1_orders_2(X2)|~v3_waybel_3(X2)|~v1_lattice3(X2)|~v1_yellow_0(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_30781, negated_conjecture, (k1_card_1(X1)=X1|~m1_subset_1(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_14257, negated_conjecture, (v1_finset_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0))|~v1_cantor_1(esk6_2(k1_card_1(esk11_1(esk2_0)),esk1_0),esk1_0))).#trainneg
cnf(i_0_2089, plain, (v4_orders_2(g1_orders_2(k1_card_1(X2),u1_orders_2(X1)))|~l1_orders_2(X1)|~v4_orders_2(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2090, plain, (v5_orders_2(g1_orders_2(k1_card_1(X2),u1_orders_2(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_30790, negated_conjecture, (v1_finset_1(X1)|r2_hidden(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~m1_subset_1(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_30888, negated_conjecture, (r1_tarski(a_1_0_waybel31(esk2_0),X1)|~v1_finset_1(esk3_2(a_1_0_waybel31(esk2_0),X1)))).#trainneg
cnf(i_0_30796, negated_conjecture, (~m1_subset_1(X1,k4_ordinal1)|~m1_subset_1(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_13737, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|r2_hidden(X1,a_1_0_waybel23(esk1_0))|X1!=k1_card_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_tops_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0))).#trainneg
cnf(i_0_2092, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(k1_card_1(X2))))|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_26345, negated_conjecture, (~v1_xboole_0(esk4_2(esk8_1(a_1_0_waybel23(esk1_0)),esk2_0))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_26348, negated_conjecture, (~v1_xboole_0(esk7_2(esk8_1(a_1_0_waybel23(esk1_0)),esk2_0))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk8_1(a_1_0_waybel23(esk1_0)),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(esk8_1(a_1_0_waybel23(esk1_0)),esk1_0))).#trainneg
cnf(i_0_2096, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(k1_card_1(X2),k1_card_1(X2))))|~l1_orders_2(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_30868, negated_conjecture, (k1_card_1(esk3_2(a_1_0_waybel31(esk2_0),X1))=esk3_2(a_1_0_waybel31(esk2_0),X1)|r1_tarski(a_1_0_waybel31(esk2_0),X1))).#trainneg
cnf(i_0_30793, negated_conjecture, (~v2_struct_0(k2_yellow_1(X1))|~m1_subset_1(X1,a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_30765, negated_conjecture, (r1_tarski(X1,a_1_0_waybel31(esk2_0))|~m1_subset_1(esk3_2(X1,a_1_0_waybel31(esk2_0)),a_1_0_waybel31(esk2_0)))).#trainneg
cnf(i_0_30854, negated_conjecture, (m1_subset_1(esk3_2(a_1_0_waybel31(esk2_0),X1),a_1_0_waybel23(esk1_0))|r1_tarski(a_1_0_waybel31(esk2_0),X1))).#trainneg
cnf(i_0_13724, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(X1)|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0))).#trainneg
cnf(i_0_2102, plain, (v3_ordinal1(u1_orders_2(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k1_card_1(X2),k1_card_1(X2))))|~l1_orders_2(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_31002, negated_conjecture, (v1_finset_1(k11_waybel_0(X1,X2))|v1_finset_1(X2)|m1_subset_1(k1_card_1(X2),a_1_0_waybel23(esk1_0))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_waybel23(k11_waybel_0(X1,X2),esk2_0)|~v6_waybel23(k11_waybel_0(X1,X2),esk2_0))).#trainneg
cnf(i_0_31824, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X2),esk2_0))|v1_finset_1(k11_waybel_0(X1,X2))|r2_hidden(X3,a_1_0_waybel23(esk1_0))|v1_finset_1(X2)|X3!=k1_card_1(esk4_2(k1_card_1(X2),esk2_0))|~m1_subset_1(esk4_2(k1_card_1(X2),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k11_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(X1,X2),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X2),esk2_0),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_28546, negated_conjecture, (m1_subset_1(k1_card_1(esk4_2(X1,esk2_0)),a_1_0_waybel23(esk1_0))|~r2_hidden(X1,a_1_0_waybel23(esk1_0)))).#trainneg
cnf(i_0_13732, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(X1)|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|esk2_0!=X3|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(X1,esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(X1,esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_31068, plain, (k1_card_1(u1_struct_0(X1))=u1_struct_0(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_27875, negated_conjecture, (v1_tops_2(esk4_2(k1_card_1(esk11_1(esk2_0)),esk2_0),esk1_0))).#trainneg
cnf(i_0_7457, negated_conjecture, (v1_finset_1(esk4_2(k1_card_1(X1),esk2_0))|v1_finset_1(k1_card_1(X1))|r2_hidden(X2,a_1_0_waybel23(esk1_0))|X2!=k1_card_1(esk4_2(k1_card_1(X1),esk2_0))|esk2_0!=X3|~m1_subset_1(esk4_2(k1_card_1(X1),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(k1_card_1(X1),k1_zfmisc_1(u1_struct_0(X3)))|~v1_cantor_1(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_cantor_1(k1_card_1(X1),esk1_0)|~v1_tops_2(esk4_2(k1_card_1(X1),esk2_0),esk1_0)|~v1_tops_2(k1_card_1(X1),esk1_0)|~l1_orders_2(X3)|~v3_waybel_3(X3)|~v1_lattice3(X3)|~v1_yellow_0(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_2106, plain, (l1_orders_2(g1_orders_2(k1_card_1(X2),u1_orders_2(X1)))|~l1_orders_2(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2107, plain, (v1_orders_2(g1_orders_2(k1_card_1(X2),u1_orders_2(X1)))|~l1_orders_2(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_30324, negated_conjecture, (~r2_hidden(X1,a_1_0_waybel23(esk1_0))|~v2_struct_0(k2_yellow_1(esk4_2(X1,esk2_0))))).#trainneg
cnf(i_0_30052, plain, (u1_struct_0(X1)=a_1_2_waybel31(X2)|v8_struct_0(X2)|v2_struct_0(X2)|X1!=k2_yellow_1(k1_xboole_0)|~v1_orders_2(X1)|~l1_orders_2(X1)|~v1_finset_1(u1_pre_topc(X2))|~l1_pre_topc(X2)|~v6_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_26228, negated_conjecture, (m1_subset_1(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_0_waybel23(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(X1,esk1_0))).#trainneg
cnf(i_0_26227, negated_conjecture, (v1_finset_1(k11_waybel_0(esk2_0,X1))|m1_subset_1(k1_card_1(k11_waybel_0(esk2_0,X1)),a_1_0_waybel23(esk1_0))|~m1_subset_1(k11_waybel_0(esk2_0,X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~v1_cantor_1(k11_waybel_0(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_2110, plain, (u1_orders_2(X1)=X2|g1_orders_2(k1_card_1(X4),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1)|~v1_xboole_0(X4)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2111, plain, (k1_card_1(X4)=X2|g1_orders_2(k1_card_1(X4),u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1)|~v1_xboole_0(X4)|~v2_struct_0(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 124
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 263
# Removed in clause preprocessing      : 62
# Initial clauses in saturation        : 201
# Processed clauses                    : 5785
# ...of these trivial                  : 65
# ...subsumed                          : 3420
# ...remaining for further processing  : 2300
# Other redundant clauses eliminated   : 33
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 170
# Backward-rewritten                   : 359
# Generated clauses                    : 32260
# ...of the previous two non-trivial   : 29920
# Contextual simplify-reflections      : 414
# Paramodulations                      : 32142
# Factorizations                       : 0
# Equation resolutions                 : 118
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1767
#    Positive orientable unit clauses  : 122
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 51
#    Non-unit-clauses                  : 1594
# Current number of unprocessed clauses: 23518
# ...number of literals in the above   : 205764
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2831
# Proof object given clauses           : 239
# Proof search given clauses           : 2300
# Clause-clause subsumption calls (NU) : 1687391
# Rec. Clause-clause subsumption calls : 102647
# Non-unit clause-clause subsumptions  : 2982
# Unit Clause-clause subsumption calls : 38014
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 77
# BW rewrite match successes           : 35
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 914205

# -------------------------------------------------
# User time                : 57.930 s
# System time              : 1.120 s
# Total time               : 59.051 s
# Maximum resident set size: 515960 pages
57.93user 1.22system 0:59.96elapsed 98%CPU (0avgtext+0avgdata 516552maxresident)k
0inputs+1288outputs (0major+454249minor)pagefaults 0swaps
