2020-06-02 10:01:04.335596: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49
2020-06-02 10:01:04.455166: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-02 10:01:04.753263: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-02 10:01:05.523280: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-02 10:01:07.687426: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3351784 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49' loaded (query=128; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.526 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>(X2=k3_finseq_1(X1)<=>k2_finseq_1(X2)=k9_xtuple_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', d3_finseq_1)).
fof(dt_k3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', dt_k3_finseq_1)).
fof(t44_finseq_1, axiom, ![X1, X2, X3]:(((v1_relat_1(X3)&v1_funct_1(X3))&v1_finseq_1(X3))=>(X3=k10_finseq_1(X1,X2)<=>((k3_finseq_1(X3)=np__2&k1_funct_1(X3,np__1)=X1)&k1_funct_1(X3,np__2)=X2))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t44_finseq_1)).
fof(fc10_finseq_1, axiom, ![X1, X2]:v1_finseq_1(k10_finseq_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', fc10_finseq_1)).
fof(fc8_finseq_1, axiom, ![X1, X2]:(v1_relat_1(k10_finseq_1(X1,X2))&v1_funct_1(k10_finseq_1(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', fc8_finseq_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t6_boole)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', redefinition_k6_numbers)).
fof(redefinition_k4_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>k4_finseq_1(X1)=k9_xtuple_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', redefinition_k4_finseq_1)).
fof(t2_finseq_1, axiom, (k2_finseq_1(np__1)=k1_tarski(np__1)&k2_finseq_1(np__2)=k2_tarski(np__1,np__2)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t2_finseq_1)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r0, axiom, ~(r1_xxreal_0(np__2,np__0)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', rqLessOrEqual__r1_xxreal_0__r2_r0)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', spc0_boole)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', d2_tarski)).
fof(d1_grsolv_1, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>(v1_grsolv_1(X1)<=>?[X2]:((((m2_finseq_1(X2,k1_group_3(X1))&~(r1_xxreal_0(k3_finseq_1(X2),k6_numbers)))&k1_funct_1(X2,np__1)=k7_group_2(X1))&k1_funct_1(X2,k3_finseq_1(X2))=k6_group_2(X1))&![X3]:(m1_subset_1(X3,k5_numbers)=>((r2_hidden(X3,k4_finseq_1(X2))&r2_hidden(k2_nat_1(X3,np__1),k4_finseq_1(X2)))=>![X4]:((v15_algstr_0(X4)&m1_group_2(X4,X1))=>![X5]:((v15_algstr_0(X5)&m1_group_2(X5,X1))=>((X4=k1_funct_1(X2,X3)&X5=k1_funct_1(X2,k2_nat_1(X3,np__1)))=>(((v15_algstr_0(X5)&v1_group_3(X5,X4))&m1_group_6(X5,X1,X4))&![X6]:((v1_group_3(X6,X4)&m1_group_6(X6,X1,X4))=>(X6=X5=>v5_group_1(k5_group_6(X4,X6))))))))))))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', d1_grsolv_1)).
fof(abstractness_v15_algstr_0, axiom, ![X1]:(l3_algstr_0(X1)=>(v15_algstr_0(X1)=>X1=g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', abstractness_v15_algstr_0)).
fof(d8_group_2, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>k7_group_2(X1)=g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', d8_group_2)).
fof(dt_k7_group_2, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>(v15_algstr_0(k7_group_2(X1))&m1_group_2(k7_group_2(X1),X1))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', dt_k7_group_2)).
fof(t13_finseq_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>![X3]:(m1_subset_1(X3,X1)=>m2_finseq_1(k10_finseq_1(X2,X3),X1)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t13_finseq_2)).
fof(d1_group_3, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(X2=k1_group_3(X1)<=>![X3]:(r2_hidden(X3,X2)<=>(v15_algstr_0(X3)&m1_group_2(X3,X1))))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', d1_group_3)).
fof(fc1_group_3, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>~(v1_xboole_0(k1_group_3(X1)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', fc1_group_3)).
fof(redefinition_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', redefinition_k2_nat_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t1_subset)).
fof(rqRealAdd__k2_xcmplx_0__r2_r1_r3, axiom, k2_xcmplx_0(np__2,np__1)=np__3, file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', rqRealAdd__k2_xcmplx_0__r2_r1_r3)).
fof(commutativity_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_nat_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', commutativity_k2_nat_1)).
fof(redefinition_m1_group_6, axiom, ![X1, X2]:(((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&m1_group_2(X2,X1))=>![X3]:(m1_group_6(X3,X1,X2)<=>m1_group_2(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', redefinition_m1_group_6)).
fof(rqRealAdd__k2_xcmplx_0__r1_r1_r2, axiom, k2_xcmplx_0(np__1,np__1)=np__2, file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', rqRealAdd__k2_xcmplx_0__r1_r1_r2)).
fof(t77_group_6, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(((((~(v2_struct_0(X2))&v15_algstr_0(X2))&v2_group_1(X2))&v3_group_1(X2))&l3_algstr_0(X2))=>((r1_group_6(X1,X2)&v5_group_1(X1))=>v5_group_1(X2)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t77_group_6)).
fof(redefinition_r2_group_6, axiom, ![X1, X2]:((((((((((~(v2_struct_0(X1))&v15_algstr_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&~(v2_struct_0(X2)))&v15_algstr_0(X2))&v2_group_1(X2))&v3_group_1(X2))&l3_algstr_0(X2))=>(r2_group_6(X1,X2)<=>r1_group_6(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', redefinition_r2_group_6)).
fof(dt_m1_group_6, axiom, ![X1, X2]:(((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&m1_group_2(X2,X1))=>![X3]:(m1_group_6(X3,X1,X2)=>m1_group_2(X3,X1))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', dt_m1_group_6)).
fof(t71_group_6, axiom, ![X1]:(((((~(v2_struct_0(X1))&v15_algstr_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>r2_group_6(X1,k5_group_6(X1,k6_group_2(X1)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t71_group_6)).
fof(dt_k6_group_2, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>(v15_algstr_0(k6_group_2(X1))&m1_group_2(k6_group_2(X1),X1))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', dt_k6_group_2)).
fof(dt_m1_group_2, axiom, ![X1]:(((~(v2_struct_0(X1))&v2_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_group_2(X2,X1)=>((~(v2_struct_0(X2))&v2_group_1(X2))&l3_algstr_0(X2)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', dt_m1_group_2)).
fof(cc1_group_2, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_group_2(X2,X1)=>v3_group_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', cc1_group_2)).
fof(fc1_group_6, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>(v15_algstr_0(k6_group_2(X1))&v1_group_3(k6_group_2(X1),X1))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', fc1_group_6)).
fof(fc5_group_6, axiom, ![X1, X2]:((((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&v1_group_3(X2,X1))&m1_group_2(X2,X1))=>(~(v2_struct_0(k5_group_6(X1,X2)))&v15_algstr_0(k5_group_6(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', fc5_group_6)).
fof(fc6_group_6, axiom, ![X1, X2]:((((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&v1_group_3(X2,X1))&m1_group_2(X2,X1))=>(v2_group_1(k5_group_6(X1,X2))&v3_group_1(k5_group_6(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', fc6_group_6)).
fof(dt_k5_group_6, axiom, ![X1, X2]:((((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&v1_group_3(X2,X1))&m1_group_2(X2,X1))=>l3_algstr_0(k5_group_6(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', dt_k5_group_6)).
fof(t7_grsolv_1, conjecture, ![X1]:((((((~(v2_struct_0(X1))&v15_algstr_0(X1))&v2_group_1(X1))&v3_group_1(X1))&v5_group_1(X1))&l3_algstr_0(X1))=>v1_grsolv_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', t7_grsolv_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', cc8_ordinal1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', redefinition_k5_numbers)).
fof(rqLessOrEqual__r1_xxreal_0__r3_r2, axiom, ~(r1_xxreal_0(np__3,np__2)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', rqLessOrEqual__r1_xxreal_0__r3_r2)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', spc1_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r2, axiom, r1_xxreal_0(np__2,np__2), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', rqLessOrEqual__r1_xxreal_0__r2_r2)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r1, axiom, ~(r1_xxreal_0(np__2,np__1)), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', rqLessOrEqual__r1_xxreal_0__r2_r1)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r3, axiom, r1_xxreal_0(np__2,np__3), file('/scratch/mptp/allpred/preds__0.005/t7_grsolv_1', rqLessOrEqual__r1_xxreal_0__r2_r3)).
fof(c_0_44, plain, ![X22, X23]:((X23!=k3_finseq_1(X22)|k2_finseq_1(X23)=k9_xtuple_0(X22)|~m2_subset_1(X23,k1_numbers,k5_numbers)|(~v1_relat_1(X22)|~v1_funct_1(X22)|~v1_finseq_1(X22)))&(k2_finseq_1(X23)!=k9_xtuple_0(X22)|X23=k3_finseq_1(X22)|~m2_subset_1(X23,k1_numbers,k5_numbers)|(~v1_relat_1(X22)|~v1_funct_1(X22)|~v1_finseq_1(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_finseq_1])])])])).
cnf(c_0_45, plain, (k2_finseq_1(X1)=k9_xtuple_0(X2)|X1!=k3_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_46, plain, ![X118]:(~v1_relat_1(X118)|~v1_funct_1(X118)|~v1_finseq_1(X118)|m2_subset_1(k3_finseq_1(X118),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_finseq_1])])).
fof(c_0_47, plain, ![X8, X9, X10]:((((k3_finseq_1(X10)=np__2|X10!=k10_finseq_1(X8,X9)|(~v1_relat_1(X10)|~v1_funct_1(X10)|~v1_finseq_1(X10)))&(k1_funct_1(X10,np__1)=X8|X10!=k10_finseq_1(X8,X9)|(~v1_relat_1(X10)|~v1_funct_1(X10)|~v1_finseq_1(X10))))&(k1_funct_1(X10,np__2)=X9|X10!=k10_finseq_1(X8,X9)|(~v1_relat_1(X10)|~v1_funct_1(X10)|~v1_finseq_1(X10))))&(k3_finseq_1(X10)!=np__2|k1_funct_1(X10,np__1)!=X8|k1_funct_1(X10,np__2)!=X9|X10=k10_finseq_1(X8,X9)|(~v1_relat_1(X10)|~v1_funct_1(X10)|~v1_finseq_1(X10)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_finseq_1])])])).
fof(c_0_48, plain, ![X112, X113]:v1_finseq_1(k10_finseq_1(X112,X113)), inference(variable_rename,[status(thm)],[fc10_finseq_1])).
fof(c_0_49, plain, ![X97, X98]:(v1_relat_1(k10_finseq_1(X97,X98))&v1_funct_1(k10_finseq_1(X97,X98))), inference(variable_rename,[status(thm)],[fc8_finseq_1])).
fof(c_0_50, plain, ![X96]:(~v1_xboole_0(X96)|X96=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_51, plain, (k9_xtuple_0(X2)=k2_finseq_1(X1)|X1!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_53, plain, (k3_finseq_1(X1)=np__2|X1!=k10_finseq_1(X2,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (v1_finseq_1(k10_finseq_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (v1_funct_1(k10_finseq_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, plain, (v1_relat_1(k10_finseq_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_59, plain, ![X63]:(~v1_relat_1(X63)|~v1_funct_1(X63)|~v1_finseq_1(X63)|k4_finseq_1(X63)=k9_xtuple_0(X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_finseq_1])])).
cnf(c_0_60, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_51])])).
cnf(c_0_61, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (k3_finseq_1(X1)=np__2|X1!=k10_finseq_1(X2,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_63, plain, (v1_finseq_1(k10_finseq_1(X1,X2))), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_64, plain, (v1_funct_1(k10_finseq_1(X1,X2))), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (v1_relat_1(k10_finseq_1(X1,X2))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_66, plain, (k2_finseq_1(np__2)=k2_tarski(np__1,np__2)), inference(split_conjunct,[status(thm)],[t2_finseq_1])).
fof(c_0_67, plain, ~r1_xxreal_0(np__2,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r0])).
cnf(c_0_68, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_69, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_70, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
fof(c_0_71, plain, ![X30, X31, X32, X33, X34, X35, X36, X37]:(((~r2_hidden(X33,X32)|(X33=X30|X33=X31)|X32!=k2_tarski(X30,X31))&((X34!=X30|r2_hidden(X34,X32)|X32!=k2_tarski(X30,X31))&(X34!=X31|r2_hidden(X34,X32)|X32!=k2_tarski(X30,X31))))&(((esk8_3(X35,X36,X37)!=X35|~r2_hidden(esk8_3(X35,X36,X37),X37)|X37=k2_tarski(X35,X36))&(esk8_3(X35,X36,X37)!=X36|~r2_hidden(esk8_3(X35,X36,X37),X37)|X37=k2_tarski(X35,X36)))&(r2_hidden(esk8_3(X35,X36,X37),X37)|(esk8_3(X35,X36,X37)=X35|esk8_3(X35,X36,X37)=X36)|X37=k2_tarski(X35,X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_72, plain, ![X11, X13, X14, X15, X16, X17]:((((((m2_finseq_1(esk2_1(X11),k1_group_3(X11))|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(~r1_xxreal_0(k3_finseq_1(esk2_1(X11)),k6_numbers)|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(k1_funct_1(esk2_1(X11),np__1)=k7_group_2(X11)|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(k1_funct_1(esk2_1(X11),k3_finseq_1(esk2_1(X11)))=k6_group_2(X11)|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&((((v15_algstr_0(X15)|(X14!=k1_funct_1(esk2_1(X11),X13)|X15!=k1_funct_1(esk2_1(X11),k2_nat_1(X13,np__1)))|(~v15_algstr_0(X15)|~m1_group_2(X15,X11))|(~v15_algstr_0(X14)|~m1_group_2(X14,X11))|(~r2_hidden(X13,k4_finseq_1(esk2_1(X11)))|~r2_hidden(k2_nat_1(X13,np__1),k4_finseq_1(esk2_1(X11))))|~m1_subset_1(X13,k5_numbers)|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(v1_group_3(X15,X14)|(X14!=k1_funct_1(esk2_1(X11),X13)|X15!=k1_funct_1(esk2_1(X11),k2_nat_1(X13,np__1)))|(~v15_algstr_0(X15)|~m1_group_2(X15,X11))|(~v15_algstr_0(X14)|~m1_group_2(X14,X11))|(~r2_hidden(X13,k4_finseq_1(esk2_1(X11)))|~r2_hidden(k2_nat_1(X13,np__1),k4_finseq_1(esk2_1(X11))))|~m1_subset_1(X13,k5_numbers)|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(m1_group_6(X15,X11,X14)|(X14!=k1_funct_1(esk2_1(X11),X13)|X15!=k1_funct_1(esk2_1(X11),k2_nat_1(X13,np__1)))|(~v15_algstr_0(X15)|~m1_group_2(X15,X11))|(~v15_algstr_0(X14)|~m1_group_2(X14,X11))|(~r2_hidden(X13,k4_finseq_1(esk2_1(X11)))|~r2_hidden(k2_nat_1(X13,np__1),k4_finseq_1(esk2_1(X11))))|~m1_subset_1(X13,k5_numbers)|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(~v1_group_3(X16,X14)|~m1_group_6(X16,X11,X14)|(X16!=X15|v5_group_1(k5_group_6(X14,X16)))|(X14!=k1_funct_1(esk2_1(X11),X13)|X15!=k1_funct_1(esk2_1(X11),k2_nat_1(X13,np__1)))|(~v15_algstr_0(X15)|~m1_group_2(X15,X11))|(~v15_algstr_0(X14)|~m1_group_2(X14,X11))|(~r2_hidden(X13,k4_finseq_1(esk2_1(X11)))|~r2_hidden(k2_nat_1(X13,np__1),k4_finseq_1(esk2_1(X11))))|~m1_subset_1(X13,k5_numbers)|~v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))))&((m1_subset_1(esk3_2(X11,X17),k5_numbers)|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(((r2_hidden(esk3_2(X11,X17),k4_finseq_1(X17))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(r2_hidden(k2_nat_1(esk3_2(X11,X17),np__1),k4_finseq_1(X17))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(((v15_algstr_0(esk4_2(X11,X17))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(m1_group_2(esk4_2(X11,X17),X11)|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(((v15_algstr_0(esk5_2(X11,X17))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(m1_group_2(esk5_2(X11,X17),X11)|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(((esk4_2(X11,X17)=k1_funct_1(X17,esk3_2(X11,X17))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(esk5_2(X11,X17)=k1_funct_1(X17,k2_nat_1(esk3_2(X11,X17),np__1))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&(((v1_group_3(esk6_2(X11,X17),esk4_2(X11,X17))|(~v15_algstr_0(esk5_2(X11,X17))|~v1_group_3(esk5_2(X11,X17),esk4_2(X11,X17))|~m1_group_6(esk5_2(X11,X17),X11,esk4_2(X11,X17)))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(m1_group_6(esk6_2(X11,X17),X11,esk4_2(X11,X17))|(~v15_algstr_0(esk5_2(X11,X17))|~v1_group_3(esk5_2(X11,X17),esk4_2(X11,X17))|~m1_group_6(esk5_2(X11,X17),X11,esk4_2(X11,X17)))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))&((esk6_2(X11,X17)=esk5_2(X11,X17)|(~v15_algstr_0(esk5_2(X11,X17))|~v1_group_3(esk5_2(X11,X17),esk4_2(X11,X17))|~m1_group_6(esk5_2(X11,X17),X11,esk4_2(X11,X17)))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11)))&(~v5_group_1(k5_group_6(esk4_2(X11,X17),esk6_2(X11,X17)))|(~v15_algstr_0(esk5_2(X11,X17))|~v1_group_3(esk5_2(X11,X17),esk4_2(X11,X17))|~m1_group_6(esk5_2(X11,X17),X11,esk4_2(X11,X17)))|(~m2_finseq_1(X17,k1_group_3(X11))|r1_xxreal_0(k3_finseq_1(X17),k6_numbers)|k1_funct_1(X17,np__1)!=k7_group_2(X11)|k1_funct_1(X17,k3_finseq_1(X17))!=k6_group_2(X11))|v1_grsolv_1(X11)|(v2_struct_0(X11)|~v2_group_1(X11)|~v3_group_1(X11)|~l3_algstr_0(X11))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_grsolv_1])])])])])])).
cnf(c_0_73, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_74, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_75, plain, (k3_finseq_1(k10_finseq_1(X1,X2))=np__2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_62]), c_0_63]), c_0_64]), c_0_65])])])).
cnf(c_0_76, plain, (k2_finseq_1(np__2)=k2_tarski(np__1,np__2)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_77, plain, (k1_funct_1(X1,np__2)=X2|X1!=k10_finseq_1(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_78, plain, (k1_funct_1(X1,np__1)=X2|X1!=k10_finseq_1(X2,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_79, plain, (~r1_xxreal_0(np__2,np__0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_80, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_81, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_70])).
fof(c_0_82, plain, ![X65]:(~l3_algstr_0(X65)|(~v15_algstr_0(X65)|X65=g3_algstr_0(u1_struct_0(X65),u2_algstr_0(X65)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v15_algstr_0])])).
fof(c_0_83, plain, ![X43]:(v2_struct_0(X43)|~v2_group_1(X43)|~v3_group_1(X43)|~l3_algstr_0(X43)|k7_group_2(X43)=g3_algstr_0(u1_struct_0(X43),u2_algstr_0(X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_group_2])])])).
cnf(c_0_84, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_85, plain, (r2_hidden(esk3_2(X1,X2),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_86, plain, (k9_xtuple_0(X1)=k4_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_87, plain, (k9_xtuple_0(k10_finseq_1(X1,X2))=k2_tarski(np__1,np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76]), c_0_63]), c_0_64]), c_0_65])])])).
cnf(c_0_88, plain, (k1_funct_1(X1,np__2)=X2|X1!=k10_finseq_1(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_89, plain, (k1_funct_1(X1,np__1)=X2|X1!=k10_finseq_1(X2,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_90, plain, (~r1_xxreal_0(np__2,np__0)), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_91, plain, (np__0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81])])).
fof(c_0_92, plain, ![X44]:((v15_algstr_0(k7_group_2(X44))|(v2_struct_0(X44)|~v2_group_1(X44)|~v3_group_1(X44)|~l3_algstr_0(X44)))&(m1_group_2(k7_group_2(X44),X44)|(v2_struct_0(X44)|~v2_group_1(X44)|~v3_group_1(X44)|~l3_algstr_0(X44)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_group_2])])])])).
cnf(c_0_93, plain, (X1=g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))|~l3_algstr_0(X1)|~v15_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_94, plain, (v2_struct_0(X1)|k7_group_2(X1)=g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_95, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_96, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|r2_hidden(esk3_2(X1,X2),k4_finseq_1(X2))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_97, plain, (k4_finseq_1(k10_finseq_1(X1,X2))=k2_tarski(np__1,np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_63]), c_0_64]), c_0_65])])])).
cnf(c_0_98, plain, (k1_funct_1(k10_finseq_1(X1,X2),np__2)=X2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_88]), c_0_63]), c_0_64]), c_0_65])])])).
cnf(c_0_99, plain, (k1_funct_1(k10_finseq_1(X1,X2),np__1)=X1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_89]), c_0_63]), c_0_64]), c_0_65])])])).
cnf(c_0_100, plain, (~r1_xxreal_0(np__2,k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_91])])).
fof(c_0_101, plain, ![X79, X80, X81]:(v1_xboole_0(X79)|(~m1_subset_1(X80,X79)|(~m1_subset_1(X81,X79)|m2_finseq_1(k10_finseq_1(X80,X81),X79)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t13_finseq_2])])])])).
fof(c_0_102, plain, ![X24, X25, X26, X27, X28]:((((v15_algstr_0(X26)|~r2_hidden(X26,X25)|X25!=k1_group_3(X24)|(v2_struct_0(X24)|~v2_group_1(X24)|~v3_group_1(X24)|~l3_algstr_0(X24)))&(m1_group_2(X26,X24)|~r2_hidden(X26,X25)|X25!=k1_group_3(X24)|(v2_struct_0(X24)|~v2_group_1(X24)|~v3_group_1(X24)|~l3_algstr_0(X24))))&(~v15_algstr_0(X27)|~m1_group_2(X27,X24)|r2_hidden(X27,X25)|X25!=k1_group_3(X24)|(v2_struct_0(X24)|~v2_group_1(X24)|~v3_group_1(X24)|~l3_algstr_0(X24))))&((~r2_hidden(esk7_2(X24,X28),X28)|(~v15_algstr_0(esk7_2(X24,X28))|~m1_group_2(esk7_2(X24,X28),X24))|X28=k1_group_3(X24)|(v2_struct_0(X24)|~v2_group_1(X24)|~v3_group_1(X24)|~l3_algstr_0(X24)))&((v15_algstr_0(esk7_2(X24,X28))|r2_hidden(esk7_2(X24,X28),X28)|X28=k1_group_3(X24)|(v2_struct_0(X24)|~v2_group_1(X24)|~v3_group_1(X24)|~l3_algstr_0(X24)))&(m1_group_2(esk7_2(X24,X28),X24)|r2_hidden(esk7_2(X24,X28),X28)|X28=k1_group_3(X24)|(v2_struct_0(X24)|~v2_group_1(X24)|~v3_group_1(X24)|~l3_algstr_0(X24)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_group_3])])])])])])])).
cnf(c_0_103, plain, (m1_group_2(k7_group_2(X1),X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_104, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=X1|~v15_algstr_0(X1)|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_105, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=k7_group_2(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_106, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_95])])).
cnf(c_0_107, plain, (r2_hidden(esk3_2(X1,k10_finseq_1(X2,X3)),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_75]), c_0_75]), c_0_98]), c_0_99]), c_0_100])])).
cnf(c_0_108, plain, (v1_xboole_0(X1)|m2_finseq_1(k10_finseq_1(X2,X3),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
fof(c_0_109, plain, ![X78]:(v2_struct_0(X78)|~v2_group_1(X78)|~v3_group_1(X78)|~l3_algstr_0(X78)|~v1_xboole_0(k1_group_3(X78))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_group_3])])])).
cnf(c_0_110, plain, (r2_hidden(X1,X3)|v2_struct_0(X2)|~v15_algstr_0(X1)|~m1_group_2(X1,X2)|X3!=k1_group_3(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_111, plain, (v2_struct_0(X1)|m1_group_2(k7_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_112, plain, (k7_group_2(X1)=X1|v2_struct_0(X1)|~l3_algstr_0(X1)|~v15_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105])])).
fof(c_0_113, plain, ![X57, X58]:(~m1_subset_1(X57,k5_numbers)|~v7_ordinal1(X58)|k2_nat_1(X57,X58)=k2_xcmplx_0(X57,X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_1])])).
cnf(c_0_114, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__2|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107])])).
cnf(c_0_115, plain, (v1_xboole_0(X1)|m2_finseq_1(k10_finseq_1(X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_116, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_xboole_0(k1_group_3(X1))), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_117, plain, (X1=k10_finseq_1(X2,X3)|k3_finseq_1(X1)!=np__2|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_118, plain, ![X85, X86]:(~r2_hidden(X85,X86)|m1_subset_1(X85,X86)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_119, plain, (v2_struct_0(X2)|r2_hidden(X1,X3)|X3!=k1_group_3(X2)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_120, plain, (m1_group_2(X1,X1)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_121, plain, (r2_hidden(k2_nat_1(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_122, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_123, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__1|esk3_2(X1,k10_finseq_1(X2,X3))=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115])])).
cnf(c_0_124, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_xboole_0(k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_125, plain, (X1=k10_finseq_1(X2,X3)|k3_finseq_1(X1)!=np__2|k1_funct_1(X1,np__2)!=X3|k1_funct_1(X1,np__1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_126, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_127, plain, (r2_hidden(X1,X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_128, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|r2_hidden(k2_nat_1(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_129, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_130, plain, (m1_subset_1(esk3_2(X1,X2),k5_numbers)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_131, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__2|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_132, plain, (k10_finseq_1(X2,k1_funct_1(X1,np__2))=X1|k1_funct_1(X1,np__1)!=X2|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_125])])).
cnf(c_0_133, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_134, plain, (r2_hidden(X1,k1_group_3(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_127])])).
cnf(c_0_135, plain, (r2_hidden(k2_xcmplx_0(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk3_2(X1,X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129])])).
cnf(c_0_136, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|m1_subset_1(esk3_2(X1,X2),k5_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_137, plain, (esk3_2(X1,X3)=np__2|esk3_2(X1,X3)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132])])).
cnf(c_0_138, plain, (m1_subset_1(X1,k1_group_3(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134])])).
cnf(c_0_139, plain, (r2_hidden(k2_xcmplx_0(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_140, plain, (esk3_2(X1,X2)=np__1|esk3_2(X1,X2)=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|k7_group_2(X1)!=X1|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138])])).
cnf(c_0_141, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r1_r3])).
cnf(c_0_142, plain, (r2_hidden(k2_xcmplx_0(esk3_2(X1,k10_finseq_1(X2,X3)),np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_97]), c_0_75]), c_0_75]), c_0_98]), c_0_99]), c_0_100])])).
cnf(c_0_143, plain, (esk3_2(X1,X2)=np__2|esk3_2(X1,X2)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_140, c_0_112])).
cnf(c_0_144, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(evalgc,[status(thm)],[c_0_141])).
fof(c_0_145, plain, ![X66, X67]:(~m1_subset_1(X66,k5_numbers)|~v7_ordinal1(X67)|k2_nat_1(X66,X67)=k2_nat_1(X67,X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_nat_1])])).
cnf(c_0_146, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_144]), c_0_98]), c_0_99]), c_0_75]), c_0_98]), c_0_63]), c_0_64]), c_0_65])])])])).
cnf(c_0_147, plain, (esk5_2(X1,X2)=k1_funct_1(X2,k2_nat_1(esk3_2(X1,X2),np__1))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_148, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_149, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_146, c_0_112])).
cnf(c_0_150, plain, (k1_funct_1(X2,k2_nat_1(esk3_2(X1,X2),np__1))=esk5_2(X1,X2)|v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_151, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_152, plain, (esk4_2(X1,X2)=k1_funct_1(X2,esk3_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_153, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_115])])).
cnf(c_0_154, plain, (k1_funct_1(X1,k2_nat_1(np__1,esk3_2(X2,X1)))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v7_ordinal1(np__1)|~m1_subset_1(esk3_2(X2,X1),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151])])).
cnf(c_0_155, plain, (esk6_2(X1,X2)=esk5_2(X1,X2)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~v15_algstr_0(esk5_2(X1,X2))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_156, plain, (v15_algstr_0(esk5_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_157, plain, ![X39, X40, X41]:((~m1_group_6(X41,X39,X40)|m1_group_2(X41,X40)|(v2_struct_0(X39)|~v2_group_1(X39)|~v3_group_1(X39)|~l3_algstr_0(X39)|~m1_group_2(X40,X39)))&(~m1_group_2(X41,X40)|m1_group_6(X41,X39,X40)|(v2_struct_0(X39)|~v2_group_1(X39)|~v3_group_1(X39)|~l3_algstr_0(X39)|~m1_group_2(X40,X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m1_group_6])])])])])).
cnf(c_0_158, plain, (k1_funct_1(X2,esk3_2(X1,X2))=esk4_2(X1,X2)|v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_159, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_153, c_0_138]), c_0_124])).
cnf(c_0_160, plain, (k1_funct_1(X1,k2_nat_1(np__1,esk3_2(X2,X1)))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~v7_ordinal1(np__1)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(csr,[status(thm)],[c_0_154, c_0_136])).
cnf(c_0_161, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_129])])).
cnf(c_0_162, plain, (m1_group_2(esk5_2(X1,X2),X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_163, plain, (esk6_2(X1,X2)=esk5_2(X1,X2)|v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(esk5_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_164, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|v15_algstr_0(esk5_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_165, plain, (m1_group_6(X1,X3,X2)|v2_struct_0(X3)|~m1_group_2(X1,X2)|~v2_group_1(X3)|~v3_group_1(X3)|~l3_algstr_0(X3)|~m1_group_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_166, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_99]), c_0_75]), c_0_75]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_167, plain, (k1_funct_1(X1,k2_xcmplx_0(esk3_2(X2,X1),np__1))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~v7_ordinal1(np__1)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(esk3_2(X2,X1),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_168, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r1_r2])).
cnf(c_0_169, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|m1_group_2(esk5_2(X1,X2),X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_170, plain, (esk6_2(X1,X2)=esk5_2(X1,X2)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_171, plain, (v2_struct_0(X3)|m1_group_6(X1,X3,X2)|~v2_group_1(X3)|~v3_group_1(X3)|~l3_algstr_0(X3)|~m1_group_2(X2,X3)|~m1_group_2(X1,X2)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_172, plain, (m1_group_2(esk4_2(X1,X2),X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_173, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_166, c_0_112])).
cnf(c_0_174, plain, (k1_funct_1(X1,k2_xcmplx_0(esk3_2(X2,X1),np__1))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~v7_ordinal1(np__1)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(csr,[status(thm)],[c_0_167, c_0_136])).
cnf(c_0_175, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_176, plain, (m1_group_2(esk5_2(X1,k10_finseq_1(X2,X3)),X1)|r1_xxreal_0(k3_finseq_1(k10_finseq_1(X2,X3)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(k10_finseq_1(X2,X3),k3_finseq_1(k10_finseq_1(X2,X3)))!=k6_group_2(X1)|k1_funct_1(k10_finseq_1(X2,X3),np__1)!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_115])])).
cnf(c_0_177, plain, (esk6_2(X1,X2)=esk5_2(X1,X2)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(esk4_2(X1,X2),X1)|~m1_group_2(esk5_2(X1,X2),esk4_2(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171])])).
cnf(c_0_178, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|m1_group_2(esk4_2(X1,X2),X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1))), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_179, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_115])])).
cnf(c_0_180, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_159]), c_0_175]), c_0_98]), c_0_75]), c_0_75]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_181, plain, (m1_group_2(esk5_2(X1,k10_finseq_1(X2,X3)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_176, c_0_75]), c_0_75]), c_0_98]), c_0_99]), c_0_100]), c_0_124])).
cnf(c_0_182, plain, (esk6_2(X1,X2)=esk5_2(X1,X2)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_2(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_177, c_0_178])).
cnf(c_0_183, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_179, c_0_138]), c_0_124])).
cnf(c_0_184, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_180, c_0_112])).
cnf(c_0_185, plain, (m1_group_2(esk5_2(X1,X3),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_132])])).
cnf(c_0_186, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183]), c_0_75]), c_0_75]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_187, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_115])])).
cnf(c_0_188, plain, (m1_group_2(esk5_2(X1,X2),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|k7_group_2(X1)!=X1|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_138])])).
cnf(c_0_189, plain, (r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,X2),esk6_2(X1,X2)))|~v15_algstr_0(esk5_2(X1,X2))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_190, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_186, c_0_112])).
cnf(c_0_191, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_187, c_0_138]), c_0_124])).
cnf(c_0_192, plain, (m1_group_2(esk5_2(X1,X2),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_188, c_0_112])).
cnf(c_0_193, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(esk5_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))|~v5_group_1(k5_group_6(esk4_2(X1,X2),esk6_2(X1,X2)))), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_194, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_191])])).
cnf(c_0_195, plain, (m1_group_2(X2,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_191]), c_0_98]), c_0_99]), c_0_75]), c_0_98]), c_0_63]), c_0_64]), c_0_65])])])).
cnf(c_0_196, plain, (r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~m1_group_6(esk5_2(X2,X1),X2,esk4_2(X2,X1))|~v1_group_3(esk5_2(X2,X1),esk4_2(X2,X1))|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v5_group_1(k5_group_6(esk4_2(X2,X1),esk6_2(X2,X1)))|~v3_group_1(X2)|~v2_group_1(X2)), inference(csr,[status(thm)],[c_0_193, c_0_164])).
cnf(c_0_197, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_194, c_0_195])).
cnf(c_0_198, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,esk6_2(X1,k10_finseq_1(X1,X2))))|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_183]), c_0_75]), c_0_75]), c_0_98]), c_0_99])]), c_0_100])])).
cnf(c_0_199, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_115])])).
cnf(c_0_200, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,esk6_2(X1,k10_finseq_1(X1,X2))))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_198, c_0_112])).
cnf(c_0_201, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_199, c_0_138]), c_0_124])).
cnf(c_0_202, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v1_group_3(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201])])).
cnf(c_0_203, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_171])])).
cnf(c_0_204, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~v1_group_3(X2,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_203, c_0_120])).
fof(c_0_205, plain, ![X120, X121]:(v2_struct_0(X120)|~v2_group_1(X120)|~v3_group_1(X120)|~l3_algstr_0(X120)|(v2_struct_0(X121)|~v15_algstr_0(X121)|~v2_group_1(X121)|~v3_group_1(X121)|~l3_algstr_0(X121)|(~r1_group_6(X120,X121)|~v5_group_1(X120)|v5_group_1(X121)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t77_group_6])])])])).
fof(c_0_206, plain, ![X55, X56]:((~r2_group_6(X55,X56)|r1_group_6(X55,X56)|(v2_struct_0(X55)|~v15_algstr_0(X55)|~v2_group_1(X55)|~v3_group_1(X55)|~l3_algstr_0(X55)|v2_struct_0(X56)|~v15_algstr_0(X56)|~v2_group_1(X56)|~v3_group_1(X56)|~l3_algstr_0(X56)))&(~r1_group_6(X55,X56)|r2_group_6(X55,X56)|(v2_struct_0(X55)|~v15_algstr_0(X55)|~v2_group_1(X55)|~v3_group_1(X55)|~l3_algstr_0(X55)|v2_struct_0(X56)|~v15_algstr_0(X56)|~v2_group_1(X56)|~v3_group_1(X56)|~l3_algstr_0(X56)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r2_group_6])])])])).
fof(c_0_207, plain, ![X46, X47, X48]:(v2_struct_0(X46)|~v2_group_1(X46)|~v3_group_1(X46)|~l3_algstr_0(X46)|~m1_group_2(X47,X46)|(~m1_group_6(X48,X46,X47)|m1_group_2(X48,X46))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_group_6])])])])).
cnf(c_0_208, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_191])])).
cnf(c_0_209, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v5_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~r1_group_6(X1,X2)|~v5_group_1(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_210, plain, (r1_group_6(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_group_6(X1,X2)|~v15_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_206])).
fof(c_0_211, plain, ![X42]:(v2_struct_0(X42)|~v15_algstr_0(X42)|~v2_group_1(X42)|~v3_group_1(X42)|~l3_algstr_0(X42)|r2_group_6(X42,k5_group_6(X42,k6_group_2(X42)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t71_group_6])])])).
cnf(c_0_212, plain, (v2_struct_0(X1)|m1_group_2(X3,X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~m1_group_6(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_207])).
fof(c_0_213, plain, ![X45]:((v15_algstr_0(k6_group_2(X45))|(v2_struct_0(X45)|~v2_group_1(X45)|~v3_group_1(X45)|~l3_algstr_0(X45)))&(m1_group_2(k6_group_2(X45),X45)|(v2_struct_0(X45)|~v2_group_1(X45)|~v3_group_1(X45)|~l3_algstr_0(X45)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_group_2])])])])).
cnf(c_0_214, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_208, c_0_195])).
cnf(c_0_215, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_group_1(X2)|~v15_algstr_0(X2)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~v5_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1)|~r1_group_6(X1,X2)), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_216, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_group_6(X1,X2)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1)|~r2_group_6(X1,X2)), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_217, plain, (v2_struct_0(X1)|r2_group_6(X1,k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_211])).
cnf(c_0_218, plain, (v2_struct_0(X1)|m1_group_2(X3,X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~m1_group_6(X3,X1,X2)), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_219, plain, (m1_group_2(k6_group_2(X1),X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_213])).
fof(c_0_220, plain, ![X61, X62]:(((~v2_struct_0(X62)|~m1_group_2(X62,X61)|(v2_struct_0(X61)|~v2_group_1(X61)|~l3_algstr_0(X61)))&(v2_group_1(X62)|~m1_group_2(X62,X61)|(v2_struct_0(X61)|~v2_group_1(X61)|~l3_algstr_0(X61))))&(l3_algstr_0(X62)|~m1_group_2(X62,X61)|(v2_struct_0(X61)|~v2_group_1(X61)|~l3_algstr_0(X61)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_group_2])])])])])).
fof(c_0_221, plain, ![X124, X125]:(v2_struct_0(X124)|~v2_group_1(X124)|~v3_group_1(X124)|~l3_algstr_0(X124)|(~m1_group_2(X125,X124)|v3_group_1(X125))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_group_2])])])])).
cnf(c_0_222, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_115])])).
cnf(c_0_223, plain, (v5_group_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|~l3_algstr_0(X1)|~l3_algstr_0(X2)|~v5_group_1(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v2_group_1(X1)|~v2_group_1(X2)|~v15_algstr_0(X1)|~r2_group_6(X2,X1)|~v15_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216])])).
cnf(c_0_224, plain, (v2_struct_0(X1)|r2_group_6(X1,k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_217])).
fof(c_0_225, plain, ![X82]:((v15_algstr_0(k6_group_2(X82))|(v2_struct_0(X82)|~v2_group_1(X82)|~v3_group_1(X82)|~l3_algstr_0(X82)))&(v1_group_3(k6_group_2(X82),X82)|(v2_struct_0(X82)|~v2_group_1(X82)|~v3_group_1(X82)|~l3_algstr_0(X82)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_group_6])])])])).
cnf(c_0_226, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_group_2(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_171])])).
cnf(c_0_227, plain, (v2_struct_0(X1)|m1_group_2(k6_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_228, plain, (v2_group_1(X1)|v2_struct_0(X2)|~m1_group_2(X1,X2)|~v2_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_229, plain, (v2_struct_0(X1)|v3_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_221])).
cnf(c_0_230, plain, (l3_algstr_0(X1)|v2_struct_0(X2)|~m1_group_2(X1,X2)|~v2_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_231, plain, (v2_struct_0(X2)|~v2_struct_0(X1)|~m1_group_2(X1,X2)|~v2_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_232, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_222, c_0_138]), c_0_124])).
cnf(c_0_233, plain, (v5_group_1(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(X1)|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(X1)|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224])])).
cnf(c_0_234, plain, (v1_group_3(k6_group_2(X1),X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
fof(c_0_235, plain, ![X53, X54]:((~v2_struct_0(k5_group_6(X53,X54))|(v2_struct_0(X53)|~v2_group_1(X53)|~v3_group_1(X53)|~l3_algstr_0(X53)|~v1_group_3(X54,X53)|~m1_group_2(X54,X53)))&(v15_algstr_0(k5_group_6(X53,X54))|(v2_struct_0(X53)|~v2_group_1(X53)|~v3_group_1(X53)|~l3_algstr_0(X53)|~v1_group_3(X54,X53)|~m1_group_2(X54,X53)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_group_6])])])])).
cnf(c_0_236, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k6_group_2(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227])])).
cnf(c_0_237, plain, (v2_struct_0(X2)|v2_group_1(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2)), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_238, plain, (v2_struct_0(X1)|v3_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_239, plain, (v2_struct_0(X2)|l3_algstr_0(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2)), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_240, plain, (v2_struct_0(X2)|~v2_struct_0(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2)), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_241, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|~v7_ordinal1(np__1)|~v1_group_3(k6_group_2(X1),X1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v5_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233])])).
cnf(c_0_242, plain, (v2_struct_0(X1)|v1_group_3(k6_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_234])).
cnf(c_0_243, plain, (v15_algstr_0(k5_group_6(X1,X2))|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_235])).
fof(c_0_244, plain, ![X51, X52]:((v2_group_1(k5_group_6(X51,X52))|(v2_struct_0(X51)|~v2_group_1(X51)|~v3_group_1(X51)|~l3_algstr_0(X51)|~v1_group_3(X52,X51)|~m1_group_2(X52,X51)))&(v3_group_1(k5_group_6(X51,X52))|(v2_struct_0(X51)|~v2_group_1(X51)|~v3_group_1(X51)|~l3_algstr_0(X51)|~v1_group_3(X52,X51)|~m1_group_2(X52,X51)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_group_6])])])])).
fof(c_0_245, plain, ![X49, X50]:(v2_struct_0(X49)|~v2_group_1(X49)|~v3_group_1(X49)|~l3_algstr_0(X49)|~v1_group_3(X50,X49)|~m1_group_2(X50,X49)|l3_algstr_0(k5_group_6(X49,X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_group_6])])])).
cnf(c_0_246, plain, (m1_group_2(k7_group_2(k6_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_111])])).
cnf(c_0_247, plain, (v2_group_1(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_227])])).
cnf(c_0_248, plain, (v3_group_1(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_227])])).
cnf(c_0_249, plain, (l3_algstr_0(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_227])])).
cnf(c_0_250, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k6_group_2(X1))|~v3_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_227])])).
cnf(c_0_251, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(X1)|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(X1)|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1)), inference(csr,[status(thm)],[c_0_241, c_0_242])).
cnf(c_0_252, plain, (v2_struct_0(X1)|v15_algstr_0(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1)), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_253, plain, (v2_group_1(k5_group_6(X1,X2))|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_244])).
cnf(c_0_254, plain, (v3_group_1(k5_group_6(X1,X2))|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_244])).
cnf(c_0_255, plain, (v2_struct_0(X1)|l3_algstr_0(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_245])).
cnf(c_0_256, plain, (m1_group_2(k7_group_2(k6_group_2(X1)),X1)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_246, c_0_247]), c_0_248]), c_0_249]), c_0_250])).
cnf(c_0_257, plain, (v2_struct_0(X1)|~v2_struct_0(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_235])).
cnf(c_0_258, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(X1)|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(X1)|~v15_algstr_0(X1)|~v1_group_3(k6_group_2(X1),X1)|~m1_group_2(k6_group_2(X1),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252])])).
cnf(c_0_259, plain, (v2_struct_0(X1)|v2_group_1(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1)), inference(evalgc,[status(thm)],[c_0_253])).
cnf(c_0_260, plain, (v2_struct_0(X1)|v3_group_1(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1)), inference(evalgc,[status(thm)],[c_0_254])).
cnf(c_0_261, plain, (v2_struct_0(X1)|l3_algstr_0(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1)), inference(evalgc,[status(thm)],[c_0_255])).
cnf(c_0_262, plain, (r2_hidden(k7_group_2(k6_group_2(X1)),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k6_group_2(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_256])])).
cnf(c_0_263, plain, (v15_algstr_0(k6_group_2(X1))|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_213])).
cnf(c_0_264, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1)|~v2_struct_0(k5_group_6(X1,X2))), inference(evalgc,[status(thm)],[c_0_257])).
cnf(c_0_265, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_258, c_0_259]), c_0_260]), c_0_261]), c_0_227]), c_0_242])).
cnf(c_0_266, plain, (r2_hidden(k6_group_2(X1),X2)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(X1))|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_112])])).
cnf(c_0_267, plain, (v2_struct_0(X1)|v15_algstr_0(k6_group_2(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_263])).
fof(c_0_268, negated_conjecture, ~(![X1]:((((((~(v2_struct_0(X1))&v15_algstr_0(X1))&v2_group_1(X1))&v3_group_1(X1))&v5_group_1(X1))&l3_algstr_0(X1))=>v1_grsolv_1(X1))), inference(assume_negation,[status(cth)],[t7_grsolv_1])).
cnf(c_0_269, plain, (v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|~v1_group_3(k6_group_2(X1),X1)|~m1_group_2(k6_group_2(X1),X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~v5_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_265])])).
cnf(c_0_270, plain, (r2_hidden(k6_group_2(X1),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_266, c_0_267]), c_0_247]), c_0_248]), c_0_249]), c_0_250])).
fof(c_0_271, negated_conjecture, ((((((~v2_struct_0(esk1_0)&v15_algstr_0(esk1_0))&v2_group_1(esk1_0))&v3_group_1(esk1_0))&v5_group_1(esk1_0))&l3_algstr_0(esk1_0))&~v1_grsolv_1(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_268])])])])).
cnf(c_0_272, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_269, c_0_227]), c_0_242])).
cnf(c_0_273, plain, (m1_subset_1(k6_group_2(X1),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_270])])).
cnf(c_0_274, negated_conjecture, (v5_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_275, negated_conjecture, (l3_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_276, negated_conjecture, (v3_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_277, negated_conjecture, (v2_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_278, negated_conjecture, (v15_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_279, negated_conjecture, (~v1_grsolv_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_271])).
cnf(c_0_280, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_271])).
fof(c_0_281, plain, ![X186]:(~m1_subset_1(X186,k4_ordinal1)|v7_ordinal1(X186)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_282, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(np__1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_273])])).
cnf(c_0_283, negated_conjecture, (v5_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_274])).
cnf(c_0_284, negated_conjecture, (l3_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_275])).
cnf(c_0_285, negated_conjecture, (v3_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_286, negated_conjecture, (v2_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_287, negated_conjecture, (v15_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_288, negated_conjecture, (~v1_grsolv_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_289, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_280])).
cnf(c_0_290, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_291, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_292, plain, ~r1_xxreal_0(np__3,np__2), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r3_r2])).
cnf(c_0_293, negated_conjecture, (r2_hidden(np__3,k2_tarski(np__1,np__2))|~v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283]), c_0_284]), c_0_285]), c_0_286]), c_0_287])]), c_0_288]), c_0_289])])).
cnf(c_0_294, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_290])).
cnf(c_0_295, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_296, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_297, plain, (~r1_xxreal_0(np__3,np__2)), inference(split_conjunct,[status(thm)],[c_0_292])).
cnf(c_0_298, negated_conjecture, (np__3=np__2|np__3=np__1|~v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_293])])).
cnf(c_0_299, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_294, c_0_295])).
cnf(c_0_300, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_301, plain, (r1_xxreal_0(np__2,np__2)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r2])).
fof(c_0_302, plain, ~r1_xxreal_0(np__2,np__1), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r1])).
cnf(c_0_303, plain, (r1_xxreal_0(np__2,np__3)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r3])).
cnf(c_0_304, plain, (~r1_xxreal_0(np__3,np__2)), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_305, negated_conjecture, (np__3=np__1|np__3=np__2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299]), c_0_300])])])).
cnf(c_0_306, plain, (r1_xxreal_0(np__2,np__2)), inference(evalgc,[status(thm)],[c_0_301])).
cnf(c_0_307, plain, (~r1_xxreal_0(np__2,np__1)), inference(split_conjunct,[status(thm)],[c_0_302])).
cnf(c_0_308, plain, (r1_xxreal_0(np__2,np__3)), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_309, negated_conjecture, (np__3=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_305]), c_0_306])])])).
cnf(c_0_310, plain, (~r1_xxreal_0(np__2,np__1)), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_311, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_308, c_0_309]), c_0_310]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 312
# Proof object clause steps            : 232
# Proof object formula steps           : 80
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 68
# Proof object initial formulas used   : 44
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 143
# Training examples: 132 positive, 1390 negative
# Training: Positive examples begin
cnf(c_0_69, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_295, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_287, negated_conjecture, (v15_algstr_0(esk1_0))).# trainpos
cnf(c_0_286, negated_conjecture, (v2_group_1(esk1_0))).# trainpos
cnf(c_0_285, negated_conjecture, (v3_group_1(esk1_0))).# trainpos
cnf(c_0_283, negated_conjecture, (v5_group_1(esk1_0))).# trainpos
cnf(c_0_284, negated_conjecture, (l3_algstr_0(esk1_0))).# trainpos
cnf(c_0_81, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_289, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_288, negated_conjecture, (~v1_grsolv_1(esk1_0))).# trainpos
cnf(c_0_144, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).# trainpos
cnf(c_0_175, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).# trainpos
cnf(c_0_310, plain, (~r1_xxreal_0(np__2,np__1))).# trainpos
cnf(c_0_306, plain, (r1_xxreal_0(np__2,np__2))).# trainpos
cnf(c_0_68, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_90, plain, (~r1_xxreal_0(np__2,np__0))).# trainpos
cnf(c_0_304, plain, (~r1_xxreal_0(np__3,np__2))).# trainpos
cnf(c_0_308, plain, (r1_xxreal_0(np__2,np__3))).# trainpos
cnf(c_0_300, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_264, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1)|~v2_struct_0(k5_group_6(X1,X2)))).# trainpos
cnf(c_0_76, plain, (k2_finseq_1(np__2)=k2_tarski(np__1,np__2))).# trainpos
cnf(c_0_65, plain, (v1_relat_1(k10_finseq_1(X1,X2)))).# trainpos
cnf(c_0_64, plain, (v1_funct_1(k10_finseq_1(X1,X2)))).# trainpos
cnf(c_0_63, plain, (v1_finseq_1(k10_finseq_1(X1,X2)))).# trainpos
cnf(c_0_294, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_124, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_xboole_0(k1_group_3(X1)))).# trainpos
cnf(c_0_133, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_238, plain, (v2_struct_0(X1)|v3_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1))).# trainpos
cnf(c_0_240, plain, (v2_struct_0(X2)|~v2_struct_0(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2))).# trainpos
cnf(c_0_237, plain, (v2_struct_0(X2)|v2_group_1(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2))).# trainpos
cnf(c_0_239, plain, (v2_struct_0(X2)|l3_algstr_0(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2))).# trainpos
cnf(c_0_104, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=X1|~v15_algstr_0(X1)|~l3_algstr_0(X1))).# trainpos
cnf(c_0_86, plain, (k9_xtuple_0(X1)=k4_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_61, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_171, plain, (v2_struct_0(X3)|m1_group_6(X1,X3,X2)|~v2_group_1(X3)|~v3_group_1(X3)|~l3_algstr_0(X3)|~m1_group_2(X2,X3)|~m1_group_2(X1,X2))).# trainpos
cnf(c_0_267, plain, (v2_struct_0(X1)|v15_algstr_0(k6_group_2(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).# trainpos
cnf(c_0_95, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_151, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_129, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_62, plain, (k3_finseq_1(X1)=np__2|X1!=k10_finseq_1(X2,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_216, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_group_6(X1,X2)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1)|~r2_group_6(X1,X2))).# trainpos
cnf(c_0_111, plain, (v2_struct_0(X1)|m1_group_2(k7_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).# trainpos
cnf(c_0_227, plain, (v2_struct_0(X1)|m1_group_2(k6_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).# trainpos
cnf(c_0_242, plain, (v2_struct_0(X1)|v1_group_3(k6_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).# trainpos
cnf(c_0_88, plain, (k1_funct_1(X1,np__2)=X2|X1!=k10_finseq_1(X3,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_89, plain, (k1_funct_1(X1,np__1)=X2|X1!=k10_finseq_1(X2,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_105, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=k7_group_2(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).# trainpos
cnf(c_0_51, plain, (k9_xtuple_0(X2)=k2_finseq_1(X1)|X1!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_215, plain, (v2_struct_0(X2)|v2_struct_0(X1)|v5_group_1(X2)|~v15_algstr_0(X2)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~v5_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1)|~r1_group_6(X1,X2))).# trainpos
cnf(c_0_115, plain, (v1_xboole_0(X1)|m2_finseq_1(k10_finseq_1(X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_224, plain, (v2_struct_0(X1)|r2_group_6(X1,k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).# trainpos
cnf(c_0_125, plain, (X1=k10_finseq_1(X2,X3)|k3_finseq_1(X1)!=np__2|k1_funct_1(X1,np__2)!=X3|k1_funct_1(X1,np__1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_252, plain, (v2_struct_0(X1)|v15_algstr_0(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1))).# trainpos
cnf(c_0_259, plain, (v2_struct_0(X1)|v2_group_1(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1))).# trainpos
cnf(c_0_260, plain, (v2_struct_0(X1)|v3_group_1(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1))).# trainpos
cnf(c_0_261, plain, (v2_struct_0(X1)|l3_algstr_0(k5_group_6(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~v1_group_3(X2,X1))).# trainpos
cnf(c_0_119, plain, (v2_struct_0(X2)|r2_hidden(X1,X3)|X3!=k1_group_3(X2)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2))).# trainpos
cnf(c_0_218, plain, (v2_struct_0(X1)|m1_group_2(X3,X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1)|~m1_group_6(X3,X1,X2))).# trainpos
cnf(c_0_164, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|v15_algstr_0(esk5_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_136, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|m1_subset_1(esk3_2(X1,X2),k5_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_178, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|m1_group_2(esk4_2(X1,X2),X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_169, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|m1_group_2(esk5_2(X1,X2),X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_96, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|r2_hidden(esk3_2(X1,X2),k4_finseq_1(X2))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_158, plain, (k1_funct_1(X2,esk3_2(X1,X2))=esk4_2(X1,X2)|v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_128, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|r2_hidden(k2_nat_1(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_150, plain, (k1_funct_1(X2,k2_nat_1(esk3_2(X1,X2),np__1))=esk5_2(X1,X2)|v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_163, plain, (esk6_2(X1,X2)=esk5_2(X1,X2)|v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(esk5_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2)))).# trainpos
cnf(c_0_193, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(esk5_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))|~v5_group_1(k5_group_6(esk4_2(X1,X2),esk6_2(X1,X2))))).# trainpos
cnf(c_0_91, plain, (np__0=k6_numbers)).# trainpos
cnf(c_0_99, plain, (k1_funct_1(k10_finseq_1(X1,X2),np__1)=X1)).# trainpos
cnf(c_0_226, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_group_2(X1,X3))).# trainpos
cnf(c_0_98, plain, (k1_funct_1(k10_finseq_1(X1,X2),np__2)=X2)).# trainpos
cnf(c_0_75, plain, (k3_finseq_1(k10_finseq_1(X1,X2))=np__2)).# trainpos
cnf(c_0_177, plain, (esk6_2(X1,X2)=esk5_2(X1,X2)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(esk4_2(X1,X2),X1)|~m1_group_2(esk5_2(X1,X2),esk4_2(X1,X2)))).# trainpos
cnf(c_0_223, plain, (v5_group_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|~l3_algstr_0(X1)|~l3_algstr_0(X2)|~v5_group_1(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v2_group_1(X1)|~v2_group_1(X2)|~v15_algstr_0(X1)|~r2_group_6(X2,X1)|~v15_algstr_0(X2))).# trainpos
cnf(c_0_100, plain, (~r1_xxreal_0(np__2,k6_numbers))).# trainpos
cnf(c_0_106, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_112, plain, (k7_group_2(X1)=X1|v2_struct_0(X1)|~l3_algstr_0(X1)|~v15_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).# trainpos
cnf(c_0_132, plain, (k10_finseq_1(X2,k1_funct_1(X1,np__2))=X1|k1_funct_1(X1,np__1)!=X2|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_60, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_250, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k6_group_2(X1))|~v3_group_1(X1))).# trainpos
cnf(c_0_154, plain, (k1_funct_1(X1,k2_nat_1(np__1,esk3_2(X2,X1)))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v7_ordinal1(np__1)|~m1_subset_1(esk3_2(X2,X1),k5_numbers))).# trainpos
cnf(c_0_135, plain, (r2_hidden(k2_xcmplx_0(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk3_2(X1,X2),k5_numbers))).# trainpos
cnf(c_0_176, plain, (m1_group_2(esk5_2(X1,k10_finseq_1(X2,X3)),X1)|r1_xxreal_0(k3_finseq_1(k10_finseq_1(X2,X3)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(k10_finseq_1(X2,X3),k3_finseq_1(k10_finseq_1(X2,X3)))!=k6_group_2(X1)|k1_funct_1(k10_finseq_1(X2,X3),np__1)!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_236, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k6_group_2(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).# trainpos
cnf(c_0_120, plain, (m1_group_2(X1,X1)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_233, plain, (v5_group_1(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(X1)|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(X1)|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1))).# trainpos
cnf(c_0_87, plain, (k9_xtuple_0(k10_finseq_1(X1,X2))=k2_tarski(np__1,np__2))).# trainpos
cnf(c_0_247, plain, (v2_group_1(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).# trainpos
cnf(c_0_97, plain, (k4_finseq_1(k10_finseq_1(X1,X2))=k2_tarski(np__1,np__2))).# trainpos
cnf(c_0_127, plain, (r2_hidden(X1,X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_249, plain, (l3_algstr_0(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).# trainpos
cnf(c_0_248, plain, (v3_group_1(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).# trainpos
cnf(c_0_246, plain, (m1_group_2(k7_group_2(k6_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).# trainpos
cnf(c_0_185, plain, (m1_group_2(esk5_2(X1,X3),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).# trainpos
cnf(c_0_134, plain, (r2_hidden(X1,k1_group_3(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_107, plain, (r2_hidden(esk3_2(X1,k10_finseq_1(X2,X3)),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).# trainpos
cnf(c_0_161, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_114, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__2|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).# trainpos
cnf(c_0_142, plain, (r2_hidden(k2_xcmplx_0(esk3_2(X1,k10_finseq_1(X2,X3)),np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).# trainpos
cnf(c_0_123, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__1|esk3_2(X1,k10_finseq_1(X2,X3))=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).# trainpos
cnf(c_0_167, plain, (k1_funct_1(X1,k2_xcmplx_0(esk3_2(X2,X1),np__1))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~v7_ordinal1(np__1)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(esk3_2(X2,X1),k5_numbers))).# trainpos
cnf(c_0_137, plain, (esk3_2(X1,X3)=np__2|esk3_2(X1,X3)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).# trainpos
cnf(c_0_262, plain, (r2_hidden(k7_group_2(k6_group_2(X1)),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k6_group_2(X1))))).# trainpos
cnf(c_0_266, plain, (r2_hidden(k6_group_2(X1),X2)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(X1))|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).# trainpos
cnf(c_0_138, plain, (m1_subset_1(X1,k1_group_3(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_140, plain, (esk3_2(X1,X2)=np__1|esk3_2(X1,X2)=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|k7_group_2(X1)!=X1|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_146, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).# trainpos
cnf(c_0_153, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).# trainpos
cnf(c_0_166, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).# trainpos
cnf(c_0_180, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).# trainpos
cnf(c_0_179, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).# trainpos
cnf(c_0_186, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).# trainpos
cnf(c_0_187, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).# trainpos
cnf(c_0_188, plain, (m1_group_2(esk5_2(X1,X2),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|k7_group_2(X1)!=X1|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_195, plain, (m1_group_2(X2,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).# trainpos
cnf(c_0_198, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,esk6_2(X1,k10_finseq_1(X1,X2))))|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).# trainpos
cnf(c_0_273, plain, (m1_subset_1(k6_group_2(X1),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).# trainpos
cnf(c_0_194, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_199, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).# trainpos
cnf(c_0_202, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v1_group_3(X2,X1))).# trainpos
cnf(c_0_203, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1))).# trainpos
cnf(c_0_208, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_222, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).# trainpos
cnf(c_0_241, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|~v7_ordinal1(np__1)|~v1_group_3(k6_group_2(X1),X1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v5_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1))))).# trainpos
cnf(c_0_258, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(X1)|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(X1)|~v15_algstr_0(X1)|~v1_group_3(k6_group_2(X1),X1)|~m1_group_2(k6_group_2(X1),X1))).# trainpos
cnf(c_0_269, plain, (v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|~v1_group_3(k6_group_2(X1),X1)|~m1_group_2(k6_group_2(X1),X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k6_group_2(X1),k1_group_3(X1))|~v5_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_282, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(np__1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).# trainpos
cnf(c_0_293, negated_conjecture, (r2_hidden(np__3,k2_tarski(np__1,np__2))|~v7_ordinal1(np__1))).# trainpos
cnf(c_0_298, negated_conjecture, (np__3=np__2|np__3=np__1|~v7_ordinal1(np__1))).# trainpos
cnf(c_0_305, negated_conjecture, (np__3=np__1|np__3=np__2)).# trainpos
cnf(c_0_309, negated_conjecture, (np__3=np__1)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_87, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_188, plain, (v1_xboole_0(esk12_0))).#trainneg
cnf(i_0_229, plain, (v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_187, plain, (v1_xcmplx_0(esk12_0))).#trainneg
cnf(i_0_186, plain, (v1_xxreal_0(esk12_0))).#trainneg
cnf(i_0_33, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_174, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_173, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_234, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_224, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_68, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_238, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_218, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_185, plain, (v1_xreal_0(esk12_0))).#trainneg
cnf(i_0_211, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_212, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_36, plain, (k1_tarski(np__1)=k2_finseq_1(np__1))).#trainneg
cnf(i_0_161, plain, (v1_relat_1(esk10_1(X1)))).#trainneg
cnf(i_0_158, plain, (v1_funct_1(esk10_1(X1)))).#trainneg
cnf(i_0_155, plain, (v1_finseq_1(esk10_1(X1)))).#trainneg
cnf(i_0_157, plain, (v1_xboole_0(esk10_1(X1)))).#trainneg
cnf(i_0_120, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_156, plain, (v1_finset_1(esk10_1(X1)))).#trainneg
cnf(i_0_112, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).#trainneg
cnf(i_0_131, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_135, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_217, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_223, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_226, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_125, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_121, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_182, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_177, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_181, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_183, plain, (r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_221, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_180, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_169, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_191, plain, (r1_xxreal_0(np__0,np__3))).#trainneg
cnf(i_0_214, plain, (r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_172, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_171, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_164, plain, (~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_167, plain, (~r1_xxreal_0(np__3,np__0))).#trainneg
cnf(i_0_215, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_258, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_129, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_128, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_208, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_213, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_192, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_154, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_143, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_222, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_160, plain, (v4_relat_1(esk10_1(X1),k5_numbers))).#trainneg
cnf(i_0_127, plain, (m1_subset_1(esk9_1(X1),X1))).#trainneg
cnf(i_0_162, plain, (m1_finseq_1(esk10_1(X1),X1))).#trainneg
cnf(i_0_159, plain, (v5_relat_1(esk10_1(X1),X1))).#trainneg
cnf(i_0_195, plain, (v1_relat_1(esk13_2(X1,X2)))).#trainneg
cnf(i_0_196, plain, (v1_xboole_0(esk13_2(X1,X2)))).#trainneg
cnf(i_0_236, plain, (~v1_xboole_0(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_92, plain, (v2_struct_0(X1)|v5_group_1(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_gr_cy_1(X1))).#trainneg
cnf(i_0_130, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_122, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_256, plain, (v1_xboole_0(X1)|v1_relat_1(esk17_1(X1)))).#trainneg
cnf(i_0_227, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_176, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_84, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_134, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_253, plain, (v1_xboole_0(X1)|v1_funct_1(esk17_1(X1)))).#trainneg
cnf(i_0_250, plain, (v1_xboole_0(X1)|v1_finseq_1(esk17_1(X1)))).#trainneg
cnf(i_0_153, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).#trainneg
cnf(i_0_240, plain, (v3_card_1(k10_finseq_1(X1,X2),np__2))).#trainneg
cnf(i_0_193, plain, (v5_relat_1(esk13_2(X1,X2),X2))).#trainneg
cnf(i_0_194, plain, (v4_relat_1(esk13_2(X1,X2),X1))).#trainneg
cnf(i_0_251, plain, (v1_xboole_0(X1)|v1_finset_1(esk17_1(X1)))).#trainneg
cnf(i_0_249, plain, (v1_xboole_0(X1)|v2_finseq_1(esk17_1(X1)))).#trainneg
cnf(i_0_116, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_247, plain, (v1_relat_1(esk16_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_139, plain, (v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_245, plain, (v1_funct_1(esk16_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_242, plain, (v1_finseq_1(esk16_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_235, plain, (v1_finset_1(k1_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_31, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_grsolv_1(X1)|~r1_xxreal_0(k3_finseq_1(esk2_1(X1)),k6_numbers))).#trainneg
cnf(i_0_255, plain, (v1_xboole_0(X1)|v4_relat_1(esk17_1(X1),k5_numbers))).#trainneg
cnf(i_0_81, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_178, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_225, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_252, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk17_1(X1)))).#trainneg
cnf(i_0_244, plain, (v1_finset_1(esk16_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_231, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_109, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_241, plain, (v2_finseq_1(esk16_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_91, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_119, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_90, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_118, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_89, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_117, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_257, plain, (v1_xboole_0(X1)|m1_finseq_1(esk17_1(X1),X1))).#trainneg
cnf(i_0_254, plain, (v1_xboole_0(X1)|v5_relat_1(esk17_1(X1),X1))).#trainneg
cnf(i_0_246, plain, (v4_relat_1(esk16_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_202, plain, (v3_card_1(k1_finseq_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_243, plain, (v3_card_1(esk16_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_83, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_220, plain, (v2_struct_0(X1)|v5_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_261, plain, (m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_110, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_189, plain, (v1_xboole_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(k1_finseq_1(X1)))).#trainneg
cnf(i_0_260, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_168, plain, (v1_xboole_0(k1_finseq_1(X1))|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_133, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_146, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_147, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_197, plain, (m1_subset_1(esk13_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_230, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_46, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2))).#trainneg
cnf(i_0_47, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4))).#trainneg
cnf(i_0_126, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_205, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_88, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_151, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_152, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_54, plain, (v2_struct_0(X1)|v15_algstr_0(k7_group_2(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_237, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_44, plain, (X3=k2_tarski(X1,X2)|esk8_3(X1,X2,X3)!=X2|~r2_hidden(esk8_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_82, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_228, plain, (m1_subset_1(k4_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_170, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_group_6(X1,X1)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_209, plain, (k3_finseq_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_111, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_200, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_201, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_248, plain, (v1_relat_1(k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))))).#trainneg
cnf(i_0_184, plain, (v2_struct_0(X1)|m1_group_2(esk11_1(X1),X1)|~v2_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_141, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_259, plain, (v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_219, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_45, plain, (X3=k2_tarski(X1,X2)|esk8_3(X1,X2,X3)!=X1|~r2_hidden(esk8_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_203, plain, (v2_xxreal_0(X2)|~v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_165, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_group_6(X1,X1)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_79, plain, (v2_struct_0(X1)|v1_gr_cy_1(k6_group_2(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_86, plain, (v2_struct_0(X1)|v7_struct_0(k6_group_2(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_138, plain, (v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_105, plain, (v2_struct_0(X1)|v8_struct_0(k6_group_2(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_70, plain, (v2_struct_0(X1)|v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_69, plain, (v2_struct_0(X1)|v3_group_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_137, plain, (m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_107, plain, (k7_group_2(k7_group_2(X1))=k7_group_2(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_75, plain, (v2_struct_0(X1)|v1_group_3(k7_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_42, plain, (v2_struct_0(X3)|v15_algstr_0(X1)|X2!=k1_group_3(X3)|~v2_group_1(X3)|~v3_group_1(X3)|~l3_algstr_0(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_30, plain, (k1_funct_1(esk2_1(X1),np__1)=k7_group_2(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_grsolv_1(X1))).#trainneg
cnf(i_0_144, plain, (X1=X2|g3_algstr_0(X3,X1)!=g3_algstr_0(X4,X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k2_zfmisc_1(X3,X3),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X3,X3),X3))))).#trainneg
cnf(i_0_150, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_123, plain, (v2_struct_0(X1)|v2_group_1(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_64, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_group_6(X1,X2)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1)|~r1_group_6(X1,X2))).#trainneg
cnf(i_0_32, plain, (v2_struct_0(X1)|m2_finseq_1(esk2_1(X1),k1_group_3(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_grsolv_1(X1))).#trainneg
cnf(i_0_41, plain, (v2_struct_0(X2)|m1_group_2(X1,X2)|X3!=k1_group_3(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_239, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_149, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_179, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_43, plain, (X3=k2_tarski(X1,X2)|esk8_3(X1,X2,X3)=X2|esk8_3(X1,X2,X3)=X1|r2_hidden(esk8_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_29, plain, (k1_funct_1(esk2_1(X1),k3_finseq_1(esk2_1(X1)))=k6_group_2(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v1_grsolv_1(X1))).#trainneg
cnf(i_0_199, plain, (v2_struct_0(X1)|m1_group_6(esk14_2(X1,X2),X1,X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_101, plain, (v15_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_100, plain, (l3_algstr_0(g3_algstr_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_145, plain, (X1=X2|g3_algstr_0(X1,X3)!=g3_algstr_0(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_50, plain, (v2_struct_0(X2)|m1_group_2(X1,X3)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X3,X2)|~m1_group_6(X1,X2,X3))).#trainneg
cnf(i_0_39, plain, (X2=k1_group_3(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(esk7_2(X1,X2))|~r2_hidden(esk7_2(X1,X2),X2)|~m1_group_2(esk7_2(X1,X2),X1))).#trainneg
cnf(i_0_106, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_group_6(X2,X1)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X2)|~v3_group_1(X1)|~l3_algstr_0(X2)|~l3_algstr_0(X1)|~r2_group_6(X1,X2))).#trainneg
cnf(i_0_38, plain, (X2=k1_group_3(X1)|v2_struct_0(X1)|v15_algstr_0(esk7_2(X1,X2))|r2_hidden(esk7_2(X1,X2),X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_37, plain, (X2=k1_group_3(X1)|v2_struct_0(X1)|r2_hidden(esk7_2(X1,X2),X2)|m1_group_2(esk7_2(X1,X2),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_21, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|v15_algstr_0(esk4_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m2_finseq_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_27, plain, (v2_struct_0(X3)|v1_group_3(X1,X2)|X2!=k1_funct_1(esk2_1(X3),X4)|X1!=k1_funct_1(esk2_1(X3),k2_nat_1(X4,np__1))|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v2_group_1(X3)|~v3_group_1(X3)|~l3_algstr_0(X3)|~v1_grsolv_1(X3)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~m1_subset_1(X4,k5_numbers)|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~r2_hidden(k2_nat_1(X4,np__1),k4_finseq_1(esk2_1(X3))))).#trainneg
cnf(i_0_26, plain, (v2_struct_0(X2)|m1_group_6(X1,X2,X3)|X3!=k1_funct_1(esk2_1(X2),X4)|X1!=k1_funct_1(esk2_1(X2),k2_nat_1(X4,np__1))|~v15_algstr_0(X3)|~v15_algstr_0(X1)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~v1_grsolv_1(X2)|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~m1_subset_1(X4,k5_numbers)|~r2_hidden(X4,k4_finseq_1(esk2_1(X2)))|~r2_hidden(k2_nat_1(X4,np__1),k4_finseq_1(esk2_1(X2))))).#trainneg
cnf(i_0_15, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_group_3(esk6_2(X1,X2),esk4_2(X1,X2))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(esk5_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2)))).#trainneg
cnf(i_0_25, plain, (v2_struct_0(X3)|v5_group_1(k5_group_6(X2,X1))|X1!=X4|X2!=k1_funct_1(esk2_1(X3),X5)|X4!=k1_funct_1(esk2_1(X3),k2_nat_1(X5,np__1))|~v15_algstr_0(X4)|~v15_algstr_0(X2)|~v2_group_1(X3)|~v3_group_1(X3)|~l3_algstr_0(X3)|~v1_grsolv_1(X3)|~m1_group_2(X4,X3)|~m1_group_2(X2,X3)|~v1_group_3(X1,X2)|~m1_subset_1(X5,k5_numbers)|~m1_group_6(X1,X3,X2)|~r2_hidden(X5,k4_finseq_1(esk2_1(X3)))|~r2_hidden(k2_nat_1(X5,np__1),k4_finseq_1(esk2_1(X3))))).#trainneg
cnf(i_0_14, plain, (v2_struct_0(X1)|v1_grsolv_1(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|m1_group_6(esk6_2(X1,X2),X1,esk4_2(X1,X2))|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v15_algstr_0(esk5_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2)))).#trainneg
cnf(i_0_46, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1))).#trainneg
cnf(i_0_263, plain, (esk15_0=k6_numbers)).#trainneg
cnf(i_0_264, plain, (esk12_0=k6_numbers)).#trainneg
cnf(i_0_262, plain, (~v1_xboole_0(k2_finseq_1(np__1)))).#trainneg
cnf(i_0_316, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_339, plain, (v5_relat_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_309, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_265, plain, (esk10_1(X1)=k6_numbers)).#trainneg
cnf(i_0_405, negated_conjecture, (r2_group_6(X1,X1)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_47, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3))).#trainneg
cnf(i_0_377, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(esk11_1(X1)))).#trainneg
cnf(i_0_425, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_319, plain, (v3_ordinal1(k6_numbers))).#trainneg
cnf(i_0_135, plain, (~r1_xxreal_0(np__1,k6_numbers))).#trainneg
cnf(i_0_269, plain, (esk13_2(X1,X2)=k6_numbers)).#trainneg
cnf(i_0_160, plain, (v4_relat_1(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_167, plain, (~r1_xxreal_0(np__3,k6_numbers))).#trainneg
cnf(i_0_552, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_286, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_274, plain, (v1_finset_1(k2_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_275, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_548, plain, (r2_hidden(X1,k2_tarski(X2,X1)))).#trainneg
cnf(i_0_520, plain, (m1_group_2(esk3_2(X1,X2),X3)|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_503, plain, (r2_group_6(k5_group_6(X1,k6_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(X1)|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(X1)|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_459, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_xboole_0(X3)|X3!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_379, plain, (l3_algstr_0(esk11_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_532, plain, (k1_funct_1(X1,k2_nat_1(np__1,esk3_2(X2,X1)))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v7_ordinal1(esk3_2(X2,X1)))).#trainneg
cnf(i_0_521, plain, (m1_subset_1(esk3_2(X1,X2),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_383, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_321, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_323, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_25, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|X2!=k1_funct_1(esk2_1(X3),k2_nat_1(X4,np__1))|X1!=k1_funct_1(esk2_1(X3),X4)|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(k2_nat_1(X4,np__1),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_191, plain, (r1_xxreal_0(k6_numbers,np__3))).#trainneg
cnf(i_0_187, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_185, plain, (v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_284, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_529, plain, (r2_hidden(k2_nat_1(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(esk3_2(X1,X2)))).#trainneg
cnf(i_0_439, plain, (v15_algstr_0(X1)|v2_struct_0(X2)|v1_xboole_0(X3)|X3!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_524, plain, (m1_group_2(k2_nat_1(esk3_2(X1,X2),np__1),X3)|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_331, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_429, plain, (v5_group_1(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_378, plain, (v2_group_1(esk11_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_527, plain, (r2_hidden(k2_nat_1(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk3_2(X1,X2),k5_numbers))).#trainneg
cnf(i_0_380, plain, (v3_group_1(esk11_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_525, plain, (m1_subset_1(k2_nat_1(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_435, plain, (v5_group_1(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_186, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_156, plain, (v1_finset_1(k6_numbers))).#trainneg
cnf(i_0_155, plain, (v1_finseq_1(k6_numbers))).#trainneg
cnf(i_0_182, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_169, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_158, plain, (v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_161, plain, (v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_294, plain, (v3_card_1(k2_finseq_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_301, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_546, plain, (m1_group_2(esk6_2(X1,X2),esk4_2(X1,X2))|v2_struct_0(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|~m1_group_2(esk4_2(X1,X2),X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_417, plain, (v1_relat_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_522, plain, (r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|~v1_xboole_0(k4_finseq_1(X1))|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_561, plain, (r2_hidden(X1,k2_tarski(X1,X2)))).#trainneg
cnf(i_0_180, plain, (r1_xxreal_0(k6_numbers,np__1))).#trainneg
cnf(i_0_121, plain, (k2_xcmplx_0(k6_numbers,np__1)=np__1)).#trainneg
cnf(i_0_295, plain, (m2_finseq_1(esk17_1(X1),X1)|v1_xboole_0(X1))).#trainneg
cnf(i_0_221, plain, (r1_xxreal_0(k6_numbers,np__2))).#trainneg
cnf(i_0_640, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_630, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_300, plain, (m1_subset_1(k2_tarski(np__1,np__2),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_302, plain, (v1_xboole_0(X1)|~v7_ordinal1(X1)|~v1_xboole_0(k2_finseq_1(X1)))).#trainneg
cnf(i_0_590, plain, (v1_group_3(X1,X1)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_556, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k7_group_2(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_644, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_580, plain, (m1_group_2(esk3_2(X1,k10_finseq_1(X2,X3)),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|k4_finseq_1(k10_finseq_1(X2,X3))!=k1_group_3(X4)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1))).#trainneg
cnf(i_0_622, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k9_xtuple_0(X1)|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_593, plain, (v5_group_1(X1)|v2_struct_0(X1)|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v5_group_1(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(X1)|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(X1)|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(X1)|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1))))).#trainneg
cnf(i_0_426, plain, (v2_group_1(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_427, plain, (l3_algstr_0(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_128, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_125, plain, (k2_xcmplx_0(np__1,k6_numbers)=np__1)).#trainneg
cnf(i_0_303, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v1_finseq_1(u1_struct_0(X1))|~v1_funct_1(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_388, plain, (v1_funct_1(k5_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_332, plain, (v1_xreal_0(np__3)|~v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_371, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_305, plain, (X1=k1_finseq_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_304, plain, (k1_finseq_1(X1)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_307, plain, (v1_xboole_0(k2_finseq_1(X1))|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_650, plain, (r2_hidden(k2_xcmplx_0(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_564, plain, (m1_subset_1(esk3_2(X1,k10_finseq_1(X2,X3)),k5_numbers)|r1_xxreal_0(np__2,k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_626, plain, (k1_funct_1(X1,k2_xcmplx_0(np__1,esk3_2(X2,X1)))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~v7_ordinal1(esk3_2(X2,X1))|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_708, plain, (r2_hidden(esk5_2(X1,k10_finseq_1(X2,X3)),X4)|v2_struct_0(X1)|v1_grsolv_1(X1)|X4!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk5_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_699, plain, (m1_group_2(esk6_2(X1,X2),esk4_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(esk4_2(X1,X2),X1)|~m1_group_2(esk5_2(X1,X2),esk4_2(X1,X2)))).#trainneg
cnf(i_0_703, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_grsolv_1(X2)|~m1_group_2(X1,esk5_2(X2,k10_finseq_1(X3,X4)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X4!=k6_group_2(X2)|X3!=k7_group_2(X2)|~m1_subset_1(X4,k1_group_3(X2))|~m1_subset_1(X3,k1_group_3(X2)))).#trainneg
cnf(i_0_667, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v1_xxreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_428, plain, (v3_group_1(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_399, plain, (v2_xxreal_0(np__3)|~v1_xreal_0(np__3)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_197, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_310, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_402, plain, (v2_xxreal_0(np__3)|~v1_xreal_0(np__3)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_637, plain, (v2_xxreal_0(np__3)|~v1_xreal_0(np__3)|~v1_xreal_0(k6_numbers)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_573, plain, (v1_finset_1(k2_tarski(np__1,np__2))|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_697, plain, (X1=u1_struct_0(X2)|~v1_xboole_0(X1)|~l3_algstr_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_342, plain, (v4_relat_1(X1,k5_numbers)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_314, plain, (v3_ordinal1(esk9_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_317, plain, (v3_ordinal1(k2_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_820, plain, (esk5_2(X1,k10_finseq_1(X2,X3))=X4|esk5_2(X1,k10_finseq_1(X2,X3))=X5|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_tarski(X5,X4)!=k1_group_3(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk5_2(X1,k10_finseq_1(X2,X3))))).#trainneg
cnf(i_0_773, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k4_finseq_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_822, plain, (m1_group_2(esk5_2(X1,k10_finseq_1(X2,X3)),X4)|v2_struct_0(X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|X5!=k1_group_3(X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|X5!=k1_group_3(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk5_2(X1,k10_finseq_1(X2,X3))))).#trainneg
cnf(i_0_834, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xreal_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_802, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_159, plain, (v5_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_327, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_666, plain, (k2_xcmplx_0(np__2,k6_numbers)=np__2|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_733, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_735, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_334, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_338, plain, (v5_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_815, plain, (k3_finseq_1(X1)=np__2|k9_xtuple_0(X1)!=k2_tarski(np__1,np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_689, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_707, plain, (v5_group_1(esk5_2(X1,k10_finseq_1(X2,X3)))|v2_struct_0(X1)|v1_grsolv_1(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_652, plain, (v1_xboole_0(k4_finseq_1(X1))|v15_algstr_0(esk3_2(X2,X1))|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k4_finseq_1(X1)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_799, plain, (k2_finseq_1(X1)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_734, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_853, plain, (v1_relat_1(k9_xtuple_0(k6_numbers)))).#trainneg
cnf(i_0_162, plain, (m1_finseq_1(k6_numbers,X1))).#trainneg
cnf(i_0_341, plain, (v4_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_786, plain, (k2_xcmplx_0(np__3,k6_numbers)=np__3|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_801, plain, (X1=k2_finseq_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_343, plain, (X1=k2_tarski(X2,X3)|v1_xboole_0(X1)|esk8_3(X2,X3,X1)!=X3|~m1_subset_1(esk8_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_344, plain, (X1=k3_finseq_1(X2)|k4_finseq_1(X2)!=k2_finseq_1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_819, plain, (k9_xtuple_0(X2)=k2_tarski(np__1,np__2)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_757, plain, (m1_group_2(k7_group_2(k7_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_645, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_873, plain, (u1_struct_0(X1)=k6_numbers|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_824, plain, (v1_grsolv_1(X2)|v2_struct_0(X2)|~v1_xboole_0(X1)|X1!=k1_group_3(X2)|X4!=k6_group_2(X2)|X3!=k7_group_2(X2)|~m1_subset_1(X4,k1_group_3(X2))|~m1_subset_1(X3,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(esk5_2(X2,k10_finseq_1(X3,X4))))).#trainneg
cnf(i_0_566, plain, (k9_xtuple_0(k10_finseq_1(X1,X2))=k2_finseq_1(X3)|X3!=np__2|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_711, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|k2_tarski(X1,X3)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_596, plain, (v1_xboole_0(X1)|m2_finseq_1(X3,X1)|~m1_subset_1(k1_funct_1(X3,np__2),X1)|~m1_subset_1(X2,X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_825, plain, (r2_hidden(esk5_2(X1,X3),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|X4!=k1_group_3(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk5_2(X1,X3))|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_381, plain, (v5_group_1(esk11_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_194, plain, (v4_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_345, plain, (v3_ordinal1(k4_finseq_1(X1))|~v5_ordinal1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_738, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_694, plain, (v3_card_1(k2_tarski(np__1,np__2),np__2)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_357, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_631, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_346, plain, (k3_finseq_1(k1_card_1(X1))=k1_card_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_347, plain, (k1_card_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_959, plain, (v5_group_1(k7_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_936, plain, (m2_finseq_1(k6_numbers,X1))).#trainneg
cnf(i_0_892, plain, (k3_finseq_1(X1)=np__2|k2_tarski(np__1,np__2)!=k4_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1015, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_883, plain, (r2_hidden(k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3))),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_954, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k6_group_2(X1)))|~v3_group_1(X1))).#trainneg
cnf(i_0_965, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k7_group_2(X1)))|~v3_group_1(X1))).#trainneg
cnf(i_0_1025, plain, (k1_card_1(np__2)=np__2)).#trainneg
cnf(i_0_580, plain, (m1_group_2(esk3_2(X1,k10_finseq_1(X2,X3)),X4)|v1_grsolv_1(X1)|v2_struct_0(X4)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1))).#trainneg
cnf(i_0_348, plain, (v1_card_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_905, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_979, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_349, plain, (v3_ordinal1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_352, plain, (k3_finseq_1(k3_finseq_1(esk17_1(X1)))=k3_finseq_1(esk17_1(X1))|v1_xboole_0(X1)|~v1_funct_1(esk17_1(X1))|~v1_relat_1(esk17_1(X1)))).#trainneg
cnf(i_0_970, plain, (v5_group_1(k7_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_881, plain, (r2_hidden(k2_nat_1(esk3_2(X1,k10_finseq_1(X2,X3)),np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_893, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k4_finseq_1(X2)!=k4_finseq_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_888, plain, (k4_finseq_1(X2)=k2_tarski(np__1,np__2)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_947, plain, (X1=np__2|k2_finseq_1(X1)!=k2_tarski(np__1,np__2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_955, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k7_group_2(k6_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_998, plain, (~v1_xboole_0(X1)|X1!=np__2|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_910, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_885, plain, (r2_hidden(k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3))),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_932, plain, (v1_relat_1(k4_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_602, plain, (~v1_xboole_0(X2)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_353, plain, (k3_finseq_1(k3_finseq_1(esk16_1(X1)))=k3_finseq_1(esk16_1(X1))|~v1_funct_1(esk16_1(X1))|~v1_relat_1(esk16_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_367, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_900, plain, (k2_xcmplx_0(np__2,k6_numbers)=np__2)).#trainneg
cnf(i_0_938, plain, (k2_xcmplx_0(np__3,k6_numbers)=np__3)).#trainneg
cnf(i_0_911, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_906, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_355, plain, (k3_finseq_1(k3_finseq_1(k10_finseq_1(X1,X2)))=k3_finseq_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_1049, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1098, plain, (k2_nat_1(esk3_2(X1,k10_finseq_1(X2,X3)),np__1)=np__2|k2_nat_1(esk3_2(X1,k10_finseq_1(X2,X3)),np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1106, plain, (k4_finseq_1(X1)=k2_tarski(np__1,np__2)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1128, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1183, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1064, plain, (v1_card_1(np__2))).#trainneg
cnf(i_0_1122, plain, (k6_numbers!=np__2)).#trainneg
cnf(i_0_1078, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1157, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_364, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk9_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_365, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_887, plain, (r2_hidden(k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3))),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_896, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_856, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_956, plain, (v2_group_1(k7_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_1111, plain, (X1=np__2|k2_tarski(np__1,np__2)!=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_366, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1080, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v3_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1079, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_368, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_370, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1192, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1193, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1235, plain, (r2_hidden(esk3_2(X1,X2),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1186, plain, (esk3_2(X1,X3)=np__1|esk3_2(X1,X3)=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1219, plain, (r2_hidden(esk3_2(X1,k10_finseq_1(X2,X3)),k4_finseq_1(X4))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X4)!=np__2|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_1133, plain, (k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_992, plain, (k9_xtuple_0(X1)=k2_tarski(np__1,np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1209, plain, (r2_hidden(np__1,k4_finseq_1(X1))|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_957, plain, (l3_algstr_0(k7_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_1248, plain, (k2_xcmplx_0(k6_numbers,np__3)=np__3|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_373, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_739, plain, (v1_funct_1(k2_tarski(np__1,np__2))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_879, plain, (m1_subset_1(k2_tarski(np__1,np__2),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1270, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_916, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_376, plain, (v1_relat_1(k9_xtuple_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))))).#trainneg
cnf(i_0_385, plain, (v1_funct_1(esk9_1(k1_zfmisc_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1379, plain, (esk4_2(X3,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X3)|v2_struct_0(X3)|esk3_2(X3,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_1402, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1385, plain, (m1_group_2(np__2,X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_1210, plain, (r2_hidden(np__2,k4_finseq_1(X1))|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_574, plain, (v15_algstr_0(X1)|v2_struct_0(X2)|k2_tarski(X3,X1)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_958, plain, (v3_group_1(k7_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1223, plain, (k3_finseq_1(X1)=np__2|k4_finseq_1(X2)!=k4_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_386, plain, (v1_funct_1(k2_finseq_1(X1))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1408, plain, (k2_xcmplx_0(k6_numbers,np__3)=np__3|~v7_ordinal1(k6_numbers)|~m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_1450, plain, (v1_relat_1(k2_tarski(np__1,np__2))|k3_finseq_1(k6_numbers)!=np__2)).#trainneg
cnf(i_0_930, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1245, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_387, plain, (v1_funct_1(k4_finseq_1(X1))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1308, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_tarski(np__1,np__2)))).#trainneg
cnf(i_0_1074, plain, (v5_group_1(esk3_2(X1,k10_finseq_1(X2,X3)))|v2_struct_0(X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|~l3_algstr_0(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v2_group_1(X4)|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1212, plain, (r2_hidden(np__1,X1)|X1!=k4_finseq_1(X2)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1136, plain, (k3_finseq_1(X1)!=np__2|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1205, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X3)|v2_struct_0(X3)|k2_nat_1(esk3_2(X3,k10_finseq_1(X1,X2)),np__1)=np__1|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_1054, plain, (k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1027, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|k1_group_3(X1)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_710, plain, (v15_algstr_0(X1)|v2_struct_0(X2)|k2_tarski(X1,X3)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_967, plain, (v2_group_1(k7_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_1240, plain, (X1=k3_finseq_1(X2)|k2_finseq_1(X1)!=k2_tarski(np__1,np__2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=np__2)).#trainneg
cnf(i_0_390, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1486, plain, (v1_funct_1(k2_tarski(np__1,np__2))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_690, plain, (k3_finseq_1(k3_finseq_1(k6_numbers))=k3_finseq_1(k6_numbers)|~v1_funct_1(k6_numbers)|~v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_391, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1594, plain, (k3_finseq_1(k6_numbers)!=np__2)).#trainneg
cnf(i_0_1606, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X3)|v2_struct_0(X3)|esk5_2(X3,k10_finseq_1(X1,X2))=X2|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_1570, plain, (m1_subset_1(k7_group_2(k6_group_2(X1)),k5_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_968, plain, (l3_algstr_0(k7_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_1668, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_393, plain, (X1=k2_tarski(X2,X3)|v1_xboole_0(X1)|esk8_3(X2,X3,X1)!=X2|~m1_subset_1(esk8_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_1670, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1653, plain, (v1_card_1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_912, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1018, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_408, plain, (v5_group_1(k6_group_2(X1))|v2_struct_0(k6_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_411, plain, (r1_group_6(X1,X1)|v2_struct_0(X1)|v2_struct_0(g3_algstr_0(u1_struct_0(X2),u2_algstr_0(X2)))|v2_struct_0(X2)|~l3_algstr_0(g3_algstr_0(u1_struct_0(X2),u2_algstr_0(X2)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(g3_algstr_0(u1_struct_0(X2),u2_algstr_0(X2)))|~v2_group_1(X1)|~l3_algstr_0(X2)|~v3_group_1(X2))).#trainneg
cnf(i_0_1076, plain, (m1_group_2(esk3_2(X1,X3),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|k2_tarski(np__1,np__2)!=k1_group_3(X4)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1418, plain, (r2_hidden(esk3_2(X1,np__2),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=k1_funct_1(np__2,np__2)|k7_group_2(X1)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1199, plain, (k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|k2_nat_1(esk3_2(X1,k10_finseq_1(X2,X3)),np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1665, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1236, plain, (r2_hidden(k2_nat_1(esk3_2(X1,X2),np__1),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_969, plain, (v3_group_1(k7_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1167, plain, (k9_xtuple_0(np__2)=k2_tarski(np__1,np__2)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_413, plain, (v3_ordinal1(u2_algstr_0(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1659, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers)|~v7_ordinal1(k6_numbers)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_1158, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1492, plain, (r1_xxreal_0(X1,k1_finseq_1(X2))|r1_xxreal_0(k1_finseq_1(X2),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_415, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2)),u1_struct_0(X2)))|~r2_hidden(X1,u2_algstr_0(X2))|~l3_algstr_0(X2))).#trainneg
cnf(i_0_1035, plain, (r2_hidden(esk3_2(X1,X3),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1407, plain, (r2_hidden(k2_xcmplx_0(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(esk3_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_1505, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__1|esk4_2(X1,k10_finseq_1(X2,X3))=X3|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1811, plain, (esk4_2(X3,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X3)|v2_struct_0(X3)|esk4_2(X3,k10_finseq_1(X1,X2))=X2|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_1789, plain, (r2_hidden(esk3_2(X1,X2),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k3_finseq_1(X2)!=np__2|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1699, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(np__3,k5_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_418, plain, (v1_xboole_0(u2_algstr_0(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1494, plain, (r1_xxreal_0(X1,k2_finseq_1(X2))|r1_xxreal_0(k2_finseq_1(X2),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_419, plain, (v4_relat_1(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_420, plain, (v1_xboole_0(u2_algstr_0(X1))|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1719, plain, (m1_group_2(esk3_2(X1,X2),X3)|v1_grsolv_1(X1)|v2_struct_0(X3)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k3_finseq_1(X2)!=np__2|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1))).#trainneg
cnf(i_0_1692, plain, (m1_subset_1(k6_group_2(X1),k5_numbers)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(X1))|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_1680, plain, (esk5_2(X1,k10_finseq_1(X2,X3))=X3|esk5_2(X1,k10_finseq_1(X2,X3))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_1688, plain, (m1_group_2(k6_group_2(X1),X2)|v2_struct_0(X2)|v2_struct_0(X1)|X3!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X3!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1691, plain, (v2_struct_0(X2)|~v1_xboole_0(X1)|X1!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_1683, plain, (k6_group_2(X1)=X2|k6_group_2(X1)=X3|v2_struct_0(X1)|k2_tarski(X3,X2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1801, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_1728, plain, (esk3_2(X1,np__2)=np__2|esk3_2(X1,np__2)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=k1_funct_1(np__2,np__2)|k7_group_2(X1)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1693, plain, (m1_subset_1(k7_group_2(k6_group_2(X1)),k5_numbers)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_457, plain, (v1_funct_1(esk2_1(X1))|v2_struct_0(X1)|~v1_grsolv_1(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1700, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(np__3,k1_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_421, plain, (v5_relat_1(u2_algstr_0(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1708, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1703, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(np__3,k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1493, plain, (r1_xxreal_0(X1,u1_struct_0(X2))|r1_xxreal_0(u1_struct_0(X2),X1)|~v1_xboole_0(X1)|~l3_algstr_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_422, plain, (v1_relat_1(k9_xtuple_0(u2_algstr_0(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1742, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X3)|v2_struct_0(X3)|k2_nat_1(np__1,esk3_2(X3,k10_finseq_1(X1,X2)))=np__2|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v7_ordinal1(esk3_2(X3,k10_finseq_1(X1,X2))))).#trainneg
cnf(i_0_1724, plain, (m1_group_2(esk3_2(X1,np__2),X2)|v2_struct_0(X2)|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k6_group_2(X1)!=k1_funct_1(np__2,np__2)|k7_group_2(X1)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1242, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k2_tarski(np__1,np__2)!=k4_finseq_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|k3_finseq_1(X2)!=np__2)).#trainneg
cnf(i_0_1000, plain, (X1=k3_finseq_1(X2)|k4_finseq_1(X2)!=k2_tarski(np__1,np__2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|X1!=np__2)).#trainneg
cnf(i_0_1462, plain, (~v1_xboole_0(k4_finseq_1(X1))|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_458, plain, (v1_finseq_1(esk2_1(X1))|v2_struct_0(X1)|~v1_grsolv_1(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1704, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(np__3,k1_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1711, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1764, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v5_ordinal1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2)|~v1_finseq_1(np__2))).#trainneg
cnf(i_0_1771, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(np__3,k1_numbers)|~v7_ordinal1(k6_numbers)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_629, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_908, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_430, plain, (m1_group_2(k7_group_2(X1),k7_group_2(X1))|v2_struct_0(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_436, plain, (v1_group_3(k7_group_2(X1),k7_group_2(X1))|v2_struct_0(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1839, plain, (esk3_2(X1,X2)=np__2|esk3_2(X1,X2)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1833, plain, (esk4_2(X1,k10_finseq_1(X2,X3))=X3|esk4_2(X1,k10_finseq_1(X2,X3))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1840, plain, (m1_subset_1(esk3_2(X1,X2),k5_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1938, plain, (esk5_2(X1,k10_finseq_1(X2,X1))=X2|esk5_2(X1,k10_finseq_1(X2,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|X2!=k7_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1937, plain, (v3_ordinal1(X1)|v2_struct_0(X1)|~v3_ordinal1(k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_571, plain, (v1_relat_1(esk2_1(X1))|v2_struct_0(X1)|~v1_grsolv_1(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1845, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_441, plain, (esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))=X2|g3_algstr_0(X1,esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))!=g3_algstr_0(X3,X2)|~v1_funct_2(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))).#trainneg
cnf(i_0_1895, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1893, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_1969, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_876, plain, (v4_relat_1(k10_finseq_1(X1,X2),k5_numbers)|v1_xboole_0(X3)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_442, plain, (u2_algstr_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X3,X2)|~v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_462, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2057, plain, (esk5_2(X1,k10_finseq_1(X1,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1999, plain, (esk4_2(X1,k10_finseq_1(X2,X1))=X2|esk4_2(X1,k10_finseq_1(X2,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|X2!=k7_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2010, plain, (m1_group_2(np__2,X1)|v1_grsolv_1(X2)|v2_struct_0(X2)|v2_struct_0(X1)|esk3_2(X2,k10_finseq_1(X2,X3))=np__1|k2_tarski(np__1,np__2)!=k1_group_3(X1)|X3!=k6_group_2(X2)|k7_group_2(X2)!=X2|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X2))|~l3_algstr_0(X1)|~l3_algstr_0(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v2_group_1(X1)|~v2_group_1(X2)|~m1_subset_1(X3,k1_group_3(X2))|~v15_algstr_0(X2))).#trainneg
cnf(i_0_1875, plain, (v5_group_1(esk3_2(X1,X2))|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|~l3_algstr_0(X3)|~v5_group_1(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2012, plain, (r2_hidden(k2_nat_1(np__2,np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1841, plain, (r2_hidden(esk3_2(X1,X2),k4_finseq_1(X3))|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2100, plain, (u2_algstr_0(X1)=X2|X1!=g3_algstr_0(X3,X2)|~l3_algstr_0(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_859, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2060, plain, (v3_ordinal1(np__3)|v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_1949, plain, (k3_finseq_1(X1)=np__1|k4_finseq_1(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_463, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1960, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1650, plain, (m2_subset_1(k3_finseq_1(k6_numbers),k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_1975, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2070, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_467, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_468, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__2,X1))=k2_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1871, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_grsolv_1(X3)|v2_struct_0(X3)|~m1_group_2(X1,esk3_2(X3,X4))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X2)|k1_funct_1(X4,np__2)!=k6_group_2(X3)|k1_funct_1(X4,np__1)!=k7_group_2(X3)|k3_finseq_1(X4)!=np__2|~m2_finseq_1(X4,k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_2154, plain, (esk4_2(X1,k10_finseq_1(X1,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2142, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2144, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2143, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_797, plain, (v3_card_1(k6_numbers,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1894, plain, (v5_relat_1(u2_algstr_0(X1),k6_numbers)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2262, plain, (k2_xcmplx_0(np__2,np__2)=k2_xcmplx_0(np__1,np__3)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1944, plain, (k3_finseq_1(X1)=np__3|k4_finseq_1(X1)!=k2_tarski(np__1,np__2)|np__3!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_469, plain, (k2_xcmplx_0(np__2,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_2280, plain, (k2_xcmplx_0(np__3,np__1)=k2_xcmplx_0(np__1,np__3)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1763, plain, (k2_tarski(np__1,np__2)=k4_finseq_1(np__2)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_2240, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2243, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_477, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_478, plain, (v15_algstr_0(esk8_3(X1,X2,X3))|v2_struct_0(X4)|esk8_3(X1,X2,X3)=X1|esk8_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2)|X3!=k1_group_3(X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4))).#trainneg
cnf(i_0_1908, plain, (k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|esk5_2(X1,k10_finseq_1(X2,X3))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1872, plain, (v2_group_1(esk3_2(X1,X2))|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|~l3_algstr_0(X3)|~v2_group_1(X3)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1874, plain, (v3_group_1(esk3_2(X1,X2))|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2145, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2156, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|m1_group_2(np__2,X3)|v1_grsolv_1(X1)|v2_struct_0(X3)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X3)|X2!=k6_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_2005, plain, (k2_nat_1(np__2,np__1)=np__1|esk5_2(X1,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1870, plain, (v2_struct_0(X1)|v1_grsolv_1(X2)|v2_struct_0(X2)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(esk3_2(X2,X3))|k2_tarski(np__1,np__2)!=k1_group_3(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X2)|k1_funct_1(X3,np__1)!=k7_group_2(X2)|k3_finseq_1(X3)!=np__2|~m2_finseq_1(X3,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_2016, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1953, plain, (X1=k3_finseq_1(X2)|k4_finseq_1(X2)!=k2_tarski(np__1,np__2)|X1!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2300, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,esk6_2(X1,k10_finseq_1(X1,X1))))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2296, plain, (m1_group_6(esk6_2(X1,k10_finseq_1(X1,X1)),X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2297, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=esk5_2(X1,k10_finseq_1(X1,X1))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2342, plain, (v15_algstr_0(np__2)|v2_struct_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_875, plain, (v4_relat_1(esk2_1(X1),k5_numbers)|v2_struct_0(X1)|~v1_grsolv_1(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2312, plain, (k2_xcmplx_0(np__2,np__2)=k2_xcmplx_0(np__1,np__3)|~v1_xcmplx_0(np__1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_2352, plain, (k3_finseq_1(X1)=np__1|k4_finseq_1(X1)!=k4_finseq_1(np__2)|np__2!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_479, plain, (m1_group_2(esk8_3(X1,X2,X3),X4)|v2_struct_0(X4)|esk8_3(X1,X2,X3)=X1|esk8_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2)|X3!=k1_group_3(X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4))).#trainneg
cnf(i_0_2335, plain, (k2_xcmplx_0(np__3,np__1)=k2_xcmplx_0(np__2,np__2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_2338, plain, (~v1_xboole_0(k4_finseq_1(np__2))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_2372, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_480, plain, (m1_subset_1(esk8_3(X1,X2,X3),X3)|esk8_3(X1,X2,X3)=X1|esk8_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2))).#trainneg
cnf(i_0_486, plain, (v15_algstr_0(g3_algstr_0(X1,esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))|~v1_funct_2(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))).#trainneg
cnf(i_0_2303, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(X1,X1,X1)|~v1_group_3(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2306, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(X1,X1)|~m1_group_2(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2308, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(X1,X1,X1)|~v1_group_3(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2295, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2294, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=esk5_2(X1,k10_finseq_1(X1,X1))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2299, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2423, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(X1,X1,X1)|~v1_group_3(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2298, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2434, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_2403, plain, (np__3=np__2|esk5_2(X1,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2405, plain, (np__3=np__2|esk5_2(X1,k10_finseq_1(X2,X3))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_2345, plain, (v15_algstr_0(np__1)|v2_struct_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_576, plain, (m1_subset_1(X1,k2_tarski(X2,X1)))).#trainneg
cnf(i_0_2336, plain, (k2_xcmplx_0(np__3,np__2)=k2_xcmplx_0(np__2,np__3)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_2366, plain, (k3_finseq_1(X1)=np__2|k4_finseq_1(np__2)!=k4_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_487, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_2337, plain, (k2_xcmplx_0(np__3,np__1)=k2_xcmplx_0(np__1,np__3)|~v1_xcmplx_0(np__1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_2371, plain, (v1_finset_1(k4_finseq_1(np__2))|~v7_ordinal1(np__2)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_1340, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X3)|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_489, plain, (l3_algstr_0(g3_algstr_0(X1,esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))|~v1_funct_2(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))).#trainneg
cnf(i_0_490, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1)))|~v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_2301, plain, (esk4_2(k1_funct_1(X1,np__2),X1)=k1_funct_1(X1,np__2)|v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~l3_algstr_0(k1_funct_1(X1,np__2))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2494, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_2498, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),X1))).#trainneg
cnf(i_0_2503, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(X1,X1)|~m1_group_2(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2573, plain, (v15_algstr_0(g3_algstr_0(k6_numbers,u2_algstr_0(X1)))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2445, plain, (k2_xcmplx_0(np__2,np__2)=k2_xcmplx_0(np__1,np__3)|~v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2461, plain, (k2_xcmplx_0(np__3,np__1)=k2_xcmplx_0(np__2,np__2)|~v1_xcmplx_0(np__1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_2339, plain, (r2_hidden(np__1,k4_finseq_1(np__2))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_2340, plain, (r2_hidden(np__2,k4_finseq_1(np__2))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_2462, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_494, plain, (X1=X2|g3_algstr_0(X1,esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))!=g3_algstr_0(X2,X3)|~v1_funct_2(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))))).#trainneg
cnf(i_0_495, plain, (u1_struct_0(X1)=X2|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X2,X3)|~v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_2238, plain, (r1_xxreal_0(X1,k2_xcmplx_0(k6_numbers,X2))|r1_xxreal_0(k2_xcmplx_0(k6_numbers,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1834, plain, (v15_algstr_0(esk3_2(X1,X2))|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1016, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2242, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),k6_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2017, plain, (k2_nat_1(np__2,np__1)=np__1|k2_nat_1(np__2,np__1)=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2563, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_tarski(X2,X1)))).#trainneg
cnf(i_0_2582, plain, (l3_algstr_0(g3_algstr_0(k6_numbers,u2_algstr_0(X1)))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2569, plain, (k2_xcmplx_0(np__3,np__2)=k2_xcmplx_0(np__2,np__3)|~v1_xcmplx_0(np__1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_2351, plain, (k3_finseq_1(X1)=np__3|k4_finseq_1(X1)!=k4_finseq_1(np__2)|np__3!=np__2|~v7_ordinal1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_496, plain, (r2_hidden(esk11_1(X1),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk11_1(X1)))).#trainneg
cnf(i_0_2576, plain, (k2_xcmplx_0(np__3,np__1)=k2_xcmplx_0(np__1,np__3)|~v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_498, plain, (r2_hidden(k7_group_2(X1),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(X1)))).#trainneg
cnf(i_0_994, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_tarski(np__1,np__2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|X1!=np__2)).#trainneg
cnf(i_0_2600, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2594, plain, (m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2599, plain, (m1_group_6(X1,X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2595, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X1)),X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X1)),X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2744, plain, (m1_group_2(k7_group_2(X1),X2)|v2_struct_0(X2)|v2_struct_0(X1)|X3!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X3!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2715, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_712, plain, (m1_subset_1(X1,k2_tarski(X1,X2)))).#trainneg
cnf(i_0_2620, plain, (k2_xcmplx_0(np__2,np__2)=k2_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_499, plain, (m1_group_2(esk14_2(X1,X2),X2)|v2_struct_0(X1)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2720, plain, (k2_xcmplx_0(np__3,np__2)=k2_xcmplx_0(np__2,np__3)|~v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2634, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_860, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_501, plain, (m1_group_2(esk14_2(X1,X2),X1)|v2_struct_0(X1)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_504, plain, (v15_algstr_0(esk7_2(X1,X2))|v2_struct_0(X3)|X2=k1_group_3(X1)|v2_struct_0(X1)|X2!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2684, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2596, plain, (m1_group_2(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2616, plain, (r2_hidden(esk6_2(X1,k10_finseq_1(X1,X1)),X2)|v2_struct_0(X1)|v1_grsolv_1(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk6_2(X1,k10_finseq_1(X1,X1)))|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2641, plain, (u1_struct_0(X1)=X2|v2_struct_0(X1)|k7_group_2(X1)!=g3_algstr_0(X2,X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2640, plain, (u1_struct_0(X1)=X2|X1!=g3_algstr_0(X2,X3)|~l3_algstr_0(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1468, plain, (k1_funct_1(X1,k2_xcmplx_0(np__1,esk3_2(X2,X1)))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~v7_ordinal1(np__1)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(esk3_2(X2,X1),k1_numbers))).#trainneg
cnf(i_0_1873, plain, (l3_algstr_0(esk3_2(X1,X2))|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|~l3_algstr_0(X3)|~v2_group_1(X3)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1734, plain, (k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__1|k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2848, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(esk14_2(X1,X2))|~m1_group_2(X2,X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_861, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2736, plain, (k2_xcmplx_0(np__3,np__1)=k2_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_2796, plain, (v1_xreal_0(k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_2565, plain, (m1_subset_1(np__2,k4_finseq_1(np__2))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_2837, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2841, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_506, plain, (m1_subset_1(esk7_2(X1,X2),X2)|X2=k1_group_3(X1)|v15_algstr_0(esk7_2(X1,X2))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_509, plain, (X1=k1_group_3(k5_group_6(X2,X3))|m1_group_2(esk7_2(k5_group_6(X2,X3),X1),k5_group_6(X2,X3))|r2_hidden(esk7_2(k5_group_6(X2,X3),X1),X1)|v2_struct_0(k5_group_6(X2,X3))|v2_struct_0(X2)|~v3_group_1(k5_group_6(X2,X3))|~v2_group_1(k5_group_6(X2,X3))|~v1_group_3(X3,X2)|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_2778, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X1)),X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X1)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1))))).#trainneg
cnf(i_0_2764, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(X1,X1,X1)|~v1_group_3(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2765, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),X1))).#trainneg
cnf(i_0_2891, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X3)|v2_struct_0(X3)|esk4_2(X3,k10_finseq_1(X1,X2))=X2|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_2773, plain, (m1_group_6(X1,X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),X1))).#trainneg
cnf(i_0_2777, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X1)),X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2770, plain, (m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X1)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1))))).#trainneg
cnf(i_0_2774, plain, (v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~m1_group_6(esk5_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2),esk4_2(k1_funct_1(X1,np__2),X1))|~v1_group_3(esk5_2(k1_funct_1(X1,np__2),X1),esk4_2(k1_funct_1(X1,np__2),X1))|~m2_finseq_1(X1,k1_group_3(k1_funct_1(X1,np__2)))|~l3_algstr_0(k1_funct_1(X1,np__2))|~v5_group_1(k5_group_6(esk4_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2)))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_865, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_2822, plain, (k2_xcmplx_0(np__3,np__2)=k2_xcmplx_0(np__2,np__3))).#trainneg
cnf(i_0_511, plain, (r1_xxreal_0(k3_finseq_1(k10_finseq_1(X1,X2)),k6_numbers)|v1_grsolv_1(X3)|v15_algstr_0(esk4_2(X3,k10_finseq_1(X1,X2)))|v2_struct_0(X3)|v1_xboole_0(k1_group_3(X3))|k1_funct_1(k10_finseq_1(X1,X2),k3_finseq_1(k10_finseq_1(X1,X2)))!=k6_group_2(X3)|k1_funct_1(k10_finseq_1(X1,X2),np__1)!=k7_group_2(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_2625, plain, (m1_subset_1(np__1,k4_finseq_1(np__2))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_864, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_513, plain, (r1_xxreal_0(k3_finseq_1(k10_finseq_1(X1,X2)),k6_numbers)|v1_grsolv_1(X3)|v15_algstr_0(esk5_2(X3,k10_finseq_1(X1,X2)))|v2_struct_0(X3)|v1_xboole_0(k1_group_3(X3))|k1_funct_1(k10_finseq_1(X1,X2),k3_finseq_1(k10_finseq_1(X1,X2)))!=k6_group_2(X3)|k1_funct_1(k10_finseq_1(X1,X2),np__1)!=k7_group_2(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_514, plain, (m1_subset_1(esk3_2(X1,esk2_1(X2)),k5_numbers)|r1_xxreal_0(k3_finseq_1(esk2_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X2)|k6_group_2(X2)!=k6_group_2(X1)|k1_funct_1(esk2_1(X2),np__1)!=k7_group_2(X1)|~m2_finseq_1(esk2_1(X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_2763, plain, (v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~m1_group_6(esk5_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2),k1_funct_1(X1,np__2))|~v1_group_3(esk5_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2))|~m2_finseq_1(X1,k1_group_3(k1_funct_1(X1,np__2)))|~l3_algstr_0(k1_funct_1(X1,np__2))|~v5_group_1(k5_group_6(k1_funct_1(X1,np__2),k1_funct_1(X1,np__2)))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2868, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=X2|esk6_2(X1,k10_finseq_1(X1,X1))=X3|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_tarski(X3,X2)!=k1_group_3(X1)|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk6_2(X1,k10_finseq_1(X1,X1)))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2887, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v2_struct_0(X2)|X1!=k7_group_2(X2)|~l3_algstr_0(X1)|~v15_algstr_0(X1)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_2873, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X1)),X2)|v2_struct_0(X2)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X3!=k1_group_3(X1)|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk6_2(X1,k10_finseq_1(X1,X1)))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2885, plain, (u1_struct_0(g3_algstr_0(X1,X2))=X1|~l3_algstr_0(g3_algstr_0(X1,X2))|~v15_algstr_0(g3_algstr_0(X1,X2)))).#trainneg
cnf(i_0_2757, plain, (k3_finseq_1(X1)=np__1|k9_xtuple_0(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_516, plain, (m1_group_2(esk4_2(X1,k10_finseq_1(X2,X3)),X1)|r1_xxreal_0(k3_finseq_1(k10_finseq_1(X2,X3)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(k10_finseq_1(X2,X3),k3_finseq_1(k10_finseq_1(X2,X3)))!=k6_group_2(X1)|k1_funct_1(k10_finseq_1(X2,X3),np__1)!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_3021, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3025, plain, (r1_xxreal_0(X1,k2_zfmisc_1(X2,X3))|r1_xxreal_0(k2_zfmisc_1(X2,X3),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_523, plain, (v15_algstr_0(k2_nat_1(esk3_2(X1,X2),np__1))|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2984, plain, (m1_group_6(X1,X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(X1,X1)|~m1_group_2(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2980, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(X1,X1)|~m1_group_2(X1,X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3044, plain, (u1_struct_0(k7_group_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~l3_algstr_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(X1)))).#trainneg
cnf(i_0_2976, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X1)),X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1854, plain, (v4_relat_1(u2_algstr_0(X1),k2_zfmisc_1(k6_numbers,k6_numbers))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2801, plain, (m2_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers,k5_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_2752, plain, (k3_finseq_1(X1)=np__3|k9_xtuple_0(X1)!=k2_tarski(np__1,np__2)|np__3!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1672, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(k6_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2353, plain, (v3_card_1(k4_finseq_1(np__2),np__2)|~v7_ordinal1(np__2)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_923, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_533, plain, (v1_group_3(X1,X2)|v2_struct_0(X3)|X1!=k1_funct_1(esk2_1(X3),k2_nat_1(np__1,X4))|X2!=k1_funct_1(esk2_1(X3),X4)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(k2_nat_1(np__1,X4),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_534, plain, (v1_group_3(X1,X2)|v2_struct_0(X3)|X1!=k1_funct_1(esk2_1(X3),k2_xcmplx_0(X4,np__1))|X2!=k1_funct_1(esk2_1(X3),X4)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(k2_xcmplx_0(X4,np__1),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2982, plain, (esk4_2(X1,k10_finseq_1(X2,X3))=X3|esk5_2(X1,k10_finseq_1(X2,X3))=X3|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3118, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_3081, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_grsolv_1(X2)|~m1_group_2(X1,esk4_2(X2,k10_finseq_1(X3,X4)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X4!=k6_group_2(X2)|X3!=k7_group_2(X2)|~m1_subset_1(X4,k1_group_3(X2))|~m1_subset_1(X3,k1_group_3(X2)))).#trainneg
cnf(i_0_3116, plain, (m1_group_6(X1,X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_3013, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3057, plain, (v2_struct_0(g3_algstr_0(X1,X2))|~v1_xboole_0(X1)|~l3_algstr_0(g3_algstr_0(X1,X2))|~v15_algstr_0(g3_algstr_0(X1,X2)))).#trainneg
cnf(i_0_3027, plain, (esk5_2(X1,k10_finseq_1(X2,X3))=X4|esk5_2(X1,k10_finseq_1(X2,X3))=X5|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_tarski(X4,X5)!=k1_group_3(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3164, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3142, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__1,np__3),k5_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_2751, plain, (k3_finseq_1(X1)=np__3|k9_xtuple_0(X1)!=k2_tarski(np__1,np__2)|np__3!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_535, plain, (v1_group_3(X1,X2)|v2_struct_0(X3)|X1!=k1_funct_1(esk2_1(X3),k2_nat_1(np__1,X4))|X2!=k1_funct_1(esk2_1(X3),X4)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(k2_nat_1(np__1,X4),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_907, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_3145, plain, (k2_xcmplx_0(np__1,np__3)!=np__2|~v1_xboole_0(k2_xcmplx_0(np__1,np__3))|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_2919, plain, (m2_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers,k5_numbers)|~v7_ordinal1(np__1)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_1345, plain, (X1=esk9_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1481, plain, (X1=esk9_1(k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_536, plain, (v1_group_3(X1,X2)|v2_struct_0(X3)|v1_xboole_0(k4_finseq_1(esk2_1(X3)))|X1!=k1_funct_1(esk2_1(X3),k2_nat_1(X4,np__1))|X2!=k1_funct_1(esk2_1(X3),X4)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~m1_subset_1(k2_nat_1(X4,np__1),k4_finseq_1(esk2_1(X3))))).#trainneg
cnf(i_0_537, plain, (v1_group_3(X1,X2)|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(esk2_1(X3)),k6_numbers)|v1_grsolv_1(X4)|v2_struct_0(X4)|X1!=k1_funct_1(esk2_1(X3),k2_nat_1(esk3_2(X4,esk2_1(X3)),np__1))|X2!=k1_funct_1(esk2_1(X3),esk3_2(X4,esk2_1(X3)))|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(esk3_2(X4,esk2_1(X3)),k4_finseq_1(esk2_1(X3)))|~m1_subset_1(esk3_2(X4,esk2_1(X3)),k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|k1_funct_1(esk2_1(X3),k3_finseq_1(esk2_1(X3)))!=k6_group_2(X4)|k1_funct_1(esk2_1(X3),np__1)!=k7_group_2(X4)|~m2_finseq_1(esk2_1(X3),k1_group_3(X4))|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4))).#trainneg
cnf(i_0_3295, plain, (m1_group_6(esk6_2(X1,k10_finseq_1(X1,X2)),X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3294, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3278, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3258, plain, (v4_relat_1(u2_algstr_0(X1),k6_numbers)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_3324, plain, (v3_ordinal1(k2_xcmplx_0(np__1,np__3))|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_2750, plain, (k3_finseq_1(X1)=np__3|k9_xtuple_0(X1)!=k2_tarski(np__1,np__2)|np__3!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(k6_numbers)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_538, plain, (m1_group_6(X1,X2,X3)|v2_struct_0(X2)|X1!=k1_funct_1(esk2_1(X2),k2_nat_1(np__1,X4))|X3!=k1_funct_1(esk2_1(X2),X4)|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~r2_hidden(k2_nat_1(np__1,X4),k4_finseq_1(esk2_1(X2)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X2)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X3)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_3320, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_3347, plain, (k2_xcmplx_0(np__1,np__3)!=np__2|~v1_xboole_0(k2_xcmplx_0(np__1,np__3))|~v7_ordinal1(np__1)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_3002, plain, (m2_subset_1(k2_xcmplx_0(np__2,np__3),k1_numbers,k5_numbers)|~v7_ordinal1(np__2)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_3343, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3256, plain, (m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_539, plain, (m1_group_6(X1,X2,X3)|v2_struct_0(X2)|X1!=k1_funct_1(esk2_1(X2),k2_xcmplx_0(X4,np__1))|X3!=k1_funct_1(esk2_1(X2),X4)|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~r2_hidden(k2_xcmplx_0(X4,np__1),k4_finseq_1(esk2_1(X2)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X2)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X3)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_540, plain, (m1_group_6(X1,X2,X3)|v2_struct_0(X2)|X1!=k1_funct_1(esk2_1(X2),k2_nat_1(np__1,X4))|X3!=k1_funct_1(esk2_1(X2),X4)|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~r2_hidden(k2_nat_1(np__1,X4),k4_finseq_1(esk2_1(X2)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X2)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X3)|~v15_algstr_0(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_3293, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1070, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_grsolv_1(X3)|v2_struct_0(X3)|~m1_group_2(X1,esk3_2(X3,k10_finseq_1(X4,X5)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X5!=k6_group_2(X3)|X4!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X4,X5),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_1214, plain, (X1=np__1|X1=np__2|X2!=k4_finseq_1(X3)|~r2_hidden(X1,X2)|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3383, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|~m1_group_2(X1,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3283, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3143, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_541, plain, (m1_group_6(X1,X2,X3)|v2_struct_0(X2)|v1_xboole_0(k4_finseq_1(esk2_1(X2)))|X1!=k1_funct_1(esk2_1(X2),k2_nat_1(X4,np__1))|X3!=k1_funct_1(esk2_1(X2),X4)|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~r2_hidden(X4,k4_finseq_1(esk2_1(X2)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X3)|~v15_algstr_0(X1)|~m1_subset_1(k2_nat_1(X4,np__1),k4_finseq_1(esk2_1(X2))))).#trainneg
cnf(i_0_3421, plain, (k2_xcmplx_0(np__2,np__3)!=np__2|~v1_xboole_0(k2_xcmplx_0(np__2,np__3))|~v7_ordinal1(np__2)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_2929, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,np__3),k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_3423, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3266, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k6_numbers)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_542, plain, (m1_group_6(X1,X2,X3)|v2_struct_0(X2)|r1_xxreal_0(k3_finseq_1(esk2_1(X2)),k6_numbers)|v1_grsolv_1(X4)|v2_struct_0(X4)|X1!=k1_funct_1(esk2_1(X2),k2_nat_1(esk3_2(X4,esk2_1(X2)),np__1))|X3!=k1_funct_1(esk2_1(X2),esk3_2(X4,esk2_1(X2)))|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~r2_hidden(esk3_2(X4,esk2_1(X2)),k4_finseq_1(esk2_1(X2)))|~m1_subset_1(esk3_2(X4,esk2_1(X2)),k5_numbers)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X3)|~v15_algstr_0(X1)|k1_funct_1(esk2_1(X2),k3_finseq_1(esk2_1(X2)))!=k6_group_2(X4)|k1_funct_1(esk2_1(X2),np__1)!=k7_group_2(X4)|~m2_finseq_1(esk2_1(X2),k1_group_3(X4))|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4))).#trainneg
cnf(i_0_544, plain, (v1_group_3(esk6_2(X1,X2),esk4_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(esk4_2(X1,X2),X1)|~m1_group_2(esk5_2(X1,X2),esk4_2(X1,X2)))).#trainneg
cnf(i_0_3389, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1))).#trainneg
cnf(i_0_3481, plain, (m1_group_2(k7_group_2(esk3_2(X1,k10_finseq_1(X2,X3))),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(esk3_2(X1,k10_finseq_1(X2,X3)))|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~l3_algstr_0(esk3_2(X1,k10_finseq_1(X2,X3)))|~v3_group_1(esk3_2(X1,k10_finseq_1(X2,X3)))|~v2_group_1(esk3_2(X1,k10_finseq_1(X2,X3))))).#trainneg
cnf(i_0_3478, plain, (m1_group_2(k6_group_2(esk3_2(X1,k10_finseq_1(X2,X3))),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(esk3_2(X1,k10_finseq_1(X2,X3)))|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~l3_algstr_0(esk3_2(X1,k10_finseq_1(X2,X3)))|~v3_group_1(esk3_2(X1,k10_finseq_1(X2,X3)))|~v2_group_1(esk3_2(X1,k10_finseq_1(X2,X3))))).#trainneg
cnf(i_0_3493, plain, (esk3_2(X1,X2)=np__2|esk3_2(X1,X2)=np__1|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X3)!=np__2|k4_finseq_1(X2)!=k4_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3463, plain, (m1_group_2(X1,X2)|v1_grsolv_1(X3)|v2_struct_0(X3)|v2_struct_0(X2)|esk3_2(X3,k10_finseq_1(X3,X4))=np__1|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X4!=k6_group_2(X3)|k7_group_2(X3)!=X3|~m1_group_2(X1,np__2)|~m2_finseq_1(k10_finseq_1(X3,X4),k1_group_3(X3))|~l3_algstr_0(X2)|~l3_algstr_0(X3)|~v3_group_1(X2)|~v3_group_1(X3)|~v2_group_1(X2)|~v2_group_1(X3)|~m1_subset_1(X4,k1_group_3(X3))|~v15_algstr_0(X3))).#trainneg
cnf(i_0_2639, plain, (k6_numbers=X2|g3_algstr_0(k6_numbers,u2_algstr_0(X1))!=g3_algstr_0(X2,X3)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_792, plain, (u1_struct_0(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2928, plain, (k2_xcmplx_0(k2_xcmplx_0(np__1,np__3),k6_numbers)=k2_xcmplx_0(np__1,np__3)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_547, plain, (m1_group_2(esk6_2(X1,X2),X1)|v2_struct_0(X1)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|~m1_group_2(esk4_2(X1,X2),X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m1_group_6(esk5_2(X1,X2),X1,esk4_2(X1,X2))|~v1_group_3(esk5_2(X1,X2),esk4_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_3550, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k6_numbers))))).#trainneg
cnf(i_0_554, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,esk11_1(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_575, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|k2_tarski(X3,X1)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_3563, plain, (v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3503, plain, (k2_nat_1(esk3_2(X1,X2),np__1)=np__2|k2_nat_1(esk3_2(X1,X2),np__1)=np__1|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X3)!=np__2|k4_finseq_1(X2)!=k4_finseq_1(X3)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3467, plain, (m1_group_2(X1,X2)|v1_grsolv_1(X3)|v2_struct_0(X3)|v2_struct_0(X2)|esk3_2(X3,k10_finseq_1(X5,X4))=np__1|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X4!=k6_group_2(X3)|X5!=k7_group_2(X3)|~m1_group_2(X1,np__2)|~m2_finseq_1(k10_finseq_1(X5,X4),k1_group_3(X3))|~l3_algstr_0(X2)|~l3_algstr_0(X3)|~v3_group_1(X2)|~v3_group_1(X3)|~v2_group_1(X2)|~v2_group_1(X3)|~m1_subset_1(X4,k1_group_3(X3))|~m1_subset_1(X5,k1_group_3(X3)))).#trainneg
cnf(i_0_3603, plain, (m1_group_2(k6_group_2(esk3_2(X1,X3)),X4)|v1_grsolv_1(X1)|v2_struct_0(esk3_2(X1,X3))|v2_struct_0(X4)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X4)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(esk3_2(X1,X3))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(esk3_2(X1,X3))|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(esk3_2(X1,X3))|~v2_group_1(X4)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3468, plain, (m1_group_2(esk14_2(X1,esk3_2(X2,k10_finseq_1(X3,X4))),X5)|v1_grsolv_1(X2)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X5)|X4!=k6_group_2(X2)|X3!=k7_group_2(X2)|~m2_finseq_1(k10_finseq_1(X3,X4),k1_group_3(X2))|~l3_algstr_0(X5)|~l3_algstr_0(X2)|~v3_group_1(X5)|~v3_group_1(X2)|~v2_group_1(X5)|~v2_group_1(X2)|~m1_group_2(esk3_2(X2,k10_finseq_1(X3,X4)),X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3462, plain, (m1_group_2(X1,X2)|v1_grsolv_1(X3)|v2_struct_0(X3)|v2_struct_0(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X2)|k1_funct_1(X4,np__2)!=k6_group_2(X3)|X5!=k7_group_2(X3)|~m1_group_2(X1,esk3_2(X3,X4))|~m2_finseq_1(X4,k1_group_3(X3))|~l3_algstr_0(X2)|~l3_algstr_0(X3)|~v3_group_1(X2)|~v3_group_1(X3)|~v2_group_1(X2)|~v2_group_1(X3)|k1_funct_1(X4,np__1)!=X5|k3_finseq_1(X4)!=np__2|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_3581, plain, (m1_group_2(k7_group_2(esk3_2(X1,X3)),X4)|v1_grsolv_1(X1)|v2_struct_0(esk3_2(X1,X3))|v2_struct_0(X4)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X4)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(esk3_2(X1,X3))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(esk3_2(X1,X3))|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(esk3_2(X1,X3))|~v2_group_1(X4)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_940, plain, (u1_struct_0(X1)=k2_finseq_1(X2)|~v1_xboole_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1652, plain, (k1_card_1(k3_finseq_1(k6_numbers))=k3_finseq_1(k6_numbers)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_584, plain, (esk8_3(X1,X2,k2_tarski(X3,X4))=X4|esk8_3(X1,X2,k2_tarski(X3,X4))=X3|esk8_3(X1,X2,k2_tarski(X3,X4))=X1|esk8_3(X1,X2,k2_tarski(X3,X4))=X2|k2_tarski(X3,X4)=k2_tarski(X1,X2))).#trainneg
cnf(i_0_2231, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k3_finseq_1(k6_numbers),k5_numbers)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_581, plain, (m1_group_2(esk3_2(X1,esk2_1(X2)),X3)|r1_xxreal_0(k3_finseq_1(esk2_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|k6_group_2(X2)!=k6_group_2(X1)|k1_funct_1(esk2_1(X2),np__1)!=k7_group_2(X1)|k4_finseq_1(esk2_1(X2))!=k1_group_3(X3)|~m2_finseq_1(esk2_1(X2),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_585, plain, (esk7_2(X1,k2_tarski(X2,X3))=X3|esk7_2(X1,k2_tarski(X2,X3))=X2|k2_tarski(X2,X3)=k1_group_3(X1)|v15_algstr_0(esk7_2(X1,k2_tarski(X2,X3)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3666, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X2,X3)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X2,X3))=X3|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X2,X3)),X1,X3)|~v1_group_3(esk5_2(X1,k10_finseq_1(X2,X3)),X3)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_3704, plain, (m1_group_2(k6_group_2(esk3_2(X1,X2)),X3)|v1_grsolv_1(X1)|v2_struct_0(esk3_2(X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k3_finseq_1(X2)!=np__2|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(esk3_2(X1,X2))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(esk3_2(X1,X2))|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(esk3_2(X1,X2))|~v2_group_1(X3)|~v2_group_1(X1))).#trainneg
cnf(i_0_3757, plain, (m1_group_2(k7_group_2(esk3_2(X1,X2)),X3)|v1_grsolv_1(X1)|v2_struct_0(esk3_2(X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k3_finseq_1(X2)!=np__2|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(esk3_2(X1,X2))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(esk3_2(X1,X2))|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(esk3_2(X1,X2))|~v2_group_1(X3)|~v2_group_1(X1))).#trainneg
cnf(i_0_871, plain, (u1_struct_0(X1)=u1_struct_0(X2)|~l3_algstr_0(X2)|~v2_struct_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2800, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,np__3),k6_numbers)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(np__1,np__3))|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1943, plain, (k3_finseq_1(X1)=np__3|k4_finseq_1(X1)!=k2_tarski(np__1,np__2)|np__3!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_601, plain, (v3_card_1(X2,np__2)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_931, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2232, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k3_finseq_1(k6_numbers),k1_numbers)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_3344, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__1,np__3),k5_numbers)|~v7_ordinal1(np__1)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_610, plain, (v1_xboole_0(X1)|m1_group_2(esk9_1(X1),X2)|v2_struct_0(X2)|X1!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_611, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|m1_group_2(u2_algstr_0(X1),X2)|v2_struct_0(X2)|k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_3902, plain, (esk5_2(X1,k10_finseq_1(X2,X3))=X3|m1_group_2(esk6_2(X1,k10_finseq_1(X2,X3)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X2,X3)),X3)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X3,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X2,X3)),X3))).#trainneg
cnf(i_0_3506, plain, (k6_group_2(X1)=np__2|k6_group_2(X1)=np__1|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|X3!=k4_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|X3!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3469, plain, (m1_group_2(esk14_2(esk3_2(X1,k10_finseq_1(X2,X3)),X4),X5)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(esk3_2(X1,k10_finseq_1(X2,X3)))|k2_tarski(np__1,np__2)!=k1_group_3(X5)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X5)|~l3_algstr_0(X1)|~v3_group_1(X5)|~v3_group_1(X1)|~v2_group_1(X5)|~v2_group_1(X1)|~m1_group_2(X4,esk3_2(X1,k10_finseq_1(X2,X3)))|~l3_algstr_0(esk3_2(X1,k10_finseq_1(X2,X3)))|~v3_group_1(esk3_2(X1,k10_finseq_1(X2,X3)))|~v2_group_1(esk3_2(X1,k10_finseq_1(X2,X3))))).#trainneg
cnf(i_0_3597, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_grsolv_1(X3)|v2_struct_0(esk3_2(X3,k10_finseq_1(X4,X5)))|v2_struct_0(X3)|~m1_group_2(X1,k6_group_2(esk3_2(X3,k10_finseq_1(X4,X5))))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X5!=k6_group_2(X3)|X4!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X4,X5),k1_group_3(X3))|~l3_algstr_0(esk3_2(X3,k10_finseq_1(X4,X5)))|~l3_algstr_0(X3)|~v3_group_1(esk3_2(X3,k10_finseq_1(X4,X5)))|~v3_group_1(X3)|~v2_group_1(esk3_2(X3,k10_finseq_1(X4,X5)))|~v2_group_1(X3))).#trainneg
cnf(i_0_3732, plain, (m1_group_2(X1,X2)|v1_grsolv_1(X3)|v2_struct_0(X2)|v2_struct_0(X3)|esk4_2(X3,k10_finseq_1(X4,X5))=X5|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X5!=k6_group_2(X3)|X4!=k7_group_2(X3)|~m1_group_2(X1,np__1)|~m2_finseq_1(k10_finseq_1(X4,X5),k1_group_3(X3))|~l3_algstr_0(X2)|~l3_algstr_0(X3)|~v3_group_1(X2)|~v3_group_1(X3)|~v2_group_1(X2)|~v2_group_1(X3)|~m1_subset_1(X5,k1_group_3(X3))|~m1_subset_1(X4,k1_group_3(X3)))).#trainneg
cnf(i_0_2932, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,np__3),np__2)|r1_xxreal_0(np__2,k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_612, plain, (v1_xboole_0(k1_zfmisc_1(k5_numbers))|m1_group_2(k2_finseq_1(X1),X2)|v2_struct_0(X2)|k1_group_3(X2)!=k1_zfmisc_1(k5_numbers)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1017, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4021, plain, (v3_ordinal1(k2_xcmplx_0(np__1,np__3))|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_4008, plain, (r1_xxreal_0(k6_numbers,esk9_1(k5_numbers))|r1_xxreal_0(esk9_1(k5_numbers),k6_numbers))).#trainneg
cnf(i_0_3620, plain, (v1_finset_1(u1_struct_0(X2))|~v7_ordinal1(X1)|~v1_xboole_0(X1)|~l3_algstr_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_613, plain, (v1_xboole_0(k1_zfmisc_1(k5_numbers))|m1_group_2(k4_finseq_1(X1),X2)|v2_struct_0(X2)|k1_group_3(X2)!=k1_zfmisc_1(k5_numbers)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_614, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|m1_group_2(k5_numbers,X1)|v2_struct_0(X1)|k1_group_3(X1)!=k1_zfmisc_1(k1_numbers)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4065, plain, (k6_group_2(X1)=np__1|k6_group_2(X1)=np__2|v2_struct_0(X1)|X4!=k2_tarski(np__1,np__2)|X4!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4011, plain, (r1_xxreal_0(k6_numbers,esk3_2(X1,k10_finseq_1(X2,X3)))|r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,X3)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1211, plain, (r2_hidden(np__2,X1)|X1!=k4_finseq_1(X2)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4064, plain, (k6_group_2(X1)=np__1|k6_group_2(X1)=np__2|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|k4_finseq_1(X2)!=k1_group_3(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1247, plain, (r2_hidden(k2_xcmplx_0(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(esk3_2(X1,X2),k5_numbers))).#trainneg
cnf(i_0_2748, plain, (r2_hidden(k7_group_2(X1),X2)|v2_struct_0(k7_group_2(X1))|v2_struct_0(X1)|X2!=k1_group_3(k7_group_2(X1))|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2443, plain, (esk6_2(k1_funct_1(X1,np__2),X1)=esk5_2(k1_funct_1(X1,np__2),X1)|v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~v1_group_3(esk5_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2))|~m1_group_2(esk5_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2))|~m2_finseq_1(X1,k1_group_3(k1_funct_1(X1,np__2)))|~l3_algstr_0(k1_funct_1(X1,np__2))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_866, plain, (v1_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_2933, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,np__3),k6_numbers)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_4001, plain, (k3_finseq_1(X1)=np__3|k4_finseq_1(X1)!=k4_finseq_1(np__2)|np__3!=np__2|~v7_ordinal1(np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_616, plain, (v1_xboole_0(k1_numbers)|m1_group_2(np__1,X1)|v2_struct_0(X1)|k1_group_3(X1)!=k1_numbers|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1246, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3345, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers)|~v7_ordinal1(np__1)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_3418, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__2,np__3),k5_numbers)|~v7_ordinal1(np__2)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_618, plain, (v1_xboole_0(k1_numbers)|m1_group_2(np__2,X1)|v2_struct_0(X1)|k1_group_3(X1)!=k1_numbers|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_620, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|m1_group_2(X2,X3)|v2_struct_0(X3)|k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_1679, plain, (esk5_2(X1,X3)=k1_funct_1(X3,np__2)|esk5_2(X1,X3)=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2615, plain, (v5_group_1(esk6_2(X1,k10_finseq_1(X1,X1)))|v2_struct_0(X1)|v1_grsolv_1(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1288, plain, (k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1748, plain, (X1!=np__2|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4235, plain, (r2_hidden(k7_group_2(X1),k1_group_3(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4161, plain, (v1_finset_1(u1_struct_0(X1))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_4310, plain, (u1_struct_0(X1)!=np__2|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_621, plain, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1305, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4269, plain, (v3_ordinal1(k2_xcmplx_0(np__2,np__3))|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_623, plain, (v1_finset_1(k9_xtuple_0(X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_625, plain, (X1=k1_group_3(esk11_1(X2))|m1_group_2(esk7_2(esk11_1(X2),X1),esk11_1(X2))|r2_hidden(esk7_2(esk11_1(X2),X1),X1)|v2_struct_0(esk11_1(X2))|v2_struct_0(X2)|~v3_group_1(esk11_1(X2))|~v2_group_1(esk11_1(X2))|~l3_algstr_0(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_4253, plain, (r1_xxreal_0(np__2,esk3_2(X1,k10_finseq_1(X2,X3)))|r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,X3)),np__2)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4254, plain, (r1_xxreal_0(np__2,esk3_2(X1,X2))|r1_xxreal_0(esk3_2(X1,X2),np__2)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4293, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X3)|v2_struct_0(X3)|k2_xcmplx_0(np__1,esk3_2(X3,k10_finseq_1(X1,X2)))=np__1|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~v7_ordinal1(esk3_2(X3,k10_finseq_1(X1,X2)))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_4302, plain, (np__3=np__1|np__3=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_4298, plain, (np__3=np__1|np__3=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_4258, plain, (r1_xxreal_0(np__2,k7_group_2(k6_group_2(X1)))|r1_xxreal_0(k7_group_2(k6_group_2(X1)),np__2)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_986, plain, (v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(k6_numbers,k6_numbers),k6_numbers)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_4145, plain, (r1_xxreal_0(esk9_1(k1_numbers),k6_numbers)|r1_xxreal_0(k6_numbers,esk9_1(k1_numbers)))).#trainneg
cnf(i_0_4309, plain, (u2_algstr_0(X1)!=np__2|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_628, plain, (v3_ordinal1(esk3_2(X1,X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~v3_ordinal1(k4_finseq_1(X2))|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4152, plain, (v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,esk9_1(k5_numbers))|~v1_xreal_0(esk9_1(k5_numbers))|~v2_xxreal_0(esk9_1(k5_numbers)))).#trainneg
cnf(i_0_4250, plain, (r1_xxreal_0(np__2,esk9_1(k5_numbers))|r1_xxreal_0(esk9_1(k5_numbers),np__2))).#trainneg
cnf(i_0_632, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|X2!=k1_funct_1(esk2_1(X3),k2_nat_1(np__1,X4))|X1!=k1_funct_1(esk2_1(X3),X4)|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X1,X3)|~r2_hidden(k2_nat_1(np__1,X4),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_633, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|X2!=k1_funct_1(esk2_1(X3),k2_xcmplx_0(X4,np__1))|X1!=k1_funct_1(esk2_1(X3),X4)|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X1,X3)|~r2_hidden(k2_xcmplx_0(X4,np__1),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1436, plain, (r2_hidden(esk3_2(X1,X3),k4_finseq_1(X4))|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X4)!=np__2|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2738, plain, (m1_subset_1(k7_group_2(X1),k5_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3289, plain, (esk5_2(X1,k10_finseq_1(X2,X3))=np__2|esk5_2(X1,k10_finseq_1(X2,X3))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_4400, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|np__3=np__2|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_4403, plain, (v1_funct_2(u2_algstr_0(X1),k6_numbers,k6_numbers)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_634, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|X2!=k1_funct_1(esk2_1(X3),k2_nat_1(np__1,X4))|X1!=k1_funct_1(esk2_1(X3),X4)|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X1,X3)|~r2_hidden(k2_nat_1(np__1,X4),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_4404, plain, (v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,esk9_1(k1_numbers))|~v1_xreal_0(esk9_1(k1_numbers))|~v2_xxreal_0(esk9_1(k1_numbers)))).#trainneg
cnf(i_0_4326, plain, (v1_finset_1(k2_finseq_1(X2))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4327, plain, (v1_finset_1(k1_finseq_1(X2))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_635, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|v1_xboole_0(k4_finseq_1(esk2_1(X3)))|X2!=k1_funct_1(esk2_1(X3),k2_nat_1(X4,np__1))|X1!=k1_funct_1(esk2_1(X3),X4)|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X1,X3)|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~m1_subset_1(k2_nat_1(X4,np__1),k4_finseq_1(esk2_1(X3))))).#trainneg
cnf(i_0_636, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(esk2_1(X3)),k6_numbers)|v1_grsolv_1(X4)|v2_struct_0(X4)|X2!=k1_funct_1(esk2_1(X3),k2_nat_1(esk3_2(X4,esk2_1(X3)),np__1))|X1!=k1_funct_1(esk2_1(X3),esk3_2(X4,esk2_1(X3)))|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X1,X3)|~r2_hidden(esk3_2(X4,esk2_1(X3)),k4_finseq_1(esk2_1(X3)))|~m1_subset_1(esk3_2(X4,esk2_1(X3)),k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|k1_funct_1(esk2_1(X3),k3_finseq_1(esk2_1(X3)))!=k6_group_2(X4)|k1_funct_1(esk2_1(X3),np__1)!=k7_group_2(X4)|~m2_finseq_1(esk2_1(X3),k1_group_3(X4))|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4))).#trainneg
cnf(i_0_4391, plain, (np__3=np__1|esk5_2(X1,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_4398, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__1|np__3=np__2|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_966, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k7_group_2(k7_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_4562, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_4317, plain, (m1_subset_1(k7_group_2(X1),k1_group_3(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4350, plain, (r1_xxreal_0(np__2,esk9_1(k1_numbers))|r1_xxreal_0(esk9_1(k1_numbers),np__2))).#trainneg
cnf(i_0_639, plain, (k2_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1896, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3419, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__2,np__3),k1_numbers)|~v7_ordinal1(np__2)|~m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_3655, plain, (m2_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers,k5_numbers)|~v7_ordinal1(k6_numbers)|~m1_subset_1(k2_xcmplx_0(np__1,np__3),k5_numbers)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_3645, plain, (k1_finseq_1(X3)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_642, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k6_numbers)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_646, plain, (v15_algstr_0(k2_nat_1(np__1,esk3_2(X1,X2)))|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4581, plain, (esk4_2(X3,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X3)|v2_struct_0(X3)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_3470, plain, (m1_group_2(esk11_1(esk3_2(X1,k10_finseq_1(X2,X3))),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(esk3_2(X1,k10_finseq_1(X2,X3)))|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~l3_algstr_0(esk3_2(X1,k10_finseq_1(X2,X3)))|~v2_group_1(esk3_2(X1,k10_finseq_1(X2,X3))))).#trainneg
cnf(i_0_4665, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1867, plain, (m1_group_2(esk3_2(X1,X2),X3)|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_group_3(k6_group_2(X3))!=k2_tarski(np__1,np__2)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(k6_group_2(X3))|~l3_algstr_0(X1)|~v3_group_1(k6_group_2(X3))|~v3_group_1(X1)|~v2_group_1(k6_group_2(X3))|~v2_group_1(X1))).#trainneg
cnf(i_0_4015, plain, (r1_xxreal_0(k6_numbers,k7_group_2(k6_group_2(X1)))|r1_xxreal_0(k7_group_2(k6_group_2(X1)),k6_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2850, plain, (v2_group_1(esk14_2(X1,X2))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~m1_group_2(X2,X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_4718, plain, (r1_xxreal_0(esk9_1(k1_numbers),np__3)|r1_xxreal_0(np__3,esk9_1(k1_numbers)))).#trainneg
cnf(i_0_4306, plain, (k1_finseq_1(X1)!=np__2|~v1_xboole_0(X1))).#trainneg
cnf(i_0_647, plain, (m1_group_2(k2_nat_1(np__1,esk3_2(X1,X2)),X3)|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3424, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3790, plain, (k2_finseq_1(X3)=k2_finseq_1(X2)|~v1_xboole_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3799, plain, (k2_finseq_1(X3)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_648, plain, (m1_subset_1(k2_nat_1(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_651, plain, (v1_xboole_0(X1)|v15_algstr_0(esk9_1(X1))|v2_struct_0(X2)|X1!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_1459, plain, (m1_group_2(np__1,X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1255, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4402, plain, (r1_xxreal_0(k6_group_2(X1),np__2)|r1_xxreal_0(np__2,k6_group_2(X1))|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1))|~v15_algstr_0(k6_group_2(X1)))).#trainneg
cnf(i_0_3508, plain, (k7_group_2(X1)=np__2|k7_group_2(X1)=np__1|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|X3!=k4_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|X3!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3487, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=np__2|esk6_2(X1,k10_finseq_1(X1,X1))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|X3!=k4_finseq_1(X2)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|X3!=k1_group_3(X1)|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk6_2(X1,k10_finseq_1(X1,X1)))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3686, plain, (m1_group_2(np__2,X1)|v2_struct_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_1238, plain, (r2_hidden(k2_nat_1(np__1,esk3_2(X1,X2)),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2235, plain, (r1_xxreal_0(X1,k2_xcmplx_0(np__3,X2))|r1_xxreal_0(k2_xcmplx_0(np__3,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2851, plain, (l3_algstr_0(esk14_2(X1,X2))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~m1_group_2(X2,X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_4311, plain, (k2_finseq_1(X1)!=np__2|~v1_xboole_0(X1))).#trainneg
cnf(i_0_653, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|v15_algstr_0(u2_algstr_0(X1))|v2_struct_0(X2)|k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1)))!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_4849, plain, (v2_xxreal_0(np__3)|r1_xxreal_0(np__3,esk9_1(k1_numbers))|~v1_xreal_0(np__3)|~v1_xreal_0(esk9_1(k1_numbers))|~v2_xxreal_0(esk9_1(k1_numbers)))).#trainneg
cnf(i_0_654, plain, (v1_xboole_0(k1_zfmisc_1(k5_numbers))|v15_algstr_0(k2_finseq_1(X1))|v2_struct_0(X2)|k1_group_3(X2)!=k1_zfmisc_1(k5_numbers)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_655, plain, (v1_xboole_0(k1_zfmisc_1(k5_numbers))|v15_algstr_0(k4_finseq_1(X1))|v2_struct_0(X2)|k1_group_3(X2)!=k1_zfmisc_1(k5_numbers)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4916, plain, (k7_group_2(X1)=np__1|k7_group_2(X1)=np__2|v2_struct_0(X1)|X4!=k2_tarski(np__1,np__2)|X4!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4848, plain, (r1_xxreal_0(k6_group_2(X1),k6_numbers)|r1_xxreal_0(k6_numbers,k6_group_2(X1))|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1))|~v15_algstr_0(k6_group_2(X1)))).#trainneg
cnf(i_0_4820, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|esk4_2(X1,k10_finseq_1(X1,X2))=X1|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_4845, plain, (m1_group_2(np__2,X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|k1_group_3(k6_group_2(X4))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_4862, plain, (r1_xxreal_0(np__3,esk3_2(X1,k10_finseq_1(X2,X3)))|r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,X3)),np__3)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2852, plain, (v3_group_1(esk14_2(X1,X2))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_4859, plain, (r1_xxreal_0(np__3,esk9_1(k5_numbers))|r1_xxreal_0(esk9_1(k5_numbers),np__3))).#trainneg
cnf(i_0_656, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|v15_algstr_0(k5_numbers)|v2_struct_0(X1)|k1_group_3(X1)!=k1_zfmisc_1(k1_numbers)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1654, plain, (k9_xtuple_0(k3_finseq_1(k6_numbers))=k2_finseq_1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_4734, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers)|~v1_xreal_0(np__2)|~v7_ordinal1(k6_numbers)|~m1_subset_1(k2_xcmplx_0(np__1,np__3),k5_numbers))).#trainneg
cnf(i_0_658, plain, (v1_xboole_0(k1_numbers)|v15_algstr_0(np__1)|v2_struct_0(X1)|k1_group_3(X1)!=k1_numbers|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_660, plain, (v1_xboole_0(k1_numbers)|v15_algstr_0(np__2)|v2_struct_0(X1)|k1_group_3(X1)!=k1_numbers|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4915, plain, (k7_group_2(X1)=np__1|k7_group_2(X1)=np__2|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|k4_finseq_1(X2)!=k1_group_3(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1415, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_4497, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3290, plain, (esk5_2(X1,k10_finseq_1(X2,X3))=np__2|esk5_2(X1,k10_finseq_1(X2,X3))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X4)!=k1_group_3(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X4)!=np__2|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_3943, plain, (m1_group_2(k7_group_2(np__2),X4)|v1_grsolv_1(X1)|v2_struct_0(X4)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_4841, plain, (m1_group_2(np__1,X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|r2_hidden(np__3,k2_tarski(np__1,np__2))|k1_group_3(k6_group_2(X3))!=k2_tarski(np__1,np__2)|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1215, plain, (X1=np__2|X1=np__1|~r2_hidden(X1,k4_finseq_1(X2))|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2808, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__1,np__3))=k2_xcmplx_0(np__3,np__2)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_662, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|v15_algstr_0(X2)|v2_struct_0(X3)|k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_787, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_4000, plain, (v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(k2_xcmplx_0(np__1,np__3))|~v2_xxreal_0(k2_xcmplx_0(np__1,np__3))|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_4246, plain, (v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(k2_xcmplx_0(np__1,np__3))|~v2_xxreal_0(k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_664, plain, (m1_group_2(k2_nat_1(esk3_2(X1,esk2_1(X2)),np__1),X3)|r1_xxreal_0(k3_finseq_1(esk2_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|k6_group_2(X2)!=k6_group_2(X1)|k1_funct_1(esk2_1(X2),np__1)!=k7_group_2(X1)|k4_finseq_1(esk2_1(X2))!=k1_group_3(X3)|~m2_finseq_1(esk2_1(X2),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_671, plain, (v15_algstr_0(k2_nat_1(np__1,esk3_2(X1,X2)))|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4958, plain, (k7_group_2(X1)=np__2|k7_group_2(X1)=np__1|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5058, plain, (esk3_2(X1,X2)=np__1|esk3_2(X1,X2)=np__2|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5063, plain, (k2_nat_1(esk3_2(X1,X2),np__1)=np__1|k2_nat_1(esk3_2(X1,X2),np__1)=np__2|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4966, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_4963, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_4965, plain, (m1_group_6(esk6_2(X1,k10_finseq_1(X1,X2)),X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5055, plain, (esk6_2(X1,k10_finseq_1(X1,X1))=np__1|esk6_2(X1,k10_finseq_1(X1,X1))=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k4_finseq_1(X2)!=k1_group_3(X1)|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(esk6_2(X1,k10_finseq_1(X1,X1)))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2853, plain, (v5_group_1(esk14_2(X1,X2))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_5011, plain, (k2_finseq_1(k3_finseq_1(k6_numbers))!=k2_tarski(np__1,np__2)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_672, plain, (m1_group_2(k2_nat_1(np__1,esk3_2(X1,X2)),X3)|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5083, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v1_xreal_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5070, plain, (v1_xreal_0(k2_xcmplx_0(np__2,np__3))|~v1_xreal_0(k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(np__1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_5005, plain, (v3_ordinal1(k2_finseq_1(k3_finseq_1(k6_numbers)))|~v5_ordinal1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_673, plain, (m1_subset_1(k2_nat_1(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4971, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,esk6_2(X1,k10_finseq_1(X1,X2))))|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_4967, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_4964, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5043, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|np__3=np__2|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_4968, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5065, plain, (k2_nat_1(np__1,esk3_2(X1,X2))=np__1|k2_nat_1(np__1,esk3_2(X1,X2))=np__2|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3257, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k6_numbers))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3267, plain, (v5_relat_1(esk9_1(k1_zfmisc_1(k6_numbers)),X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5004, plain, (k2_finseq_1(k3_finseq_1(k6_numbers))=k4_finseq_1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5232, plain, (k4_finseq_1(k3_finseq_1(k6_numbers))!=np__2|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_677, plain, (v15_algstr_0(k2_xcmplx_0(esk3_2(X1,X2),np__1))|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5141, plain, (k9_xtuple_0(k6_numbers)!=k2_tarski(np__1,np__2)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5249, plain, (k4_finseq_1(k3_finseq_1(k6_numbers))!=k2_tarski(np__1,np__2)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5149, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_678, plain, (m1_group_2(k2_xcmplx_0(esk3_2(X1,X2),np__1),X3)|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_679, plain, (m1_subset_1(k2_xcmplx_0(esk3_2(X1,X2),np__1),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5200, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5197, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=esk5_2(X1,k10_finseq_1(X1,X2))|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5199, plain, (m1_group_6(esk6_2(X1,k10_finseq_1(X1,X2)),X1,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5135, plain, (X1=np__2|X1=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X1)|k3_finseq_1(X2)!=np__2|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5223, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|m1_group_2(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1))).#trainneg
cnf(i_0_2946, plain, (u1_struct_0(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_5247, plain, (k4_finseq_1(k3_finseq_1(k6_numbers))=k9_xtuple_0(k6_numbers)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5139, plain, (k2_tarski(np__1,np__2)!=k6_numbers|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_xboole_0(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5379, plain, (v1_finset_1(k2_zfmisc_1(X2,X3))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_683, plain, (v3_ordinal1(k2_nat_1(esk3_2(X1,X2),np__1))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|~v3_ordinal1(k4_finseq_1(X2))|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5115, plain, (esk3_2(X1,esk17_1(k1_group_3(X1)))=np__2|esk3_2(X1,esk17_1(k1_group_3(X1)))=np__1|r1_xxreal_0(k3_finseq_1(esk17_1(k1_group_3(X1))),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(esk17_1(k1_group_3(X1)),k3_finseq_1(esk17_1(k1_group_3(X1))))!=k6_group_2(X1)|k1_funct_1(esk17_1(k1_group_3(X1)),np__1)!=k7_group_2(X1)|k3_finseq_1(esk17_1(k1_group_3(X1)))!=np__2|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4095, plain, (m1_group_2(X1,X2)|v1_grsolv_1(X3)|v2_struct_0(np__2)|v2_struct_0(X3)|v2_struct_0(X2)|esk3_2(X3,k10_finseq_1(X4,X5))=np__1|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X5!=k6_group_2(X3)|X4!=k7_group_2(X3)|~m1_group_2(X1,k6_group_2(np__2))|~m2_finseq_1(k10_finseq_1(X4,X5),k1_group_3(X3))|~l3_algstr_0(np__2)|~l3_algstr_0(X2)|~l3_algstr_0(X3)|~v3_group_1(np__2)|~v3_group_1(X2)|~v3_group_1(X3)|~v2_group_1(np__2)|~v2_group_1(X2)|~v2_group_1(X3)|~m1_subset_1(X5,k1_group_3(X3))|~m1_subset_1(X4,k1_group_3(X3)))).#trainneg
cnf(i_0_5217, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|X1=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~m1_subset_1(X1,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5354, plain, (v1_relat_1(u1_struct_0(X3))|~v1_xboole_0(X1)|~l3_algstr_0(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_3097, plain, (u1_struct_0(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_5261, plain, (r1_xxreal_0(k2_xcmplx_0(np__1,np__3),np__3)|r1_xxreal_0(np__3,k2_xcmplx_0(np__1,np__3))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_5450, plain, (k9_xtuple_0(k6_numbers)!=np__2|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5257, plain, (k2_tarski(np__1,np__2)!=k4_finseq_1(k6_numbers)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5558, plain, (v1_finset_1(k2_zfmisc_1(X2,X3))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4314, plain, (m1_group_2(k7_group_2(X1),X2)|v2_struct_0(X2)|v2_struct_0(X1)|k1_group_3(k7_group_2(X1))!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_971, plain, (r2_hidden(k7_group_2(k7_group_2(X1)),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k7_group_2(X1))))).#trainneg
cnf(i_0_1469, plain, (k1_funct_1(X1,k2_xcmplx_0(np__1,esk3_2(X2,X1)))=esk5_2(X2,X1)|r1_xxreal_0(k3_finseq_1(X1),k6_numbers)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X1,k3_finseq_1(X1))!=k6_group_2(X2)|k1_funct_1(X1,np__1)!=k7_group_2(X2)|~v7_ordinal1(np__1)|~m2_finseq_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(esk3_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_3464, plain, (m1_group_2(X1,X2)|v1_grsolv_1(X3)|v2_struct_0(X3)|v2_struct_0(X2)|r2_hidden(np__3,k2_tarski(np__1,np__2))|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X4!=k6_group_2(X3)|k7_group_2(X3)!=X3|~m1_group_2(X1,np__1)|~m2_finseq_1(k10_finseq_1(X3,X4),k1_group_3(X3))|~l3_algstr_0(X2)|~l3_algstr_0(X3)|~v3_group_1(X2)|~v3_group_1(X3)|~v2_group_1(X2)|~v2_group_1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X4,k1_group_3(X3))|~v15_algstr_0(X3))).#trainneg
cnf(i_0_1198, plain, (k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|k2_nat_1(esk3_2(X1,k10_finseq_1(X2,X3)),np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5489, plain, (v1_relat_1(u1_struct_0(X1))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_5411, plain, (k9_xtuple_0(k6_numbers)!=k4_finseq_1(np__2)|~v1_finseq_1(np__2)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(np__2)|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(np__2)|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_692, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(k6_numbers,X1))=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_695, plain, (v3_card_1(k9_xtuple_0(X1),k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_716, plain, (m1_group_2(esk11_1(k6_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_5205, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,esk6_2(X1,k10_finseq_1(X1,X2))))|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1616, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X3)|v2_struct_0(X3)|k2_nat_1(np__1,esk3_2(X3,k10_finseq_1(X1,X2)))=np__1|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_5198, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_5707, plain, (v5_group_1(esk11_1(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5702, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(esk11_1(k6_group_2(X1)))|~v3_group_1(X1))).#trainneg
cnf(i_0_5704, plain, (v2_group_1(esk11_1(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_718, plain, (m1_group_2(k6_group_2(k6_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_5582, plain, (k4_finseq_1(k6_numbers)!=np__2|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5650, plain, (v1_relat_1(k2_finseq_1(X2))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5651, plain, (v1_relat_1(k1_finseq_1(X2))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_722, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_723, plain, (m1_group_2(esk4_2(X1,esk17_1(k1_group_3(X1))),X1)|r1_xxreal_0(k3_finseq_1(esk17_1(k1_group_3(X1))),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(esk17_1(k1_group_3(X1)),k3_finseq_1(esk17_1(k1_group_3(X1))))!=k6_group_2(X1)|k1_funct_1(esk17_1(k1_group_3(X1)),np__1)!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5743, plain, (np__3=np__2|np__3=np__1|v1_group_3(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(X1,X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),X1))).#trainneg
cnf(i_0_5590, plain, (m1_subset_1(k7_group_2(k7_group_2(X1)),k5_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k7_group_2(X1))))).#trainneg
cnf(i_0_5705, plain, (l3_algstr_0(esk11_1(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_5706, plain, (v3_group_1(esk11_1(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_724, plain, (m1_group_2(esk5_2(X1,esk17_1(k1_group_3(X1))),X1)|r1_xxreal_0(k3_finseq_1(esk17_1(k1_group_3(X1))),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(esk17_1(k1_group_3(X1)),k3_finseq_1(esk17_1(k1_group_3(X1))))!=k6_group_2(X1)|k1_funct_1(esk17_1(k1_group_3(X1)),np__1)!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5452, plain, (k4_finseq_1(np__2)!=k6_numbers|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_5584, plain, (k4_finseq_1(k6_numbers)!=k4_finseq_1(np__2)|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_725, plain, (r1_xxreal_0(k3_finseq_1(esk17_1(k1_group_3(X1))),k6_numbers)|v1_grsolv_1(X1)|v15_algstr_0(esk4_2(X1,esk17_1(k1_group_3(X1))))|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(esk17_1(k1_group_3(X1)),k3_finseq_1(esk17_1(k1_group_3(X1))))!=k6_group_2(X1)|k1_funct_1(esk17_1(k1_group_3(X1)),np__1)!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_726, plain, (r1_xxreal_0(k3_finseq_1(esk17_1(k1_group_3(X1))),k6_numbers)|v1_grsolv_1(X1)|v15_algstr_0(esk5_2(X1,esk17_1(k1_group_3(X1))))|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k1_funct_1(esk17_1(k1_group_3(X1)),k3_finseq_1(esk17_1(k1_group_3(X1))))!=k6_group_2(X1)|k1_funct_1(esk17_1(k1_group_3(X1)),np__1)!=k7_group_2(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5625, plain, (k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__1|k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2739, plain, (k7_group_2(X1)=X2|k7_group_2(X1)=X3|v2_struct_0(X1)|k2_tarski(X3,X2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2811, plain, (m1_group_2(esk14_2(X1,k6_group_2(X2)),X2)|v2_struct_0(X2)|v2_struct_0(X1)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_group_2(k6_group_2(X2),X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1224, plain, (k4_finseq_1(X1)=k4_finseq_1(X2)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5758, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k6_group_2(k6_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_5757, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k6_group_2(k6_group_2(X1)))|~v3_group_1(X1))).#trainneg
cnf(i_0_5759, plain, (v2_group_1(k6_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_5789, plain, (k2_xcmplx_0(k6_numbers,np__2)=np__2|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_743, plain, (v1_xboole_0(k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_xboole_0(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5790, plain, (k2_xcmplx_0(k6_numbers,np__3)=np__3|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_5358, plain, (v5_relat_1(u1_struct_0(X3),X2)|~v1_xboole_0(X1)|~l3_algstr_0(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_754, plain, (m1_group_2(esk11_1(k7_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_756, plain, (m1_group_2(k6_group_2(k7_group_2(X1)),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_5762, plain, (v5_group_1(k6_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5760, plain, (l3_algstr_0(k6_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_5761, plain, (v3_group_1(k6_group_2(k6_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5681, plain, (v3_card_1(k9_xtuple_0(np__2),np__2)|~v7_ordinal1(np__2)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_759, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_6047, plain, (k2_xcmplx_0(k6_numbers,np__2)=np__2|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_6060, plain, (v5_relat_1(u1_struct_0(X1),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~l3_algstr_0(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_5534, plain, (v4_relat_1(u1_struct_0(X3),X1)|~v1_xboole_0(X2)|~l3_algstr_0(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_761, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k6_numbers,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_778, plain, (X1=k1_group_3(k7_group_2(X2))|m1_group_2(esk7_2(k7_group_2(X2),X1),k7_group_2(X2))|r2_hidden(esk7_2(k7_group_2(X2),X1),X1)|v2_struct_0(k7_group_2(X2))|v2_struct_0(X2)|~v3_group_1(k7_group_2(X2))|~v2_group_1(k7_group_2(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_6101, plain, (v5_group_1(k6_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6102, plain, (r2_hidden(k6_group_2(k7_group_2(X1)),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k7_group_2(X1))))).#trainneg
cnf(i_0_6031, plain, (m1_group_2(esk3_2(X1,k10_finseq_1(X2,X3)),X4)|v2_struct_0(X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(k6_group_2(X4)))|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|k1_group_3(k6_group_2(k6_group_2(X4)))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(k6_group_2(k6_group_2(X4)))|~l3_algstr_0(X1)|~v3_group_1(k6_group_2(k6_group_2(X4)))|~v3_group_1(X1)|~v2_group_1(k6_group_2(k6_group_2(X4)))|~v2_group_1(X1))).#trainneg
cnf(i_0_6097, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k6_group_2(k7_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_5987, plain, (r2_hidden(esk3_2(X1,X2),k4_finseq_1(X3))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_5763, plain, (r2_hidden(k6_group_2(k6_group_2(X1)),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_6075, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(esk11_1(k7_group_2(X1)))|~v3_group_1(X1))).#trainneg
cnf(i_0_6096, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k6_group_2(k7_group_2(X1)))|~v3_group_1(X1))).#trainneg
cnf(i_0_6151, plain, (k2_xcmplx_0(k6_numbers,np__2)=np__2)).#trainneg
cnf(i_0_781, plain, (v1_xboole_0(k1_numbers)|m1_group_2(k6_numbers,X1)|v2_struct_0(X1)|k1_group_3(X1)!=k1_numbers|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6053, plain, (k2_xcmplx_0(k6_numbers,np__3)=np__3|~v1_xcmplx_0(np__1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_5168, plain, (v3_ordinal1(k9_xtuple_0(k6_numbers))|~v5_ordinal1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_6158, plain, (v4_relat_1(u1_struct_0(X1),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~l3_algstr_0(X3)|~v2_struct_0(X3))).#trainneg
cnf(i_0_782, plain, (v1_xboole_0(k1_numbers)|v15_algstr_0(k6_numbers)|v2_struct_0(X1)|k1_group_3(X1)!=k1_numbers|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_785, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(k6_numbers,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_6207, plain, (m1_subset_1(k6_group_2(k7_group_2(X1)),k5_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k7_group_2(X1))))).#trainneg
cnf(i_0_6088, plain, (m1_group_2(k6_group_2(k7_group_2(esk3_2(X1,X2))),X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(esk3_2(X1,X2))|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~l3_algstr_0(esk3_2(X1,X2))|~v3_group_1(esk3_2(X1,X2))|~v2_group_1(esk3_2(X1,X2)))).#trainneg
cnf(i_0_6077, plain, (v2_group_1(esk11_1(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_6078, plain, (l3_algstr_0(esk11_1(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_791, plain, (k1_finseq_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5413, plain, (v3_ordinal1(k9_xtuple_0(k6_numbers))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5414, plain, (v1_funct_1(k9_xtuple_0(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k5_numbers)|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k5_numbers)|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_2013, plain, (r2_hidden(k2_nat_1(np__1,np__2),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3212, plain, (m1_group_2(X3,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X2,X3))=X3|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_6436, plain, (m1_group_2(X3,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2712, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|k2_nat_1(np__2,np__1)=np__2|k2_nat_1(np__2,np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_6325, plain, (k2_xcmplx_0(k6_numbers,np__3)=np__3|~v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_800, plain, (k2_finseq_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_995, plain, (v1_finset_1(k2_tarski(np__1,np__2))|~v7_ordinal1(X1)|X1!=np__2|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6326, plain, (v3_ordinal1(k4_finseq_1(k6_numbers))|~v5_ordinal1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_6406, plain, (v3_ordinal1(k4_finseq_1(k6_numbers))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_805, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_806, plain, (v15_algstr_0(k2_xcmplx_0(np__1,esk3_2(X1,X2)))|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5611, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X3)|v2_struct_0(X3)|np__3=np__1|np__3=np__2|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_3195, plain, (esk6_2(X1,k10_finseq_1(X2,X3))=esk5_2(X1,k10_finseq_1(X2,X3))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X2,X3))=X3|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X2,X3)),X1,X3)|~v1_group_3(esk5_2(X1,k10_finseq_1(X2,X3)),X3)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_1207, plain, (k2_nat_1(esk3_2(X1,X3),np__1)=np__1|k2_nat_1(esk3_2(X1,X3),np__1)=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1868, plain, (m1_group_2(esk3_2(X1,X2),X3)|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(k6_group_2(X3)))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_group_3(k7_group_2(k6_group_2(X3)))!=k2_tarski(np__1,np__2)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(k7_group_2(k6_group_2(X3)))|~l3_algstr_0(X1)|~v3_group_1(k7_group_2(k6_group_2(X3)))|~v3_group_1(X1)|~v2_group_1(k7_group_2(k6_group_2(X3)))|~v2_group_1(X1))).#trainneg
cnf(i_0_6079, plain, (v3_group_1(esk11_1(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6513, plain, (k2_xcmplx_0(k6_numbers,np__3)=np__3)).#trainneg
cnf(i_0_807, plain, (m1_group_2(k2_xcmplx_0(np__1,esk3_2(X1,X2)),X3)|v2_struct_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(X2)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6548, plain, (v1_finset_1(k2_tarski(np__1,np__2))|X1!=np__2|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6408, plain, (v1_funct_1(k4_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k5_numbers)|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_808, plain, (m1_subset_1(k2_xcmplx_0(np__1,esk3_2(X1,X2)),k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_817, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v5_ordinal1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_1749, plain, (r2_hidden(k2_nat_1(esk3_2(X1,np__2),np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=k1_funct_1(np__2,np__2)|k7_group_2(X1)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6033, plain, (m1_group_2(esk3_2(X1,X2),X3)|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(k6_group_2(X3)))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_group_3(k6_group_2(k6_group_2(X3)))!=k2_tarski(np__1,np__2)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(k6_group_2(k6_group_2(X3)))|~l3_algstr_0(X1)|~v3_group_1(k6_group_2(k6_group_2(X3)))|~v3_group_1(X1)|~v2_group_1(k6_group_2(k6_group_2(X3)))|~v2_group_1(X1))).#trainneg
cnf(i_0_1729, plain, (m1_subset_1(esk3_2(X1,np__2),k5_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=k1_funct_1(np__2,np__2)|k7_group_2(X1)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2602, plain, (esk6_2(k1_funct_1(X1,np__2),X1)=k1_funct_1(X1,np__2)|v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~l3_algstr_0(k1_funct_1(X1,np__2))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5304, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v2_struct_0(X1)|np__3=np__2|np__3=np__1|v1_grsolv_1(X1)|~m1_group_2(X1,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_6098, plain, (v2_group_1(k6_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_6099, plain, (l3_algstr_0(k6_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_5234, plain, (v1_finset_1(k4_finseq_1(k3_finseq_1(k6_numbers)))|~v7_ordinal1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_5236, plain, (v1_xboole_0(k3_finseq_1(k6_numbers))|~v7_ordinal1(k3_finseq_1(k6_numbers))|~v1_xboole_0(k4_finseq_1(k3_finseq_1(k6_numbers)))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_838, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|m1_group_2(k6_numbers,X3)|v2_struct_0(X3)|k1_zfmisc_1(k2_zfmisc_1(X1,X2))!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_839, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|v15_algstr_0(k6_numbers)|v2_struct_0(X3)|k1_zfmisc_1(k2_zfmisc_1(X1,X2))!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_2166, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|r2_hidden(k2_nat_1(np__2,np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_3575, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_grsolv_1(X3)|v2_struct_0(esk3_2(X3,k10_finseq_1(X4,X5)))|v2_struct_0(X3)|~m1_group_2(X1,k7_group_2(esk3_2(X3,k10_finseq_1(X4,X5))))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X2)|X5!=k6_group_2(X3)|X4!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X4,X5),k1_group_3(X3))|~l3_algstr_0(esk3_2(X3,k10_finseq_1(X4,X5)))|~l3_algstr_0(X3)|~v3_group_1(esk3_2(X3,k10_finseq_1(X4,X5)))|~v3_group_1(X3)|~v2_group_1(esk3_2(X3,k10_finseq_1(X4,X5)))|~v2_group_1(X3))).#trainneg
cnf(i_0_6621, plain, (k2_nat_1(esk3_2(X1,X2),np__1)=np__2|k2_nat_1(esk3_2(X1,X2),np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k3_finseq_1(X2)!=np__2|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6740, plain, (esk3_2(X1,k10_finseq_1(X1,X1))=np__1|r2_hidden(k2_nat_1(np__2,np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_6687, plain, (k2_nat_1(esk3_2(X1,np__2),np__1)=np__2|k2_nat_1(esk3_2(X1,np__2),np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=k1_funct_1(np__2,np__2)|k7_group_2(X1)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6100, plain, (v3_group_1(k6_group_2(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6656, plain, (v1_finset_1(k2_tarski(np__1,np__2)))).#trainneg
cnf(i_0_842, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_6726, plain, (v1_finset_1(k9_xtuple_0(k6_numbers))|~v7_ordinal1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_6730, plain, (v1_xboole_0(k3_finseq_1(k6_numbers))|~v7_ordinal1(k3_finseq_1(k6_numbers))|~v1_xboole_0(k9_xtuple_0(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_844, plain, (v15_algstr_0(g3_algstr_0(X1,k6_numbers))|~v1_funct_2(k6_numbers,k2_zfmisc_1(X1,X1),X1))).#trainneg
cnf(i_0_845, plain, (l3_algstr_0(g3_algstr_0(X1,k6_numbers))|~v1_funct_2(k6_numbers,k2_zfmisc_1(X1,X1),X1))).#trainneg
cnf(i_0_3215, plain, (m1_group_2(X1,X2)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X3,np__2)!=k6_group_2(X2)|X4!=k7_group_2(X2)|~m1_group_2(X1,esk4_2(X2,X3))|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X2))|~m1_subset_1(X4,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k1_funct_1(X3,np__1)!=X4|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1600, plain, (k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))=np__1|esk5_2(X1,k10_finseq_1(X2,X3))=X3|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3562, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X2,X3)),X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X2,X3))=X3|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X2,X3)),X3)|~m1_group_2(esk5_2(X1,k10_finseq_1(X2,X3)),X3)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_6812, plain, (r2_hidden(np__2,k4_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_nat_1(esk3_2(X1,X2),np__1)=np__1|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k3_finseq_1(X2)!=np__2)).#trainneg
cnf(i_0_6080, plain, (v5_group_1(esk11_1(k7_group_2(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6062, plain, (v5_relat_1(u1_struct_0(X1),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_6845, plain, (v1_finset_1(k4_finseq_1(np__2))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_849, plain, (k6_numbers=X1|g3_algstr_0(X2,k6_numbers)!=g3_algstr_0(X3,X1)|~v1_funct_2(k6_numbers,k2_zfmisc_1(X2,X2),X2))).#trainneg
cnf(i_0_6856, plain, (v1_finset_1(k4_finseq_1(k6_numbers))|~v7_ordinal1(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_6858, plain, (v1_xboole_0(k3_finseq_1(k6_numbers))|~v7_ordinal1(k3_finseq_1(k6_numbers))|~v1_xboole_0(k4_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_850, plain, (X1=X2|g3_algstr_0(X1,k6_numbers)!=g3_algstr_0(X2,X3)|~v1_funct_2(k6_numbers,k2_zfmisc_1(X1,X1),X1))).#trainneg
cnf(i_0_854, plain, (v1_xboole_0(k1_zfmisc_1(X1))|m1_group_2(X1,X2)|v2_struct_0(X2)|k1_zfmisc_1(X1)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_6796, plain, (k2_nat_1(np__1,esk3_2(X1,X2))=np__2|k2_nat_1(esk3_2(X1,X2),np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,X2))|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6794, plain, (k2_xcmplx_0(np__1,esk3_2(X1,X2))=np__2|k2_nat_1(esk3_2(X1,X2),np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|~v7_ordinal1(esk3_2(X1,X2))|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1241, plain, (r2_hidden(k2_nat_1(np__1,esk3_2(X1,X2)),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6294, plain, (m1_subset_1(k6_group_2(k6_group_2(X1)),k5_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_6851, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_6160, plain, (v4_relat_1(u1_struct_0(X1),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_855, plain, (v1_xboole_0(k1_zfmisc_1(X1))|v15_algstr_0(X1)|v2_struct_0(X2)|k1_zfmisc_1(X1)!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_5237, plain, (v1_xboole_0(k4_finseq_1(k3_finseq_1(k6_numbers)))|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_6387, plain, (v1_finset_1(k1_finseq_1(X2))|~v7_ordinal1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_857, plain, (v15_algstr_0(g3_algstr_0(X1,k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))).#trainneg
cnf(i_0_858, plain, (l3_algstr_0(g3_algstr_0(X1,k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))|~v1_funct_2(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))).#trainneg
cnf(i_0_7042, plain, (m1_subset_1(k6_group_2(k6_group_2(X1)),k5_numbers)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_4297, plain, (k2_xcmplx_0(np__1,esk3_2(X1,X3))=np__1|k2_xcmplx_0(np__1,esk3_2(X1,X3))=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X3))|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2239, plain, (r1_xxreal_0(X1,k2_xcmplx_0(np__2,X2))|r1_xxreal_0(k2_xcmplx_0(np__2,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_7043, plain, (v3_ordinal1(X1)|~r2_hidden(X1,k6_numbers)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_7091, plain, (v1_finset_1(k1_finseq_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7057, plain, (k4_finseq_1(k3_finseq_1(k6_numbers))=k6_numbers|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_862, plain, (k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)=X2|g3_algstr_0(X1,k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))!=g3_algstr_0(X3,X2)|~v1_funct_2(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))).#trainneg
cnf(i_0_6959, plain, (v5_relat_1(k2_finseq_1(X3),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_6960, plain, (v5_relat_1(k1_finseq_1(X3),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_863, plain, (X1=X2|g3_algstr_0(X1,k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))!=g3_algstr_0(X2,X3)|~v1_funct_2(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1),k2_zfmisc_1(X1,X1),X1)|~v1_funct_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1)))).#trainneg
cnf(i_0_1876, plain, (r2_hidden(esk3_2(X1,X2),X3)|v2_struct_0(X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k1_group_3(X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~v15_algstr_0(esk3_2(X1,X2))|k2_tarski(np__1,np__2)!=k1_group_3(X4)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2428, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|esk5_2(X1,k10_finseq_1(X1,X2))=X2|k2_nat_1(np__2,np__1)=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_7114, plain, (r1_xxreal_0(np__2,k6_group_2(k6_group_2(X1)))|r1_xxreal_0(k6_group_2(k6_group_2(X1)),np__2)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7116, plain, (r1_xxreal_0(k6_numbers,k6_group_2(k6_group_2(X1)))|r1_xxreal_0(k6_group_2(k6_group_2(X1)),k6_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5062, plain, (k2_xcmplx_0(np__1,esk3_2(X1,X2))=np__1|k2_xcmplx_0(np__1,esk3_2(X1,X2))=np__2|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2378, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__2)|r1_xxreal_0(np__2,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1746, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4318, plain, (v2_struct_0(X1)|~v1_xboole_0(k1_group_3(k7_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_952, plain, (m1_group_2(k7_group_2(k6_group_2(k6_group_2(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_7133, plain, (v3_ordinal1(X1)|~r2_hidden(X1,k6_numbers))).#trainneg
cnf(i_0_3641, plain, (v5_relat_1(u2_algstr_0(X1),k1_finseq_1(X2))|~l3_algstr_0(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_7196, plain, (k9_xtuple_0(k6_numbers)=k6_numbers|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_xboole_0(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_878, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2973, plain, (v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~v1_group_3(esk5_2(k1_funct_1(X1,np__2),X1),esk4_2(k1_funct_1(X1,np__2),X1))|~m1_group_2(esk5_2(k1_funct_1(X1,np__2),X1),esk4_2(k1_funct_1(X1,np__2),X1))|~m1_group_2(esk4_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2))|~m2_finseq_1(X1,k1_group_3(k1_funct_1(X1,np__2)))|~l3_algstr_0(k1_funct_1(X1,np__2))|~v5_group_1(k5_group_6(esk4_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2)))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_917, plain, (v1_grsolv_1(X1)|v5_group_1(esk5_2(X1,X3))|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4669, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_2307, plain, (v1_group_3(esk6_2(k1_funct_1(X1,np__2),X1),esk4_2(k1_funct_1(X1,np__2),X1))|v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~m1_group_6(k1_funct_1(X1,np__2),k1_funct_1(X1,np__2),esk4_2(k1_funct_1(X1,np__2),X1))|~v1_group_3(k1_funct_1(X1,np__2),esk4_2(k1_funct_1(X1,np__2),X1))|~m2_finseq_1(X1,k1_group_3(k1_funct_1(X1,np__2)))|~l3_algstr_0(k1_funct_1(X1,np__2))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3795, plain, (v5_relat_1(u2_algstr_0(X1),k2_finseq_1(X2))|~l3_algstr_0(X1)|~v1_xboole_0(X2)|~v2_struct_0(X1))).#trainneg
cnf(i_0_7080, plain, (v1_xboole_0(k9_xtuple_0(k6_numbers))|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_897, plain, (X1=k1_group_3(k6_group_2(X2))|m1_group_2(esk7_2(k6_group_2(X2),X1),k6_group_2(X2))|r2_hidden(esk7_2(k6_group_2(X2),X1),X1)|v2_struct_0(k6_group_2(X2))|v2_struct_0(X2)|~v3_group_1(k6_group_2(X2))|~v2_group_1(k6_group_2(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_7139, plain, (v1_finset_1(k2_finseq_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7050, plain, (v4_relat_1(k2_finseq_1(X3),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_902, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_922, plain, (k9_xtuple_0(X1)=k6_numbers|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(k3_finseq_1(X1)))).#trainneg
cnf(i_0_7362, plain, (v5_group_1(k7_group_2(k6_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7358, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k7_group_2(k6_group_2(k6_group_2(X2))))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_7357, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k6_group_2(k6_group_2(X1))))|~v3_group_1(X1))).#trainneg
cnf(i_0_7359, plain, (v2_group_1(k7_group_2(k6_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_924, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_941, plain, (k2_finseq_1(X1)=k2_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7428, plain, (v1_group_3(esk6_2(k1_funct_1(X1,np__2),X1),k1_funct_1(X1,np__2))|v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~m1_group_6(k1_funct_1(X1,np__2),k1_funct_1(X1,np__2),k1_funct_1(X1,np__2))|~v1_group_3(k1_funct_1(X1,np__2),k1_funct_1(X1,np__2))|~m2_finseq_1(X1,k1_group_3(k1_funct_1(X1,np__2)))|~l3_algstr_0(k1_funct_1(X1,np__2))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(k1_funct_1(X1,np__2))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7414, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X1))=np__1|k6_group_2(X1)!=X1|~v1_group_3(np__2,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(np__2,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X1)),X1))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|k7_group_2(X1)!=X1|~m1_subset_1(X1,k1_group_3(X1))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_7424, plain, (esk5_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_7727, plain, (v1_group_3(esk6_2(X1,k10_finseq_1(X1,X1)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_7664, plain, (k3_finseq_1(X1)=k6_numbers|k9_xtuple_0(X1)!=k6_numbers|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7740, plain, (m1_group_2(X2,X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|X2!=k6_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_7368, plain, (v3_ordinal1(X1)|~r2_hidden(k1_zfmisc_1(X1),k6_numbers))).#trainneg
cnf(i_0_7360, plain, (l3_algstr_0(k7_group_2(k6_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_943, plain, (k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))=k2_tarski(X2,X3)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|esk8_3(X2,X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))!=X3|~m2_finseq_1(esk8_3(X2,X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))),X1))).#trainneg
cnf(i_0_946, plain, (X1=k3_finseq_1(X2)|k4_finseq_1(X2)!=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_953, plain, (m1_group_2(k7_group_2(k6_group_2(k7_group_2(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_7790, plain, (m1_group_2(k6_group_2(X1),X2)|v1_grsolv_1(X2)|v2_struct_0(X2)|v2_struct_0(X1)|k6_group_2(X1)!=k6_group_2(X2)|X3!=k7_group_2(X2)|~v7_ordinal1(np__1)|~m1_subset_1(X3,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k1_group_3(X2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7747, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|np__3=np__2|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_7792, plain, (esk5_2(X1,k10_finseq_1(X2,k6_group_2(X3)))=X2|esk5_2(X1,k10_finseq_1(X2,k6_group_2(X3)))=k6_group_2(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|k6_group_2(X3)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_group_3(X1)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_7794, plain, (esk4_2(X1,k10_finseq_1(X2,k6_group_2(X3)))=X2|esk4_2(X1,k10_finseq_1(X2,k6_group_2(X3)))=k6_group_2(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|k6_group_2(X3)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_group_3(X1)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_7752, plain, (esk6_2(X1,k10_finseq_1(X1,X2))=X2|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_group_2(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1372, plain, (r2_hidden(k2_nat_1(np__2,np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_1239, plain, (r2_hidden(k2_xcmplx_0(esk3_2(X1,X2),np__1),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_7369, plain, (v3_ordinal1(X1)|~r2_hidden(k2_tarski(X2,X1),k6_numbers))).#trainneg
cnf(i_0_7361, plain, (v3_group_1(k7_group_2(k6_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7465, plain, (v1_xboole_0(k4_finseq_1(k6_numbers))|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_963, plain, (m1_group_2(k7_group_2(k7_group_2(k6_group_2(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_7811, plain, (k2_tarski(np__1,np__2)!=k6_numbers)).#trainneg
cnf(i_0_964, plain, (m1_group_2(k7_group_2(k7_group_2(k7_group_2(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_7812, plain, (k3_finseq_1(X1)=k6_numbers|k4_finseq_1(X1)!=k6_numbers|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7959, plain, (v5_group_1(k7_group_2(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6026, plain, (m1_group_2(esk6_2(k6_group_2(k6_group_2(X1)),k10_finseq_1(k6_group_2(k6_group_2(X1)),k6_group_2(k6_group_2(X1)))),X1)|v2_struct_0(X1)|v1_grsolv_1(k6_group_2(k6_group_2(X1)))|v2_struct_0(k6_group_2(k6_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k6_group_2(k6_group_2(k6_group_2(X1)))!=k6_group_2(k6_group_2(X1))|~m2_finseq_1(k10_finseq_1(k6_group_2(k6_group_2(X1)),k6_group_2(k6_group_2(X1))),k1_group_3(k6_group_2(k6_group_2(X1))))|~l3_algstr_0(k6_group_2(k6_group_2(X1)))|~v3_group_1(k6_group_2(k6_group_2(X1)))|~v2_group_1(k6_group_2(k6_group_2(X1)))|~v15_algstr_0(k6_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_6566, plain, (k9_xtuple_0(X1)!=np__2|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3197, plain, (m1_group_6(esk6_2(X1,k10_finseq_1(X2,X3)),X1,X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X2,X3))=X3|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X2,X3)),X1,X3)|~v1_group_3(esk5_2(X1,k10_finseq_1(X2,X3)),X3)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_7869, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k6_group_2(k7_group_2(X1))))|~v3_group_1(X1))).#trainneg
cnf(i_0_7954, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k7_group_2(k6_group_2(X1))))|~v3_group_1(X1))).#trainneg
cnf(i_0_7381, plain, (k4_finseq_1(k6_numbers)=k6_numbers|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_7810, plain, (k3_finseq_1(k6_numbers)=k6_numbers|k2_finseq_1(k3_finseq_1(k6_numbers))!=k6_numbers|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_976, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_978, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k6_numbers,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_981, plain, (g3_algstr_0(k6_numbers,u2_algstr_0(X1))=X1|~l3_algstr_0(X1)|~v15_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_7965, plain, (k4_finseq_1(X1)!=k6_numbers|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8037, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X2,X3)),X3)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X2,X3))=X3|v1_grsolv_1(X1)|~m1_group_2(X3,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X2,X3)),X1,X3)|~v1_group_3(esk5_2(X1,k10_finseq_1(X2,X3)),X3)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1)))).#trainneg
cnf(i_0_8034, plain, (k4_finseq_1(X1)!=np__2|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7964, plain, (k4_finseq_1(np__2)!=k6_numbers|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_7982, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k7_group_2(k7_group_2(X1))))|~v3_group_1(X1))).#trainneg
cnf(i_0_7871, plain, (v2_group_1(k7_group_2(k6_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_987, plain, (m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(k6_numbers,k6_numbers),k6_numbers)))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_996, plain, (v3_card_1(k2_tarski(np__1,np__2),X1)|~v7_ordinal1(X1)|X1!=np__2|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1392, plain, (k2_nat_1(np__2,np__1)=np__1|k2_nat_1(np__2,np__1)=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_2146, plain, (m1_group_2(esk6_2(X1,k10_finseq_1(X1,X1)),esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_6242, plain, (m1_group_2(np__2,X1)|v1_grsolv_1(X2)|v2_struct_0(X2)|v2_struct_0(X1)|esk3_2(X2,k10_finseq_1(X2,X3))=np__1|k1_group_3(k6_group_2(k6_group_2(X1)))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X2)|k7_group_2(X2)!=X2|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X2))|~l3_algstr_0(X1)|~l3_algstr_0(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v2_group_1(X1)|~v2_group_1(X2)|~m1_subset_1(X3,k1_group_3(X2))|~v15_algstr_0(X2))).#trainneg
cnf(i_0_7370, plain, (v3_ordinal1(X1)|~r2_hidden(k2_tarski(X1,X2),k6_numbers))).#trainneg
cnf(i_0_7872, plain, (l3_algstr_0(k7_group_2(k6_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_8056, plain, (k3_finseq_1(k6_numbers)=k6_numbers|k4_finseq_1(k3_finseq_1(k6_numbers))!=k6_numbers|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_7051, plain, (v4_relat_1(k1_finseq_1(X3),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3615, plain, (esk3_2(X1,k10_finseq_1(X1,X2))=np__1|m1_group_2(X3,X4)|v1_grsolv_1(X1)|v2_struct_0(X4)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X2!=k6_group_2(X1)|~m1_group_2(X3,np__2)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_1869, plain, (m1_group_2(esk3_2(X1,X2),X3)|v2_struct_0(X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k1_group_3(k7_group_2(X3))!=k2_tarski(np__1,np__2)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(k7_group_2(X3))|~l3_algstr_0(X1)|~v3_group_1(k7_group_2(X3))|~v3_group_1(X1)|~v2_group_1(k7_group_2(X3))|~v2_group_1(X1))).#trainneg
cnf(i_0_1774, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_6957, plain, (v5_relat_1(k2_zfmisc_1(X3,X4),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1005, plain, (v1_xboole_0(X1)|m2_finseq_1(X2,X1)|k3_finseq_1(X2)!=np__2|~m1_subset_1(k1_funct_1(X2,np__2),X1)|~m1_subset_1(k1_funct_1(X2,np__1),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1006, plain, (v1_xboole_0(X1)|m2_finseq_1(esk2_1(X2),X1)|v2_struct_0(X2)|k7_group_2(X2)!=X3|k3_finseq_1(esk2_1(X2))!=np__2|~m1_subset_1(k1_funct_1(esk2_1(X2),np__2),X1)|~m1_subset_1(X3,X1)|~v1_finseq_1(esk2_1(X2))|~v1_funct_1(esk2_1(X2))|~v1_relat_1(esk2_1(X2))|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_8270, plain, (r1_xxreal_0(X1,esk3_2(X2,k10_finseq_1(X3,X4)))|r1_xxreal_0(esk3_2(X2,k10_finseq_1(X3,X4)),X1)|v1_grsolv_1(X2)|v2_struct_0(X2)|~m1_subset_1(X1,k5_numbers)|X4!=k6_group_2(X2)|X3!=k7_group_2(X2)|~m2_finseq_1(k10_finseq_1(X3,X4),k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_8183, plain, (v1_xboole_0(u2_algstr_0(X1))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_8280, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6958, plain, (v5_relat_1(k2_zfmisc_1(X3,X4),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_7048, plain, (v4_relat_1(k2_zfmisc_1(X3,X4),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1011, plain, (k2_xcmplx_0(X1,np__3)=k2_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1012, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8363, plain, (r1_xxreal_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3)))|r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,X3)),np__1)|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8336, plain, (r1_xxreal_0(esk3_2(X1,X3),X4)|r1_xxreal_0(X4,esk3_2(X1,X3))|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X4,k5_numbers)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_8364, plain, (r1_xxreal_0(np__1,esk3_2(X1,X2))|r1_xxreal_0(esk3_2(X1,X2),np__1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|k3_finseq_1(X2)!=np__2|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8370, plain, (r1_xxreal_0(np__1,k6_group_2(k6_group_2(X1)))|r1_xxreal_0(k6_group_2(k6_group_2(X1)),np__1)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7823, plain, (v3_ordinal1(X1)|~r2_hidden(k1_zfmisc_1(k1_zfmisc_1(X1)),k6_numbers))).#trainneg
cnf(i_0_8359, plain, (r1_xxreal_0(np__1,esk9_1(k5_numbers))|r1_xxreal_0(esk9_1(k5_numbers),np__1))).#trainneg
cnf(i_0_7995, plain, (k3_finseq_1(k6_numbers)=k6_numbers|k9_xtuple_0(k6_numbers)!=k6_numbers|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_7049, plain, (v4_relat_1(k2_zfmisc_1(X3,X4),X2)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1014, plain, (k2_xcmplx_0(X1,np__2)=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_8200, plain, (r2_hidden(esk6_2(X1,k10_finseq_1(X1,X1)),X2)|v2_struct_0(esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k1_group_3(esk4_2(X1,k10_finseq_1(X1,X1)))|~l3_algstr_0(esk4_2(X1,k10_finseq_1(X1,X1)))|~v3_group_1(esk4_2(X1,k10_finseq_1(X1,X1)))|~v2_group_1(esk4_2(X1,k10_finseq_1(X1,X1)))|~v15_algstr_0(esk6_2(X1,k10_finseq_1(X1,X1)))|k6_group_2(X1)!=X1|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_8371, plain, (r1_xxreal_0(np__1,k7_group_2(k6_group_2(X1)))|r1_xxreal_0(k7_group_2(k6_group_2(X1)),np__1)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8337, plain, (r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,X3)),X4)|r1_xxreal_0(X4,esk3_2(X1,k10_finseq_1(X2,X3)))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X4,k5_numbers)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_7960, plain, (r2_hidden(k7_group_2(k7_group_2(k6_group_2(X1))),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k7_group_2(k6_group_2(X1)))))).#trainneg
cnf(i_0_4842, plain, (m1_group_2(np__1,X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|r2_hidden(np__3,k2_tarski(np__1,np__2))|k1_group_3(k6_group_2(X3))!=k2_tarski(np__1,np__2)|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_3042, plain, (esk6_2(k1_funct_1(X1,np__2),X1)=X2|esk6_2(k1_funct_1(X1,np__2),X1)=X3|v1_grsolv_1(k1_funct_1(X1,np__2))|v2_struct_0(k1_funct_1(X1,np__2))|k2_tarski(X2,X3)!=k1_group_3(k1_funct_1(X1,np__2))|k6_group_2(k1_funct_1(X1,np__2))!=k1_funct_1(X1,np__2)|~m2_finseq_1(X1,k1_group_3(k1_funct_1(X1,np__2)))|~l3_algstr_0(k1_funct_1(X1,np__2))|~v3_group_1(k1_funct_1(X1,np__2))|~v2_group_1(k1_funct_1(X1,np__2))|~v15_algstr_0(esk6_2(k1_funct_1(X1,np__2),X1))|~v15_algstr_0(k1_funct_1(X1,np__2))|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__1)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6248, plain, (m1_group_2(np__1,X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|esk4_2(X1,k10_finseq_1(X2,X3))=X3|k1_group_3(k6_group_2(k6_group_2(X4)))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_7363, plain, (r2_hidden(k7_group_2(k6_group_2(k6_group_2(X1))),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k6_group_2(k6_group_2(X1)))))).#trainneg
cnf(i_0_6245, plain, (m1_group_2(np__1,X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|r2_hidden(np__3,k2_tarski(np__1,np__2))|k1_group_3(k6_group_2(k6_group_2(X3)))!=k2_tarski(np__1,np__2)|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_6089, plain, (m1_group_2(k6_group_2(k7_group_2(esk3_2(X1,X2))),X3)|v1_grsolv_1(X1)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(esk3_2(X1,X2))|k2_tarski(np__1,np__2)!=k1_group_3(X3)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X4|k3_finseq_1(X2)!=np__2|X4!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~l3_algstr_0(esk3_2(X1,X2))|~v3_group_1(esk3_2(X1,X2))|~v2_group_1(esk3_2(X1,X2)))).#trainneg
cnf(i_0_7874, plain, (v5_group_1(k7_group_2(k6_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1887, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2643, plain, (r2_hidden(k2_nat_1(np__1,np__1),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_7873, plain, (v3_group_1(k7_group_2(k6_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1037, plain, (m1_group_2(esk4_2(X1,k6_numbers),X1)|r1_xxreal_0(k3_finseq_1(k6_numbers),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(k6_numbers,k3_finseq_1(k6_numbers))!=k6_group_2(X1)|k7_group_2(X1)!=k1_funct_1(k6_numbers,np__1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1038, plain, (m1_group_2(esk5_2(X1,k6_numbers),X1)|r1_xxreal_0(k3_finseq_1(k6_numbers),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(k6_numbers,k3_finseq_1(k6_numbers))!=k6_group_2(X1)|k7_group_2(X1)!=k1_funct_1(k6_numbers,np__1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8549, plain, (r1_xxreal_0(k6_group_2(X1),np__1)|r1_xxreal_0(np__1,k6_group_2(X1))|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1))|~v15_algstr_0(k6_group_2(X1)))).#trainneg
cnf(i_0_8634, plain, (esk4_2(X3,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X3)|v2_struct_0(X3)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_8518, plain, (r1_xxreal_0(X1,esk3_2(X2,X3))|r1_xxreal_0(esk3_2(X2,X3),X1)|v1_grsolv_1(X2)|v2_struct_0(X2)|k1_funct_1(X3,np__2)!=k6_group_2(X2)|k3_finseq_1(X3)!=np__2|k1_funct_1(X3,np__1)!=k7_group_2(X2)|~m1_subset_1(X1,k5_numbers)|~m2_finseq_1(X3,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_7956, plain, (v2_group_1(k7_group_2(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_4023, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(np__1,np__3))|r1_xxreal_0(k2_xcmplx_0(np__1,np__3),k6_numbers)|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1039, plain, (r1_xxreal_0(k3_finseq_1(k6_numbers),k6_numbers)|v1_grsolv_1(X1)|v15_algstr_0(esk4_2(X1,k6_numbers))|v2_struct_0(X1)|k1_funct_1(k6_numbers,k3_finseq_1(k6_numbers))!=k6_group_2(X1)|k7_group_2(X1)!=k1_funct_1(k6_numbers,np__1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8353, plain, (X1=u2_algstr_0(X2)|~v1_xboole_0(X1)|~l3_algstr_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1040, plain, (r1_xxreal_0(k3_finseq_1(k6_numbers),k6_numbers)|v1_grsolv_1(X1)|v15_algstr_0(esk5_2(X1,k6_numbers))|v2_struct_0(X1)|k1_funct_1(k6_numbers,k3_finseq_1(k6_numbers))!=k6_group_2(X1)|k7_group_2(X1)!=k1_funct_1(k6_numbers,np__1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1048, plain, (k2_xcmplx_0(k6_numbers,X1)=k2_xcmplx_0(X1,k6_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8636, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=X2|v1_grsolv_1(X3)|v2_struct_0(X3)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_8554, plain, (r1_xxreal_0(X1,esk3_2(X2,k10_finseq_1(X3,X2)))|r1_xxreal_0(esk3_2(X2,k10_finseq_1(X3,X2)),X1)|v1_grsolv_1(X2)|v2_struct_0(X2)|k6_group_2(X2)!=X2|X3!=k7_group_2(X2)|~m1_subset_1(X3,k1_group_3(X2))|~m1_subset_1(X1,k5_numbers)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2))).#trainneg
cnf(i_0_8507, plain, (r1_xxreal_0(esk3_2(X1,X3),np__1)|r1_xxreal_0(np__1,esk3_2(X1,X3))|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_7957, plain, (l3_algstr_0(k7_group_2(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_1081, plain, (m1_subset_1(esk3_2(X1,k3_finseq_1(esk17_1(X2))),k5_numbers)|r1_xxreal_0(k3_finseq_1(esk17_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(X2)|k1_funct_1(k3_finseq_1(esk17_1(X2)),k3_finseq_1(esk17_1(X2)))!=k6_group_2(X1)|k1_funct_1(k3_finseq_1(esk17_1(X2)),np__1)!=k7_group_2(X1)|~m2_finseq_1(k3_finseq_1(esk17_1(X2)),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1082, plain, (m1_group_2(esk3_2(X1,k3_finseq_1(esk17_1(X2))),X3)|r1_xxreal_0(k3_finseq_1(esk17_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|v1_xboole_0(X2)|k1_funct_1(k3_finseq_1(esk17_1(X2)),k3_finseq_1(esk17_1(X2)))!=k6_group_2(X1)|k1_funct_1(k3_finseq_1(esk17_1(X2)),np__1)!=k7_group_2(X1)|k4_finseq_1(k3_finseq_1(esk17_1(X2)))!=k1_group_3(X3)|~m2_finseq_1(k3_finseq_1(esk17_1(X2)),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1))).#trainneg
cnf(i_0_7748, plain, (np__3=np__1|np__3=np__2|m1_group_2(esk6_2(X1,k10_finseq_1(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__2)|~m1_group_6(X2,X1,X1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_831, plain, (m1_group_2(k7_group_2(esk5_2(X1,k10_finseq_1(X2,X3))),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(esk5_2(X1,k10_finseq_1(X2,X3)))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(esk5_2(X1,k10_finseq_1(X2,X3)))|~v3_group_1(esk5_2(X1,k10_finseq_1(X2,X3)))|~v2_group_1(esk5_2(X1,k10_finseq_1(X2,X3))))).#trainneg
cnf(i_0_8819, plain, (r1_xxreal_0(esk3_2(X1,k10_finseq_1(X1,X1)),X2)|r1_xxreal_0(X2,esk3_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k7_group_2(X1)!=X1|~m1_subset_1(X2,k5_numbers)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_8758, plain, (u2_algstr_0(X1)=k6_numbers|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_7371, plain, (v3_ordinal1(X1)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~r2_hidden(k1_group_3(X1),k6_numbers))).#trainneg
cnf(i_0_7958, plain, (v3_group_1(k7_group_2(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4256, plain, (r1_xxreal_0(np__2,k2_xcmplx_0(np__1,np__3))|r1_xxreal_0(k2_xcmplx_0(np__1,np__3),np__2)|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1084, plain, (m2_subset_1(k3_finseq_1(esk17_1(X1)),k1_numbers,k5_numbers)|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_1(esk17_1(X1)))|~v1_funct_1(k3_finseq_1(esk17_1(X1)))|~v1_relat_1(k3_finseq_1(esk17_1(X1))))).#trainneg
cnf(i_0_5401, plain, (m1_subset_1(k9_xtuple_0(k6_numbers),k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_1083, plain, (m1_group_2(k2_nat_1(esk3_2(X1,k3_finseq_1(esk17_1(X2))),np__1),X3)|r1_xxreal_0(k3_finseq_1(esk17_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|v1_xboole_0(X2)|k1_funct_1(k3_finseq_1(esk17_1(X2)),k3_finseq_1(esk17_1(X2)))!=k6_group_2(X1)|k1_funct_1(k3_finseq_1(esk17_1(X2)),np__1)!=k7_group_2(X1)|k4_finseq_1(k3_finseq_1(esk17_1(X2)))!=k1_group_3(X3)|~m2_finseq_1(k3_finseq_1(esk17_1(X2)),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1))).#trainneg
cnf(i_0_1085, plain, (k9_xtuple_0(k3_finseq_1(esk17_1(X1)))=k2_finseq_1(X2)|v1_xboole_0(X1)|X2!=k3_finseq_1(esk17_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_1(esk17_1(X1)))|~v1_funct_1(k3_finseq_1(esk17_1(X1)))|~v1_relat_1(k3_finseq_1(esk17_1(X1))))).#trainneg
cnf(i_0_4857, plain, (m1_group_2(k2_nat_1(np__1,esk3_2(X1,k10_finseq_1(X2,X3))),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X4)|~v7_ordinal1(esk3_2(X1,k10_finseq_1(X2,X3)))|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1))).#trainneg
cnf(i_0_5613, plain, (esk5_2(X3,k10_finseq_1(X1,X2))=k1_funct_1(k10_finseq_1(X1,X2),np__3)|v1_grsolv_1(X3)|v2_struct_0(X3)|esk3_2(X3,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X3)|X1!=k7_group_2(X3)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,k1_group_3(X3))|~m1_subset_1(X1,k1_group_3(X3)))).#trainneg
cnf(i_0_8677, plain, (esk3_2(X1,X3)=np__1|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k1_funct_1(X3,np__2),k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4216, plain, (r2_hidden(X2,X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X3!=k1_group_3(X1)|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_subset_1(X2,k1_group_3(X1))|~m1_subset_1(X1,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X2)|~v7_ordinal1(np__1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1674, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8899, plain, (r1_xxreal_0(np__1,esk3_2(X1,k10_finseq_1(X1,X1)))|r1_xxreal_0(esk3_2(X1,k10_finseq_1(X1,X1)),np__1)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_8910, plain, (v15_algstr_0(g3_algstr_0(k6_numbers,k6_numbers))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_8928, plain, (l3_algstr_0(g3_algstr_0(u1_struct_0(X1),k6_numbers))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1086, plain, (k1_card_1(k3_finseq_1(esk17_1(X1)))=k3_finseq_1(esk17_1(X1))|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_1(esk17_1(X1)))|~v1_funct_1(k3_finseq_1(esk17_1(X1)))|~v1_relat_1(k3_finseq_1(esk17_1(X1))))).#trainneg
cnf(i_0_1087, plain, (v1_card_1(k3_finseq_1(esk17_1(X1)))|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_1(esk17_1(X1)))|~v1_funct_1(k3_finseq_1(esk17_1(X1)))|~v1_relat_1(k3_finseq_1(esk17_1(X1))))).#trainneg
cnf(i_0_1088, plain, (k9_xtuple_0(k3_finseq_1(esk17_1(X1)))=k2_finseq_1(k3_finseq_1(esk17_1(X1)))|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_1(esk17_1(X1)))|~v1_funct_1(k3_finseq_1(esk17_1(X1)))|~v1_relat_1(k3_finseq_1(esk17_1(X1))))).#trainneg
cnf(i_0_8896, plain, (r1_xxreal_0(k6_group_2(k6_group_2(X1)),esk3_2(X2,k10_finseq_1(X2,X2)))|r1_xxreal_0(esk3_2(X2,k10_finseq_1(X2,X2)),k6_group_2(k6_group_2(X1)))|v1_grsolv_1(X2)|v2_struct_0(X2)|v2_struct_0(X1)|k6_group_2(X2)!=X2|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8889, plain, (r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,X3)),esk3_2(X4,k10_finseq_1(X4,X4)))|r1_xxreal_0(esk3_2(X4,k10_finseq_1(X4,X4)),esk3_2(X1,k10_finseq_1(X2,X3)))|v1_grsolv_1(X4)|v2_struct_0(X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X4)!=X4|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~v15_algstr_0(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8900, plain, (r1_xxreal_0(np__2,esk3_2(X1,k10_finseq_1(X1,X1)))|r1_xxreal_0(esk3_2(X1,k10_finseq_1(X1,X1)),np__2)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_8897, plain, (r1_xxreal_0(k7_group_2(k6_group_2(X1)),esk3_2(X2,k10_finseq_1(X2,X2)))|r1_xxreal_0(esk3_2(X2,k10_finseq_1(X2,X2)),k7_group_2(k6_group_2(X1)))|v1_grsolv_1(X2)|v2_struct_0(X2)|v2_struct_0(X1)|k6_group_2(X2)!=X2|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8898, plain, (r1_xxreal_0(k6_numbers,esk3_2(X1,k10_finseq_1(X1,X1)))|r1_xxreal_0(esk3_2(X1,k10_finseq_1(X1,X1)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_8911, plain, (l3_algstr_0(g3_algstr_0(k6_numbers,k6_numbers))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_8929, plain, (v15_algstr_0(g3_algstr_0(u1_struct_0(X1),k6_numbers))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_9030, plain, (m2_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers,k5_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_9031, plain, (m2_subset_1(k2_xcmplx_0(np__2,np__3),k1_numbers,k5_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1113, plain, (m1_group_2(esk11_1(k7_group_2(k6_group_2(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(k6_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(k6_group_2(X1)))|~v2_group_1(k7_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_1116, plain, (m1_group_2(k6_group_2(k7_group_2(k6_group_2(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(k6_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(k6_group_2(X1)))|~v3_group_1(k7_group_2(k6_group_2(X1)))|~v2_group_1(k7_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_9225, plain, (v5_group_1(k6_group_2(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9129, plain, (r1_xxreal_0(esk3_2(X1,k10_finseq_1(X1,X1)),esk3_2(X2,k10_finseq_1(X3,X4)))|r1_xxreal_0(esk3_2(X2,k10_finseq_1(X3,X4)),esk3_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X2)|v1_grsolv_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X2))|k6_group_2(X1)!=X1|X4!=k6_group_2(X2)|X3!=k7_group_2(X2)|~l3_algstr_0(X1)|~l3_algstr_0(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v2_group_1(X1)|~v2_group_1(X2)|~v15_algstr_0(X1)|~m1_subset_1(X4,k1_group_3(X2))|~m1_subset_1(X3,k1_group_3(X2)))).#trainneg
cnf(i_0_9252, plain, (m1_group_6(X2,X1,X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,X1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v1_group_3(X2,X1))).#trainneg
cnf(i_0_9230, plain, (v1_group_3(X2,esk4_2(X1,k10_finseq_1(X1,X2)))|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),esk4_2(X1,k10_finseq_1(X1,X2)))|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),esk4_2(X1,k10_finseq_1(X1,X2)))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9201, plain, (v5_group_1(esk11_1(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9231, plain, (v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,esk4_2(X1,k10_finseq_1(X1,X2)))|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),esk4_2(X1,k10_finseq_1(X1,X2)))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X2)),X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9229, plain, (m1_group_6(X2,X1,esk4_2(X1,k10_finseq_1(X1,X2)))|v1_grsolv_1(X1)|v2_struct_0(X1)|r2_hidden(np__3,k2_tarski(np__1,np__2))|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_group_6(esk5_2(X1,k10_finseq_1(X1,X2)),X1,esk4_2(X1,k10_finseq_1(X1,X2)))|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),esk4_2(X1,k10_finseq_1(X1,X2)))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9196, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(esk11_1(k7_group_2(k6_group_2(X1))))|~v3_group_1(X1))).#trainneg
cnf(i_0_9220, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k6_group_2(k7_group_2(k6_group_2(X1))))|~v3_group_1(X1))).#trainneg
cnf(i_0_9168, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__1,np__3),k5_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1117, plain, (m1_group_2(k7_group_2(k6_group_2(k7_group_2(k6_group_2(X1)))),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(k6_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(k6_group_2(X1)))|~v3_group_1(k7_group_2(k6_group_2(X1)))|~v2_group_1(k7_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_8913, plain, (v1_funct_2(k6_numbers,k6_numbers,k6_numbers)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1118, plain, (m1_group_2(k7_group_2(k7_group_2(k7_group_2(k6_group_2(X1)))),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(k6_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(k6_group_2(X1)))|~v3_group_1(k7_group_2(k6_group_2(X1)))|~v2_group_1(k7_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_9221, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k6_group_2(k7_group_2(k6_group_2(X2))))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_9270, plain, (m1_group_6(X1,X2,X2)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X2)|v2_struct_0(X2)|X1!=k6_group_2(X2)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X2,k10_finseq_1(X2,X1)),X2)|~v1_group_3(X1,X2)|~m1_subset_1(X1,k1_group_3(X2))|~m2_finseq_1(k10_finseq_1(X2,X1),k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2)|~m1_group_2(X2,X2)|~m1_group_2(esk5_2(X2,k10_finseq_1(X2,X1)),X2))).#trainneg
cnf(i_0_9302, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X1,k10_finseq_1(X1,X2)),esk4_2(X1,k10_finseq_1(X1,X2)))|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X2)),X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_group_2(esk5_2(X1,k10_finseq_1(X1,X2)),esk4_2(X1,k10_finseq_1(X1,X2))))).#trainneg
cnf(i_0_9272, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(X2,X1,X1)|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(X1,X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9314, plain, (m1_group_6(X1,X2,esk4_2(X2,k10_finseq_1(X2,X1)))|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X2)|v2_struct_0(X2)|X1!=k6_group_2(X2)|~v7_ordinal1(np__1)|~v1_group_3(esk5_2(X2,k10_finseq_1(X2,X1)),esk4_2(X2,k10_finseq_1(X2,X1)))|~v1_group_3(X1,X2)|~m1_subset_1(X1,k1_group_3(X2))|~m2_finseq_1(k10_finseq_1(X2,X1),k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2)|~m1_group_2(esk4_2(X2,k10_finseq_1(X2,X1)),X2)|~m1_group_2(esk5_2(X2,k10_finseq_1(X2,X1)),esk4_2(X2,k10_finseq_1(X2,X1))))).#trainneg
cnf(i_0_9264, plain, (r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,X1)),esk3_2(X3,k10_finseq_1(X3,X3)))|r1_xxreal_0(esk3_2(X3,k10_finseq_1(X3,X3)),esk3_2(X1,k10_finseq_1(X2,X1)))|v1_grsolv_1(X3)|v1_grsolv_1(X1)|v2_struct_0(X3)|v2_struct_0(X1)|k6_group_2(X3)!=X3|k6_group_2(X1)!=X1|X2!=k7_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v15_algstr_0(X3)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9344, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k6_group_2(k7_group_2(k6_group_2(X1)))))|~v3_group_1(X1))).#trainneg
cnf(i_0_9369, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(k7_group_2(k7_group_2(k6_group_2(X1)))))|~v3_group_1(X1))).#trainneg
cnf(i_0_9320, plain, (v3_ordinal1(k2_xcmplx_0(np__1,np__3))|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_8536, plain, (k3_finseq_1(k6_numbers)=k6_numbers|k4_finseq_1(k6_numbers)!=k6_numbers|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_1140, plain, (m1_subset_1(esk3_2(X1,k3_finseq_1(esk16_1(X2))),k5_numbers)|r1_xxreal_0(k3_finseq_1(esk16_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(k3_finseq_1(esk16_1(X2)),k3_finseq_1(esk16_1(X2)))!=k6_group_2(X1)|k1_funct_1(k3_finseq_1(esk16_1(X2)),np__1)!=k7_group_2(X1)|~m2_finseq_1(k3_finseq_1(esk16_1(X2)),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1141, plain, (m1_group_2(esk3_2(X1,k3_finseq_1(esk16_1(X2))),X3)|r1_xxreal_0(k3_finseq_1(esk16_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|k1_funct_1(k3_finseq_1(esk16_1(X2)),k3_finseq_1(esk16_1(X2)))!=k6_group_2(X1)|k1_funct_1(k3_finseq_1(esk16_1(X2)),np__1)!=k7_group_2(X1)|k4_finseq_1(k3_finseq_1(esk16_1(X2)))!=k1_group_3(X3)|~m2_finseq_1(k3_finseq_1(esk16_1(X2)),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9262, plain, (r1_xxreal_0(esk3_2(X1,k10_finseq_1(X2,k6_group_2(X3))),esk3_2(X4,k10_finseq_1(X4,X4)))|r1_xxreal_0(esk3_2(X4,k10_finseq_1(X4,X4)),esk3_2(X1,k10_finseq_1(X2,k6_group_2(X3))))|v1_grsolv_1(X4)|v1_grsolv_1(X1)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|k6_group_2(X4)!=X4|k6_group_2(X3)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~v15_algstr_0(X4)|k1_group_3(X1)!=k1_group_3(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_9023, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9298, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~m1_group_6(X2,X1,esk4_2(X1,k10_finseq_1(X1,X2)))|~v1_group_3(X2,esk4_2(X1,k10_finseq_1(X1,X2)))|~v1_group_3(X2,X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X2)),X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9226, plain, (r2_hidden(k6_group_2(k7_group_2(k6_group_2(X1))),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k7_group_2(k6_group_2(X1)))))).#trainneg
cnf(i_0_7984, plain, (v2_group_1(k7_group_2(k7_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_9169, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__1,np__3),k1_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1143, plain, (m2_subset_1(k3_finseq_1(esk16_1(X1)),k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_1(esk16_1(X1)))|~v1_funct_1(k3_finseq_1(esk16_1(X1)))|~v1_relat_1(k3_finseq_1(esk16_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5394, plain, (k2_zfmisc_1(X2,X3)=k6_numbers|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1142, plain, (m1_group_2(k2_nat_1(esk3_2(X1,k3_finseq_1(esk16_1(X2))),np__1),X3)|r1_xxreal_0(k3_finseq_1(esk16_1(X2)),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X3)|k1_funct_1(k3_finseq_1(esk16_1(X2)),k3_finseq_1(esk16_1(X2)))!=k6_group_2(X1)|k1_funct_1(k3_finseq_1(esk16_1(X2)),np__1)!=k7_group_2(X1)|k4_finseq_1(k3_finseq_1(esk16_1(X2)))!=k1_group_3(X3)|~m2_finseq_1(k3_finseq_1(esk16_1(X2)),k1_group_3(X1))|~l3_algstr_0(X3)|~l3_algstr_0(X1)|~v3_group_1(X3)|~v3_group_1(X1)|~v2_group_1(X3)|~v2_group_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1144, plain, (k9_xtuple_0(k3_finseq_1(esk16_1(X1)))=k2_finseq_1(X2)|X2!=k3_finseq_1(esk16_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_1(esk16_1(X1)))|~v1_funct_1(k3_finseq_1(esk16_1(X1)))|~v1_relat_1(k3_finseq_1(esk16_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9466, plain, (r1_xxreal_0(esk3_2(X1,k10_finseq_1(X1,X1)),esk3_2(X2,k10_finseq_1(X2,X2)))|r1_xxreal_0(esk3_2(X2,k10_finseq_1(X2,X2)),esk3_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X2)|v1_grsolv_1(X1)|v2_struct_0(X2)|v2_struct_0(X1)|k6_group_2(X1)!=X1|k6_group_2(X2)!=X2|k7_group_2(X2)!=X2|~l3_algstr_0(X1)|~l3_algstr_0(X2)|~v3_group_1(X1)|~v3_group_1(X2)|~v2_group_1(X1)|~v2_group_1(X2)|~v15_algstr_0(X1)|~v15_algstr_0(X2))).#trainneg
cnf(i_0_9445, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X2!=k6_group_2(X1)|~v7_ordinal1(np__1)|~v1_group_3(X2,esk4_2(X1,k10_finseq_1(X1,X2)))|~v1_group_3(X2,X1)|~m1_group_2(X2,esk4_2(X1,k10_finseq_1(X1,X2)))|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X2)),X1)|~m1_subset_1(X2,k1_group_3(X1))|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v5_group_1(k5_group_6(esk4_2(X1,k10_finseq_1(X1,X2)),X2))|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9434, plain, (m1_group_6(X1,X2,X2)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X2)|v2_struct_0(X2)|X1!=k6_group_2(X2)|~v7_ordinal1(np__1)|~v1_group_3(X1,X2)|~m1_group_2(X1,X2)|~m1_subset_1(X1,k1_group_3(X2))|~m2_finseq_1(k10_finseq_1(X2,X1),k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2))).#trainneg
cnf(i_0_1145, plain, (k1_card_1(k3_finseq_1(esk16_1(X1)))=k3_finseq_1(esk16_1(X1))|~v1_finseq_1(k3_finseq_1(esk16_1(X1)))|~v1_funct_1(k3_finseq_1(esk16_1(X1)))|~v1_relat_1(k3_finseq_1(esk16_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8959, plain, (m1_subset_1(k4_finseq_1(k6_numbers),k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_1146, plain, (v1_card_1(k3_finseq_1(esk16_1(X1)))|~v1_finseq_1(k3_finseq_1(esk16_1(X1)))|~v1_funct_1(k3_finseq_1(esk16_1(X1)))|~v1_relat_1(k3_finseq_1(esk16_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1147, plain, (k9_xtuple_0(k3_finseq_1(esk16_1(X1)))=k2_finseq_1(k3_finseq_1(esk16_1(X1)))|~v1_finseq_1(k3_finseq_1(esk16_1(X1)))|~v1_funct_1(k3_finseq_1(esk16_1(X1)))|~v1_relat_1(k3_finseq_1(esk16_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6903, plain, (m1_group_2(k6_group_2(esk4_2(X1,X2)),X1)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(esk4_2(X1,X2))|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X3|k3_finseq_1(X2)!=np__2|X3!=k7_group_2(X1)|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~m1_subset_1(X3,k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(esk4_2(X1,X2))|~v3_group_1(esk4_2(X1,X2))|~v2_group_1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_6700, plain, (r1_xxreal_0(np__2,esk3_2(X1,np__2))|r1_xxreal_0(esk3_2(X1,np__2),np__2)|v1_grsolv_1(X1)|v2_struct_0(X1)|k6_group_2(X1)!=k1_funct_1(np__2,np__2)|k7_group_2(X1)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4639, plain, (m1_group_2(k6_group_2(k7_group_2(k7_group_2(X1))),X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(k7_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(k7_group_2(X1)))|~v3_group_1(k7_group_2(k7_group_2(X1)))|~v2_group_1(k7_group_2(k7_group_2(X1))))).#trainneg
cnf(i_0_1220, plain, (r2_hidden(k2_nat_1(esk3_2(X1,k10_finseq_1(X2,X3)),np__1),k4_finseq_1(X4))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X4)!=np__2|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_1571, plain, (k7_group_2(k6_group_2(X1))=X2|k7_group_2(k6_group_2(X1))=X3|v2_struct_0(X1)|k2_tarski(X3,X2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_6106, plain, (v1_grsolv_1(X1)|v5_group_1(esk3_2(X1,k10_finseq_1(X2,X3)))|v2_struct_0(X1)|v2_struct_0(k6_group_2(k6_group_2(X4)))|v2_struct_0(X4)|k1_group_3(k6_group_2(k6_group_2(X4)))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(k6_group_2(k6_group_2(X4)))|~l3_algstr_0(X1)|~v3_group_1(k6_group_2(k6_group_2(X4)))|~v3_group_1(X1)|~v2_group_1(k6_group_2(k6_group_2(X4)))|~v2_group_1(X1)|~l3_algstr_0(X4)|~v5_group_1(X4)|~v3_group_1(X4)|~v2_group_1(X4))).#trainneg
cnf(i_0_5820, plain, (k2_xcmplx_0(np__1,esk3_2(X1,X3))=np__2|k2_nat_1(np__1,esk3_2(X1,X3))=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X3,np__2)!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X3))|~m2_finseq_1(X3,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_funct_1(X3,np__1)!=X2|k3_finseq_1(X3)!=np__2|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_7985, plain, (l3_algstr_0(k7_group_2(k7_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_9174, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__2,np__3),k5_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1149, plain, (m1_subset_1(esk8_3(X1,X2,k5_numbers),k1_numbers)|esk8_3(X1,X2,k5_numbers)=X1|esk8_3(X1,X2,k5_numbers)=X2|k2_tarski(X1,X2)=k5_numbers)).#trainneg
cnf(i_0_8925, plain, (v1_funct_2(k6_numbers,k2_zfmisc_1(k6_numbers,k6_numbers),k6_numbers)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_5503, plain, (v1_relat_1(k9_xtuple_0(u1_struct_0(X4)))|~v1_xboole_0(X3)|~l3_algstr_0(X4)|~v2_struct_0(X4))).#trainneg
cnf(i_0_1150, plain, (m1_subset_1(esk7_2(X1,k5_numbers),k1_numbers)|k1_group_3(X1)=k5_numbers|v15_algstr_0(esk7_2(X1,k5_numbers))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1154, plain, (k2_xcmplx_0(np__2,k2_xcmplx_0(k6_numbers,X1))=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_9654, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|~m1_group_2(X1,k6_group_2(k7_group_2(k7_group_2(X2))))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_9658, plain, (v5_group_1(k6_group_2(k7_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9706, plain, (k7_group_2(k6_group_2(X1))=X2|k7_group_2(k6_group_2(X1))=X3|v2_struct_0(X1)|v2_struct_0(k6_group_2(X1))|k2_tarski(X2,X3)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k6_group_2(X1))|~v3_group_1(k6_group_2(X1))|~v2_group_1(k6_group_2(X1)))).#trainneg
cnf(i_0_9592, plain, (m1_group_6(X1,X2,X2)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X2)|v2_struct_0(X2)|v1_xboole_0(k1_group_3(X2))|X1!=k6_group_2(X2)|~v7_ordinal1(np__1)|~v1_group_3(X1,X2)|~m1_subset_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2)|~m1_subset_1(X2,k1_group_3(X2)))).#trainneg
cnf(i_0_9659, plain, (r2_hidden(k6_group_2(k7_group_2(k7_group_2(X1))),X2)|v2_struct_0(X1)|X2!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k7_group_2(k7_group_2(X1)))))).#trainneg
cnf(i_0_9653, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k6_group_2(k7_group_2(k7_group_2(X1))))|~v3_group_1(X1))).#trainneg
cnf(i_0_9724, plain, (v3_ordinal1(k2_xcmplx_0(np__2,np__3))|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_8349, plain, (u2_algstr_0(X1)!=np__2|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1156, plain, (k2_xcmplx_0(np__3,k2_xcmplx_0(k6_numbers,X1))=k2_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_9746, plain, (v1_relat_1(k9_xtuple_0(u1_struct_0(X1)))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~l3_algstr_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1164, plain, (k9_xtuple_0(np__2)=k2_finseq_1(X1)|X1!=np__2|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_9847, plain, (m1_group_6(X1,X2,X2)|r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(X2)|v2_struct_0(X2)|X1!=k6_group_2(X2)|~v1_group_3(X1,X2)|~m1_subset_1(X1,k1_group_3(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X2))).#trainneg
cnf(i_0_9753, plain, (v1_relat_1(k9_xtuple_0(u1_struct_0(X1)))|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_9175, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(np__2,np__3),k1_numbers)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1179, plain, (v1_group_3(X1,X2)|v2_struct_0(X3)|X1!=k1_funct_1(esk2_1(X3),k2_xcmplx_0(np__1,X4))|X2!=k1_funct_1(esk2_1(X3),X4)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(k2_xcmplx_0(np__1,X4),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_9842, plain, (k7_group_2(k6_group_2(X1))=np__2|k7_group_2(k6_group_2(X1))=np__1|v2_struct_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_9849, plain, (m1_subset_1(k6_group_2(k7_group_2(k7_group_2(X1))),k5_numbers)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k7_group_2(k7_group_2(X1)))))).#trainneg
cnf(i_0_6208, plain, (k6_group_2(k7_group_2(X1))=X2|k6_group_2(k7_group_2(X1))=X3|v2_struct_0(X1)|k2_tarski(X3,X2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k6_group_2(k7_group_2(X1))))).#trainneg
cnf(i_0_7493, plain, (v1_finset_1(k2_finseq_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7986, plain, (v3_group_1(k7_group_2(k7_group_2(k7_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1180, plain, (m1_group_6(X1,X2,X3)|v2_struct_0(X2)|X1!=k1_funct_1(esk2_1(X2),k2_xcmplx_0(np__1,X4))|X3!=k1_funct_1(esk2_1(X2),X4)|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~r2_hidden(k2_xcmplx_0(np__1,X4),k4_finseq_1(esk2_1(X2)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X2)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X3)|~v15_algstr_0(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_9955, plain, (v1_relat_1(k9_xtuple_0(k2_finseq_1(X2)))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9956, plain, (v1_relat_1(k9_xtuple_0(k1_finseq_1(X2)))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1181, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|X2!=k1_funct_1(esk2_1(X3),k2_xcmplx_0(np__1,X4))|X1!=k1_funct_1(esk2_1(X3),X4)|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X1,X3)|~r2_hidden(k2_xcmplx_0(np__1,X4),k4_finseq_1(esk2_1(X3)))|~r2_hidden(X4,k4_finseq_1(esk2_1(X3)))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_5309, plain, (X1=np__1|X1=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|k4_finseq_1(X2)!=k1_group_3(X1)|k3_finseq_1(X2)!=np__2|k6_group_2(X1)!=X1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~m1_subset_1(X1,k1_group_3(X1)))).#trainneg
cnf(i_0_1067, plain, (m1_group_2(esk3_2(X1,k10_finseq_1(X2,X3)),X4)|v2_struct_0(X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k6_group_2(X4))|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|k1_group_3(k6_group_2(X4))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(k6_group_2(X4))|~l3_algstr_0(X1)|~v3_group_1(k6_group_2(X4))|~v3_group_1(X1)|~v2_group_1(k6_group_2(X4))|~v2_group_1(X1))).#trainneg
cnf(i_0_10216, plain, (X1=np__2|X1=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_tarski(np__1,np__2)!=k1_group_3(X1)|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_10137, plain, (k6_group_2(k7_group_2(X1))=X2|k6_group_2(k7_group_2(X1))=X3|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|k2_tarski(X2,X3)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_7245, plain, (esk4_2(X1,k10_finseq_1(X1,X2))=X1|v1_grsolv_1(X1)|v2_struct_0(X1)|k2_nat_1(np__2,np__1)=np__1|esk5_2(X1,k10_finseq_1(X1,X2))=X2|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1244, plain, (r2_hidden(k2_xcmplx_0(np__1,esk3_2(X1,X2)),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~v7_ordinal1(esk3_2(X1,X2))|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4257, plain, (r1_xxreal_0(np__2,k2_xcmplx_0(np__1,np__3))|r1_xxreal_0(k2_xcmplx_0(np__1,np__3),np__2)|v1_xboole_0(k1_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1188, plain, (k2_xcmplx_0(np__3,X1)=k2_xcmplx_0(X1,np__3)|~v7_ordinal1(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_10139, plain, (v1_finset_1(k4_finseq_1(k3_finseq_1(k6_numbers)))|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_1189, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7824, plain, (v3_ordinal1(X1)|~r2_hidden(k1_zfmisc_1(k2_tarski(X2,X1)),k6_numbers))).#trainneg
cnf(i_0_9198, plain, (v2_group_1(esk11_1(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_10279, plain, (v1_finset_1(k9_xtuple_0(k6_numbers))|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers))|~v1_funct_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_1191, plain, (k2_xcmplx_0(np__2,X1)=k2_xcmplx_0(X1,np__2)|~v7_ordinal1(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_5000, plain, (v2_xxreal_0(np__3)|r1_xxreal_0(np__3,esk9_1(k5_numbers))|~v1_xreal_0(np__3)|~v1_xreal_0(esk9_1(k5_numbers))|~v2_xxreal_0(esk9_1(k5_numbers)))).#trainneg
cnf(i_0_1216, plain, (v3_card_1(k4_finseq_1(X1),np__2)|~v7_ordinal1(np__2)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1226, plain, (v1_group_3(X1,X2)|v2_struct_0(X3)|X1!=k1_funct_1(esk2_1(X3),k2_nat_1(X4,np__1))|X2!=k1_funct_1(esk2_1(X3),X4)|~m1_group_2(X2,X3)|~m1_group_2(X1,X3)|~r2_hidden(k2_nat_1(X4,np__1),k2_tarski(np__1,np__2))|~r2_hidden(X4,k2_tarski(np__1,np__2))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|k3_finseq_1(esk2_1(X3))!=np__2|~v1_finseq_1(esk2_1(X3))|~v1_funct_1(esk2_1(X3))|~v1_relat_1(esk2_1(X3)))).#trainneg
cnf(i_0_10258, plain, (X1=np__1|X1=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|k6_group_2(X1)!=X1|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_10261, plain, (k6_group_2(k7_group_2(X1))=np__2|k6_group_2(k7_group_2(X1))=np__1|v2_struct_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_10436, plain, (r2_hidden(np__3,k2_tarski(np__1,np__2))|v1_grsolv_1(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(k5_group_6(X1,k6_group_2(X1)))|v2_struct_0(X1)|~v7_ordinal1(np__1)|~l3_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~v3_group_1(k5_group_6(X1,k6_group_2(X1)))|~v2_group_1(k5_group_6(X1,k6_group_2(X1)))|~v15_algstr_0(k5_group_6(X1,k6_group_2(X1)))|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_10460, plain, (k2_xcmplx_0(np__2,X1)=k2_xcmplx_0(X1,np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2015, plain, (r2_hidden(k2_nat_1(np__1,np__2),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X1,X2))=np__1|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X1,X2),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k1_group_3(X1))|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1714, plain, (v1_grsolv_1(X1)|v5_group_1(esk3_2(X1,k10_finseq_1(X2,X3)))|v2_struct_0(X1)|v2_struct_0(k6_group_2(X4))|v2_struct_0(X4)|k1_group_3(k6_group_2(X4))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(k6_group_2(X4))|~l3_algstr_0(X1)|~v3_group_1(k6_group_2(X4))|~v3_group_1(X1)|~v2_group_1(k6_group_2(X4))|~v2_group_1(X1)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4))).#trainneg
cnf(i_0_3942, plain, (m1_group_2(k7_group_2(np__1),X4)|v1_grsolv_1(X1)|v2_struct_0(X4)|v2_struct_0(X1)|esk4_2(X1,k10_finseq_1(X2,X3))=X3|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_7825, plain, (v3_ordinal1(X1)|~r2_hidden(k1_zfmisc_1(k2_tarski(X1,X2)),k6_numbers))).#trainneg
cnf(i_0_9199, plain, (l3_algstr_0(esk11_1(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_9935, plain, (k2_finseq_1(X1)!=k6_numbers|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2)|X1!=np__2|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10455, plain, (v1_finset_1(k4_finseq_1(k6_numbers))|~v1_xboole_0(k3_finseq_1(k6_numbers))|~v1_finseq_1(k3_finseq_1(k6_numbers))|~v1_relat_1(k3_finseq_1(k6_numbers)))).#trainneg
cnf(i_0_9953, plain, (v1_relat_1(k9_xtuple_0(k2_zfmisc_1(X2,X3)))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9954, plain, (v1_relat_1(k9_xtuple_0(k2_zfmisc_1(X2,X3)))|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1227, plain, (v5_group_1(k5_group_6(X1,X2))|v2_struct_0(X3)|X2!=k1_funct_1(esk2_1(X3),k2_nat_1(X4,np__1))|X1!=k1_funct_1(esk2_1(X3),X4)|~m1_group_6(X2,X3,X1)|~v1_group_3(X2,X1)|~m1_group_2(X1,X3)|~r2_hidden(k2_nat_1(X4,np__1),k2_tarski(np__1,np__2))|~r2_hidden(X4,k2_tarski(np__1,np__2))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v15_algstr_0(X2)|~v15_algstr_0(X1)|k3_finseq_1(esk2_1(X3))!=np__2|~v1_finseq_1(esk2_1(X3))|~v1_funct_1(esk2_1(X3))|~v1_relat_1(esk2_1(X3)))).#trainneg
cnf(i_0_1228, plain, (m1_group_6(X1,X2,X3)|v2_struct_0(X2)|X1!=k1_funct_1(esk2_1(X2),k2_nat_1(X4,np__1))|X3!=k1_funct_1(esk2_1(X2),X4)|~m1_group_2(X3,X2)|~m1_group_2(X1,X2)|~r2_hidden(k2_nat_1(X4,np__1),k2_tarski(np__1,np__2))|~r2_hidden(X4,k2_tarski(np__1,np__2))|~m1_subset_1(X4,k5_numbers)|~v1_grsolv_1(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v15_algstr_0(X3)|~v15_algstr_0(X1)|k3_finseq_1(esk2_1(X2))!=np__2|~v1_finseq_1(esk2_1(X2))|~v1_funct_1(esk2_1(X2))|~v1_relat_1(esk2_1(X2)))).#trainneg
cnf(i_0_4227, plain, (r2_hidden(k2_xcmplx_0(np__1,esk3_2(X1,k10_finseq_1(X2,X3))),k2_tarski(np__1,np__2))|v1_grsolv_1(X1)|v2_struct_0(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4466, plain, (m1_group_6(X1,X1,esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X1))=np__1|k6_group_2(X1)!=X1|~v1_group_3(np__2,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(np__2,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X1)),X1)|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|k7_group_2(X1)!=X1|~m1_subset_1(X1,k1_group_3(X1))|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_7422, plain, (v1_grsolv_1(X1)|v5_group_1(esk5_2(X1,X2))|v2_struct_0(X1)|k1_funct_1(X2,np__2)!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=X1|k3_finseq_1(X2)!=np__2|k7_group_2(X1)!=X1|~m1_subset_1(k1_funct_1(X2,np__2),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_10782, plain, (m1_group_2(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|v1_grsolv_1(X1)|v2_struct_0(X1)|esk5_2(X1,k10_finseq_1(X1,X1))=np__1|k6_group_2(X1)!=X1|~v1_group_3(X1,esk4_2(X1,k10_finseq_1(X1,X1)))|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1)|k4_finseq_1(np__2)!=k1_group_3(X1)|~v1_group_3(np__2,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(np__2,esk4_2(X1,k10_finseq_1(X1,X1)))|~m1_group_2(esk4_2(X1,k10_finseq_1(X1,X1)),X1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_10697, plain, (esk4_2(X1,k10_finseq_1(X2,X3))=X3|m1_group_2(k7_group_2(np__1),X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(X4)|v1_xboole_0(k1_group_3(X1))|k2_tarski(np__1,np__2)!=k1_group_3(X4)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X4)|~l3_algstr_0(X1)|~v3_group_1(X4)|~v3_group_1(X1)|~v2_group_1(X4)|~v2_group_1(X1))).#trainneg
cnf(i_0_3085, plain, (v5_group_1(esk4_2(X1,k10_finseq_1(X2,X3)))|v2_struct_0(X1)|v1_grsolv_1(X1)|~l3_algstr_0(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_1576, plain, (m1_group_2(k7_group_2(k6_group_2(X1)),X2)|v2_struct_0(X2)|v2_struct_0(X1)|X3!=k1_group_3(X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X3!=k1_group_3(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(k7_group_2(k6_group_2(X1))))).#trainneg
cnf(i_0_4395, plain, (np__3=np__1|esk5_2(X1,k10_finseq_1(X2,X3))=X3|v1_grsolv_1(X1)|v2_struct_0(X1)|esk3_2(X1,k10_finseq_1(X2,X3))=np__1|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__2)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1)))).#trainneg
cnf(i_0_9200, plain, (v3_group_1(esk11_1(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_10690, negated_conjecture, (m1_subset_1(np__3,k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_10715, plain, (k9_xtuple_0(np__2)!=k6_numbers|X1!=np__2|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_1230, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v5_ordinal1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2)).#trainneg
cnf(i_0_2341, plain, (r2_hidden(np__2,X1)|X1!=k4_finseq_1(np__2)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_1232, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_tarski(np__1,np__2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2)).#trainneg
cnf(i_0_4919, plain, (X1=np__1|X1=np__2|v1_grsolv_1(X1)|v2_struct_0(X1)|k3_finseq_1(X2)!=np__2|X3!=k4_finseq_1(X2)|X3!=k1_group_3(X1)|k6_group_2(X1)!=X1|~m2_finseq_1(k10_finseq_1(X1,X1),k1_group_3(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_1926, plain, (m1_group_2(X1,X2)|v2_struct_0(X2)|v1_grsolv_1(X3)|v2_struct_0(X3)|~m1_group_2(X1,esk3_2(X3,np__2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k2_tarski(np__1,np__2)!=k1_group_3(X2)|k6_group_2(X3)!=k1_funct_1(np__2,np__2)|k7_group_2(X3)!=k1_funct_1(np__2,np__1)|~m2_finseq_1(np__2,k1_group_3(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_1115, plain, (m1_group_2(esk3_2(X1,k10_finseq_1(X2,X3)),X4)|v2_struct_0(X4)|v1_grsolv_1(X1)|v2_struct_0(X1)|v2_struct_0(k7_group_2(k6_group_2(X4)))|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|k1_group_3(k7_group_2(k6_group_2(X4)))!=k2_tarski(np__1,np__2)|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~m2_finseq_1(k10_finseq_1(X2,X3),k1_group_3(X1))|~l3_algstr_0(k7_group_2(k6_group_2(X4)))|~l3_algstr_0(X1)|~v3_group_1(k7_group_2(k6_group_2(X4)))|~v3_group_1(X1)|~v2_group_1(k7_group_2(k6_group_2(X4)))|~v2_group_1(X1))).#trainneg
cnf(i_0_7764, plain, (r2_hidden(X2,X3)|v1_grsolv_1(X1)|v2_struct_0(X1)|np__3=np__1|X3!=k1_group_3(X1)|X2!=k6_group_2(X1)|k7_group_2(X1)!=X1|~m1_subset_1(X2,k1_group_3(X1))|~m1_subset_1(X1,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X2)|~v7_ordinal1(np__1)|~v7_ordinal1(np__2)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_9222, plain, (v2_group_1(k6_group_2(k7_group_2(k6_group_2(X1))))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_10820, negated_conjecture, (m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_10687, negated_conjecture, (m1_subset_1(np__3,k2_tarski(np__1,np__2))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_10691, negated_conjecture, (r2_hidden(np__3,k4_finseq_1(np__2))|~v7_ordinal1(np__1)|~v1_finseq_1(np__2)|~v1_funct_1(np__2)|~v1_relat_1(np__2))).#trainneg
cnf(i_0_1237, plain, (m1_subset_1(k2_nat_1(esk3_2(X1,X2),np__1),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1243, plain, (m1_subset_1(esk3_2(X1,X2),k2_tarski(np__1,np__2))|r1_xxreal_0(k3_finseq_1(X2),k6_numbers)|v1_grsolv_1(X1)|v2_struct_0(X1)|k1_funct_1(X2,k3_finseq_1(X2))!=k6_group_2(X1)|k1_funct_1(X2,np__1)!=k7_group_2(X1)|~m2_finseq_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k3_finseq_1(X2)!=np__2|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10813, plain, (esk3_2(X1,k10_finseq_1(X2,X3))=np__1|esk5_2(X1,k10_finseq_1(X2,X3))=X3|np__3=np__1|v1_grsolv_1(X1)|v2_struct_0(X1)|v1_xboole_0(k1_group_3(X1))|X3!=k6_group_2(X1)|X2!=k7_group_2(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X3,k1_group_3(X1))|~m1_subset_1(X2,k1_group_3(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 150
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 261
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 255
# Processed clauses                    : 2619
# ...of these trivial                  : 25
# ...subsumed                          : 1072
# ...remaining for further processing  : 1522
# Other redundant clauses eliminated   : 132
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 168
# Backward-rewritten                   : 167
# Generated clauses                    : 10868
# ...of the previous two non-trivial   : 9541
# Contextual simplify-reflections      : 585
# Paramodulations                      : 10650
# Factorizations                       : 24
# Equation resolutions                 : 194
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1184
#    Positive orientable unit clauses  : 92
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 1076
# Current number of unprocessed clauses: 6996
# ...number of literals in the above   : 85504
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1857
# Proof object given clauses           : 132
# Proof search given clauses           : 1522
# Clause-clause subsumption calls (NU) : 1157410
# Rec. Clause-clause subsumption calls : 43763
# Non-unit clause-clause subsumptions  : 1734
# Unit Clause-clause subsumption calls : 10597
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 45
# BW rewrite match successes           : 24
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 373197

# -------------------------------------------------
# User time                : 30.885 s
# System time              : 0.653 s
# Total time               : 31.538 s
# Maximum resident set size: 362276 pages
30.88user 0.70system 0:32.04elapsed 98%CPU (0avgtext+0avgdata 362276maxresident)k
0inputs+832outputs (0major+194609minor)pagefaults 0swaps
