# ENIGMA: LightGBM model 'Enigma/mizar40-all-T10/mega-VHSLCAXPh2e15/lgb-t150-d30-l6400-e0.15/model.lgb' loaded. (hash_base: 32768; conj_feats: 52; version: 991; iters: 150)
# Preprocessing time       : 1.386 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t81_group_5, conjecture, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_struct_0(k10_group_5(X1),X2)<=>k7_group_3(X1,X2)=k6_domain_1(u1_struct_0(X1),X2)))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t81_group_5)).
fof(d3_group_1, axiom, ![X1]:(l3_algstr_0(X1)=>(v3_group_1(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4))))))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', d3_group_1)).
fof(d2_group_3, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_group_3(X1,X2,X3)=k6_algstr_0(X1,k6_algstr_0(X1,k2_group_1(X1,X3),X2),X3)))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', d2_group_3)).
fof(t77_group_5, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_struct_0(k10_group_5(X1),X2)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2))))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t77_group_5)).
fof(dt_k2_group_1, axiom, ![X1, X2]:(((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', dt_k2_group_1)).
fof(t13_group_1, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(k6_algstr_0(X1,X2,X3)=X4<=>X3=k6_algstr_0(X1,k2_group_1(X1,X2),X4)))))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t13_group_1)).
fof(t6_group_1, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X2,X4)|k6_algstr_0(X1,X3,X2)=k6_algstr_0(X1,X4,X2))=>X3=X4))))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t6_group_1)).
fof(dt_k6_algstr_0, axiom, ![X1, X2, X3]:(((l3_algstr_0(X1)&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', dt_k6_algstr_0)).
fof(involutiveness_k2_group_1, axiom, ![X1, X2]:(((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>k2_group_1(X1,k2_group_1(X1,X2))=X2), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', involutiveness_k2_group_1)).
fof(d7_group_3, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_group_3(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&X2=k2_group_3(X1,X3,X4)))))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', d7_group_3)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t4_subset)).
fof(dt_k7_group_3, axiom, ![X1, X2]:(((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', dt_k7_group_3)).
fof(t82_group_3, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r2_hidden(k2_group_3(X1,X2,X3),k7_group_3(X1,X2))))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t82_group_3)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', d3_tarski)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t3_subset)).
fof(symmetry_r2_group_3, axiom, ![X1, X2, X3]:((((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>(r2_group_3(X1,X2,X3)=>r2_group_3(X1,X3,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', symmetry_r2_group_3)).
fof(t80_group_3, axiom, ![X1, X2]:((((~(v2_struct_0(X2))&v2_group_1(X2))&v3_group_1(X2))&l3_algstr_0(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>(r2_hidden(X1,k7_group_3(X2,X3))<=>?[X4]:((m1_subset_1(X4,u1_struct_0(X2))&X4=X1)&r2_group_3(X2,X3,X4))))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t80_group_3)).
fof(t83_group_3, axiom, ![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r2_hidden(X2,k7_group_3(X1,X2)))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t83_group_3)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', d1_tarski)).
fof(redefinition_r2_group_3, axiom, ![X1, X2, X3]:((((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>(r2_group_3(X1,X2,X3)<=>r1_group_3(X1,X2,X3))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', redefinition_r2_group_3)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t2_subset)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', redefinition_k6_domain_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', d10_xboole_0)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t7_boole)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', fc2_xboole_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/home/yan/atp/benchmarks/mizar40/all/t81_group_5', t1_subset)).
fof(c_0_26, negated_conjecture, ~(![X1]:((((~(v2_struct_0(X1))&v2_group_1(X1))&v3_group_1(X1))&l3_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_struct_0(k10_group_5(X1),X2)<=>k7_group_3(X1,X2)=k6_domain_1(u1_struct_0(X1),X2))))), inference(assume_negation,[status(cth)],[t81_group_5])).
fof(c_0_27, plain, ![X99, X100, X101, X102]:((~v3_group_1(X99)|(~m1_subset_1(X100,u1_struct_0(X99))|(~m1_subset_1(X101,u1_struct_0(X99))|(~m1_subset_1(X102,u1_struct_0(X99))|k6_algstr_0(X99,k6_algstr_0(X99,X100,X101),X102)=k6_algstr_0(X99,X100,k6_algstr_0(X99,X101,X102)))))|~l3_algstr_0(X99))&((m1_subset_1(esk4_1(X99),u1_struct_0(X99))|v3_group_1(X99)|~l3_algstr_0(X99))&((m1_subset_1(esk5_1(X99),u1_struct_0(X99))|v3_group_1(X99)|~l3_algstr_0(X99))&((m1_subset_1(esk6_1(X99),u1_struct_0(X99))|v3_group_1(X99)|~l3_algstr_0(X99))&(k6_algstr_0(X99,k6_algstr_0(X99,esk4_1(X99),esk5_1(X99)),esk6_1(X99))!=k6_algstr_0(X99,esk4_1(X99),k6_algstr_0(X99,esk5_1(X99),esk6_1(X99)))|v3_group_1(X99)|~l3_algstr_0(X99)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_group_1])])])])])])])).
fof(c_0_28, plain, ![X96, X97, X98]:(v2_struct_0(X96)|~v2_group_1(X96)|~v3_group_1(X96)|~l3_algstr_0(X96)|(~m1_subset_1(X97,u1_struct_0(X96))|(~m1_subset_1(X98,u1_struct_0(X96))|k2_group_3(X96,X97,X98)=k6_algstr_0(X96,k6_algstr_0(X96,k2_group_1(X96,X98),X97),X98)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_group_3])])])])])])).
fof(c_0_29, plain, ![X307, X308, X309]:((~r1_struct_0(k10_group_5(X307),X308)|(~m1_subset_1(X309,u1_struct_0(X307))|k6_algstr_0(X307,X308,X309)=k6_algstr_0(X307,X309,X308))|~m1_subset_1(X308,u1_struct_0(X307))|(v2_struct_0(X307)|~v2_group_1(X307)|~v3_group_1(X307)|~l3_algstr_0(X307)))&((m1_subset_1(esk57_2(X307,X308),u1_struct_0(X307))|r1_struct_0(k10_group_5(X307),X308)|~m1_subset_1(X308,u1_struct_0(X307))|(v2_struct_0(X307)|~v2_group_1(X307)|~v3_group_1(X307)|~l3_algstr_0(X307)))&(k6_algstr_0(X307,X308,esk57_2(X307,X308))!=k6_algstr_0(X307,esk57_2(X307,X308),X308)|r1_struct_0(k10_group_5(X307),X308)|~m1_subset_1(X308,u1_struct_0(X307))|(v2_struct_0(X307)|~v2_group_1(X307)|~v3_group_1(X307)|~l3_algstr_0(X307))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t77_group_5])])])])])])])])).
fof(c_0_30, negated_conjecture, ((((~v2_struct_0(esk1_0)&v2_group_1(esk1_0))&v3_group_1(esk1_0))&l3_algstr_0(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~r1_struct_0(k10_group_5(esk1_0),esk2_0)|k7_group_3(esk1_0,esk2_0)!=k6_domain_1(u1_struct_0(esk1_0),esk2_0))&(r1_struct_0(k10_group_5(esk1_0),esk2_0)|k7_group_3(esk1_0,esk2_0)=k6_domain_1(u1_struct_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_26])])])])])])).
cnf(c_0_31, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4))|~v3_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (v2_struct_0(X1)|k2_group_3(X1,X2,X3)=k6_algstr_0(X1,k6_algstr_0(X1,k2_group_1(X1,X3),X2),X3)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_33, plain, ![X129, X130]:(v2_struct_0(X129)|~v2_group_1(X129)|~v3_group_1(X129)|~l3_algstr_0(X129)|~m1_subset_1(X130,u1_struct_0(X129))|m1_subset_1(k2_group_1(X129,X130),u1_struct_0(X129))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_group_1])])])).
cnf(c_0_34, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~r1_struct_0(k10_group_5(X1),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk2_0)|k7_group_3(esk1_0,esk2_0)=k6_domain_1(u1_struct_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, negated_conjecture, (l3_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, negated_conjecture, (v3_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, negated_conjecture, (v2_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_41, plain, ![X280, X281, X282, X283]:((k6_algstr_0(X280,X281,X282)!=X283|X282=k6_algstr_0(X280,k2_group_1(X280,X281),X283)|~m1_subset_1(X283,u1_struct_0(X280))|~m1_subset_1(X282,u1_struct_0(X280))|~m1_subset_1(X281,u1_struct_0(X280))|(v2_struct_0(X280)|~v2_group_1(X280)|~v3_group_1(X280)|~l3_algstr_0(X280)))&(X282!=k6_algstr_0(X280,k2_group_1(X280,X281),X283)|k6_algstr_0(X280,X281,X282)=X283|~m1_subset_1(X283,u1_struct_0(X280))|~m1_subset_1(X282,u1_struct_0(X280))|~m1_subset_1(X281,u1_struct_0(X280))|(v2_struct_0(X280)|~v2_group_1(X280)|~v3_group_1(X280)|~l3_algstr_0(X280)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t13_group_1])])])])])])])).
cnf(c_0_42, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4))|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (k6_algstr_0(X1,k6_algstr_0(X1,k2_group_1(X1,X3),X2),X3)=k2_group_3(X1,X2,X3)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_32])).
cnf(c_0_44, plain, (v2_struct_0(X1)|m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_45, plain, ![X303, X304, X305, X306]:((k6_algstr_0(X303,X304,X305)!=k6_algstr_0(X303,X304,X306)|X305=X306|~m1_subset_1(X306,u1_struct_0(X303))|~m1_subset_1(X305,u1_struct_0(X303))|~m1_subset_1(X304,u1_struct_0(X303))|(v2_struct_0(X303)|~v2_group_1(X303)|~v3_group_1(X303)|~l3_algstr_0(X303)))&(k6_algstr_0(X303,X305,X304)!=k6_algstr_0(X303,X306,X304)|X305=X306|~m1_subset_1(X306,u1_struct_0(X303))|~m1_subset_1(X305,u1_struct_0(X303))|~m1_subset_1(X304,u1_struct_0(X303))|(v2_struct_0(X303)|~v2_group_1(X303)|~v3_group_1(X303)|~l3_algstr_0(X303)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_group_1])])])])])])])).
cnf(c_0_46, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_struct_0(k10_group_5(X1),X2)), inference(evalgc,[status(thm)],[c_0_34])).
cnf(c_0_47, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|r1_struct_0(k10_group_5(esk1_0),esk2_0)), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_49, negated_conjecture, (l3_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_37])).
cnf(c_0_50, negated_conjecture, (v3_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_38])).
cnf(c_0_51, negated_conjecture, (v2_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_39])).
cnf(c_0_52, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_53, plain, (k6_algstr_0(X2,X3,X1)=X4|v2_struct_0(X2)|X1!=k6_algstr_0(X2,k2_group_1(X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_54, plain, (k6_algstr_0(X1,k2_group_1(X1,X2),k6_algstr_0(X1,X3,X2))=k2_group_3(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43])])).
cnf(c_0_55, plain, (v2_struct_0(X1)|m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_44])).
fof(c_0_56, plain, ![X140, X141, X142]:(~l3_algstr_0(X140)|~m1_subset_1(X141,u1_struct_0(X140))|~m1_subset_1(X142,u1_struct_0(X140))|m1_subset_1(k6_algstr_0(X140,X141,X142),u1_struct_0(X140))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k6_algstr_0])])).
cnf(c_0_57, plain, (X3=X4|v2_struct_0(X1)|k6_algstr_0(X1,X2,X3)!=k6_algstr_0(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_58, negated_conjecture, (k6_algstr_0(esk1_0,esk2_0,X1)=k6_algstr_0(esk1_0,X1,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
fof(c_0_59, plain, ![X199, X200]:(v2_struct_0(X199)|~v2_group_1(X199)|~v3_group_1(X199)|~l3_algstr_0(X199)|~m1_subset_1(X200,u1_struct_0(X199))|k2_group_1(X199,k2_group_1(X199,X200))=X200), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[involutiveness_k2_group_1])])])).
cnf(c_0_60, plain, (k6_algstr_0(X2,X3,X1)=X4|v2_struct_0(X2)|X1!=k6_algstr_0(X2,k2_group_1(X2,X3),X4)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_61, plain, (k6_algstr_0(X1,k2_group_1(X1,X2),k6_algstr_0(X1,X3,X2))=k2_group_3(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_62, plain, (m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_63, plain, ![X116, X117, X118, X120]:(((m1_subset_1(esk8_3(X116,X117,X118),u1_struct_0(X116))|~r1_group_3(X116,X117,X118)|~m1_subset_1(X118,u1_struct_0(X116))|~m1_subset_1(X117,u1_struct_0(X116))|(v2_struct_0(X116)|~v2_group_1(X116)|~v3_group_1(X116)|~l3_algstr_0(X116)))&(X117=k2_group_3(X116,X118,esk8_3(X116,X117,X118))|~r1_group_3(X116,X117,X118)|~m1_subset_1(X118,u1_struct_0(X116))|~m1_subset_1(X117,u1_struct_0(X116))|(v2_struct_0(X116)|~v2_group_1(X116)|~v3_group_1(X116)|~l3_algstr_0(X116))))&(~m1_subset_1(X120,u1_struct_0(X116))|X117!=k2_group_3(X116,X118,X120)|r1_group_3(X116,X117,X118)|~m1_subset_1(X118,u1_struct_0(X116))|~m1_subset_1(X117,u1_struct_0(X116))|(v2_struct_0(X116)|~v2_group_1(X116)|~v3_group_1(X116)|~l3_algstr_0(X116)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_group_3])])])])])])])])).
fof(c_0_64, plain, ![X296, X297, X298]:(~r2_hidden(X296,X297)|~m1_subset_1(X297,k1_zfmisc_1(X298))|m1_subset_1(X296,X298)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_65, plain, ![X146, X147]:(v2_struct_0(X146)|~v2_group_1(X146)|~v3_group_1(X146)|~l3_algstr_0(X146)|~m1_subset_1(X147,u1_struct_0(X146))|m1_subset_1(k7_group_3(X146,X147),k1_zfmisc_1(u1_struct_0(X146)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_group_3])])])).
cnf(c_0_66, plain, (X3=X4|v2_struct_0(X1)|k6_algstr_0(X1,X2,X3)!=k6_algstr_0(X1,X2,X4)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_67, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,k2_group_1(esk1_0,X1),esk2_0)=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_55]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_68, plain, (v2_struct_0(X1)|k2_group_1(X1,k2_group_1(X1,X2))=X2|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X4,X2)|v2_struct_0(X1)|X3!=k2_group_3(X1,X4,X2)|~m1_subset_1(k6_algstr_0(X1,X4,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_70, plain, (m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_71, plain, (X1=k2_group_3(X2,X3,esk8_3(X2,X1,X3))|v2_struct_0(X2)|~r1_group_3(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_72, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, plain, (v2_struct_0(X1)|m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_74, plain, ![X318, X319, X320]:(v2_struct_0(X318)|~v2_group_1(X318)|~v3_group_1(X318)|~l3_algstr_0(X318)|(~m1_subset_1(X319,u1_struct_0(X318))|(~m1_subset_1(X320,u1_struct_0(X318))|r2_hidden(k2_group_3(X318,X319,X320),k7_group_3(X318,X319))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t82_group_3])])])])])])).
fof(c_0_75, plain, ![X109, X110, X111, X109, X110]:((~r1_tarski(X109,X110)|(~r2_hidden(X111,X109)|r2_hidden(X111,X110)))&((r2_hidden(esk7_2(X109,X110),X109)|r1_tarski(X109,X110))&(~r2_hidden(esk7_2(X109,X110),X110)|r1_tarski(X109,X110)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])])).
fof(c_0_76, plain, ![X294, X295, X294, X295]:((~m1_subset_1(X294,k1_zfmisc_1(X295))|r1_tarski(X294,X295))&(~r1_tarski(X294,X295)|m1_subset_1(X294,k1_zfmisc_1(X295)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])])])).
cnf(c_0_77, negated_conjecture, (esk2_0=X1|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,k2_group_1(esk1_0,X2),X1)!=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_78, plain, (k2_group_1(X1,k2_group_1(X1,X2))=X2|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_79, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X4,X2)|v2_struct_0(X1)|X3!=k2_group_3(X1,X4,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_80, plain, (k2_group_3(X2,X3,esk8_3(X2,X1,X3))=X1|v2_struct_0(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~r1_group_3(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_81, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_82, plain, ![X277, X278, X279]:(v2_struct_0(X277)|~v2_group_1(X277)|~v3_group_1(X277)|~l3_algstr_0(X277)|~m1_subset_1(X278,u1_struct_0(X277))|~m1_subset_1(X279,u1_struct_0(X277))|(~r2_group_3(X277,X278,X279)|r2_group_3(X277,X279,X278))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[symmetry_r2_group_3])])])).
fof(c_0_83, plain, ![X313, X314, X315, X317]:((((m1_subset_1(esk58_3(X313,X314,X315),u1_struct_0(X314))|~r2_hidden(X313,k7_group_3(X314,X315))|~m1_subset_1(X315,u1_struct_0(X314))|(v2_struct_0(X314)|~v2_group_1(X314)|~v3_group_1(X314)|~l3_algstr_0(X314)))&(esk58_3(X313,X314,X315)=X313|~r2_hidden(X313,k7_group_3(X314,X315))|~m1_subset_1(X315,u1_struct_0(X314))|(v2_struct_0(X314)|~v2_group_1(X314)|~v3_group_1(X314)|~l3_algstr_0(X314))))&(r2_group_3(X314,X315,esk58_3(X313,X314,X315))|~r2_hidden(X313,k7_group_3(X314,X315))|~m1_subset_1(X315,u1_struct_0(X314))|(v2_struct_0(X314)|~v2_group_1(X314)|~v3_group_1(X314)|~l3_algstr_0(X314))))&(~m1_subset_1(X317,u1_struct_0(X314))|X317!=X313|~r2_group_3(X314,X315,X317)|r2_hidden(X313,k7_group_3(X314,X315))|~m1_subset_1(X315,u1_struct_0(X314))|(v2_struct_0(X314)|~v2_group_1(X314)|~v3_group_1(X314)|~l3_algstr_0(X314)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t80_group_3])])])])])])])])).
cnf(c_0_84, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_85, plain, (v2_struct_0(X1)|m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_86, plain, (v2_struct_0(X1)|r2_hidden(k2_group_3(X1,X2,X3),k7_group_3(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_87, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_88, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_89, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X2,X1)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,X2),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_90, plain, (k6_algstr_0(X1,esk8_3(X1,X2,X3),X2)=k6_algstr_0(X1,X3,esk8_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80])])])).
cnf(c_0_91, plain, (v2_struct_0(X1)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_92, plain, (v2_struct_0(X1)|r2_group_3(X1,X3,X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_group_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_93, plain, (r2_group_3(X1,X2,esk58_3(X3,X1,X2))|v2_struct_0(X1)|~r2_hidden(X3,k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_94, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85])])).
cnf(c_0_95, plain, (v2_struct_0(X1)|r2_hidden(k2_group_3(X1,X2,X3),k7_group_3(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_96, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_97, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_88])).
fof(c_0_98, plain, ![X321, X322]:(v2_struct_0(X321)|~v2_group_1(X321)|~v3_group_1(X321)|~l3_algstr_0(X321)|(~m1_subset_1(X322,u1_struct_0(X321))|r2_hidden(X322,k7_group_3(X321,X322)))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t83_group_3])])])])])])).
fof(c_0_99, plain, ![X92, X93, X94, X94, X92, X93]:(((~r2_hidden(X94,X93)|X94=X92|X93!=k1_tarski(X92))&(X94!=X92|r2_hidden(X94,X93)|X93!=k1_tarski(X92)))&((~r2_hidden(esk3_2(X92,X93),X93)|esk3_2(X92,X93)!=X92|X93=k1_tarski(X92))&(r2_hidden(esk3_2(X92,X93),X93)|esk3_2(X92,X93)=X92|X93=k1_tarski(X92)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])])).
cnf(c_0_100, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X2,X1)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_55]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_101, plain, (k6_algstr_0(X1,esk8_3(X1,X2,X3),X2)=k6_algstr_0(X1,X3,esk8_3(X1,X2,X3))|v2_struct_0(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_102, plain, (v2_struct_0(X1)|r2_group_3(X1,X3,X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~r2_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_103, plain, (v2_struct_0(X1)|r2_group_3(X1,X2,esk58_3(X3,X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,k7_group_3(X1,X2))), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_104, plain, (m1_subset_1(esk58_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_105, plain, (esk58_3(X1,X2,X3)=X1|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_106, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_107, plain, (k6_algstr_0(X1,X2,k2_group_3(X1,X3,X2))=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_79])])).
cnf(c_0_108, plain, (m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_109, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97])])).
cnf(c_0_110, plain, (v2_struct_0(X1)|r2_hidden(X2,k7_group_3(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_111, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_112, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X2,esk8_3(esk1_0,X1,X2))!=k6_algstr_0(esk1_0,esk2_0,esk8_3(esk1_0,X1,X2))|~m1_subset_1(esk8_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_group_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
fof(c_0_113, plain, ![X269, X270, X271]:((~r2_group_3(X269,X270,X271)|r1_group_3(X269,X270,X271)|(v2_struct_0(X269)|~v2_group_1(X269)|~v3_group_1(X269)|~l3_algstr_0(X269)|~m1_subset_1(X270,u1_struct_0(X269))|~m1_subset_1(X271,u1_struct_0(X269))))&(~r1_group_3(X269,X270,X271)|r2_group_3(X269,X270,X271)|(v2_struct_0(X269)|~v2_group_1(X269)|~v3_group_1(X269)|~l3_algstr_0(X269)|~m1_subset_1(X270,u1_struct_0(X269))|~m1_subset_1(X271,u1_struct_0(X269))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r2_group_3])])])])).
cnf(c_0_114, plain, (r2_group_3(X1,esk58_3(X2,X1,X3),X3)|v2_struct_0(X1)|~m1_subset_1(esk58_3(X2,X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,k7_group_3(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_115, plain, (v2_struct_0(X2)|m1_subset_1(esk58_3(X1,X2,X3),u1_struct_0(X2))|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,k7_group_3(X2,X3))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_116, plain, (esk58_3(X1,X2,X3)=X1|v2_struct_0(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,k7_group_3(X2,X3))), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_117, plain, (r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_118, plain, (k6_algstr_0(X1,X2,k2_group_3(X1,X3,X2))=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_119, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_120, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_85])])).
cnf(c_0_121, plain, (v2_struct_0(X1)|r2_hidden(X2,k7_group_3(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_122, plain, (X1=X3|X2!=k1_tarski(X3)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_111])).
fof(c_0_123, plain, ![X286, X287]:(~m1_subset_1(X286,X287)|(v1_xboole_0(X287)|r2_hidden(X286,X287))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_124, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(esk8_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_112]), c_0_48])])])).
cnf(c_0_125, plain, (r1_group_3(X1,X2,X3)|v2_struct_0(X1)|~r2_group_3(X1,X2,X3)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_126, plain, (r2_group_3(X1,esk58_3(X2,X1,X3),X3)|v2_struct_0(X1)|~r2_hidden(X2,k7_group_3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_127, plain, (esk58_3(esk7_2(k7_group_3(X1,X2),X3),X1,X2)=esk7_2(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117])])).
cnf(c_0_128, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,X1,X2)=esk2_0|k6_algstr_0(esk1_0,X1,X2)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k2_group_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_118]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_129, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_130, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_131, plain, (esk7_2(X1,X2)=X3|r1_tarski(X1,X2)|X1!=k1_tarski(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_117])])).
cnf(c_0_132, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_133, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_91]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_134, plain, (v2_struct_0(X1)|r1_group_3(X1,X2,X3)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~r2_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_135, plain, (r2_group_3(X1,esk7_2(k7_group_3(X1,X2),X3),X2)|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~r2_hidden(esk7_2(k7_group_3(X1,X2),X3),k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127])])).
fof(c_0_136, plain, ![X264, X265]:(v1_xboole_0(X264)|~m1_subset_1(X265,X264)|k6_domain_1(X264,X265)=k1_tarski(X265)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_137, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,X1,X2)=esk2_0|k6_algstr_0(esk1_0,X1,X2)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_108]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_138, plain, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(esk7_2(X1,u1_struct_0(X2)),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_139, plain, (esk7_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_131])])).
cnf(c_0_140, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_141, plain, ![X90, X91, X90, X91]:(((r1_tarski(X90,X91)|X90!=X91)&(r1_tarski(X91,X90)|X90!=X91))&(~r1_tarski(X90,X91)|~r1_tarski(X91,X90)|X90=X91)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])])])).
cnf(c_0_142, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_143, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_group_3(esk1_0,X1,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_144, plain, (r2_group_3(X1,esk7_2(k7_group_3(X1,X2),X3),X2)|r1_tarski(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(csr,[status(thm)],[c_0_135, c_0_117])).
cnf(c_0_145, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_146, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,X1,esk2_0)=esk2_0|k6_algstr_0(esk1_0,X1,esk2_0)!=k6_algstr_0(esk1_0,esk2_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_48])])).
cnf(c_0_147, plain, (r1_tarski(k1_tarski(X1),u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139])])).
cnf(c_0_148, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_149, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_150, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk7_2(X1,X2),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_142])])).
cnf(c_0_151, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk7_2(k7_group_3(esk1_0,esk2_0),X1)=esk2_0|r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|~m1_subset_1(esk7_2(k7_group_3(esk1_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_152, plain, (m1_subset_1(esk7_2(k7_group_3(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_117])])).
cnf(c_0_153, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_154, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_48])])).
fof(c_0_155, plain, ![X311, X312]:(~r2_hidden(X311,X312)|~v1_xboole_0(X312)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_156, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_147])])).
cnf(c_0_157, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_148])])).
cnf(c_0_158, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_159, plain, (r1_tarski(k1_tarski(X1),X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_139])])).
cnf(c_0_160, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk7_2(k7_group_3(esk1_0,esk2_0),X1)=esk2_0|r1_tarski(k7_group_3(esk1_0,esk2_0),X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
fof(c_0_161, plain, ![X180]:~v1_xboole_0(k1_tarski(X180)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
fof(c_0_162, plain, ![X284, X285]:(~r2_hidden(X284,X285)|m1_subset_1(X284,X285)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_163, negated_conjecture, (k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_48])])])).
cnf(c_0_164, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_165, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k1_tarski(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_166, plain, (r2_hidden(X1,k1_tarski(X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_157])])).
cnf(c_0_167, plain, (X1=k1_tarski(X2)|v1_xboole_0(X1)|~r1_tarski(X1,k1_tarski(X2))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159])])).
cnf(c_0_168, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~r2_hidden(esk2_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_160])])).
cnf(c_0_169, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_170, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_171, negated_conjecture, (r2_hidden(esk2_0,k7_group_3(esk1_0,esk2_0))|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_163]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_172, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_173, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166])])).
cnf(c_0_174, negated_conjecture, (k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(k7_group_3(esk1_0,esk2_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,k7_group_3(esk1_0,esk2_0))|~r2_hidden(esk2_0,k1_tarski(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168])])).
cnf(c_0_175, plain, (~v1_xboole_0(k1_tarski(X1))), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_176, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_177, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,k7_group_3(esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_171, c_0_157])).
cnf(c_0_178, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173])])).
cnf(c_0_179, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(k7_group_3(esk1_0,esk2_0))|~m1_subset_1(X1,k7_group_3(esk1_0,esk2_0))|~m1_subset_1(esk2_0,k1_tarski(X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_142]), c_0_175])])).
cnf(c_0_180, plain, (m1_subset_1(X1,k7_group_3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_121])])).
cnf(c_0_181, plain, (m1_subset_1(X1,k1_tarski(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_166])])).
cnf(c_0_182, negated_conjecture, (r2_hidden(esk2_0,k7_group_3(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_177, c_0_178])])).
cnf(c_0_183, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(k7_group_3(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_181]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_184, negated_conjecture, (~v1_xboole_0(k7_group_3(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_182])])).
cnf(c_0_185, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)), inference(sr,[status(thm)],[c_0_183, c_0_184])).
cnf(c_0_186, plain, (k2_group_3(X1,X2,X3)=X4|v2_struct_0(X1)|k7_group_3(X1,X2)!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_95])])).
cnf(c_0_187, negated_conjecture, (k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_185]), c_0_48])]), c_0_178])])).
cnf(c_0_188, negated_conjecture, (k2_group_3(X1,X2,X3)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187])])).
cnf(c_0_189, plain, (r1_struct_0(k10_group_5(X1),X2)|v2_struct_0(X1)|k6_algstr_0(X1,X2,esk57_2(X1,X2))!=k6_algstr_0(X1,esk57_2(X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_190, negated_conjecture, (k6_algstr_0(X1,X2,esk2_0)=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X3)!=k7_group_3(esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_188])])).
cnf(c_0_191, plain, (m1_subset_1(esk57_2(X1,X2),u1_struct_0(X1))|r1_struct_0(k10_group_5(X1),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_192, plain, (v2_struct_0(X1)|r1_struct_0(k10_group_5(X1),X2)|k6_algstr_0(X1,esk57_2(X1,X2),X2)!=k6_algstr_0(X1,X2,esk57_2(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_193, negated_conjecture, (k6_algstr_0(esk1_0,X1,esk2_0)=k6_algstr_0(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_190]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_194, plain, (v2_struct_0(X1)|r1_struct_0(k10_group_5(X1),X2)|m1_subset_1(esk57_2(X1,X2),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_195, negated_conjecture, (~r1_struct_0(k10_group_5(esk1_0),esk2_0)|k7_group_3(esk1_0,esk2_0)!=k6_domain_1(u1_struct_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_196, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk2_0)|~m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193]), c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_197, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk2_0)|m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])])).
cnf(c_0_198, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|~r1_struct_0(k10_group_5(esk1_0),esk2_0)), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_199, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk2_0)), inference(csr,[status(thm)],[c_0_196, c_0_197])).
cnf(c_0_200, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_198, c_0_199])])])).
cnf(c_0_201, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_153]), c_0_187]), c_0_48])]), c_0_178]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 202
# Proof object clause steps            : 149
# Proof object formula steps           : 53
# Proof object conjectures             : 54
# Proof object clause conjectures      : 51
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 40
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 57
# Proof object simplifying inferences  : 127
# Training examples: 99 positive, 1170 negative
# Training: Positive examples begin
cnf(c_0_52, negated_conjecture, (~v2_struct_0(esk1_0))).#trainpos
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_148, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2))).#trainpos
cnf(c_0_51, negated_conjecture, (v2_group_1(esk1_0))).#trainpos
cnf(c_0_50, negated_conjecture, (v3_group_1(esk1_0))).#trainpos
cnf(c_0_157, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1))).#trainpos
cnf(c_0_166, plain, (r2_hidden(X1,k1_tarski(X1)))).#trainpos
cnf(c_0_49, negated_conjecture, (l3_algstr_0(esk1_0))).#trainpos
cnf(c_0_172, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainpos
cnf(c_0_122, plain, (X1=X3|X2!=k1_tarski(X3)|~r2_hidden(X1,X2))).#trainpos
cnf(c_0_176, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainpos
cnf(c_0_142, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainpos
cnf(c_0_181, plain, (m1_subset_1(X1,k1_tarski(X1)))).#trainpos
cnf(c_0_198, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|~r1_struct_0(k10_group_5(esk1_0),esk2_0))).#trainpos
cnf(c_0_47, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|r1_struct_0(k10_group_5(esk1_0),esk2_0))).#trainpos
cnf(c_0_153, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainpos
cnf(c_0_97, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainpos
cnf(c_0_158, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2))).#trainpos
cnf(c_0_117, plain, (r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X1))).#trainpos
cnf(c_0_96, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2))).#trainpos
cnf(c_0_84, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainpos
cnf(c_0_129, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2))).#trainpos
cnf(c_0_121, plain, (v2_struct_0(X1)|r2_hidden(X2,k7_group_3(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_78, plain, (k2_group_1(X1,k2_group_1(X1,X2))=X2|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_55, plain, (v2_struct_0(X1)|m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_85, plain, (v2_struct_0(X1)|m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_194, plain, (v2_struct_0(X1)|r1_struct_0(k10_group_5(X1),X2)|m1_subset_1(esk57_2(X1,X2),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_180, plain, (m1_subset_1(X1,k7_group_3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainpos
cnf(c_0_131, plain, (esk7_2(X1,X2)=X3|r1_tarski(X1,X2)|X1!=k1_tarski(X3))).#trainpos
cnf(c_0_116, plain, (esk58_3(X1,X2,X3)=X1|v2_struct_0(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,k7_group_3(X2,X3)))).#trainpos
cnf(c_0_134, plain, (v2_struct_0(X1)|r1_group_3(X1,X2,X3)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~r2_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_102, plain, (v2_struct_0(X1)|r2_group_3(X1,X3,X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~r2_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_46, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_struct_0(k10_group_5(X1),X2))).#trainpos
cnf(c_0_95, plain, (v2_struct_0(X1)|r2_hidden(k2_group_3(X1,X2,X3),k7_group_3(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_115, plain, (v2_struct_0(X2)|m1_subset_1(esk58_3(X1,X2,X3),u1_struct_0(X2))|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,k7_group_3(X2,X3)))).#trainpos
cnf(c_0_103, plain, (v2_struct_0(X1)|r2_group_3(X1,X2,esk58_3(X3,X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X3,k7_group_3(X1,X2)))).#trainpos
cnf(c_0_66, plain, (X3=X4|v2_struct_0(X1)|k6_algstr_0(X1,X2,X3)!=k6_algstr_0(X1,X2,X4)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_91, plain, (v2_struct_0(X1)|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_60, plain, (k6_algstr_0(X2,X3,X1)=X4|v2_struct_0(X2)|X1!=k6_algstr_0(X2,k2_group_1(X2,X3),X4)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2)))).#trainpos
cnf(c_0_192, plain, (v2_struct_0(X1)|r1_struct_0(k10_group_5(X1),X2)|k6_algstr_0(X1,esk57_2(X1,X2),X2)!=k6_algstr_0(X1,X2,esk57_2(X1,X2))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_80, plain, (k2_group_3(X2,X3,esk8_3(X2,X1,X3))=X1|v2_struct_0(X2)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~r1_group_3(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2)))).#trainpos
cnf(c_0_58, negated_conjecture, (k6_algstr_0(esk1_0,esk2_0,X1)=k6_algstr_0(esk1_0,X1,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_109, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainpos
cnf(c_0_42, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4))|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_150, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk7_2(X1,X2),X2))).#trainpos
cnf(c_0_94, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainpos
cnf(c_0_197, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk2_0)|m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_139, plain, (esk7_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2))).#trainpos
cnf(c_0_127, plain, (esk58_3(esk7_2(k7_group_3(X1,X2),X3),X1,X2)=esk7_2(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainpos
cnf(c_0_186, plain, (k2_group_3(X1,X2,X3)=X4|v2_struct_0(X1)|k7_group_3(X1,X2)!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainpos
cnf(c_0_114, plain, (r2_group_3(X1,esk58_3(X2,X1,X3),X3)|v2_struct_0(X1)|~m1_subset_1(esk58_3(X2,X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,k7_group_3(X1,X3)))).#trainpos
cnf(c_0_67, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,k2_group_1(esk1_0,X1),esk2_0)=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_120, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainpos
cnf(c_0_77, negated_conjecture, (esk2_0=X1|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,k2_group_1(esk1_0,X2),X1)!=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_130, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainpos
cnf(c_0_152, plain, (m1_subset_1(esk7_2(k7_group_3(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainpos
cnf(c_0_159, plain, (r1_tarski(k1_tarski(X1),X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2))).#trainpos
cnf(c_0_135, plain, (r2_group_3(X1,esk7_2(k7_group_3(X1,X2),X3),X2)|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~r2_hidden(esk7_2(k7_group_3(X1,X2),X3),k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainpos
cnf(c_0_89, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X2,X1)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,X2),u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_167, plain, (X1=k1_tarski(X2)|v1_xboole_0(X1)|~r1_tarski(X1,k1_tarski(X2))|~m1_subset_1(X2,X1))).#trainpos
cnf(c_0_108, plain, (m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainpos
cnf(c_0_100, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X2,X1)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_175, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainpos
cnf(c_0_70, plain, (m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_43, plain, (k6_algstr_0(X1,k6_algstr_0(X1,k2_group_1(X1,X3),X2),X3)=k2_group_3(X1,X2,X3)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainpos
cnf(c_0_54, plain, (k6_algstr_0(X1,k2_group_1(X1,X2),k6_algstr_0(X1,X3,X2))=k2_group_3(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainpos
cnf(c_0_69, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X4,X2)|v2_struct_0(X1)|X3!=k2_group_3(X1,X4,X2)|~m1_subset_1(k6_algstr_0(X1,X4,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainpos
cnf(c_0_107, plain, (k6_algstr_0(X1,X2,k2_group_3(X1,X3,X2))=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainpos
cnf(c_0_90, plain, (k6_algstr_0(X1,esk8_3(X1,X2,X3),X2)=k6_algstr_0(X1,X3,esk8_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3))).#trainpos
cnf(c_0_128, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,X1,X2)=esk2_0|k6_algstr_0(esk1_0,X1,X2)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k2_group_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_112, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X2,esk8_3(esk1_0,X1,X2))!=k6_algstr_0(esk1_0,esk2_0,esk8_3(esk1_0,X1,X2))|~m1_subset_1(esk8_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_group_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_137, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,X1,X2)=esk2_0|k6_algstr_0(esk1_0,X1,X2)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_124, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(esk8_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_146, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,X1,esk2_0)=esk2_0|k6_algstr_0(esk1_0,X1,esk2_0)!=k6_algstr_0(esk1_0,esk2_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_133, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_154, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0)).#trainpos
cnf(c_0_143, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_group_3(esk1_0,X1,esk2_0))).#trainpos
cnf(c_0_163, negated_conjecture, (k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0)).#trainpos
cnf(c_0_138, plain, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(esk7_2(X1,u1_struct_0(X2)),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainpos
cnf(c_0_151, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk7_2(k7_group_3(esk1_0,esk2_0),X1)=esk2_0|r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|~m1_subset_1(esk7_2(k7_group_3(esk1_0,esk2_0),X1),u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_171, negated_conjecture, (r2_hidden(esk2_0,k7_group_3(esk1_0,esk2_0))|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_147, plain, (r1_tarski(k1_tarski(X1),u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainpos
cnf(c_0_160, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk7_2(k7_group_3(esk1_0,esk2_0),X1)=esk2_0|r1_tarski(k7_group_3(esk1_0,esk2_0),X1))).#trainpos
cnf(c_0_156, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainpos
cnf(c_0_168, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~r2_hidden(esk2_0,X1))).#trainpos
cnf(c_0_165, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k1_tarski(esk2_0)))).#trainpos
cnf(c_0_174, negated_conjecture, (k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(k7_group_3(esk1_0,esk2_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,k7_group_3(esk1_0,esk2_0))|~r2_hidden(esk2_0,k1_tarski(X1)))).#trainpos
cnf(c_0_173, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_178, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_179, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(k7_group_3(esk1_0,esk2_0))|~m1_subset_1(X1,k7_group_3(esk1_0,esk2_0))|~m1_subset_1(esk2_0,k1_tarski(X1)))).#trainpos
cnf(c_0_182, negated_conjecture, (r2_hidden(esk2_0,k7_group_3(esk1_0,esk2_0)))).#trainpos
cnf(c_0_184, negated_conjecture, (~v1_xboole_0(k7_group_3(esk1_0,esk2_0)))).#trainpos
cnf(c_0_183, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(k7_group_3(esk1_0,esk2_0)))).#trainpos
cnf(c_0_187, negated_conjecture, (k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0))).#trainpos
cnf(c_0_188, negated_conjecture, (k2_group_3(X1,X2,X3)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainpos
cnf(c_0_190, negated_conjecture, (k6_algstr_0(X1,X2,esk2_0)=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X3)!=k7_group_3(esk1_0,esk2_0))).#trainpos
cnf(c_0_193, negated_conjecture, (k6_algstr_0(esk1_0,X1,esk2_0)=k6_algstr_0(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_196, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk2_0)|~m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainpos
cnf(c_0_200, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_201, plain, (~v2_struct_0(esk18_0))).#trainneg
cnf(i_0_264, plain, (~v2_struct_0(esk38_0))).#trainneg
cnf(i_0_247, plain, (~v1_xboole_0(esk32_0))).#trainneg
cnf(i_0_250, plain, (~v1_xboole_0(esk34_0))).#trainneg
cnf(i_0_277, plain, (~v1_xboole_0(esk42_0))).#trainneg
cnf(i_0_292, plain, (~v1_xboole_0(esk48_0))).#trainneg
cnf(i_0_311, plain, (~v1_xboole_0(esk54_0))).#trainneg
cnf(i_0_255, plain, (~v1_finset_1(esk36_0))).#trainneg
cnf(i_0_165, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_319, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_95, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_327, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_158, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_134, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_362, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_199, plain, (v2_group_1(esk18_0))).#trainneg
cnf(i_0_174, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_180, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_157, plain, (v2_struct_0(X1)|v1_funct_1(u2_algstr_0(X1))|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_164, plain, (v2_struct_0(X1)|v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1)|~v1_group_1(X1))).#trainneg
cnf(i_0_115, plain, (v2_struct_0(X1)|v15_algstr_0(k10_group_5(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_130, plain, (v2_struct_0(X1)|v15_algstr_0(k7_group_2(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_8, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=X1|~l3_algstr_0(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_364, plain, (X1=X2|k1_tarski(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_260, plain, (v2_group_1(esk38_0))).#trainneg
cnf(i_0_368, plain, (X1=X2|v1_xboole_0(X3)|X3!=k1_tarski(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_370, negated_conjecture, (esk2_0=X1|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_173, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_177, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_79, plain, (v3_finseq_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_198, plain, (v3_group_1(esk18_0))).#trainneg
cnf(i_0_259, plain, (v3_group_1(esk38_0))).#trainneg
cnf(i_0_342, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_340, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_136, plain, (v2_struct_0(X2)|v2_group_1(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2))).#trainneg
cnf(i_0_135, plain, (v2_struct_0(X2)|l3_algstr_0(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2))).#trainneg
cnf(i_0_114, plain, (v2_struct_0(X1)|m1_group_2(k10_group_5(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_129, plain, (v2_struct_0(X1)|m1_group_2(k7_group_2(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_341, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2))).#trainneg
cnf(i_0_374, plain, (l3_algstr_0(k10_group_5(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_375, plain, (v2_group_1(k10_group_5(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_146, plain, (l3_algstr_0(esk10_0))).#trainneg
cnf(i_0_207, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v2_struct_0(esk19_1(X1)))).#trainneg
cnf(i_0_373, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_382, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_175, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_202, plain, (l3_algstr_0(esk18_0))).#trainneg
cnf(i_0_376, plain, (l3_algstr_0(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_377, plain, (v2_group_1(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_379, plain, (r2_hidden(X1,X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_381, plain, (X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_109, plain, (g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))=k7_group_2(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_111, plain, (k8_group_2(X1,X2)=u1_struct_0(X2)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_128, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_132, plain, (v2_struct_0(X1)|m1_subset_1(k8_group_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_395, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_tarski(X3))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_265, plain, (l3_algstr_0(esk38_0))).#trainneg
cnf(i_0_388, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_383, plain, (m1_subset_1(esk7_2(X1,X2),X1)|r1_tarski(X1,X2))).#trainneg
cnf(i_0_400, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_200, plain, (v15_algstr_0(esk18_0))).#trainneg
cnf(i_0_262, plain, (v15_algstr_0(esk38_0))).#trainneg
cnf(i_0_397, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k6_domain_1(X2,X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_386, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(X3))|k1_zfmisc_1(X3)!=k1_tarski(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_156, plain, (v2_struct_0(X1)|v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_163, plain, (v2_struct_0(X1)|v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l3_algstr_0(X1)|~v1_group_1(X1))).#trainneg
cnf(i_0_326, plain, (v2_struct_0(X1)|r2_group_3(X1,X2,X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_110, plain, (k5_group_3(X1,k8_group_2(X1,k7_group_2(X1)),X2)=k7_group_3(X1,X2)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_353, plain, (v2_struct_0(X2)|r2_hidden(X3,k7_group_3(X2,X4))|X1!=X3|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~r2_group_3(X2,X4,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2)))).#trainneg
cnf(i_0_323, plain, (v2_struct_0(X1)|r2_group_3(X1,X2,X3)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_425, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_161, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_353, plain, (r2_hidden(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_387, plain, (k1_tarski(X1)=X2|~m1_subset_1(X2,k1_zfmisc_1(k1_tarski(X1)))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_393, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_tarski(X3))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_222, plain, (v1_xboole_0(esk24_0))).#trainneg
cnf(i_0_436, plain, (r2_group_3(X1,X2,X2)|v2_struct_0(X1)|r1_tarski(u1_struct_0(X1),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_390, plain, (k1_tarski(X1)=k1_tarski(X2)|~r2_hidden(X2,k1_tarski(X1))|~r2_hidden(X1,k1_tarski(X2)))).#trainneg
cnf(i_0_101, plain, (k3_group_3(X1,X2,X3)=a_3_0_group_3(X1,X2,X3)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_105, plain, (k3_group_3(X1,k6_domain_1(u1_struct_0(X1),X3),X2)=k5_group_3(X1,X2,X3)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_106, plain, (v2_struct_0(X2)|r1_group_3(X2,X3,X4)|X3!=k2_group_3(X2,X4,X1)|~v2_group_1(X2)|~v3_group_1(X2)|~l3_algstr_0(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2)))).#trainneg
cnf(i_0_334, plain, (X3=k6_algstr_0(X1,k2_group_1(X1,X2),X4)|v2_struct_0(X1)|k6_algstr_0(X1,X2,X3)!=X4|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_347, plain, (X2=X4|v2_struct_0(X1)|k6_algstr_0(X1,X2,X3)!=k6_algstr_0(X1,X4,X3)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_419, plain, (k7_group_3(X1,X2)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k7_group_3(X1,X2)))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_392, plain, (k7_group_2(X1)=X1|v2_struct_0(X1)|~v15_algstr_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_289, plain, (v1_xboole_0(esk46_0))).#trainneg
cnf(i_0_491, plain, (m1_group_2(X1,X1)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_310, plain, (v3_finseq_1(esk54_0))).#trainneg
cnf(i_0_275, plain, (v1_finseq_1(esk42_0))).#trainneg
cnf(i_0_509, plain, (m1_group_2(k10_group_5(X1),k10_group_5(X1))|v2_struct_0(k10_group_5(X1))|v2_struct_0(X1)|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_503, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k6_domain_1(X2,X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_396, plain, (k6_domain_1(X1,X2)=X3|v1_xboole_0(k1_zfmisc_1(X1))|v1_xboole_0(X1)|k1_zfmisc_1(X1)!=k1_tarski(X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_398, plain, (v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_188, plain, (X1=X2|g3_algstr_0(X1,X3)!=g3_algstr_0(X2,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,k2_zfmisc_1(X1,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X1),X1))))).#trainneg
cnf(i_0_384, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_403, plain, (X1=X2|v2_struct_0(X3)|k7_group_3(X3,X1)!=k1_tarski(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_408, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_412, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,X2))|m1_subset_1(esk57_2(X1,k2_group_1(X1,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_416, plain, (k8_group_2(X1,X2)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k8_group_2(X1,X2)))|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_551, plain, (u1_struct_0(X2)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_group_2(X2,X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_536, plain, (m1_subset_1(u1_struct_0(k7_group_2(X1)),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_145, plain, (l1_struct_0(esk9_0))).#trainneg
cnf(i_0_531, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_404, plain, (k2_group_1(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|u1_struct_0(X1)!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_559, plain, (v2_struct_0(esk9_0)|~v1_xboole_0(u1_struct_0(esk9_0)))).#trainneg
cnf(i_0_170, plain, (v1_card_1(k4_ordinal1))).#trainneg
cnf(i_0_191, plain, (v1_card_1(esk16_0))).#trainneg
cnf(i_0_430, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|v1_xboole_0(k6_domain_1(X1,X3))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,k6_domain_1(X1,X3)))).#trainneg
cnf(i_0_406, plain, (k8_group_2(X1,X2)=X3|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X3)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_407, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k8_group_2(X2,X3))|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_437, negated_conjecture, (r2_group_3(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_439, plain, (k5_group_3(X1,u1_struct_0(k7_group_2(X1)),X2)=k7_group_3(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(k7_group_2(X1),X1))).#trainneg
cnf(i_0_441, plain, (esk58_3(X1,X2,X3)=X1|v2_struct_0(X2)|v1_xboole_0(k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k7_group_3(X2,X3)))).#trainneg
cnf(i_0_442, plain, (esk58_3(X1,X2,X1)=X1|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_446, plain, (m1_subset_1(k2_group_3(X1,X2,X3),k7_group_3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_450, plain, (esk58_3(k2_group_3(X1,X2,X3),X1,X2)=k2_group_3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_454, plain, (r1_tarski(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X3,esk7_2(X1,k7_group_3(X2,X3)))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(esk7_2(X1,k7_group_3(X2,X3)),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_455, plain, (m1_subset_1(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_468, plain, (r2_group_3(X1,k2_group_1(X1,X2),k2_group_1(X1,X2))|r1_tarski(u1_struct_0(X1),X3)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_557, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(k7_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_538, plain, (m1_subset_1(u1_struct_0(k10_group_5(X1)),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_231, plain, (v1_card_1(esk29_0))).#trainneg
cnf(i_0_241, plain, (v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v2_struct_0(esk30_1(X1)))).#trainneg
cnf(i_0_409, plain, (k7_group_3(X1,X2)=X3|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_172, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_214, plain, (v3_ordinal1(esk21_0))).#trainneg
cnf(i_0_569, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_group_2(X3,X2)|~r2_hidden(X1,u1_struct_0(X3))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_413, plain, (m1_subset_1(esk7_2(k1_tarski(X1),X2),X3)|r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_414, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k1_tarski(X3))|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_tarski(X3)))).#trainneg
cnf(i_0_470, plain, (k1_tarski(X1)=k1_tarski(X2)|v1_xboole_0(k1_tarski(X1))|~r2_hidden(X1,k1_tarski(X2))|~m1_subset_1(X2,k1_tarski(X1)))).#trainneg
cnf(i_0_477, plain, (r1_group_3(X1,k2_group_3(X1,X2,X3),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_479, plain, (k6_algstr_0(X2,k2_group_1(X2,X3),k6_algstr_0(X2,X3,X1))=X1|v2_struct_0(X2)|~m1_subset_1(k6_algstr_0(X2,X3,X1),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_486, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,k2_group_1(X1,X2),X3))=X3|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,k2_group_1(X1,X2),X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_488, plain, (r2_hidden(X3,k7_group_3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(esk8_3(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_492, plain, (k5_group_3(X1,k8_group_2(X1,X1),X2)=k7_group_3(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_522, plain, (X1=k6_algstr_0(X2,X3,X4)|v2_struct_0(X2)|k6_algstr_0(X2,X1,X5)!=k6_algstr_0(X2,X3,k6_algstr_0(X2,X4,X5))|~m1_subset_1(k6_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_604, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(k10_group_5(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_233, plain, (v3_ordinal1(esk29_0))).#trainneg
cnf(i_0_418, plain, (k6_domain_1(X1,X2)=X1|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_domain_1(X1,X2)))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_244, plain, (v3_ordinal1(esk32_0))).#trainneg
cnf(i_0_258, plain, (v1_relat_1(esk37_0))).#trainneg
cnf(i_0_606, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X3))|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_420, plain, (esk7_2(k1_zfmisc_1(X1),X2)=X1|r1_tarski(k1_zfmisc_1(X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(esk7_2(k1_zfmisc_1(X1),X2))))).#trainneg
cnf(i_0_421, plain, (m1_subset_1(X1,X2)|r1_tarski(k1_zfmisc_1(X2),X3)|~r2_hidden(X1,esk7_2(k1_zfmisc_1(X2),X3)))).#trainneg
cnf(i_0_544, negated_conjecture, (k6_algstr_0(esk1_0,esk2_0,X1)=k6_algstr_0(esk1_0,X1,esk2_0)|m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_545, negated_conjecture, (m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_549, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),k2_group_1(esk1_0,esk2_0))|m1_subset_1(esk57_2(esk1_0,k2_group_1(esk1_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_556, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(k7_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_570, plain, (m1_subset_1(esk7_2(k8_group_2(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k8_group_2(X1,X2),X3)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_519, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k1_tarski(X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_577, plain, (k5_group_3(X1,u1_struct_0(X1),X2)=k7_group_3(X1,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_280, plain, (v1_relat_1(esk42_0))).#trainneg
cnf(i_0_422, plain, (r1_struct_0(k10_group_5(X1),esk7_2(u1_struct_0(X1),X2))|m1_subset_1(esk57_2(X1,esk7_2(u1_struct_0(X1),X2)),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(u1_struct_0(X1),X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_261, plain, (v1_group_1(esk38_0))).#trainneg
cnf(i_0_235, plain, (v1_ordinal1(esk29_0))).#trainneg
cnf(i_0_715, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m1_subset_1(esk2_0,u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(esk1_0,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_730, plain, (r1_tarski(k7_group_3(X1,X2),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_578, plain, (r2_group_3(X1,X2,esk7_2(k7_group_3(X1,X2),X3))|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~r2_hidden(esk7_2(k7_group_3(X1,X2),X3),k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_580, plain, (r2_group_3(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(k7_group_3(X1,X2))|~r2_hidden(X3,k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,k7_group_3(X1,X2)))).#trainneg
cnf(i_0_582, plain, (r2_group_3(X1,X2,X2)|v2_struct_0(X1)|~r2_hidden(X2,k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_585, plain, (m1_subset_1(X3,k7_group_3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(esk8_3(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_588, plain, (esk58_3(X3,X1,X2)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X3,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_589, plain, (r1_tarski(k1_tarski(X1),k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_590, negated_conjecture, (m1_subset_1(X1,k7_group_3(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_773, plain, (k7_group_3(X1,X2)=u1_struct_0(X1)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_246, plain, (v1_ordinal1(esk32_0))).#trainneg
cnf(i_0_674, negated_conjecture, (k6_algstr_0(esk1_0,X1,X2)=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|X2!=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_429, plain, (v1_xboole_0(X1)|m1_subset_1(esk7_2(k6_domain_1(X1,X2),X3),X1)|r1_tarski(k6_domain_1(X1,X2),X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_234, plain, (v2_ordinal1(esk29_0))).#trainneg
cnf(i_0_245, plain, (v2_ordinal1(esk32_0))).#trainneg
cnf(i_0_821, plain, (r1_tarski(k6_domain_1(X1,X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_432, plain, (k1_tarski(esk7_2(X1,X2))=X3|v1_xboole_0(k1_zfmisc_1(X1))|r1_tarski(X1,X2)|k1_zfmisc_1(X1)!=k1_tarski(X3))).#trainneg
cnf(i_0_433, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(k1_tarski(X1)))|k1_zfmisc_1(k1_tarski(X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_603, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(k10_group_5(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_627, plain, (r2_group_3(X1,k2_group_3(X1,X2,X3),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_632, plain, (r2_group_3(X1,k2_group_3(X1,X2,X3),X2)|v2_struct_0(X1)|~r2_hidden(k2_group_3(X1,X2,X3),k7_group_3(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_634, plain, (r2_group_3(X1,X2,X3)|v2_struct_0(X1)|v1_xboole_0(k7_group_3(X1,X3))|~r2_hidden(X2,k7_group_3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,k7_group_3(X1,X3)))).#trainneg
cnf(i_0_659, plain, (r1_tarski(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r1_group_3(X2,esk7_2(X1,k7_group_3(X2,X3)),X3)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(esk7_2(X1,k7_group_3(X2,X3)),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_210, plain, (v1_xreal_0(esk20_0))).#trainneg
cnf(i_0_669, negated_conjecture, (X1=k2_group_1(esk1_0,X2)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,X1,esk2_0)!=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X2))|~m1_subset_1(k2_group_1(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_434, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(k7_group_3(X3,X1)))|v2_struct_0(X3)|k1_zfmisc_1(k7_group_3(X3,X1))!=k1_tarski(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_512, plain, (v1_xboole_0(X1)|r2_hidden(X2,X1)|~r2_hidden(X2,k1_tarski(X3))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_223, plain, (v1_xreal_0(esk25_0))).#trainneg
cnf(i_0_251, plain, (v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_435, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(X3))|v1_xboole_0(X3)|k1_zfmisc_1(X3)!=k1_tarski(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_449, plain, (k1_tarski(k2_group_3(X1,X2,X3))=X4|v1_xboole_0(k1_zfmisc_1(k7_group_3(X1,X2)))|v2_struct_0(X1)|k1_zfmisc_1(k7_group_3(X1,X2))!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_677, plain, (r2_hidden(esk7_2(k7_group_3(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_678, plain, (r2_hidden(k2_group_3(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_737, negated_conjecture, (esk57_2(esk1_0,esk2_0)=X1|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=u1_struct_0(esk1_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_552, plain, (u1_struct_0(k7_group_2(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k7_group_2(X1))))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_286, plain, (v1_xreal_0(esk46_0))).#trainneg
cnf(i_0_452, plain, (X1=X2|v1_xboole_0(k1_zfmisc_1(X1))|k1_zfmisc_1(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_257, plain, (v1_funct_1(esk37_0))).#trainneg
cnf(i_0_278, plain, (v1_funct_1(esk42_0))).#trainneg
cnf(i_0_514, plain, (v1_xboole_0(X1)|r2_hidden(X2,X1)|v1_xboole_0(k6_domain_1(X1,X3))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,k6_domain_1(X1,X3)))).#trainneg
cnf(i_0_457, plain, (X1=X2|v2_struct_0(X3)|k7_group_3(X3,X4)!=k1_tarski(X2)|~r2_group_3(X3,X4,X1)|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_458, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(k7_group_3(X3,X4)))|v2_struct_0(X3)|k1_zfmisc_1(k7_group_3(X3,X4))!=k1_tarski(X2)|~r2_group_3(X3,X4,X1)|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_785, plain, (r2_group_3(X1,X2,k2_group_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_788, plain, (r2_group_3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_793, plain, (m1_subset_1(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r2_group_3(X2,X1,X3))).#trainneg
cnf(i_0_800, plain, (r2_hidden(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(X4))|~r2_group_3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_553, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k10_group_5(X1))))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_232, plain, (v1_finset_1(esk29_0))).#trainneg
cnf(i_0_907, negated_conjecture, (X1=X2|k7_group_3(esk1_0,X1)!=k1_tarski(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_459, plain, (esk58_3(X1,X2,X3)=X1|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r2_group_3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2)))).#trainneg
cnf(i_0_276, plain, (v1_finset_1(esk42_0))).#trainneg
cnf(i_0_281, plain, (v7_ordinal1(esk43_0))).#trainneg
cnf(i_0_826, plain, (k6_domain_1(X1,X2)=X1|v1_xboole_0(X1)|~r1_tarski(X1,k6_domain_1(X1,X2))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_460, plain, (esk7_2(k1_zfmisc_1(k1_tarski(X1)),X2)=k1_tarski(X1)|r1_tarski(k1_zfmisc_1(k1_tarski(X1)),X2)|~r2_hidden(X1,esk7_2(k1_zfmisc_1(k1_tarski(X1)),X2)))).#trainneg
cnf(i_0_846, plain, (r1_tarski(k1_tarski(X1),k7_group_3(X2,X3))|v2_struct_0(X2)|~r1_group_3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_847, plain, (r1_tarski(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~m1_subset_1(esk7_2(X1,k7_group_3(X2,X3)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r2_group_3(X2,esk7_2(X1,k7_group_3(X2,X3)),X3))).#trainneg
cnf(i_0_291, plain, (v7_ordinal1(esk48_0))).#trainneg
cnf(i_0_462, plain, (k6_domain_1(k1_tarski(X1),X2)=k1_tarski(X1)|v1_xboole_0(k1_tarski(X1))|~r2_hidden(X1,k6_domain_1(k1_tarski(X1),X2))|~m1_subset_1(X2,k1_tarski(X1)))).#trainneg
cnf(i_0_209, plain, (v1_int_1(esk20_0))).#trainneg
cnf(i_0_243, plain, (v1_int_1(esk31_0))).#trainneg
cnf(i_0_695, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(k10_group_5(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k10_group_5(X2))))).#trainneg
cnf(i_0_830, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(k10_group_5(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k10_group_5(X2))))).#trainneg
cnf(i_0_463, plain, (r2_hidden(esk7_2(k1_tarski(X1),X2),X3)|r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_464, plain, (r2_hidden(X1,X2)|v1_xboole_0(k1_tarski(X3))|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_tarski(X3)))).#trainneg
cnf(i_0_263, plain, (v8_struct_0(esk38_0))).#trainneg
cnf(i_0_469, plain, (k1_tarski(esk7_2(k1_tarski(X1),X2))=k1_tarski(X1)|r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,k1_tarski(esk7_2(k1_tarski(X1),X2))))).#trainneg
cnf(i_0_862, negated_conjecture, (k1_tarski(esk2_0)=X1|v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_211, plain, (v1_xcmplx_0(esk20_0))).#trainneg
cnf(i_0_254, plain, (v1_xcmplx_0(esk35_0))).#trainneg
cnf(i_0_471, plain, (esk58_3(X1,X2,X3)=X4|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|u1_struct_0(X2)!=k1_tarski(X4)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_472, plain, (r1_struct_0(k10_group_5(X1),esk58_3(X2,X1,X3))|m1_subset_1(esk57_2(X1,esk58_3(X2,X1,X3)),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,k7_group_3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_943, plain, (r2_hidden(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(X4))|~r1_group_3(X2,X4,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_945, plain, (r1_tarski(k1_tarski(X1),k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_994, plain, (esk58_3(esk7_2(k1_tarski(X1),X2),X3,X4)=esk7_2(k1_tarski(X1),X2)|v2_struct_0(X3)|r1_tarski(k1_tarski(X1),X2)|~m1_subset_1(X4,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~r2_hidden(X1,k7_group_3(X3,X4)))).#trainneg
cnf(i_0_1036, plain, (r2_hidden(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(X4))|~r2_group_3(X2,X4,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_1037, plain, (r2_hidden(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_706, plain, (X1=X2|v2_struct_0(X3)|u1_struct_0(X3)!=k1_tarski(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_527, plain, (u1_struct_0(X1)=X2|v2_struct_0(X1)|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X2,X3)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_288, plain, (v1_xcmplx_0(esk46_0))).#trainneg
cnf(i_0_1066, negated_conjecture, (esk2_0=X1|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_475, plain, (a_3_0_group_3(X1,k6_domain_1(u1_struct_0(X1),X2),X3)=k5_group_3(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_567, plain, (v1_xboole_0(k1_tarski(X2))|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m1_subset_1(X3,k1_tarski(X2))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_212, plain, (v1_xxreal_0(esk20_0))).#trainneg
cnf(i_0_253, plain, (v1_xxreal_0(esk35_0))).#trainneg
cnf(i_0_710, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_541, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k7_group_3(X2,X3)))).#trainneg
cnf(i_0_476, plain, (k3_group_3(X1,k1_tarski(X2),X3)=k5_group_3(X1,X3,X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_482, plain, (esk8_3(X1,X2,X3)=X4|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|u1_struct_0(X1)!=k1_tarski(X4)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1082, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(esk1_0,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_287, plain, (v1_xxreal_0(esk46_0))).#trainneg
cnf(i_0_256, plain, (v2_finseq_1(esk37_0))).#trainneg
cnf(i_0_274, plain, (v2_finseq_1(esk42_0))).#trainneg
cnf(i_0_171, plain, (v2_card_1(k4_ordinal1))).#trainneg
cnf(i_0_331, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_252, plain, (v2_xxreal_0(esk35_0))).#trainneg
cnf(i_0_346, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_248, plain, (v1_xboole_0(esk33_1(X1)))).#trainneg
cnf(i_0_192, plain, (v1_finseq_1(esk17_1(X1)))).#trainneg
cnf(i_0_313, plain, (v1_finseq_1(esk55_1(X1)))).#trainneg
cnf(i_0_1092, plain, (esk46_0=k1_xboole_0)).#trainneg
cnf(i_0_1093, plain, (esk24_0=k1_xboole_0)).#trainneg
cnf(i_0_286, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_483, plain, (r1_struct_0(k10_group_5(X1),esk8_3(X1,X2,X3))|m1_subset_1(esk57_2(X1,esk8_3(X1,X2,X3)),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_287, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_288, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_328, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_679, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k7_group_3(X2,X3)))).#trainneg
cnf(i_0_500, plain, (r2_hidden(X1,X2)|r1_tarski(k1_zfmisc_1(X2),X3)|~r2_hidden(X1,esk7_2(k1_zfmisc_1(X2),X3)))).#trainneg
cnf(i_0_196, plain, (v1_relat_1(esk17_1(X1)))).#trainneg
cnf(i_0_269, plain, (v1_relat_1(esk39_1(X1)))).#trainneg
cnf(i_0_318, plain, (v1_relat_1(esk55_1(X1)))).#trainneg
cnf(i_0_193, plain, (v1_funct_1(esk17_1(X1)))).#trainneg
cnf(i_0_267, plain, (v1_funct_1(esk39_1(X1)))).#trainneg
cnf(i_0_315, plain, (v1_funct_1(esk55_1(X1)))).#trainneg
cnf(i_0_314, plain, (v1_finset_1(esk55_1(X1)))).#trainneg
cnf(i_0_266, plain, (v5_ordinal1(esk39_1(X1)))).#trainneg
cnf(i_0_312, plain, (v2_finseq_1(esk55_1(X1)))).#trainneg
cnf(i_0_329, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_335, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_213, plain, (m1_subset_1(esk20_0,k1_numbers))).#trainneg
cnf(i_0_776, plain, (k1_tarski(X1)=X2|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_tarski(X1)))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_1094, plain, (esk33_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_279, plain, (v4_relat_1(esk42_0,k5_numbers))).#trainneg
cnf(i_0_1106, plain, (np__1=X1|k1_tarski(X1)!=k1_numbers)).#trainneg
cnf(i_0_501, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k8_group_2(X2,X3))|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_506, plain, (m1_group_2(esk38_0,esk38_0))).#trainneg
cnf(i_0_507, plain, (m1_group_2(esk18_0,esk18_0))).#trainneg
cnf(i_0_1129, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_group_2(X3,X2)|~r2_hidden(X1,u1_struct_0(X3))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_1143, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X3))|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(X3)))).#trainneg
cnf(i_0_508, plain, (m1_group_2(k7_group_2(X1),k7_group_2(X1))|v2_struct_0(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_513, plain, (v1_xboole_0(X1)|r2_hidden(esk7_2(k6_domain_1(X1,X2),X3),X1)|r1_tarski(k6_domain_1(X1,X2),X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1131, plain, (r2_hidden(esk7_2(k8_group_2(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k8_group_2(X1,X2),X3)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1205, plain, (r1_tarski(k8_group_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1211, plain, (r1_tarski(u1_struct_0(X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_85, plain, (v6_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_29, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1214, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_126, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1117, plain, (np__1=X1|k1_tarski(X1)!=k5_numbers)).#trainneg
cnf(i_0_1221, plain, (k1_numbers=k5_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_330, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_521, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),X5)=k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X4,X5))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_40, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_43, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_18, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_69, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_23, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_39, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_42, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_56, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_76, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_33, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1224, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1248, plain, (v1_xreal_0(esk31_0))).#trainneg
cnf(i_0_1250, plain, (v1_finset_1(esk48_0))).#trainneg
cnf(i_0_1121, plain, (esk20_0=X1|k1_tarski(X1)!=k1_numbers)).#trainneg
cnf(i_0_528, plain, (u1_struct_0(X1)=X2|v2_struct_0(X1)|g3_algstr_0(u1_struct_0(X1),u2_algstr_0(X1))!=g3_algstr_0(X2,X3)|~v1_funct_1(u2_algstr_0(X1))|~m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~v1_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1251, plain, (v1_finset_1(esk43_0))).#trainneg
cnf(i_0_738, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk57_2(esk1_0,esk2_0))|m1_subset_1(esk57_2(esk1_0,esk57_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_742, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),k2_group_1(esk1_0,esk57_2(esk1_0,esk2_0)))|m1_subset_1(esk57_2(esk1_0,k2_group_1(esk1_0,esk57_2(esk1_0,esk2_0))),u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_1252, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_523, plain, (X1=X2|v2_struct_0(X3)|k6_algstr_0(X3,X4,k6_algstr_0(X3,X5,X1))!=k6_algstr_0(X3,k6_algstr_0(X3,X4,X5),X2)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(k6_algstr_0(X3,X4,X5),u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X5,u1_struct_0(X3))|~m1_subset_1(X4,u1_struct_0(X3)))).#trainneg
cnf(i_0_524, plain, (X1=k6_algstr_0(X2,k2_group_1(X2,k6_algstr_0(X2,X3,X4)),X5)|v2_struct_0(X2)|k6_algstr_0(X2,X3,k6_algstr_0(X2,X4,X1))!=X5|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(k6_algstr_0(X2,X3,X4),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_1260, plain, (~v1_xboole_0(esk36_0))).#trainneg
cnf(i_0_50, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_55, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_84, plain, (v4_funct_1(X1)|~v3_finseq_1(X1))).#trainneg
cnf(i_0_293, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk49_1(X1)))).#trainneg
cnf(i_0_359, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_160, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_304, plain, (v1_finseq_1(esk53_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_230, plain, (l1_struct_0(esk28_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_298, plain, (v1_relat_1(esk50_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_1223, plain, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1059, negated_conjecture, (esk57_2(esk1_0,esk2_0)=X1|k1_tarski(X1)!=u1_struct_0(esk1_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_546, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,esk58_3(X2,X1,X3)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,esk58_3(X2,X1,X3))),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,k7_group_3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1153, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(esk1_0,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1218, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X1,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_547, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,esk8_3(X1,X2,X3)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,esk8_3(X1,X2,X3))),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1307, plain, (u1_struct_0(k7_group_2(X1))=u1_struct_0(X1)|v2_struct_0(k7_group_2(X1))|v2_struct_0(X1)|~m1_group_2(X1,k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v3_group_1(k7_group_2(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_1308, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v2_struct_0(k10_group_5(X1))|v2_struct_0(X1)|~m1_group_2(X1,k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v3_group_1(k10_group_5(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_309, plain, (v1_relat_1(esk53_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_144, plain, (v1_funct_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_297, plain, (v1_funct_1(esk50_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_307, plain, (v1_funct_1(esk53_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_306, plain, (v1_finset_1(esk53_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_303, plain, (v2_finseq_1(esk53_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_88, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_89, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_19, plain, (v1_finseq_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1292, plain, (v2_struct_0(esk28_1(X1))|~v1_xboole_0(u1_struct_0(esk28_1(X1)))|~v1_card_1(X1))).#trainneg
cnf(i_0_1293, plain, (r1_tarski(X1,k1_numbers)|~r2_hidden(esk7_2(X1,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1326, plain, (r1_tarski(k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1119, plain, (k1_xboole_0=X1|k1_tarski(X1)!=k5_numbers)).#trainneg
cnf(i_0_548, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,esk7_2(u1_struct_0(X1),X2)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,esk7_2(u1_struct_0(X1),X2))),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(u1_struct_0(X1),X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1329, plain, (k1_numbers=k5_numbers|~r1_tarski(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1144, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_709, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(k10_group_5(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k10_group_5(X2)))|~l3_algstr_0(k10_group_5(X2))|~v3_group_1(k10_group_5(X2))|~v2_group_1(k10_group_5(X2)))).#trainneg
cnf(i_0_550, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,k2_group_1(X1,X2)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,k2_group_1(X1,X2))),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_554, plain, (u1_struct_0(k7_group_2(X1))=X2|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1347, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(esk1_0,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_17, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_34, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_21, plain, (v1_group_1(X1)|~v2_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_65, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_52, plain, (v8_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_70, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_153, plain, (v3_card_1(k1_tarski(X1),np__1))).#trainneg
cnf(i_0_1374, plain, (v7_ordinal1(esk29_0))).#trainneg
cnf(i_0_1373, plain, (v7_ordinal1(esk32_0)|~v1_finset_1(esk32_0))).#trainneg
cnf(i_0_1375, plain, (v7_ordinal1(esk21_0)|~v1_finset_1(esk21_0))).#trainneg
cnf(i_0_1297, plain, (k1_tarski(X1)=X2|k1_tarski(X2)!=k1_zfmisc_1(k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1367, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),k2_group_1(esk1_0,k2_group_1(esk1_0,esk2_0)))|m1_subset_1(esk57_2(esk1_0,k2_group_1(esk1_0,k2_group_1(esk1_0,esk2_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1379, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_571, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(k8_group_2(X2,X3))|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k8_group_2(X2,X3)))).#trainneg
cnf(i_0_195, plain, (v4_relat_1(esk17_1(X1),k5_numbers))).#trainneg
cnf(i_0_317, plain, (v4_relat_1(esk55_1(X1),k5_numbers))).#trainneg
cnf(i_0_78, plain, (v3_card_1(X1,k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_194, plain, (v5_relat_1(esk17_1(X1),X1))).#trainneg
cnf(i_0_268, plain, (v5_relat_1(esk39_1(X1),X1))).#trainneg
cnf(i_0_316, plain, (v5_relat_1(esk55_1(X1),X1))).#trainneg
cnf(i_0_220, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk23_1(X1)))).#trainneg
cnf(i_0_294, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk49_1(X1)))).#trainneg
cnf(i_0_299, plain, (v1_zfmisc_1(X1)|~v1_zfmisc_1(esk51_1(X1)))).#trainneg
cnf(i_0_179, plain, (v1_finset_1(X1)|~v1_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1403, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1386, plain, (v7_ordinal1(esk16_0)|~v1_finset_1(esk16_0))).#trainneg
cnf(i_0_850, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_1(esk1_0,X1)=k2_group_1(esk1_0,X2)|k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1))!=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X2))|~m1_subset_1(k2_group_1(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1409, negated_conjecture, (k6_algstr_0(esk1_0,esk12_1(u1_struct_0(esk1_0)),esk2_0)=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1408, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk12_1(u1_struct_0(esk1_0)),esk2_0)=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_591, plain, (m1_subset_1(esk58_3(X1,X2,X3),k7_group_3(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(esk58_3(X1,X2,X3),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r2_hidden(X1,k7_group_3(X2,X3)))).#trainneg
cnf(i_0_272, plain, (v1_finset_1(X1)|~v1_finset_1(esk41_1(X1)))).#trainneg
cnf(i_0_224, plain, (v2_struct_0(X1)|v1_zfmisc_1(esk26_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_301, plain, (v1_xboole_0(X1)|v3_card_1(esk52_1(X1),np__1))).#trainneg
cnf(i_0_10, plain, (v1_xboole_0(X1)|v3_card_1(X1,np__1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_71, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_82, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_1475, plain, (v1_zfmisc_1(k1_tarski(X1)))).#trainneg
cnf(i_0_1322, plain, (r1_tarski(k1_tarski(X1),k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1360, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),k2_group_1(esk1_0,k2_group_1(esk1_0,esk57_2(esk1_0,esk2_0))))|m1_subset_1(esk57_2(esk1_0,k2_group_1(esk1_0,k2_group_1(esk1_0,esk57_2(esk1_0,esk2_0)))),u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_1452, negated_conjecture, (esk2_0=X1|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk12_1(u1_struct_0(esk1_0)),X1)!=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_592, plain, (r2_group_3(X1,k2_group_1(X1,esk58_3(X2,X1,X3)),k2_group_1(X1,esk58_3(X2,X1,X3)))|r1_tarski(u1_struct_0(X1),X4)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,k7_group_3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1436, negated_conjecture, (k6_algstr_0(esk1_0,k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0))),k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))))=esk2_0|m1_subset_1(esk57_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_831, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(k10_group_5(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k10_group_5(X2)))|~l3_algstr_0(k10_group_5(X2))|~v3_group_1(k10_group_5(X2))|~v2_group_1(k10_group_5(X2)))).#trainneg
cnf(i_0_593, plain, (r2_group_3(X1,k2_group_1(X1,esk8_3(X1,X2,X3)),k2_group_1(X1,esk8_3(X1,X2,X3)))|r1_tarski(u1_struct_0(X1),X4)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_594, plain, (r2_group_3(X1,k2_group_1(X1,esk7_2(u1_struct_0(X1),X2)),k2_group_1(X1,esk7_2(u1_struct_0(X1),X2)))|r1_tarski(u1_struct_0(X1),X3)|v2_struct_0(X1)|r1_tarski(u1_struct_0(X1),X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_28, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_47, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_80, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_159, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_176, plain, (v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_178, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v8_struct_0(X1))).#trainneg
cnf(i_0_284, plain, (v1_xboole_0(X1)|v1_subset_1(esk45_1(X1),X1))).#trainneg
cnf(i_0_30, plain, (v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_66, plain, (v2_finseq_1(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_308, plain, (v4_relat_1(esk53_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_249, plain, (m1_subset_1(esk33_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_271, plain, (m1_subset_1(esk40_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1570, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1578, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk40_1(X2)))).#trainneg
cnf(i_0_1549, plain, (v1_xboole_0(u1_struct_0(esk9_0))|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_596, plain, (r2_group_3(X1,k2_group_1(X1,k2_group_1(X1,X2)),k2_group_1(X1,k2_group_1(X1,X2)))|r1_tarski(u1_struct_0(X1),X3)|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1587, plain, (u1_struct_0(esk9_0)=k1_xboole_0|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_1543, plain, (~m1_subset_1(esk36_0,k5_numbers))).#trainneg
cnf(i_0_1544, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_1550, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1553, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_598, plain, (m1_subset_1(esk7_2(u1_struct_0(k7_group_2(X1)),X2),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(u1_struct_0(k7_group_2(X1)),X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_599, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(k7_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k7_group_2(X2))))).#trainneg
cnf(i_0_81, plain, (v13_struct_0(X1,k1_xboole_0)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_83, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_290, plain, (v3_card_1(esk47_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_296, plain, (v3_card_1(esk50_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_305, plain, (v3_card_1(esk53_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_229, plain, (v13_struct_0(esk28_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_225, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk26_1(X1)))).#trainneg
cnf(i_0_282, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk44_1(X1)))).#trainneg
cnf(i_0_227, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(esk27_1(X1)))).#trainneg
cnf(i_0_218, plain, (v1_xboole_0(esk22_2(X1,X2)))).#trainneg
cnf(i_0_217, plain, (v1_relat_1(esk22_2(X1,X2)))).#trainneg
cnf(i_0_221, plain, (v1_xboole_0(X1)|m1_subset_1(esk23_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1559, plain, (v1_finset_1(esk17_1(X1)))).#trainneg
cnf(i_0_1552, plain, (v1_finset_1(u1_struct_0(esk38_0))|~l1_struct_0(esk38_0))).#trainneg
cnf(i_0_1451, negated_conjecture, (X1=esk12_1(u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,X1,esk2_0)!=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_600, plain, (u1_struct_0(k10_group_5(X1))=X2|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1569, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1545, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1648, plain, (~v1_xboole_0(esk47_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_1554, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_605, plain, (m1_subset_1(esk7_2(u1_struct_0(X1),X2),u1_struct_0(X3))|v2_struct_0(X3)|r1_tarski(u1_struct_0(X1),X2)|~m1_group_2(X1,X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_607, plain, (k1_tarski(X1)=esk7_2(k1_tarski(X2),X3)|r1_tarski(k1_tarski(X2),X3)|~r2_hidden(X1,esk7_2(k1_tarski(X2),X3))|~r2_hidden(X2,k1_zfmisc_1(k1_tarski(X1))))).#trainneg
cnf(i_0_285, plain, (v1_xboole_0(X1)|m1_subset_1(esk45_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_295, plain, (v1_xboole_0(X1)|m1_subset_1(esk49_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_302, plain, (v1_xboole_0(X1)|m1_subset_1(esk52_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_300, plain, (v1_zfmisc_1(X1)|m1_subset_1(esk51_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_273, plain, (v1_finset_1(X1)|m1_subset_1(esk41_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13, plain, (v2_struct_0(X1)|v13_struct_0(X1,np__1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_86, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,k1_xboole_0))).#trainneg
cnf(i_0_15, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_270, plain, (~v1_subset_1(esk40_1(X1),X1))).#trainneg
cnf(i_0_11, plain, (v1_finseq_1(X2)|~v3_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_61, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_59, plain, (v4_relat_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1812, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_1815, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_1795, plain, (v3_ordinal1(esk20_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_608, plain, (X1=esk7_2(k1_tarski(X2),X3)|r1_tarski(k1_tarski(X2),X3)|~m1_subset_1(X1,k1_zfmisc_1(esk7_2(k1_tarski(X2),X3)))|~r2_hidden(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1755, plain, (v1_finseq_1(esk20_0)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_1651, plain, (~v1_xboole_0(esk50_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_1654, plain, (~v1_xboole_0(esk53_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1604, plain, (X1=u1_struct_0(X2)|~v1_xboole_0(X1)|~l3_algstr_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_610, plain, (r2_hidden(X1,X2)|r1_tarski(k1_tarski(X3),X4)|~r2_hidden(X1,esk7_2(k1_tarski(X3),X4))|~r2_hidden(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_611, plain, (m1_subset_1(X1,X2)|r1_tarski(k1_tarski(X3),X4)|~r2_hidden(X1,esk7_2(k1_tarski(X3),X4))|~r2_hidden(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_74, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_finseq_1(X1))).#trainneg
cnf(i_0_16, plain, (~v2_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_204, plain, (v2_struct_0(X1)|v2_group_1(esk19_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_237, plain, (v2_struct_0(X1)|v2_group_1(esk30_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_203, plain, (v2_struct_0(X1)|v3_group_1(esk19_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_236, plain, (v2_struct_0(X1)|v3_group_1(esk30_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_206, plain, (v2_struct_0(X1)|v15_algstr_0(esk19_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_239, plain, (v2_struct_0(X1)|v15_algstr_0(esk30_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_205, plain, (v2_struct_0(X1)|v1_group_1(esk19_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_238, plain, (v2_struct_0(X1)|v1_group_1(esk30_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_240, plain, (v2_struct_0(X1)|v8_struct_0(esk30_1(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_99, plain, (v3_group_1(X1)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1847, plain, (u1_struct_0(esk9_0)=u1_struct_0(X1)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_1862, plain, (v4_relat_1(esk37_0,k5_numbers))).#trainneg
cnf(i_0_1825, plain, (v7_ordinal1(k1_xboole_0)|~v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_1460, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk12_1(u1_struct_0(esk1_0))=k2_group_1(esk1_0,X1)|k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))!=k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1))|~m1_subset_1(k2_group_1(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_612, plain, (v1_xboole_0(k6_domain_1(X1,X2))|v1_xboole_0(X1)|m1_subset_1(esk7_2(k1_tarski(X3),X4),X1)|r1_tarski(k1_tarski(X3),X4)|~m1_subset_1(X2,X1)|~r2_hidden(X3,k6_domain_1(X1,X2)))).#trainneg
cnf(i_0_1772, plain, (v1_finseq_1(k1_xboole_0)|~v3_finseq_1(k5_numbers))).#trainneg
cnf(i_0_1864, plain, (~l1_struct_0(esk28_1(np__1))|~v2_struct_0(esk28_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_1775, plain, (v1_finseq_1(np__1)|~v3_finseq_1(k5_numbers))).#trainneg
cnf(i_0_1848, plain, (u1_struct_0(X1)=u1_struct_0(X2)|~l3_algstr_0(X2)|~v2_struct_0(X2)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_614, plain, (r1_struct_0(k10_group_5(X1),esk7_2(k1_tarski(X2),X3))|m1_subset_1(esk57_2(X1,esk7_2(k1_tarski(X2),X3)),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k1_tarski(X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_616, plain, (r2_group_3(X1,k2_group_1(X1,esk7_2(k1_tarski(X2),X3)),k2_group_1(X1,esk7_2(k1_tarski(X2),X3)))|r1_tarski(u1_struct_0(X1),X4)|v2_struct_0(X1)|r1_tarski(k1_tarski(X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_98, plain, (v3_group_1(X1)|m1_subset_1(esk5_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_97, plain, (v3_group_1(X1)|m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_147, plain, (v2_struct_0(X1)|m1_group_2(esk11_1(X1),X1)|~v2_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_216, plain, (v4_relat_1(esk22_2(X1,X2),X1))).#trainneg
cnf(i_0_215, plain, (v5_relat_1(esk22_2(X1,X2),X2))).#trainneg
cnf(i_0_190, plain, (k7_group_2(k7_group_2(X1))=k7_group_2(X1)|v2_struct_0(X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_27, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_64, plain, (v1_zfmisc_1(X2)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_14, plain, (v3_finseq_1(X2)|~v3_finseq_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12, plain, (v6_ordinal1(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1980, plain, (l3_algstr_0(esk11_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1826, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1))).#trainneg
cnf(i_0_617, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,esk7_2(k1_tarski(X2),X3)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,esk7_2(k1_tarski(X2),X3))),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k1_tarski(X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r2_hidden(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1776, plain, (v1_finseq_1(np__1)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_1765, negated_conjecture, (v1_finseq_1(esk2_0)|~v3_finseq_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2096, negated_conjecture, (v1_finseq_1(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_619, plain, (esk7_2(k1_tarski(X1),X2)=X3|v1_xboole_0(X4)|r1_tarski(k1_tarski(X1),X2)|X4!=k1_tarski(X3)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_622, plain, (v1_xboole_0(k1_tarski(X1))|m1_subset_1(esk7_2(k1_tarski(X2),X3),X4)|r1_tarski(k1_tarski(X2),X3)|~r2_hidden(X1,X4)|~r2_hidden(X2,k1_tarski(X1)))).#trainneg
cnf(i_0_9, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2130, plain, (v2_struct_0(X2)|~r2_hidden(k7_group_3(X2,X1),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_2131, plain, (v2_struct_0(X2)|~r2_hidden(u1_struct_0(X2),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_155, plain, (v2_struct_0(X1)|v2_binop_1(u2_algstr_0(X1),u1_struct_0(X1))|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_162, plain, (v2_struct_0(X1)|v1_setwiseo(u2_algstr_0(X1),u1_struct_0(X1))|~l3_algstr_0(X1)|~v1_group_1(X1))).#trainneg
cnf(i_0_208, plain, (v2_struct_0(X1)|m1_group_2(esk19_1(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_242, plain, (v2_struct_0(X1)|m1_group_2(esk30_1(X1),X1)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_137, plain, (v2_struct_0(X2)|~v2_struct_0(X1)|~v2_group_1(X2)|~l3_algstr_0(X2)|~m1_group_2(X1,X2))).#trainneg
cnf(i_0_2163, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k7_group_2(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_2165, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(k10_group_5(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_226, plain, (v2_struct_0(X1)|m1_subset_1(esk26_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_283, plain, (v2_struct_0(X1)|m1_subset_1(esk44_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2161, plain, (v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v2_struct_0(esk11_1(X1)))).#trainneg
cnf(i_0_1561, plain, (k1_xboole_0=X1|k1_zfmisc_1(X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_2033, plain, (v1_zfmisc_1(k5_numbers)|~v1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_1805, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2192, negated_conjecture, (v3_ordinal1(esk2_0)|~v7_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_625, plain, (k1_tarski(esk7_2(k1_tarski(X1),X2))=k1_tarski(X1)|v1_xboole_0(k1_tarski(esk7_2(k1_tarski(X1),X2)))|r1_tarski(k1_tarski(X1),X2)|~m1_subset_1(X1,k1_tarski(esk7_2(k1_tarski(X1),X2))))).#trainneg
cnf(i_0_626, plain, (k1_tarski(X1)=k1_tarski(X2)|v1_xboole_0(k1_tarski(X1))|v1_xboole_0(k1_tarski(X2))|~m1_subset_1(X2,k1_tarski(X1))|~m1_subset_1(X1,k1_tarski(X2)))).#trainneg
cnf(i_0_228, plain, (v7_struct_0(X1)|m1_subset_1(esk27_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_167, plain, (v2_struct_0(X1)|v1_finseqop(u2_algstr_0(X1),u1_struct_0(X1))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_22, plain, (v2_struct_0(X1)|v3_group_1(X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_2219, plain, (v3_group_1(k7_group_2(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2221, plain, (v3_group_1(k10_group_5(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_151, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_152, plain, (v1_xboole_0(X2)|v1_finset_1(X1)|~v1_finset_1(k2_zfmisc_1(X2,X1)))).#trainneg
cnf(i_0_150, plain, (v1_xboole_0(X2)|v1_finset_1(X1)|~v1_finset_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_154, plain, (v3_card_1(u1_struct_0(X2),X1)|~l1_struct_0(X2)|~v1_card_1(X1)|~v13_struct_0(X2,X1))).#trainneg
cnf(i_0_38, plain, (v1_xboole_0(X1)|v1_subset_1(X2,X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_41, plain, (v2_struct_0(X1)|v8_struct_0(X2)|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~v8_struct_0(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_1981, plain, (v2_group_1(esk11_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1562, plain, (k1_tarski(k1_xboole_0)=X1|k1_zfmisc_1(k1_zfmisc_1(X2))!=k1_tarski(X1))).#trainneg
cnf(i_0_630, plain, (m1_subset_1(X1,k7_group_3(X2,esk58_3(X3,X2,X1)))|v2_struct_0(X2)|~m1_subset_1(esk58_3(X3,X2,X1),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r2_hidden(X3,k7_group_3(X2,X1)))).#trainneg
cnf(i_0_2054, plain, (v3_finseq_1(k5_numbers)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_2075, plain, (v6_ordinal1(k5_numbers)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2193, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2256, negated_conjecture, (v7_ordinal1(esk2_0)|~v1_finset_1(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_635, plain, (k6_algstr_0(X1,k2_group_1(X1,X2),k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4))=k6_algstr_0(X1,X3,X4)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_636, plain, (k6_algstr_0(X1,k2_group_1(X1,k2_group_1(X1,X2)),X3)=k6_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2249, plain, (m1_subset_1(X1,k7_group_3(X2,k2_group_3(X2,X1,X3)))|v2_struct_0(X2)|~r2_hidden(k2_group_3(X2,X1,X3),k7_group_3(X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_2252, plain, (m1_subset_1(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X3,k7_group_3(X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r2_group_3(X2,X1,X3)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_2324, plain, (m1_subset_1(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(esk8_3(X2,X3,X1),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r1_group_3(X2,X3,X1)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_90, plain, (X2=k1_tarski(X1)|esk3_2(X1,X2)=X1|r2_hidden(esk3_2(X1,X2),X2))).#trainneg
cnf(i_0_345, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_2149, plain, (l3_algstr_0(esk19_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_1571, plain, (esk40_1(X1)=X2|k1_zfmisc_1(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_637, plain, (X1=k2_group_1(X2,X3)|v2_struct_0(X2)|k6_algstr_0(X2,X1,k6_algstr_0(X2,X3,X4))!=X4|~m1_subset_1(k2_group_1(X2,X3),u1_struct_0(X2))|~m1_subset_1(k6_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_1827, plain, (v7_ordinal1(esk20_0)|~v1_finset_1(esk20_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1774, plain, (v1_finseq_1(k5_numbers)|~v3_finseq_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2194, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_54, plain, (~v1_xboole_0(X1)|~v1_subset_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_339, plain, (r2_hidden(esk56_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_197, plain, (m1_subset_1(esk17_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_91, plain, (X2=k1_tarski(X1)|esk3_2(X1,X2)!=X1|~r2_hidden(esk3_2(X1,X2),X2))).#trainneg
cnf(i_0_219, plain, (m1_subset_1(esk22_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_338, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk56_2(X3,X2)))).#trainneg
cnf(i_0_25, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_321, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_149, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk13_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_143, plain, (v1_funct_2(u2_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_1910, plain, (u1_struct_0(X1)=u1_struct_0(esk10_0)|~l3_algstr_0(X1)|~v2_struct_0(esk10_0)|~v2_struct_0(X1))).#trainneg
cnf(i_0_2543, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_2415, negated_conjecture, (X1=k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0)))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,X1,k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2433, plain, (v7_ordinal1(esk20_0)|~v1_finset_1(esk20_0)|~v1_card_1(k1_numbers))).#trainneg
cnf(i_0_2431, plain, (v7_ordinal1(esk20_0)|~v1_finset_1(esk20_0)|~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2240, plain, (v8_struct_0(k10_group_5(X1))|v2_struct_0(X1)|~v8_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2097, negated_conjecture, (v1_finset_1(esk2_0)|~v1_funct_1(esk2_0)|~v1_relat_1(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_643, plain, (k6_algstr_0(X1,k2_group_1(X1,k6_algstr_0(X1,X2,X3)),k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)))=X4|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_51, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_44, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_36, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_35, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_322, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_139, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_125, plain, (v2_struct_0(X1)|m1_subset_1(k5_group_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1572, plain, (k1_tarski(esk40_1(X1))=X2|k1_zfmisc_1(k1_zfmisc_1(X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_639, plain, (k6_algstr_0(X1,k2_group_1(X1,k2_group_1(X1,X2)),X3)=k6_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2195, negated_conjecture, (v7_ordinal1(esk2_0)|~v1_finset_1(esk2_0)|~v7_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2034, plain, (v1_zfmisc_1(k6_domain_1(X1,X2))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2055, plain, (v3_finseq_1(k6_domain_1(X1,X2))|v1_xboole_0(X1)|~v3_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_647, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,k2_group_1(X1,k6_algstr_0(X1,X2,X3)),X4)))=X4|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,k2_group_1(X1,k6_algstr_0(X1,X2,X3)),X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_123, plain, (v2_struct_0(X1)|m1_subset_1(k3_group_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v2_group_1(X1)|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_142, plain, (m1_subset_1(u2_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_2874, plain, (v5_relat_1(u2_algstr_0(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_2921, plain, (k2_group_3(X1,k6_algstr_0(X1,X2,X3),X2)=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_2959, plain, (k2_group_3(X1,X2,X2)=X2|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2979, plain, (esk58_3(k6_algstr_0(X1,X3,X2),X1,k6_algstr_0(X1,X2,X3))=k6_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_2983, plain, (r1_group_3(X1,X2,k6_algstr_0(X1,X3,X4))|v2_struct_0(X1)|X2!=k6_algstr_0(X1,X4,X3)|~m1_subset_1(k6_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_2990, plain, (k2_group_3(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),k6_algstr_0(X1,X2,X3))=k6_algstr_0(X1,X4,k6_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2154, plain, (l3_algstr_0(esk30_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v2_group_1(X1)|~v3_group_1(X1))).#trainneg
cnf(i_0_1646, plain, (v1_xboole_0(esk47_1(k1_xboole_0))|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_2751, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_1(esk1_0,k2_group_1(esk1_0,X1))=k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0)))|k6_algstr_0(esk1_0,X1,k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,k2_group_1(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_650, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,k2_group_1(X1,X2),X3),X4))=k6_algstr_0(X1,X3,X4)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,k2_group_1(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_2032, plain, (v1_zfmisc_1(k1_xboole_0)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_3190, plain, (v1_zfmisc_1(k1_xboole_0))).#trainneg
cnf(i_0_1647, plain, (v1_zfmisc_1(esk47_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_2076, plain, (v6_ordinal1(k6_domain_1(X1,X2))|v1_xboole_0(X1)|~v6_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_651, plain, (X1=X2|v2_struct_0(X3)|k6_algstr_0(X3,X1,k6_algstr_0(X3,k2_group_1(X3,X2),X4))!=X4|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(k6_algstr_0(X3,k2_group_1(X3,X2),X4),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X4,u1_struct_0(X3)))).#trainneg
cnf(i_0_3000, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,k6_algstr_0(X1,X2,X4),X2)|v2_struct_0(X1)|X3!=k6_algstr_0(X1,X4,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X4),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_3003, plain, (k2_group_3(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X2,X3))=k6_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X2,X3))),u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3006, plain, (k2_group_3(X1,X2,X2)=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3041, plain, (r2_group_3(X1,X2,k6_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X2!=k6_algstr_0(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_3043, plain, (r2_group_3(X1,k6_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X4!=k6_algstr_0(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3044, plain, (m1_subset_1(k6_algstr_0(X1,X2,X3),k7_group_3(X1,X4))|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|X4!=k6_algstr_0(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1567, plain, (X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3484, plain, (esk40_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_3257, negated_conjecture, (k6_algstr_0(esk1_0,esk2_0,X1)=k6_algstr_0(esk1_0,k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|X1!=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2053, plain, (v3_finseq_1(k1_xboole_0)|~v3_finseq_1(X1))).#trainneg
cnf(i_0_3493, plain, (~v1_subset_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_3497, plain, (v3_finseq_1(k1_xboole_0))).#trainneg
cnf(i_0_2037, plain, (v1_zfmisc_1(u1_struct_0(k10_group_5(X1)))|v2_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2058, plain, (v3_finseq_1(u1_struct_0(k10_group_5(X1)))|v2_struct_0(X1)|~v3_finseq_1(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_656, plain, (k6_algstr_0(X1,k2_group_1(X1,X2),k6_algstr_0(X1,X2,X3))=X3|v2_struct_0(X1)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3254, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X2)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3278, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk8_3(esk1_0,X1,X2)=esk2_0|k6_algstr_0(esk1_0,esk8_3(esk1_0,X1,X2),X1)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~r1_group_3(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3297, plain, (k2_group_3(X1,k6_algstr_0(X1,esk8_3(X1,X3,X2),X3),X2)=k6_algstr_0(X1,esk8_3(X1,X3,X2),X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X3,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3310, plain, (esk8_3(X1,X2,X3)=k6_algstr_0(X1,k2_group_1(X1,X3),X4)|v2_struct_0(X1)|k6_algstr_0(X1,esk8_3(X1,X2,X3),X2)!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3324, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X3,esk8_3(X1,X2,X3)),X4)=k6_algstr_0(X1,esk8_3(X1,X2,X3),k6_algstr_0(X1,X2,X4))|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_3483, plain, (esk41_1(k1_xboole_0)=k1_xboole_0|v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_3575, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,X1)=esk12_1(u1_struct_0(esk1_0))|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,X1,esk2_0))!=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2074, plain, (v6_ordinal1(k1_xboole_0)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_1649, plain, (v1_xboole_0(esk50_1(k1_xboole_0))|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_1650, plain, (v1_zfmisc_1(esk50_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_2079, plain, (v6_ordinal1(u1_struct_0(k10_group_5(X1)))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_662, plain, (X1=X2|v2_struct_0(X3)|k7_group_3(X3,X4)!=k1_tarski(X2)|~r1_group_3(X3,X1,X4)|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_663, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(k7_group_3(X3,X4)))|v2_struct_0(X3)|k1_zfmisc_1(k7_group_3(X3,X4))!=k1_tarski(X2)|~r1_group_3(X3,X1,X4)|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_3397, plain, (r1_group_3(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3477, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|X1=esk2_0|k6_algstr_0(esk1_0,X1,X1)!=k6_algstr_0(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2217, plain, (v3_group_1(esk11_1(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3115, plain, (esk47_1(k1_xboole_0)=k1_xboole_0|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_3805, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,esk2_0))=esk12_1(u1_struct_0(esk1_0))|k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k6_algstr_0(esk1_0,k2_group_1(esk1_0,esk2_0),esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3826, plain, (v3_card_1(k1_xboole_0,k1_xboole_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_1652, plain, (v1_xboole_0(esk53_1(k1_xboole_0))|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_1653, plain, (v1_zfmisc_1(esk53_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_3501, plain, (v7_struct_0(k10_group_5(X1))|v2_struct_0(X1)|~l1_struct_0(k10_group_5(X1))|~v1_zfmisc_1(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_684, plain, (X1=k6_algstr_0(X2,X3,k6_algstr_0(X2,X4,X5))|v2_struct_0(X2)|k6_algstr_0(X2,X1,X6)!=k6_algstr_0(X2,X3,k6_algstr_0(X2,X4,k6_algstr_0(X2,X5,X6)))|~m1_subset_1(k6_algstr_0(X2,X3,k6_algstr_0(X2,X4,X5)),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(k6_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2)))).#trainneg
cnf(i_0_686, plain, (X1=k6_algstr_0(X2,X3,X4)|v2_struct_0(X2)|k6_algstr_0(X2,X1,k6_algstr_0(X2,k2_group_1(X2,X4),X5))!=k6_algstr_0(X2,X3,X5)|~m1_subset_1(k6_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k6_algstr_0(X2,k2_group_1(X2,X4),X5),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X5,u1_struct_0(X2)))).#trainneg
cnf(i_0_3569, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,k2_group_1(X1,X2),X3))=X3|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3580, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X2)),X4)=k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X2,X4))|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_3655, plain, (k2_group_3(X1,k6_algstr_0(X1,X3,esk8_3(X1,X2,X3)),X3)=k6_algstr_0(X1,esk8_3(X1,X2,X3),X3)|v2_struct_0(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3116, plain, (esk47_1(k1_xboole_0)=u1_struct_0(X1)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_2669, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_3809, plain, (esk50_1(k1_xboole_0)=k1_xboole_0|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_3970, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,esk2_0))=esk12_1(u1_struct_0(esk1_0))|k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_694, plain, (m1_subset_1(esk7_2(u1_struct_0(k10_group_5(X1)),X2),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(u1_struct_0(k10_group_5(X1)),X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4181, plain, (v1_funct_1(k1_xboole_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_4210, plain, (v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2561, plain, (v1_finset_1(u1_struct_0(k10_group_5(X1)))|v2_struct_0(X1)|~l1_struct_0(k10_group_5(X1))|~v8_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_687, plain, (X1=k6_algstr_0(X2,k6_algstr_0(X2,X3,X4),X5)|v2_struct_0(X2)|k6_algstr_0(X2,X1,X6)!=k6_algstr_0(X2,X3,k6_algstr_0(X2,X4,k6_algstr_0(X2,X5,X6)))|~m1_subset_1(k6_algstr_0(X2,k6_algstr_0(X2,X3,X4),X5),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(k6_algstr_0(X2,X3,X4),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(k6_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_691, plain, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X4,X5)|v2_struct_0(X1)|k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X6))!=k6_algstr_0(X1,X4,k6_algstr_0(X1,X5,X6))|~m1_subset_1(k6_algstr_0(X1,X4,X5),u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_3935, plain, (X1=k6_algstr_0(X2,X3,X4)|v2_struct_0(X2)|k6_algstr_0(X2,X1,X5)!=k6_algstr_0(X2,X3,k6_algstr_0(X2,X4,X5))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(k6_algstr_0(X2,X4,X5),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_4030, plain, (k6_algstr_0(X2,X3,X1)=X1|v2_struct_0(X2)|k6_algstr_0(X2,X3,X4)!=X4|~m1_subset_1(k6_algstr_0(X2,k2_group_1(X2,X1),X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_4037, negated_conjecture, (k6_algstr_0(esk1_0,X1,k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1)))=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4118, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X2,X4))=k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,X3,X2),X4))|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_3810, plain, (esk50_1(k1_xboole_0)=u1_struct_0(X1)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_2687, plain, (v5_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_3829, plain, (esk53_1(k1_xboole_0)=k1_xboole_0|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_3802, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))=esk12_1(u1_struct_0(esk1_0))|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))))!=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4671, plain, (v3_card_1(k1_xboole_0,k1_xboole_0)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_4665, plain, (v1_finseq_1(k1_xboole_0)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_4209, plain, (r1_tarski(u1_struct_0(k10_group_5(X1)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_3832, plain, (v1_finset_1(u1_struct_0(k10_group_5(X1)))|v2_struct_0(X1)|~l1_struct_0(k10_group_5(X1))|~v1_zfmisc_1(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_700, plain, (k6_domain_1(esk7_2(k1_tarski(X1),X2),X3)=esk7_2(k1_tarski(X1),X2)|v1_xboole_0(esk7_2(k1_tarski(X1),X2))|r1_tarski(k1_tarski(X1),X2)|~m1_subset_1(X3,esk7_2(k1_tarski(X1),X2))|~r2_hidden(X1,k1_zfmisc_1(k6_domain_1(esk7_2(k1_tarski(X1),X2),X3))))).#trainneg
cnf(i_0_701, plain, (k6_domain_1(k1_tarski(X1),X2)=k1_tarski(X1)|v1_xboole_0(k1_tarski(X1))|~m1_subset_1(X2,k1_tarski(X1))|~r2_hidden(X1,k6_domain_1(k1_tarski(X1),X2)))).#trainneg
cnf(i_0_4654, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,esk2_0,X1)=esk2_0|~m1_subset_1(k2_group_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4679, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v2_struct_0(k10_group_5(X1))|v1_xboole_0(u1_struct_0(X1))|~m1_group_2(X1,k10_group_5(X1))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4684, plain, (k1_tarski(X2)=k1_tarski(X1)|~r2_hidden(X1,k1_tarski(X2))|~m1_subset_1(X2,k1_tarski(X1)))).#trainneg
cnf(i_0_4687, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0)).#trainneg
cnf(i_0_4688, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=u1_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|~r1_tarski(u1_struct_0(esk1_0),k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_4690, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0)).#trainneg
cnf(i_0_4693, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=u1_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))))).#trainneg
cnf(i_0_4702, negated_conjecture, (k3_group_3(esk1_0,k7_group_3(esk1_0,esk2_0),X1)=k5_group_3(esk1_0,X1,esk2_0)|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3830, plain, (esk53_1(k1_xboole_0)=u1_struct_0(X1)|~l3_algstr_0(X1)|~v2_struct_0(X1)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_2877, plain, (v1_relat_1(u2_algstr_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_4670, plain, (v2_finseq_1(k1_xboole_0)|~v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_3804, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1))=esk12_1(u1_struct_0(esk1_0))|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1)))!=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,k2_group_1(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_702, plain, (k6_domain_1(k1_tarski(X1),X2)=k1_tarski(X1)|v1_xboole_0(k1_tarski(X1))|v1_xboole_0(k6_domain_1(k1_tarski(X1),X2))|~m1_subset_1(X2,k1_tarski(X1))|~m1_subset_1(X1,k6_domain_1(k1_tarski(X1),X2)))).#trainneg
cnf(i_0_2783, negated_conjecture, (v7_ordinal1(esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1741, plain, (v2_struct_0(esk28_1(k1_xboole_0))|~l1_struct_0(esk28_1(k1_xboole_0))|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_1695, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|~m1_group_2(X1,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_707, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X3)))|v2_struct_0(X3)|k1_zfmisc_1(u1_struct_0(X3))!=k1_tarski(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_4705, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,esk2_0,X1)=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4706, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk58_3(X1,esk1_0,esk2_0)=esk2_0|~m1_subset_1(esk58_3(X1,esk1_0,esk2_0),u1_struct_0(esk1_0))|~r2_hidden(X1,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_4722, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4723, negated_conjecture, (r2_group_3(esk1_0,esk2_0,esk2_0)|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4725, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(esk1_0,esk2_0))|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3808, plain, (v6_ordinal1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4927, plain, (v6_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_3486, plain, (esk12_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_708, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(k7_group_2(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k7_group_2(X2)))|~l3_algstr_0(k7_group_2(X2))|~v3_group_1(k7_group_2(X2))|~v2_group_1(k7_group_2(X2)))).#trainneg
cnf(i_0_4904, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v1_xboole_0(k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_1743, plain, (v7_struct_0(esk28_1(np__1))|~l1_struct_0(esk28_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_4774, plain, (r1_tarski(u1_struct_0(k10_group_5(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_711, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk58_3(X2,X1,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~r2_hidden(X2,k7_group_3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_712, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4726, negated_conjecture, (esk58_3(esk2_0,esk1_0,esk2_0)=esk2_0|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4728, negated_conjecture, (r1_group_3(esk1_0,esk2_0,esk2_0)|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4813, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,esk2_0,k2_group_1(esk1_0,X1))=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2545, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4792, negated_conjecture, (k1_tarski(esk2_0)=X1|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_713, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk7_2(k1_tarski(X2),X3),u1_struct_0(X4))|v2_struct_0(X4)|r1_tarski(k1_tarski(X2),X3)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~r2_hidden(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4903, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(k7_group_3(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_4952, plain, (v1_finset_1(u1_struct_0(esk28_1(np__1)))|~l1_struct_0(esk28_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_716, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X3))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_group_2(X1,X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4818, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|X1=esk2_0|v1_xboole_0(k7_group_3(esk1_0,esk2_0))|~r2_hidden(X1,k7_group_3(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_4832, negated_conjecture, (esk57_2(esk1_0,esk2_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_4838, negated_conjecture, (esk2_0=X1|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|k1_tarski(X1)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_4839, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4842, negated_conjecture, (v3_card_1(k7_group_3(esk1_0,esk2_0),np__1)|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4843, negated_conjecture, (m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(X1))|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(esk2_0,X1))).#trainneg
cnf(i_0_4845, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=X1|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))))).#trainneg
cnf(i_0_4868, negated_conjecture, (X1=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1588, plain, (X1=u1_struct_0(esk9_0)|~v1_xboole_0(X1)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_4778, negated_conjecture, (k1_tarski(esk57_2(esk1_0,esk2_0))=X1|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_717, plain, (esk7_2(k1_zfmisc_1(esk7_2(k1_tarski(X1),X2)),X3)=esk7_2(k1_tarski(X1),X2)|r1_tarski(k1_zfmisc_1(esk7_2(k1_tarski(X1),X2)),X3)|r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,k1_zfmisc_1(esk7_2(k1_zfmisc_1(esk7_2(k1_tarski(X1),X2)),X3))))).#trainneg
cnf(i_0_1895, plain, (u1_struct_0(esk9_0)=u1_struct_0(esk10_0)|~v2_struct_0(esk9_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_2035, plain, (v1_zfmisc_1(k7_group_3(X1,X2))|v2_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2056, plain, (v3_finseq_1(k7_group_3(X1,X2))|v2_struct_0(X1)|~v3_finseq_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_719, plain, (esk7_2(k1_zfmisc_1(k1_tarski(X1)),X2)=k1_tarski(X1)|r1_tarski(k1_zfmisc_1(k1_tarski(X1)),X2)|v1_xboole_0(esk7_2(k1_zfmisc_1(k1_tarski(X1)),X2))|~m1_subset_1(X1,esk7_2(k1_zfmisc_1(k1_tarski(X1)),X2)))).#trainneg
cnf(i_0_4871, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(X1))|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_4874, negated_conjecture, (X1=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|k1_zfmisc_1(X1)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_4900, negated_conjecture, (esk2_0=X1|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_4902, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4905, negated_conjecture, (esk2_0=X1|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_4906, negated_conjecture, (k1_tarski(esk2_0)=X1|v1_xboole_0(u1_struct_0(esk1_0))|k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))!=k1_tarski(X1))).#trainneg
cnf(i_0_4909, negated_conjecture, (esk58_3(esk2_0,esk1_0,esk2_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4910, negated_conjecture, (esk58_3(esk2_0,esk1_0,esk2_0)=X1|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_4537, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,X1)=k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0)))|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,k6_algstr_0(esk1_0,X1,esk2_0),esk12_1(u1_struct_0(esk1_0))))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_720, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|m1_subset_1(esk7_2(esk7_2(k1_zfmisc_1(X1),X2),X3),X1)|r1_tarski(esk7_2(k1_zfmisc_1(X1),X2),X3))).#trainneg
cnf(i_0_4176, plain, (esk47_1(k1_xboole_0)=u1_struct_0(esk10_0)|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5295, plain, (v1_xboole_0(u1_struct_0(esk10_0))|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5297, plain, (v1_xboole_0(u1_struct_0(esk10_0))|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_2077, plain, (v6_ordinal1(k7_group_3(X1,X2))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_721, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|m1_subset_1(X3,X1)|v1_xboole_0(esk7_2(k1_zfmisc_1(X1),X2))|~m1_subset_1(X3,esk7_2(k1_zfmisc_1(X1),X2)))).#trainneg
cnf(i_0_722, plain, (esk7_2(k7_group_3(X1,X2),X3)=X4|v1_xboole_0(u1_struct_0(X1))|r1_tarski(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|u1_struct_0(X1)!=k1_tarski(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4911, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4913, negated_conjecture, (r2_group_3(esk1_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4986, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk58_3(X1,esk1_0,esk2_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(X1,k7_group_3(esk1_0,esk2_0))|~m1_group_2(esk1_0,esk1_0))).#trainneg
cnf(i_0_5299, plain, (u1_struct_0(esk10_0)=k1_xboole_0|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5229, negated_conjecture, (k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0)))=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),esk12_1(u1_struct_0(esk1_0))))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_724, plain, (r1_struct_0(k10_group_5(X1),esk7_2(k7_group_3(X1,X2),X3))|m1_subset_1(esk57_2(X1,esk7_2(k7_group_3(X1,X2),X3)),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5311, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(u1_struct_0(esk1_0),esk2_0))).#trainneg
cnf(i_0_5294, plain, (v3_card_1(u1_struct_0(esk10_0),k1_xboole_0)|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_1784, plain, (v1_finseq_1(X1)|v2_struct_0(X2)|~v3_finseq_1(k7_group_3(X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_726, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk7_2(k7_group_3(X1,X2),X3),u1_struct_0(X4))|v2_struct_0(X4)|r1_tarski(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_727, plain, (r2_group_3(X1,k2_group_1(X1,esk7_2(k7_group_3(X1,X2),X3)),k2_group_1(X1,esk7_2(k7_group_3(X1,X2),X3)))|r1_tarski(u1_struct_0(X1),X4)|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5035, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_5064, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(esk2_0)|~m1_subset_1(k7_group_3(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_5073, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,esk2_0,k2_group_1(X1,X2))=esk2_0|v1_xboole_0(u1_struct_0(X1))|~m1_group_2(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5102, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|X1=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_group_3(esk1_0,esk2_0,X1))).#trainneg
cnf(i_0_5105, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),k7_group_3(X1,X2))|v2_struct_0(X1)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~r2_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5106, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),k7_group_3(X1,X2))|v2_struct_0(X1)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~r2_group_3(X1,X2,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5302, plain, (X1=u1_struct_0(esk10_0)|~v1_xboole_0(X1)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_2533, plain, (v1_relat_1(esk40_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5234, negated_conjecture, (k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0)))=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,esk12_1(u1_struct_0(esk1_0)),k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_728, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,esk7_2(k7_group_3(X1,X2),X3)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,esk7_2(k7_group_3(X1,X2),X3))),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(X1,X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1824, plain, (v3_ordinal1(X1)|v2_struct_0(X2)|~v3_ordinal1(k7_group_3(X2,X1))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_748, plain, (r2_hidden(esk7_2(u1_struct_0(k7_group_2(X1)),X2),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(u1_struct_0(k7_group_2(X1)),X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_749, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(k7_group_2(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k7_group_2(X2))))).#trainneg
cnf(i_0_5107, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),k7_group_3(X1,X2))|v2_struct_0(X1)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~r1_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5108, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),u1_struct_0(X1))|v2_struct_0(X1)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_2671, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_5338, negated_conjecture, (k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0)))=k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,esk12_1(u1_struct_0(esk1_0)),esk12_1(u1_struct_0(esk1_0)))))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_750, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(k7_group_2(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,u1_struct_0(k7_group_2(X2)))|~l3_algstr_0(k7_group_2(X2))|~v3_group_1(k7_group_2(X2))|~v2_group_1(k7_group_2(X2)))).#trainneg
cnf(i_0_4661, plain, (esk50_1(k1_xboole_0)=u1_struct_0(esk10_0)|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5601, plain, (v1_relat_1(u1_struct_0(esk10_0))|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5364, plain, (v1_finseq_1(X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v1_xboole_0(k7_group_3(X2,X1)))).#trainneg
cnf(i_0_751, plain, (esk7_2(k8_group_2(X1,X2),X3)=X4|v1_xboole_0(u1_struct_0(X1))|r1_tarski(k8_group_2(X1,X2),X3)|v2_struct_0(X1)|u1_struct_0(X1)!=k1_tarski(X4)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_754, plain, (r1_struct_0(k10_group_5(X1),esk7_2(k8_group_2(X1,X2),X3))|m1_subset_1(esk57_2(X1,esk7_2(k8_group_2(X1,X2),X3)),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k8_group_2(X1,X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_5332, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk58_3(k2_group_3(esk1_0,esk2_0,X1),esk1_0,esk2_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~m1_group_2(esk1_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5334, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk58_3(X1,esk1_0,esk2_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~m1_group_2(esk1_0,esk1_0)|~r2_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5335, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk58_3(X1,esk1_0,esk2_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~m1_group_2(esk1_0,esk1_0)|~r2_group_3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5336, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk58_3(X1,esk1_0,esk2_0)=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~m1_group_2(esk1_0,esk1_0)|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5433, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_5457, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(k2_group_1(X2,X3)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_5464, negated_conjecture, (esk58_3(esk2_0,esk1_0,esk2_0)=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(esk1_0))|~m1_group_2(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1472, plain, (v1_finset_1(X1)|~v1_xboole_0(esk41_1(X1)))).#trainneg
cnf(i_0_5693, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_1825, plain, (v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_3830, plain, (esk53_1(k1_xboole_0)=u1_struct_0(X1)|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_756, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk7_2(k8_group_2(X1,X2),X3),u1_struct_0(X4))|v2_struct_0(X4)|r1_tarski(k8_group_2(X1,X2),X3)|v2_struct_0(X1)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4671, plain, (v3_card_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_4670, plain, (v2_finseq_1(k1_xboole_0))).#trainneg
cnf(i_0_4665, plain, (v1_finseq_1(k1_xboole_0))).#trainneg
cnf(i_0_5516, plain, (v3_ordinal1(X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v7_ordinal1(k7_group_3(X2,X1)))).#trainneg
cnf(i_0_5517, plain, (v3_ordinal1(X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v1_xboole_0(k7_group_3(X2,X1)))).#trainneg
cnf(i_0_757, plain, (r2_group_3(X1,k2_group_1(X1,esk7_2(k8_group_2(X1,X2),X3)),k2_group_1(X1,esk7_2(k8_group_2(X1,X2),X3)))|r1_tarski(u1_struct_0(X1),X4)|v2_struct_0(X1)|r1_tarski(k8_group_2(X1,X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_758, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,esk7_2(k8_group_2(X1,X2),X3)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,esk7_2(k8_group_2(X1,X2),X3))),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k8_group_2(X1,X2),X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X2,X1))).#trainneg
cnf(i_0_5472, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k2_group_3(esk1_0,esk2_0,X2)=esk2_0|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(X1,esk1_0)|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5477, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=k7_group_3(X1,X2)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0)))|~r2_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5479, negated_conjecture, (k7_group_3(X1,X2)=k7_group_3(esk1_0,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~r1_tarski(k7_group_3(X1,X2),k7_group_3(esk1_0,esk2_0))|~r2_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5481, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=k7_group_3(X1,X2)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0)))|~r2_group_3(X1,X2,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5483, negated_conjecture, (k7_group_3(X1,X2)=k7_group_3(esk1_0,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~r1_tarski(k7_group_3(X1,X2),k7_group_3(esk1_0,esk2_0))|~r2_group_3(X1,X2,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5562, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=k7_group_3(X1,X2)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0)))|~r1_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5564, negated_conjecture, (k7_group_3(X1,X2)=k7_group_3(esk1_0,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~r1_tarski(k7_group_3(X1,X2),k7_group_3(esk1_0,esk2_0))|~r1_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5567, negated_conjecture, (m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5568, negated_conjecture, (u1_struct_0(X1)=k7_group_3(esk1_0,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),k7_group_3(esk1_0,esk2_0))|~m1_subset_1(esk2_0,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1652, plain, (v1_xboole_0(esk53_1(k1_xboole_0)))).#trainneg
cnf(i_0_3829, plain, (esk53_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_3830, plain, (u1_struct_0(X1)=k1_xboole_0|~l3_algstr_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_760, plain, (r1_tarski(k8_group_2(X1,X2),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5606, plain, (v1_relat_1(u1_struct_0(esk10_0))|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5602, plain, (v1_funct_1(u1_struct_0(esk10_0))|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5518, plain, (v3_ordinal1(X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~v1_card_1(k7_group_3(X2,X1)))).#trainneg
cnf(i_0_762, plain, (v1_xboole_0(X1)|m1_subset_1(esk7_2(k1_tarski(X2),X3),X1)|r1_tarski(k1_tarski(X2),X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5573, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=X1|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k7_group_3(esk1_0,esk2_0)))|~r2_hidden(esk2_0,X1))).#trainneg
cnf(i_0_5575, negated_conjecture, (X1=k7_group_3(esk1_0,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|~r1_tarski(X1,k7_group_3(esk1_0,esk2_0))|~r2_hidden(esk2_0,X1))).#trainneg
cnf(i_0_5835, plain, (v1_funct_1(u1_struct_0(esk10_0))|~v2_struct_0(esk10_0))).#trainneg
cnf(i_0_5233, negated_conjecture, (k6_algstr_0(esk1_0,esk2_0,X1)=k2_group_1(esk1_0,esk12_1(u1_struct_0(esk1_0)))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,X1,k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))))!=esk2_0|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_767, plain, (r1_tarski(u1_struct_0(esk10_0),X1)|r1_struct_0(k10_group_5(esk10_0),esk7_2(u1_struct_0(esk10_0),X1))|m1_subset_1(esk57_2(esk10_0,esk7_2(u1_struct_0(esk10_0),X1)),u1_struct_0(esk10_0))|v2_struct_0(esk10_0)|~v3_group_1(esk10_0)|~v2_group_1(esk10_0))).#trainneg
cnf(i_0_769, plain, (r1_tarski(u1_struct_0(k7_group_2(X1)),X2)|r1_struct_0(k10_group_5(k7_group_2(X1)),esk7_2(u1_struct_0(k7_group_2(X1)),X2))|m1_subset_1(esk57_2(k7_group_2(X1),esk7_2(u1_struct_0(k7_group_2(X1)),X2)),u1_struct_0(k7_group_2(X1)))|v2_struct_0(k7_group_2(X1))|v2_struct_0(X1)|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_770, plain, (r1_tarski(u1_struct_0(k10_group_5(X1)),X2)|r1_struct_0(k10_group_5(k10_group_5(X1)),esk7_2(u1_struct_0(k10_group_5(X1)),X2))|m1_subset_1(esk57_2(k10_group_5(X1),esk7_2(u1_struct_0(k10_group_5(X1)),X2)),u1_struct_0(k10_group_5(X1)))|v2_struct_0(k10_group_5(X1))|v2_struct_0(X1)|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5673, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k1_tarski(X3)))).#trainneg
cnf(i_0_5678, negated_conjecture, (r2_hidden(esk7_2(k1_tarski(esk2_0),X1),u1_struct_0(esk1_0))|r1_tarski(k1_tarski(esk2_0),X1))).#trainneg
cnf(i_0_5680, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_tarski(esk2_0)))).#trainneg
cnf(i_0_5688, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k1_tarski(k2_group_1(X2,X3))))).#trainneg
cnf(i_0_5689, plain, (r2_hidden(k2_group_1(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1781, plain, (v1_finseq_1(X1)|~v3_finseq_1(k1_tarski(X1)))).#trainneg
cnf(i_0_2689, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_797, plain, (k1_tarski(X1)=k7_group_3(X2,X3)|v2_struct_0(X2)|~m1_subset_1(k7_group_3(X2,X3),k1_zfmisc_1(k1_tarski(X1)))|~r2_group_3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_5396, plain, (r1_tarski(k7_group_3(X1,X2),u1_struct_0(X3))|v2_struct_0(X3)|v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_798, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k7_group_3(X2,X3)))|v2_struct_0(X2)|~r2_group_3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_799, plain, (k7_group_3(X1,X2)=k1_tarski(X3)|v2_struct_0(X1)|~r1_tarski(k7_group_3(X1,X2),k1_tarski(X3))|~r2_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_5910, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|v1_xboole_0(k7_group_3(esk1_0,esk2_0))|~m1_subset_1(esk2_0,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_5930, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|~r2_hidden(esk2_0,k1_tarski(X1))|~r2_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5931, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|~r2_hidden(esk2_0,k1_tarski(X1))|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5952, negated_conjecture, (~r2_hidden(u1_struct_0(esk1_0),esk2_0))).#trainneg
cnf(i_0_1782, plain, (v1_finseq_1(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1756, plain, (v1_finseq_1(esk12_1(X1))|~v3_finseq_1(X1))).#trainneg
cnf(i_0_6151, negated_conjecture, (v1_xboole_0(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),esk2_0))).#trainneg
cnf(i_0_4909, negated_conjecture, (esk58_3(esk2_0,esk1_0,esk2_0)=esk2_0)).#trainneg
cnf(i_0_807, plain, (v1_xboole_0(k1_tarski(X1))|m1_subset_1(esk7_2(k6_domain_1(k1_tarski(X1),X2),X3),X4)|r1_tarski(k6_domain_1(k1_tarski(X1),X2),X3)|~r2_hidden(X1,X4)|~m1_subset_1(X2,k1_tarski(X1)))).#trainneg
cnf(i_0_5172, plain, (esk50_1(k1_xboole_0)=u1_struct_0(esk9_0)|~v2_struct_0(esk9_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_4913, negated_conjecture, (r2_group_3(esk1_0,esk2_0,esk2_0))).#trainneg
cnf(i_0_6155, negated_conjecture, (esk2_0=k1_xboole_0|~m1_subset_1(u1_struct_0(esk1_0),esk2_0))).#trainneg
cnf(i_0_4766, plain, (r2_hidden(X1,u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(X3))|~m1_group_2(X3,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_6263, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(X2,u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X3))|~m1_group_2(X3,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X3)))).#trainneg
cnf(i_0_808, plain, (esk7_2(k6_domain_1(k1_zfmisc_1(k1_tarski(X1)),X2),X3)=k1_tarski(X1)|r1_tarski(k6_domain_1(k1_zfmisc_1(k1_tarski(X1)),X2),X3)|v1_xboole_0(k1_zfmisc_1(k1_tarski(X1)))|~r2_hidden(X1,esk7_2(k6_domain_1(k1_zfmisc_1(k1_tarski(X1)),X2),X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_tarski(X1))))).#trainneg
cnf(i_0_809, plain, (esk7_2(k6_domain_1(k1_zfmisc_1(X1),X2),X3)=X1|r1_tarski(k6_domain_1(k1_zfmisc_1(X1),X2),X3)|v1_xboole_0(k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(esk7_2(k6_domain_1(k1_zfmisc_1(X1),X2),X3)))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6044, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_tarski(esk2_0)))).#trainneg
cnf(i_0_6046, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,k1_tarski(k2_group_1(X2,X3)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_6049, negated_conjecture, (m1_subset_1(esk7_2(k1_tarski(esk2_0),X1),u1_struct_0(esk1_0))|r1_tarski(k1_tarski(esk2_0),X1))).#trainneg
cnf(i_0_6061, negated_conjecture, (r1_tarski(k1_tarski(esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4687, negated_conjecture, (k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|r1_tarski(k7_group_3(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4689, negated_conjecture, (k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_1796, plain, (v3_ordinal1(esk12_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_5173, plain, (esk47_1(k1_xboole_0)=u1_struct_0(esk9_0)|~v2_struct_0(esk9_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_811, plain, (r2_hidden(X1,X2)|r1_tarski(k6_domain_1(k1_zfmisc_1(X2),X3),X4)|v1_xboole_0(k1_zfmisc_1(X2))|~r2_hidden(X1,esk7_2(k6_domain_1(k1_zfmisc_1(X2),X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_6376, negated_conjecture, (~r2_hidden(k7_group_3(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_4678, plain, (r2_hidden(X1,u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(k10_group_5(X2)))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_4767, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X2))|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X1,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_812, plain, (m1_subset_1(X1,X2)|r1_tarski(k6_domain_1(k1_zfmisc_1(X2),X3),X4)|v1_xboole_0(k1_zfmisc_1(X2))|~r2_hidden(X1,esk7_2(k6_domain_1(k1_zfmisc_1(X2),X3),X4))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_813, plain, (v1_xboole_0(k6_domain_1(X1,X2))|v1_xboole_0(X1)|m1_subset_1(esk7_2(k6_domain_1(k6_domain_1(X1,X2),X3),X4),X1)|r1_tarski(k6_domain_1(k6_domain_1(X1,X2),X3),X4)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,k6_domain_1(X1,X2)))).#trainneg
cnf(i_0_4728, negated_conjecture, (k1_tarski(esk2_0)=k7_group_3(esk1_0,esk2_0)|r1_group_3(esk1_0,esk2_0,esk2_0))).#trainneg
cnf(i_0_4902, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_6112, plain, (m1_subset_1(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,k1_tarski(X4))|~r2_group_3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_6147, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|~r2_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_tarski(X1)))).#trainneg
cnf(i_0_6149, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_tarski(X1)=k7_group_3(esk1_0,esk2_0)|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_tarski(X1)))).#trainneg
cnf(i_0_1821, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_tarski(X1)))).#trainneg
cnf(i_0_2002, plain, (v1_xboole_0(esk40_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4792, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=X1|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_816, plain, (r1_struct_0(k10_group_5(X1),esk7_2(k6_domain_1(u1_struct_0(X1),X2),X3))|m1_subset_1(esk57_2(X1,esk7_2(k6_domain_1(u1_struct_0(X1),X2),X3)),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k6_domain_1(u1_struct_0(X1),X2),X3)|v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6511, negated_conjecture, (v1_finseq_1(esk2_0)|~v3_finseq_1(k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_6422, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(k10_group_5(X1)))|v2_struct_0(X1)|~m1_group_2(X1,k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6501, plain, (r1_tarski(k6_domain_1(k6_domain_1(X1,X2),X3),X1)|v1_xboole_0(X1)|v1_xboole_0(k6_domain_1(X1,X2))|~m1_subset_1(X3,k6_domain_1(X1,X2))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_817, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk7_2(k6_domain_1(u1_struct_0(X1),X2),X3),u1_struct_0(X4))|v2_struct_0(X4)|r1_tarski(k6_domain_1(u1_struct_0(X1),X2),X3)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_818, plain, (r2_group_3(X1,k2_group_1(X1,esk7_2(k6_domain_1(u1_struct_0(X1),X2),X3)),k2_group_1(X1,esk7_2(k6_domain_1(u1_struct_0(X1),X2),X3)))|r1_tarski(u1_struct_0(X1),X4)|v2_struct_0(X1)|r1_tarski(k6_domain_1(u1_struct_0(X1),X2),X3)|v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6359, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=u1_struct_0(esk1_0)|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))))).#trainneg
cnf(i_0_6361, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=u1_struct_0(esk1_0)|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|~r1_tarski(u1_struct_0(esk1_0),k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_1822, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2023, plain, (v1_zfmisc_1(esk40_1(X1))|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_6610, negated_conjecture, (v1_zfmisc_1(k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_4906, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=X1|k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))!=k1_tarski(X1))).#trainneg
cnf(i_0_820, plain, (esk7_2(k6_domain_1(X1,X2),X3)=X4|v1_xboole_0(X1)|r1_tarski(k6_domain_1(X1,X2),X3)|X1!=k1_tarski(X4)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6512, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_5395, plain, (r1_tarski(k7_group_3(X1,X2),u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_6416, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(X2,u1_struct_0(X1))|v1_xboole_0(u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1))))).#trainneg
cnf(i_0_819, plain, (r1_struct_0(k10_group_5(X1),k2_group_1(X1,esk7_2(k6_domain_1(u1_struct_0(X1),X2),X3)))|m1_subset_1(esk57_2(X1,k2_group_1(X1,esk7_2(k6_domain_1(u1_struct_0(X1),X2),X3))),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k6_domain_1(u1_struct_0(X1),X2),X3)|v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6405, negated_conjecture, (v1_xboole_0(esk2_0)|~m1_subset_1(k7_group_3(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_6421, plain, (u1_struct_0(k7_group_2(X1))=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(k7_group_2(X1)))|v2_struct_0(X1)|~m1_group_2(X1,k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6523, plain, (m1_subset_1(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k1_tarski(X4)))).#trainneg
cnf(i_0_2044, plain, (v3_finseq_1(esk40_1(X1))|~v3_finseq_1(X1))).#trainneg
cnf(i_0_6044, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_6735, plain, (r1_tarski(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X3))|v2_struct_0(X3)|v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6264, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X3)|~m1_group_2(X3,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_829, plain, (r2_hidden(esk7_2(u1_struct_0(k10_group_5(X1)),X2),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(u1_struct_0(k10_group_5(X1)),X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6061, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6049, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|m1_subset_1(esk7_2(k7_group_3(esk1_0,esk2_0),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5680, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_5678, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|r2_hidden(esk7_2(k7_group_3(esk1_0,esk2_0),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5455, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_6573, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),k7_group_3(X1,X2))|v2_struct_0(X1)|~r2_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6574, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),k7_group_3(X1,X2))|v2_struct_0(X1)|~r1_group_3(X1,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6577, negated_conjecture, (esk57_2(esk1_0,esk2_0)=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_6589, negated_conjecture, (r2_hidden(esk2_0,X1)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_6590, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=X1|~r2_hidden(esk2_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))))).#trainneg
cnf(i_0_6592, negated_conjecture, (r2_hidden(esk2_0,X1)|~m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6593, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|~r2_hidden(esk2_0,X1))).#trainneg
cnf(i_0_2065, plain, (v6_ordinal1(esk40_1(X1))|~v6_ordinal1(X1))).#trainneg
cnf(i_0_4690, negated_conjecture, (k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_836, plain, (m1_subset_1(X1,k7_group_3(X2,esk7_2(k7_group_3(X2,X1),X3)))|v2_struct_0(X2)|r1_tarski(k7_group_3(X2,X1),X3)|~m1_subset_1(esk7_2(k7_group_3(X2,X1),X3),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_6241, plain, (v1_relat_1(u1_struct_0(esk9_0))|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_6777, negated_conjecture, (v3_ordinal1(esk2_0)|~v7_ordinal1(k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_6734, plain, (r1_tarski(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_851, plain, (v1_xboole_0(X1)|r2_hidden(esk7_2(k1_tarski(X2),X3),X1)|r1_tarski(k1_tarski(X2),X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6594, negated_conjecture, (esk7_2(k7_group_3(esk1_0,esk2_0),X1)=esk2_0|r1_tarski(k7_group_3(esk1_0,esk2_0),X1))).#trainneg
cnf(i_0_6596, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),k7_group_3(X1,X2))|v2_struct_0(X1)|~r2_group_3(X1,X2,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6606, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),X1)|v1_xboole_0(X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_6615, negated_conjecture, (r1_tarski(k7_group_3(esk1_0,esk2_0),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6616, negated_conjecture, (k1_tarski(esk57_2(esk1_0,esk2_0))=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6617, negated_conjecture, (m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6619, negated_conjecture, (X1=esk2_0|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6622, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(X1))|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_6624, negated_conjecture, (X1=k7_group_3(esk1_0,esk2_0)|~r1_tarski(X1,k7_group_3(esk1_0,esk2_0))|~r2_hidden(esk2_0,X1))).#trainneg
cnf(i_0_6626, negated_conjecture, (u1_struct_0(k10_group_5(X1))=esk2_0|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k7_group_3(esk1_0,esk2_0)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6629, negated_conjecture, (k3_group_3(X1,k7_group_3(esk1_0,esk2_0),X2)=k5_group_3(X1,X2,esk2_0)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk2_0,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6630, negated_conjecture, (u1_struct_0(k7_group_2(X1))=esk2_0|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k7_group_3(esk1_0,esk2_0)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6554, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_tarski(X1)))).#trainneg
cnf(i_0_7073, plain, (v1_relat_1(u1_struct_0(esk9_0))|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_6242, plain, (v1_funct_1(u1_struct_0(esk9_0))|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_7213, negated_conjecture, (~v1_xboole_0(esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_852, plain, (v1_xboole_0(X1)|r2_hidden(X2,X1)|v1_xboole_0(k1_tarski(X3))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_tarski(X3)))).#trainneg
cnf(i_0_7637, negated_conjecture, (k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_4955, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4765, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v1_xboole_0(u1_struct_0(X1))|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_854, plain, (k1_tarski(esk58_3(X1,X2,X3))=X4|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X2)))|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|k1_zfmisc_1(u1_struct_0(X2))!=k1_tarski(X4)|~r2_hidden(X1,k7_group_3(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_856, plain, (k1_tarski(esk8_3(X1,X2,X3))=X4|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X4)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6640, negated_conjecture, (esk7_2(X1,X2)=esk2_0|r1_tarski(X1,X2)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_6654, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=esk2_0|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6779, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_card_1(k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_6868, plain, (k6_domain_1(u1_struct_0(X1),X2)=u1_struct_0(X3)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X3)|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(k6_domain_1(u1_struct_0(X1),X2)))|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_6870, plain, (u1_struct_0(X1)=k6_domain_1(u1_struct_0(X2),X3)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),k6_domain_1(u1_struct_0(X2),X3))|~m1_group_2(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6924, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))))).#trainneg
cnf(i_0_6925, negated_conjecture, (m1_subset_1(k7_group_3(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7635, plain, (v1_funct_1(u1_struct_0(esk9_0))|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_6578, negated_conjecture, (esk2_0=np__1|k7_group_3(esk1_0,esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_857, plain, (k1_tarski(esk7_2(k8_group_2(X1,X2),X3))=X4|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|r1_tarski(k8_group_2(X1,X2),X3)|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X4)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7666, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X2))|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X1,X2)|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_858, plain, (k1_tarski(esk7_2(k7_group_3(X1,X2),X3))=X4|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|r1_tarski(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_859, plain, (k1_tarski(esk7_2(k1_tarski(X1),X2))=X3|v1_xboole_0(k1_zfmisc_1(X4))|v1_xboole_0(X4)|r1_tarski(k1_tarski(X1),X2)|k1_zfmisc_1(X4)!=k1_tarski(X3)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_6976, negated_conjecture, (r2_hidden(k2_group_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7002, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7003, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_group_3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7004, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r1_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7014, negated_conjecture, (r1_struct_0(k10_group_5(esk1_0),esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_6556, plain, (v3_ordinal1(X1)|~v1_card_1(k1_tarski(X1)))).#trainneg
cnf(i_0_2626, plain, (v1_xboole_0(esk17_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7038, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=k1_xboole_0|~r2_hidden(esk2_0,k1_xboole_0))).#trainneg
cnf(i_0_7859, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_860, plain, (k1_tarski(esk7_2(k6_domain_1(X1,X2),X3))=X4|v1_xboole_0(k1_zfmisc_1(X1))|v1_xboole_0(X1)|r1_tarski(k6_domain_1(X1,X2),X3)|k1_zfmisc_1(X1)!=k1_tarski(X4)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3117, plain, (X1=esk47_1(k1_xboole_0)|~v1_xboole_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_7939, negated_conjecture, (~r2_hidden(esk2_0,k1_xboole_0))).#trainneg
cnf(i_0_4677, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r1_tarski(u1_struct_0(X1),u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_4957, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v2_struct_0(k10_group_5(X1))|v2_struct_0(X1)|~m1_group_2(X1,k10_group_5(X1))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_861, plain, (k1_tarski(esk7_2(X1,X2))=X3|v1_xboole_0(k1_zfmisc_1(X1))|v1_xboole_0(X1)|r1_tarski(X1,X2)|k1_zfmisc_1(X1)!=k1_tarski(X3))).#trainneg
cnf(i_0_863, plain, (k1_tarski(k8_group_2(X1,X2))=X3|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_tarski(X3)|~m1_group_2(X2,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7079, plain, (u1_struct_0(X1)=k6_domain_1(u1_struct_0(X2),X3)|v1_xboole_0(u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X1))|~r1_tarski(u1_struct_0(X1),k6_domain_1(u1_struct_0(X2),X3))|~m1_group_2(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7185, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0)))|~m1_subset_1(esk2_0,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7987, negated_conjecture, (k7_group_3(esk1_0,esk2_0)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_864, plain, (k1_tarski(k2_group_1(X1,X2))=X3|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7655, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v2_struct_0(X1)|v2_struct_0(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_group_2(X1,k10_group_5(X1))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1)))).#trainneg
cnf(i_0_865, plain, (k1_tarski(k2_group_3(X1,X2,X3))=X4|v1_xboole_0(k1_zfmisc_1(k7_group_3(X1,X2)))|v1_xboole_0(k7_group_3(X1,X2))|v2_struct_0(X1)|k1_zfmisc_1(k7_group_3(X1,X2))!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_866, plain, (k1_tarski(k1_tarski(X1))=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X3)))|v1_xboole_0(k1_zfmisc_1(X3))|k1_zfmisc_1(k1_zfmisc_1(X3))!=k1_tarski(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_7209, negated_conjecture, (r2_hidden(esk57_2(esk1_0,esk2_0),esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7211, negated_conjecture, (m1_subset_1(esk57_2(esk1_0,esk2_0),esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7237, negated_conjecture, (v1_zfmisc_1(esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7240, negated_conjecture, (v1_finseq_1(esk57_2(esk1_0,esk2_0))|~v3_finseq_1(esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7241, negated_conjecture, (v3_ordinal1(esk57_2(esk1_0,esk2_0))|~v3_ordinal1(esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7476, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7485, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_6771, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_867, plain, (k1_tarski(k1_tarski(X1))=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X3)))|v1_xboole_0(k1_zfmisc_1(X3))|v1_xboole_0(X3)|k1_zfmisc_1(k1_zfmisc_1(X3))!=k1_tarski(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_7991, plain, (esk47_1(k1_xboole_0)=esk50_1(k1_xboole_0)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_868, plain, (k1_tarski(k6_domain_1(X1,X2))=X3|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X1)))|v1_xboole_0(k1_zfmisc_1(X1))|v1_xboole_0(X1)|k1_zfmisc_1(k1_zfmisc_1(X1))!=k1_tarski(X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_869, plain, (k1_tarski(k7_group_3(X1,X2))=X3|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7630, negated_conjecture, (u1_struct_0(X1)=esk2_0|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k7_group_3(esk1_0,esk2_0)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~v15_algstr_0(X1))).#trainneg
cnf(i_0_7631, negated_conjecture, (v3_ordinal1(esk57_2(esk1_0,esk2_0))|~v7_ordinal1(esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7728, negated_conjecture, (r1_tarski(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,esk2_0,X3)|~m1_subset_1(esk2_0,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7729, negated_conjecture, (r1_tarski(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r2_group_3(X2,X3,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7730, negated_conjecture, (r1_tarski(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|~r1_group_3(X2,esk2_0,X3)|~m1_subset_1(esk2_0,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7731, negated_conjecture, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(esk2_0,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7733, negated_conjecture, (r1_tarski(X1,X2)|~r2_hidden(esk2_0,X2)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7734, negated_conjecture, (r1_tarski(X1,X2)|m1_subset_1(esk2_0,X1)|X1!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_6773, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6579, negated_conjecture, (esk2_0=np__1|k7_group_3(esk1_0,esk2_0)!=k5_numbers)).#trainneg
cnf(i_0_870, plain, (k1_tarski(u1_struct_0(k7_group_2(X1)))=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_tarski(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7663, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(k10_group_5(X1)))|v1_xboole_0(u1_struct_0(X1))|~m1_group_2(X1,k10_group_5(X1))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_871, plain, (k1_tarski(u1_struct_0(k10_group_5(X1)))=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_tarski(X2)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_872, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(k1_tarski(X1)))|v1_xboole_0(k1_tarski(X1))|k1_zfmisc_1(k1_tarski(X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_7739, negated_conjecture, (k6_algstr_0(X1,X3,X2)=esk2_0|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,X3))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_7742, negated_conjecture, (r2_group_3(X1,X2,esk2_0)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7744, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7745, negated_conjecture, (esk58_3(esk2_0,X1,X2)=esk2_0|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7746, negated_conjecture, (r2_group_3(X1,esk2_0,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7747, negated_conjecture, (r1_group_3(X1,esk2_0,X2)|v2_struct_0(X1)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7750, negated_conjecture, (esk2_0=X3|v2_struct_0(X1)|~r1_group_3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(esk8_3(X1,X3,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_7751, negated_conjecture, (m1_subset_1(X1,k7_group_3(X2,esk2_0))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k7_group_3(X2,X1)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_7753, negated_conjecture, (k6_algstr_0(X1,X2,X3)=k6_algstr_0(X1,X4,X2)|v2_struct_0(X1)|X3!=esk2_0|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X4)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_6581, negated_conjecture, (esk20_0=esk2_0|k7_group_3(esk1_0,esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_873, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X1)))|v1_xboole_0(k1_zfmisc_1(X1))|k1_zfmisc_1(k1_zfmisc_1(X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_7793, plain, (u1_struct_0(k10_group_5(X1))=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(k10_group_5(X1)))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_group_2(X1,k10_group_5(X1))|~l3_algstr_0(X1)|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(X1)|~v3_group_1(k10_group_5(X1))|~v2_group_1(X1)|~v2_group_1(k10_group_5(X1)))).#trainneg
cnf(i_0_6417, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1)))).#trainneg
cnf(i_0_875, plain, (k1_tarski(X1)=X2|v1_xboole_0(k1_zfmisc_1(k7_group_3(X3,X1)))|v1_xboole_0(k7_group_3(X3,X1))|v2_struct_0(X3)|k1_zfmisc_1(k7_group_3(X3,X1))!=k1_tarski(X2)|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_878, plain, (esk7_2(k7_group_3(X1,X2),X3)=X4|r1_tarski(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|u1_struct_0(X1)!=k1_tarski(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7766, negated_conjecture, (u1_struct_0(X1)=k7_group_3(esk1_0,esk2_0)|v2_struct_0(X1)|k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0|~r1_tarski(u1_struct_0(X1),k7_group_3(esk1_0,esk2_0))|~m1_group_2(esk1_0,X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7775, negated_conjecture, (k7_group_3(esk1_0,esk2_0)=u1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))))).#trainneg
cnf(i_0_7782, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k7_group_3(esk1_0,esk2_0)))).#trainneg
cnf(i_0_7825, negated_conjecture, (k2_group_3(esk1_0,esk2_0,X1)=X2|k1_tarski(X2)!=u1_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8070, negated_conjecture, (u1_struct_0(esk38_0)=esk2_0|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk38_0)))).#trainneg
cnf(i_0_879, plain, (k1_tarski(esk7_2(k7_group_3(X1,X2),X3))=X4|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_6100, plain, (r2_hidden(X1,u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X3))|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X3,X4))|~m1_group_2(X3,X2)|~m1_subset_1(X4,u1_struct_0(X3))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_880, plain, (r2_hidden(esk7_2(k7_group_3(k7_group_2(X1),X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(k7_group_2(X1),X2),X3)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k7_group_2(X1)))|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_881, plain, (m1_subset_1(esk7_2(k7_group_3(k7_group_2(X1),X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(k7_group_2(X1),X2),X3)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k7_group_2(X1)))|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_7860, negated_conjecture, (v3_ordinal1(esk57_2(esk1_0,esk2_0))|~v1_card_1(esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8033, negated_conjecture, (k1_tarski(k1_tarski(X1))=esk2_0|k1_zfmisc_1(k1_zfmisc_1(X2))!=k7_group_3(esk1_0,esk2_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_8035, negated_conjecture, (~r2_hidden(esk2_0,esk57_2(esk1_0,esk2_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1546, plain, (v1_finset_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_8073, negated_conjecture, (u1_struct_0(esk18_0)=esk2_0|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk18_0)))).#trainneg
cnf(i_0_884, plain, (m1_subset_1(esk7_2(k7_group_3(X1,X2),X3),u1_struct_0(X4))|v2_struct_0(X4)|r1_tarski(k7_group_3(X1,X2),X3)|v2_struct_0(X1)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8566, plain, (~m1_subset_1(k1_zfmisc_1(esk36_0),k5_numbers))).#trainneg
cnf(i_0_8567, plain, (~m1_subset_1(k1_zfmisc_1(k5_numbers),k5_numbers))).#trainneg
cnf(i_0_8622, plain, (r1_tarski(k7_group_3(X1,X2),u1_struct_0(X3))|v2_struct_0(X3)|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8621, plain, (r1_tarski(k7_group_3(X1,X2),u1_struct_0(X3))|v1_xboole_0(u1_struct_0(X3))|v2_struct_0(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_882, plain, (r2_hidden(esk7_2(k7_group_3(k10_group_5(X1),X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(k10_group_5(X1),X2),X3)|v2_struct_0(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1)))).#trainneg
cnf(i_0_883, plain, (m1_subset_1(esk7_2(k7_group_3(k10_group_5(X1),X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|r1_tarski(k7_group_3(k10_group_5(X1),X2),X3)|v2_struct_0(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1)))).#trainneg
cnf(i_0_8045, negated_conjecture, (v1_zfmisc_1(esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8047, negated_conjecture, (v1_finset_1(esk57_2(esk1_0,esk2_0))|~v1_funct_1(esk57_2(esk1_0,esk2_0))|~v1_relat_1(esk57_2(esk1_0,esk2_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~v3_finseq_1(esk2_0))).#trainneg
cnf(i_0_8048, negated_conjecture, (v7_ordinal1(esk57_2(esk1_0,esk2_0))|~v1_finset_1(esk57_2(esk1_0,esk2_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~v3_ordinal1(esk2_0))).#trainneg
cnf(i_0_8065, negated_conjecture, (k1_tarski(k6_domain_1(X1,X2))=esk2_0|v1_xboole_0(X1)|k1_zfmisc_1(k1_zfmisc_1(X1))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_8076, negated_conjecture, (v7_ordinal1(esk57_2(esk1_0,esk2_0))|~v1_finset_1(esk57_2(esk1_0,esk2_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~v7_ordinal1(esk2_0))).#trainneg
cnf(i_0_8141, negated_conjecture, (k6_algstr_0(X1,X2,esk8_3(X1,X2,X3))=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,X3,esk8_3(X1,X2,X3)))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_8146, negated_conjecture, (k6_algstr_0(X1,X2,k6_algstr_0(X1,X2,X3))=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X2)))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_8158, negated_conjecture, (k6_algstr_0(X1,k2_group_3(X1,X2,X3),X3)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,X3))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8056, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_7249, negated_conjecture, (esk57_2(esk1_0,esk2_0)=k7_group_3(esk1_0,esk2_0)|k1_zfmisc_1(k7_group_3(esk1_0,esk2_0))!=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_888, plain, (k2_group_3(X1,X2,X3)=X4|v2_struct_0(X1)|u1_struct_0(X1)!=k1_tarski(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8972, plain, (v7_ordinal1(esk20_0)|~v1_finset_1(esk20_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_8568, plain, (~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_7076, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k7_group_3(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_8662, plain, (r1_tarski(k7_group_3(k10_group_5(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_889, plain, (k1_tarski(k2_group_3(X1,X2,X3))=X4|v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|k1_zfmisc_1(u1_struct_0(X1))!=k1_tarski(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_890, plain, (r2_hidden(k2_group_3(k7_group_2(X1),X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k7_group_2(X1)))|~m1_subset_1(X3,u1_struct_0(k7_group_2(X1)))|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_8159, negated_conjecture, (k6_algstr_0(X1,esk8_3(X1,X2,X3),X3)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,esk8_3(X1,X2,X3),X2))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8164, negated_conjecture, (r2_group_3(esk1_0,X1,esk2_0)|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_8167, negated_conjecture, (r2_group_3(X1,X2,esk2_0)|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8198, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(esk1_0,X1))|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_8201, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(X1,X2))|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8232, negated_conjecture, (esk58_3(esk2_0,esk1_0,X1)=esk2_0|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_8235, negated_conjecture, (esk58_3(esk2_0,X1,X2)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8266, negated_conjecture, (r2_group_3(esk1_0,esk2_0,X1)|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_8269, negated_conjecture, (r2_group_3(X1,esk2_0,X2)|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8300, negated_conjecture, (r1_group_3(esk1_0,esk2_0,X1)|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_8303, negated_conjecture, (r1_group_3(X1,esk2_0,X2)|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_1542, plain, (v1_finset_1(X1)|~m1_subset_1(esk41_1(X1),k5_numbers))).#trainneg
cnf(i_0_6580, negated_conjecture, (esk2_0=k1_xboole_0|k7_group_3(esk1_0,esk2_0)!=k5_numbers)).#trainneg
cnf(i_0_891, plain, (m1_subset_1(k2_group_3(k7_group_2(X1),X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(k7_group_2(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k7_group_2(X1)))|~m1_subset_1(X3,u1_struct_0(k7_group_2(X1)))|~l3_algstr_0(k7_group_2(X1))|~v3_group_1(k7_group_2(X1))|~v2_group_1(k7_group_2(X1)))).#trainneg
cnf(i_0_8976, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_8974, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(k1_zfmisc_1(k7_group_3(esk1_0,esk2_0)),k5_numbers))).#trainneg
cnf(i_0_8714, plain, (r1_tarski(k7_group_3(k10_group_5(X1),X2),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8631, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k7_group_3(X3,X4))|~m1_group_2(X3,X2)|~m1_subset_1(X4,u1_struct_0(X3))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_892, plain, (r2_hidden(k2_group_3(k10_group_5(X1),X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~m1_subset_1(X3,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1)))).#trainneg
cnf(i_0_893, plain, (m1_subset_1(k2_group_3(k10_group_5(X1),X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(k10_group_5(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~m1_subset_1(X3,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(k10_group_5(X1))|~v3_group_1(k10_group_5(X1))|~v2_group_1(k10_group_5(X1)))).#trainneg
cnf(i_0_8338, negated_conjecture, (m1_subset_1(X1,k7_group_3(esk1_0,esk2_0))|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_8341, negated_conjecture, (m1_subset_1(X1,k7_group_3(X2,esk2_0))|v2_struct_0(X2)|k7_group_3(X2,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2))).#trainneg
cnf(i_0_8371, negated_conjecture, (k6_algstr_0(esk1_0,X1,X2)=k6_algstr_0(esk1_0,esk2_0,X1)|X2!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8475, plain, (r1_tarski(k7_group_3(k7_group_2(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(k7_group_2(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_7634, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_tarski(X1),k5_numbers))).#trainneg
cnf(i_0_2234, plain, (v8_struct_0(esk19_1(X1))|v2_struct_0(X1)|~v8_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9548, negated_conjecture, (esk12_1(u1_struct_0(esk1_0))=k6_algstr_0(esk1_0,esk2_0,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,esk2_0,esk2_0))!=k6_algstr_0(esk1_0,esk2_0,esk2_0)|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|esk12_1(u1_struct_0(esk1_0))!=esk2_0)).#trainneg
cnf(i_0_894, plain, (m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_group_2(X1,X4)|~l3_algstr_0(X4)|~v3_group_1(X4)|~v2_group_1(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9139, negated_conjecture, (v7_ordinal1(esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_9637, plain, (m1_subset_1(X2,u1_struct_0(X3))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8536, negated_conjecture, (v7_ordinal1(esk57_2(esk1_0,esk2_0))|~v1_finset_1(esk57_2(esk1_0,esk2_0))|k6_domain_1(u1_struct_0(esk1_0),esk2_0)!=k7_group_3(esk1_0,esk2_0)|k7_group_3(esk1_0,esk2_0)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~v1_card_1(esk2_0))).#trainneg
cnf(i_0_8629, plain, (m1_subset_1(k7_group_3(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|v2_struct_0(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_8858, negated_conjecture, (k6_algstr_0(X1,X2,esk8_3(X1,X2,X3))=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,esk8_3(X1,X2,X3),X2))!=k7_group_3(esk1_0,esk2_0)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8892, negated_conjecture, (k6_algstr_0(X1,X2,k6_algstr_0(X1,X2,k2_group_1(X1,X2)))=esk2_0|v2_struct_0(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_group_1(X1,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8903, negated_conjecture, (k6_algstr_0(X1,k2_group_3(X1,X2,X3),k6_algstr_0(X1,X3,X4))=k6_algstr_0(X1,esk2_0,X4)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,X3))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~v2_group_1(X1))).#trainneg
cnf(i_0_8907, negated_conjecture, (k6_algstr_0(X1,k2_group_1(X1,k2_group_3(X1,X2,X3)),esk2_0)=X3|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,X3))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8908, negated_conjecture, (k2_group_3(X1,esk2_0,k2_group_3(X1,X2,X3))=k6_algstr_0(X1,X3,k2_group_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(k2_group_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,X3))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_8926, negated_conjecture, (k2_group_3(X1,k2_group_3(X1,X3,X2),X2)=k6_algstr_0(X1,k2_group_1(X1,X2),esk2_0)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_group_3(X1,X3,X2),u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,k6_algstr_0(X1,X3,X2))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_8940, negated_conjecture, (X1=X2|v2_struct_0(X3)|k6_algstr_0(X3,k2_group_3(X3,X4,X5),k6_algstr_0(X3,X5,X1))!=k6_algstr_0(X3,esk2_0,X2)|~m1_subset_1(esk2_0,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X5,u1_struct_0(X3))|~m1_subset_1(k2_group_3(X3,X4,X5),u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|k7_group_3(X3,k6_algstr_0(X3,X4,X5))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X4,u1_struct_0(X3)))).#trainneg
cnf(i_0_8942, negated_conjecture, (X1=k6_algstr_0(X2,X3,k2_group_3(X2,X4,X5))|v2_struct_0(X2)|k6_algstr_0(X2,X1,X5)!=k6_algstr_0(X2,X3,esk2_0)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(k2_group_3(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k7_group_3(X2,k6_algstr_0(X2,X4,X5))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X4,u1_struct_0(X2)))).#trainneg
cnf(i_0_2236, plain, (v8_struct_0(esk11_1(X1))|v2_struct_0(X1)|~v8_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9535, negated_conjecture, (esk12_1(u1_struct_0(esk1_0))=k6_algstr_0(esk1_0,esk2_0,esk2_0)|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|k6_algstr_0(esk1_0,esk2_0,esk12_1(u1_struct_0(esk1_0)))!=k6_algstr_0(esk1_0,esk2_0,k6_algstr_0(esk1_0,esk2_0,esk2_0))|~m1_subset_1(k6_algstr_0(esk1_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_group_1(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|k2_group_1(esk1_0,X1)!=esk2_0)).#trainneg
cnf(i_0_9685, plain, (v6_ordinal1(k7_group_3(X1,X2))|v2_struct_0(X3)|~v6_ordinal1(u1_struct_0(X3))|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_9686, plain, (v3_finseq_1(k7_group_3(X1,X2))|v2_struct_0(X3)|~v3_finseq_1(u1_struct_0(X3))|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_904, plain, (v1_xboole_0(k6_domain_1(X1,X2))|v1_xboole_0(X1)|r2_hidden(esk7_2(k1_tarski(X3),X4),X1)|r1_tarski(k1_tarski(X3),X4)|~m1_subset_1(X2,X1)|~r2_hidden(X3,k6_domain_1(X1,X2)))).#trainneg
cnf(i_0_8968, negated_conjecture, (k6_algstr_0(X1,esk2_0,X3)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,X3))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1)|k7_group_3(X1,X2)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_8970, negated_conjecture, (k6_algstr_0(X1,X2,X2)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,X2,X2))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8986, plain, (k7_group_3(k10_group_5(X1),X2)=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k7_group_3(k10_group_5(X1),X2)))|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8987, plain, (m1_subset_1(k7_group_3(k10_group_5(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_8988, plain, (k7_group_3(k10_group_5(X1),X2)=u1_struct_0(X1)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),k7_group_3(k10_group_5(X1),X2))|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9028, negated_conjecture, (k6_algstr_0(X1,esk8_3(X1,X2,X3),X3)=esk2_0|v2_struct_0(X1)|k7_group_3(X1,k6_algstr_0(X1,X3,esk8_3(X1,X2,X3)))!=k7_group_3(esk1_0,esk2_0)|~r1_group_3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9029, negated_conjecture, (r2_group_3(esk1_0,X1,esk2_0)|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9034, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_9035, negated_conjecture, (r1_tarski(X1,k7_group_3(X2,esk2_0))|v2_struct_0(X2)|~m1_subset_1(esk7_2(X1,k7_group_3(X2,esk2_0)),u1_struct_0(X2))|~m1_subset_1(esk2_0,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|k7_group_3(X2,esk7_2(X1,k7_group_3(X2,esk2_0)))!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_2179, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk26_1(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_2238, plain, (v8_struct_0(k7_group_2(X1))|v2_struct_0(X1)|~v8_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_10261, negated_conjecture, (k6_algstr_0(esk1_0,esk2_0,esk2_0)=esk2_0|k7_group_3(esk1_0,k6_algstr_0(esk1_0,esk2_0,esk2_0))!=k7_group_3(esk1_0,esk2_0))).#trainneg
cnf(i_0_905, plain, (v1_xboole_0(k6_domain_1(X1,X2))|v1_xboole_0(X1)|r2_hidden(esk7_2(k6_domain_1(k6_domain_1(X1,X2),X3),X4),X1)|r1_tarski(k6_domain_1(k6_domain_1(X1,X2),X3),X4)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,k6_domain_1(X1,X2)))).#trainneg
cnf(i_0_3811, plain, (X1=esk50_1(k1_xboole_0)|~v1_xboole_0(X1)|~v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_10279, plain, (v6_ordinal1(k7_group_3(k10_group_5(X1),X2))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_10280, plain, (v3_finseq_1(k7_group_3(k10_group_5(X1),X2))|v2_struct_0(X1)|~v3_finseq_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_908, plain, (X1=X2|v2_struct_0(X3)|k7_group_3(X3,esk58_3(X4,X3,X1))!=k1_tarski(X2)|~m1_subset_1(esk58_3(X4,X3,X1),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~r2_hidden(X4,k7_group_3(X3,X1)))).#trainneg
cnf(i_0_9040, negated_conjecture, (m1_subset_1(esk2_0,k7_group_3(esk1_0,X1))|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2537, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_909, plain, (X1=X2|v2_struct_0(X3)|r1_tarski(k7_group_3(X3,X1),X4)|k7_group_3(X3,esk7_2(k7_group_3(X3,X1),X4))!=k1_tarski(X2)|~m1_subset_1(esk7_2(k7_group_3(X3,X1),X4),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_10282, plain, (v1_zfmisc_1(k7_group_3(k10_group_5(X1),X2))|v2_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_910, plain, (X1=X2|v2_struct_0(X3)|k7_group_3(X3,k2_group_3(X3,X1,X4))!=k1_tarski(X2)|~m1_subset_1(k2_group_3(X3,X1,X4),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~m1_subset_1(X4,u1_struct_0(X3)))).#trainneg
cnf(i_0_912, plain, (esk58_3(X1,X2,X3)=X4|v2_struct_0(X2)|k7_group_3(X2,X3)!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(esk58_3(X1,X2,X3),u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~r2_hidden(X1,k7_group_3(X2,X3)))).#trainneg
cnf(i_0_9064, negated_conjecture, (r2_group_3(esk1_0,esk2_0,X1)|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9075, negated_conjecture, (r1_group_3(esk1_0,esk2_0,X1)|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2191, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk44_1(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_2659, plain, (v4_relat_1(esk40_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_7777, negated_conjecture, (v6_ordinal1(k7_group_3(esk1_0,esk2_0))|~v6_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7778, negated_conjecture, (v3_finseq_1(k7_group_3(esk1_0,esk2_0))|~v3_finseq_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9688, plain, (v1_zfmisc_1(k7_group_3(X1,X2))|v2_struct_0(X3)|~v1_zfmisc_1(u1_struct_0(X3))|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_9141, plain, (k7_group_3(k10_group_5(X1),X2)=u1_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k7_group_3(k10_group_5(X1),X2)))|~m1_subset_1(X2,u1_struct_0(k10_group_5(X1)))|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_9242, negated_conjecture, (m1_subset_1(X1,k7_group_3(esk1_0,esk2_0))|k7_group_3(esk1_0,X1)!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9260, negated_conjecture, (k6_algstr_0(esk1_0,k2_group_1(esk1_0,esk2_0),k6_algstr_0(esk1_0,X1,esk2_0))=X1|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9261, negated_conjecture, (k2_group_3(esk1_0,k6_algstr_0(esk1_0,X1,esk2_0),esk2_0)=k6_algstr_0(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9262, negated_conjecture, (esk58_3(k6_algstr_0(esk1_0,X1,esk2_0),esk1_0,k6_algstr_0(esk1_0,X1,esk2_0))=k6_algstr_0(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9265, negated_conjecture, (r2_group_3(esk1_0,k6_algstr_0(esk1_0,X1,esk2_0),X2)|X2!=k6_algstr_0(esk1_0,X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9266, negated_conjecture, (m1_subset_1(k6_algstr_0(esk1_0,X1,esk2_0),k7_group_3(esk1_0,X2))|X2!=k6_algstr_0(esk1_0,X1,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9271, negated_conjecture, (r1_group_3(esk1_0,X1,k6_algstr_0(esk1_0,X2,esk2_0))|X1!=k6_algstr_0(esk1_0,X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9273, negated_conjecture, (r2_group_3(esk1_0,X1,k6_algstr_0(esk1_0,X2,esk2_0))|X1!=k6_algstr_0(esk1_0,X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9275, negated_conjecture, (k6_algstr_0(esk1_0,k2_group_1(esk1_0,X1),k6_algstr_0(esk1_0,X1,esk2_0))=k2_group_3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8565, plain, (v1_finset_1(X1)|~m1_subset_1(k1_zfmisc_1(esk41_1(X1)),k5_numbers))).#trainneg
cnf(i_0_2677, plain, (v5_relat_1(esk40_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_10867, negated_conjecture, (k2_group_3(esk1_0,esk2_0,esk2_0)=esk2_0)).#trainneg
cnf(i_0_928, plain, (r2_hidden(esk7_2(k1_tarski(X1),X2),k7_group_3(X3,X4))|v2_struct_0(X3)|r1_tarski(k1_tarski(X1),X2)|~r2_group_3(X3,X4,X1)|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X4,u1_struct_0(X3))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3))).#trainneg
cnf(i_0_929, plain, (r2_hidden(X1,k7_group_3(X2,X3))|v2_struct_0(X2)|v1_xboole_0(k1_tarski(X4))|~r2_group_3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_algstr_0(X2)|~v3_group_1(X2)|~v2_group_1(X2)|~m1_subset_1(X1,k1_tarski(X4)))).#trainneg
cnf(i_0_9278, negated_conjecture, (k6_algstr_0(esk1_0,X1,esk2_0)=esk2_0|k7_group_3(esk1_0,k6_algstr_0(esk1_0,X1,esk2_0))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9280, negated_conjecture, (k6_algstr_0(esk1_0,X1,X2)=k6_algstr_0(esk1_0,k6_algstr_0(esk1_0,X1,esk2_0),X1)|X2!=k6_algstr_0(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9282, negated_conjecture, (X1=esk2_0|k6_algstr_0(esk1_0,X1,X2)!=k6_algstr_0(esk1_0,X2,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9283, negated_conjecture, (X1=X2|k6_algstr_0(esk1_0,X1,esk2_0)!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9284, negated_conjecture, (X1=k6_algstr_0(esk1_0,k2_group_1(esk1_0,esk2_0),X2)|k6_algstr_0(esk1_0,X1,esk2_0)!=X2|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9289, negated_conjecture, (X1=k6_algstr_0(esk1_0,X2,esk2_0)|k6_algstr_0(esk1_0,X1,X3)!=k6_algstr_0(esk1_0,X2,k6_algstr_0(esk1_0,X3,esk2_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9296, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X1,esk2_0)!=k6_algstr_0(esk1_0,esk2_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9297, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|X1=esk2_0|k6_algstr_0(esk1_0,X1,X1)!=k6_algstr_0(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9298, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k7_group_3(esk1_0,esk2_0)|esk2_0=X1|k6_algstr_0(esk1_0,X2,X1)!=k6_algstr_0(esk1_0,X2,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8571, plain, (v1_finset_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_4775, plain, (r1_tarski(u1_struct_0(k7_group_2(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l3_algstr_0(X1)|~v3_group_1(X1)|~v2_group_1(X1))).#trainneg
cnf(i_0_10893, negated_conjecture, (r1_group_3(esk1_0,esk2_0,esk2_0))).#trainneg
cnf(i_0_6243, plain, (v3_card_1(u1_struct_0(esk9_0),k1_xboole_0)|~v1_card_1(k1_xboole_0)|~v2_struct_0(esk9_0))).#trainneg
cnf(i_0_10505, plain, (v1_zfmisc_1(k7_group_3(X1,X2))|v2_struct_0(X3)|~m1_group_2(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_algstr_0(X3)|~v3_group_1(X3)|~v2_group_1(X3)|~v7_struct_0(X3)|~l1_struct_0(X3))).#trainneg
cnf(i_0_9328, negated_conjecture, (k6_algstr_0(esk1_0,k2_group_1(esk1_0,X1),k6_algstr_0(esk1_0,esk2_0,X1))=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9329, negated_conjecture, (k2_group_3(esk1_0,k6_algstr_0(esk1_0,esk2_0,X1),X1)=k6_algstr_0(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9330, negated_conjecture, (esk58_3(k6_algstr_0(esk1_0,esk2_0,X1),esk1_0,k6_algstr_0(esk1_0,esk2_0,X1))=k6_algstr_0(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9333, negated_conjecture, (r2_group_3(esk1_0,k6_algstr_0(esk1_0,esk2_0,X1),X2)|X2!=k6_algstr_0(esk1_0,esk2_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9334, negated_conjecture, (m1_subset_1(k6_algstr_0(esk1_0,esk2_0,X1),k7_group_3(esk1_0,X2))|X2!=k6_algstr_0(esk1_0,esk2_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9335, negated_conjecture, (k6_algstr_0(esk1_0,k6_algstr_0(esk1_0,esk2_0,X1),X1)=k6_algstr_0(esk1_0,X1,k6_algstr_0(esk1_0,esk2_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9339, negated_conjecture, (r1_group_3(esk1_0,X1,k6_algstr_0(esk1_0,esk2_0,X2))|X1!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9341, negated_conjecture, (r2_group_3(esk1_0,X1,k6_algstr_0(esk1_0,esk2_0,X2))|X1!=k6_algstr_0(esk1_0,esk2_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9344, negated_conjecture, (k6_algstr_0(esk1_0,k2_group_1(esk1_0,esk2_0),k6_algstr_0(esk1_0,esk2_0,X1))=k2_group_3(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9347, negated_conjecture, (k6_algstr_0(esk1_0,esk2_0,X1)=esk2_0|k7_group_3(esk1_0,k6_algstr_0(esk1_0,esk2_0,X1))!=k7_group_3(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8971, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_tarski(X1)),k5_numbers))).#trainneg
cnf(i_0_7995, negated_conjecture, (esk47_1(k1_xboole_0)=esk2_0|~v1_card_1(k1_xboole_0)|~m1_subset_1(u1_struct_0(esk1_0),esk2_0))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 199
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 359
# Removed in clause preprocessing      : 23
# Initial clauses in saturation        : 336
# Processed clauses                    : 1888
# ...of these trivial                  : 19
# ...subsumed                          : 600
# ...remaining for further processing  : 1269
# Other redundant clauses eliminated   : 53
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 96
# Backward-rewritten                   : 41
# Generated clauses                    : 11359
# ...of the previous two non-trivial   : 10722
# Contextual simplify-reflections      : 633
# Paramodulations                      : 11260
# Factorizations                       : 0
# Equation resolutions                 : 90
# Propositional unsat checks           : 0
# Propositional unsat check successes  : 0
# Current number of processed clauses  : 1119
#    Positive orientable unit clauses  : 143
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 33
#    Non-unit-clauses                  : 943
# Current number of unprocessed clauses: 7976
# ...number of literals in the above   : 64660
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1415
# Proof object given clauses           : 99
# Proof search given clauses           : 1269
# Clause-clause subsumption calls (NU) : 638114
# Rec. Clause-clause subsumption calls : 55511
# Non-unit clause-clause subsumptions  : 1231
# Unit Clause-clause subsumption calls : 17799
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 38
# BW rewrite match successes           : 20
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 374512

# -------------------------------------------------
# User time                : 4.286 s
# System time              : 0.272 s
# Total time               : 4.558 s
# ...preprocessing         : 1.386 s
# ...main loop             : 3.172 s
# Maximum resident set size: 171720 pages
