# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
# Preprocessing time       : 0.820 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 189 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=982, query=845, total=1829).
#TF#SERVER: Sent 243875 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7312 bytes expected).
#TF#SERVER: Received 7312 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 253 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=1102, query=772, total=1876).
#TF#SERVER: Sent 303658 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7500 bytes expected).
#TF#SERVER: Received 7500 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 373 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=1382, query=688, total=2072).
#TF#SERVER: Sent 446081 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8284 bytes expected).
#TF#SERVER: Received 8284 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 441 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=1557, query=623, total=2182).
#TF#SERVER: Sent 481894 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8724 bytes expected).
#TF#SERVER: Received 8724 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 429 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=1710, query=615, total=2327).
#TF#SERVER: Sent 532426 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9304 bytes expected).
#TF#SERVER: Received 9304 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 353 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=1850, query=672, total=2524).
#TF#SERVER: Sent 590682 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10092 bytes expected).
#TF#SERVER: Received 10092 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 640 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=1951, query=387, total=2340).
#TF#SERVER: Sent 522822 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9356 bytes expected).
#TF#SERVER: Received 9356 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 525 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=499, total=2549).
#TF#SERVER: Sent 583580 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10192 bytes expected).
#TF#SERVER: Received 10192 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 556 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=469, total=2519).
#TF#SERVER: Sent 563434 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10072 bytes expected).
#TF#SERVER: Received 10072 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 510 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=528, total=2578).
#TF#SERVER: Sent 557719 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10308 bytes expected).
#TF#SERVER: Received 10308 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 447 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=586, total=2636).
#TF#SERVER: Sent 584833 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10540 bytes expected).
#TF#SERVER: Received 10540 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 441 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=583, total=2633).
#TF#SERVER: Sent 621250 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10528 bytes expected).
#TF#SERVER: Received 10528 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 482 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=566, total=2616).
#TF#SERVER: Sent 658787 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10460 bytes expected).
#TF#SERVER: Received 10460 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 429 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=609, total=2659).
#TF#SERVER: Sent 706210 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10632 bytes expected).
#TF#SERVER: Received 10632 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 528 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=532, total=2582).
#TF#SERVER: Sent 617137 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10324 bytes expected).
#TF#SERVER: Received 10324 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 639 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=440, total=2490).
#TF#SERVER: Sent 609969 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9956 bytes expected).
#TF#SERVER: Received 9956 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 698 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=386, total=2436).
#TF#SERVER: Sent 616327 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9740 bytes expected).
#TF#SERVER: Received 9740 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 646 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=464, total=2514).
#TF#SERVER: Sent 605728 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10052 bytes expected).
#TF#SERVER: Received 10052 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 653 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=416, total=2466).
#TF#SERVER: Sent 634957 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9860 bytes expected).
#TF#SERVER: Received 9860 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 595 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=500, total=2550).
#TF#SERVER: Sent 653700 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10196 bytes expected).
#TF#SERVER: Received 10196 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 648 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=377, total=2427).
#TF#SERVER: Sent 638672 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9704 bytes expected).
#TF#SERVER: Received 9704 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 750 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=305, total=2355).
#TF#SERVER: Sent 557761 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9416 bytes expected).
#TF#SERVER: Received 9416 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 744 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=296, total=2346).
#TF#SERVER: Sent 559448 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9380 bytes expected).
#TF#SERVER: Received 9380 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 675 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=358, total=2408).
#TF#SERVER: Sent 658383 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9628 bytes expected).
#TF#SERVER: Received 9628 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 707 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=421, total=2471).
#TF#SERVER: Sent 597232 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9880 bytes expected).
#TF#SERVER: Received 9880 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 744 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=436, total=2486).
#TF#SERVER: Sent 601615 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9940 bytes expected).
#TF#SERVER: Received 9940 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 595 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=436, total=2486).
#TF#SERVER: Sent 573994 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9940 bytes expected).
#TF#SERVER: Received 9940 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 650 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=384, total=2434).
#TF#SERVER: Sent 581671 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9732 bytes expected).
#TF#SERVER: Received 9732 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 763 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=362, total=2412).
#TF#SERVER: Sent 592250 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9644 bytes expected).
#TF#SERVER: Received 9644 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 711 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=315, total=2365).
#TF#SERVER: Sent 573203 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9456 bytes expected).
#TF#SERVER: Received 9456 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 587 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=539, total=2589).
#TF#SERVER: Sent 653171 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10352 bytes expected).
#TF#SERVER: Received 10352 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 412 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=635, total=2685).
#TF#SERVER: Sent 685454 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10736 bytes expected).
#TF#SERVER: Received 10736 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 554 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=493, total=2543).
#TF#SERVER: Sent 661820 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10168 bytes expected).
#TF#SERVER: Received 10168 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 489 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=550, total=2600).
#TF#SERVER: Sent 620920 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10396 bytes expected).
#TF#SERVER: Received 10396 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 534 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=521, total=2571).
#TF#SERVER: Sent 569447 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10280 bytes expected).
#TF#SERVER: Received 10280 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 584 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=443, total=2493).
#TF#SERVER: Sent 600230 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9968 bytes expected).
#TF#SERVER: Received 9968 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 519 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=514, total=2564).
#TF#SERVER: Sent 587188 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10252 bytes expected).
#TF#SERVER: Received 10252 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 662 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=394, total=2444).
#TF#SERVER: Sent 613863 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9772 bytes expected).
#TF#SERVER: Received 9772 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 524 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=584, total=2634).
#TF#SERVER: Sent 644357 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10532 bytes expected).
#TF#SERVER: Received 10532 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 536 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=511, total=2561).
#TF#SERVER: Sent 590390 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10240 bytes expected).
#TF#SERVER: Received 10240 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 586 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=548, total=2598).
#TF#SERVER: Sent 636797 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10388 bytes expected).
#TF#SERVER: Received 10388 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 541 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=526, total=2576).
#TF#SERVER: Sent 643570 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10300 bytes expected).
#TF#SERVER: Received 10300 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 513 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=522, total=2572).
#TF#SERVER: Sent 602887 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10284 bytes expected).
#TF#SERVER: Received 10284 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 725 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=359, total=2409).
#TF#SERVER: Sent 595892 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9632 bytes expected).
#TF#SERVER: Received 9632 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 585 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=476, total=2526).
#TF#SERVER: Sent 610740 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10100 bytes expected).
#TF#SERVER: Received 10100 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 533 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=531, total=2581).
#TF#SERVER: Sent 637610 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10320 bytes expected).
#TF#SERVER: Received 10320 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 619 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=529, total=2579).
#TF#SERVER: Sent 656594 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10312 bytes expected).
#TF#SERVER: Received 10312 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 445 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=618, total=2668).
#TF#SERVER: Sent 689914 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10668 bytes expected).
#TF#SERVER: Received 10668 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 516 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=669, total=2719).
#TF#SERVER: Sent 774987 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10872 bytes expected).
#TF#SERVER: Received 10872 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 486 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=563, total=2613).
#TF#SERVER: Sent 722892 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10448 bytes expected).
#TF#SERVER: Received 10448 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 455 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=582, total=2632).
#TF#SERVER: Sent 674124 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10524 bytes expected).
#TF#SERVER: Received 10524 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 677 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=520, total=2570).
#TF#SERVER: Sent 702550 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10276 bytes expected).
#TF#SERVER: Received 10276 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 449 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=598, total=2648).
#TF#SERVER: Sent 689169 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10588 bytes expected).
#TF#SERVER: Received 10588 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 585 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=634, total=2684).
#TF#SERVER: Sent 676941 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10732 bytes expected).
#TF#SERVER: Received 10732 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 529 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=595, total=2645).
#TF#SERVER: Sent 710515 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10576 bytes expected).
#TF#SERVER: Received 10576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 569 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=504, total=2554).
#TF#SERVER: Sent 695769 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10212 bytes expected).
#TF#SERVER: Received 10212 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 610 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=422, total=2472).
#TF#SERVER: Sent 609176 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9884 bytes expected).
#TF#SERVER: Received 9884 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 563 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=561, total=2611).
#TF#SERVER: Sent 694214 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10440 bytes expected).
#TF#SERVER: Received 10440 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 374 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=665, total=2715).
#TF#SERVER: Sent 764265 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10856 bytes expected).
#TF#SERVER: Received 10856 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 423 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=674, total=2724).
#TF#SERVER: Sent 810246 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10892 bytes expected).
#TF#SERVER: Received 10892 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 491 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=551, total=2601).
#TF#SERVER: Sent 724845 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10400 bytes expected).
#TF#SERVER: Received 10400 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 560 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=652, total=2702).
#TF#SERVER: Sent 742225 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10804 bytes expected).
#TF#SERVER: Received 10804 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 586 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=659, total=2709).
#TF#SERVER: Sent 796786 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10832 bytes expected).
#TF#SERVER: Received 10832 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 572 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=598, total=2648).
#TF#SERVER: Sent 748870 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10588 bytes expected).
#TF#SERVER: Received 10588 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 465 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=622, total=2672).
#TF#SERVER: Sent 785886 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10684 bytes expected).
#TF#SERVER: Received 10684 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 620 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=488, total=2538).
#TF#SERVER: Sent 663847 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10148 bytes expected).
#TF#SERVER: Received 10148 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 511 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=590, total=2640).
#TF#SERVER: Sent 775768 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10556 bytes expected).
#TF#SERVER: Received 10556 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 729 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=471, total=2521).
#TF#SERVER: Sent 703375 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10080 bytes expected).
#TF#SERVER: Received 10080 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 607 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=473, total=2523).
#TF#SERVER: Sent 704981 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10088 bytes expected).
#TF#SERVER: Received 10088 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 783 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=365, total=2415).
#TF#SERVER: Sent 638356 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9656 bytes expected).
#TF#SERVER: Received 9656 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 739 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=422, total=2472).
#TF#SERVER: Sent 601467 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9884 bytes expected).
#TF#SERVER: Received 9884 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 687 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=357, total=2407).
#TF#SERVER: Sent 594107 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9624 bytes expected).
#TF#SERVER: Received 9624 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 637 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=436, total=2486).
#TF#SERVER: Sent 682469 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9940 bytes expected).
#TF#SERVER: Received 9940 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 700 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=427, total=2477).
#TF#SERVER: Sent 647840 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9904 bytes expected).
#TF#SERVER: Received 9904 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 451 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=633, total=2683).
#TF#SERVER: Sent 733758 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10728 bytes expected).
#TF#SERVER: Received 10728 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 692 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=472, total=2522).
#TF#SERVER: Sent 637802 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10084 bytes expected).
#TF#SERVER: Received 10084 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 630 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=595, total=2645).
#TF#SERVER: Sent 684465 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10576 bytes expected).
#TF#SERVER: Received 10576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 636 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=495, total=2545).
#TF#SERVER: Sent 676074 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10176 bytes expected).
#TF#SERVER: Received 10176 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 606 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=418, total=2468).
#TF#SERVER: Sent 627250 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9868 bytes expected).
#TF#SERVER: Received 9868 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 601 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=472, total=2522).
#TF#SERVER: Sent 666173 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10084 bytes expected).
#TF#SERVER: Received 10084 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 693 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=393, total=2443).
#TF#SERVER: Sent 647993 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9768 bytes expected).
#TF#SERVER: Received 9768 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 508 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=745, total=2795).
#TF#SERVER: Sent 849241 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11176 bytes expected).
#TF#SERVER: Received 11176 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 595 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=640, total=2690).
#TF#SERVER: Sent 795384 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10756 bytes expected).
#TF#SERVER: Received 10756 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 864 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=456, total=2506).
#TF#SERVER: Sent 663536 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10020 bytes expected).
#TF#SERVER: Received 10020 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 669 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=360, total=2410).
#TF#SERVER: Sent 635406 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9636 bytes expected).
#TF#SERVER: Received 9636 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 638 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=496, total=2546).
#TF#SERVER: Sent 745512 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10180 bytes expected).
#TF#SERVER: Received 10180 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 818 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=278, total=2328).
#TF#SERVER: Sent 605996 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9308 bytes expected).
#TF#SERVER: Received 9308 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 696 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=354, total=2404).
#TF#SERVER: Sent 627886 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9612 bytes expected).
#TF#SERVER: Received 9612 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 937 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=361, total=2411).
#TF#SERVER: Sent 624614 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9640 bytes expected).
#TF#SERVER: Received 9640 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 634 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=425, total=2475).
#TF#SERVER: Sent 594185 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9896 bytes expected).
#TF#SERVER: Received 9896 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 756 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=310, total=2360).
#TF#SERVER: Sent 569765 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9436 bytes expected).
#TF#SERVER: Received 9436 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 695 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=473, total=2523).
#TF#SERVER: Sent 725901 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10088 bytes expected).
#TF#SERVER: Received 10088 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 722 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=420, total=2470).
#TF#SERVER: Sent 606535 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9876 bytes expected).
#TF#SERVER: Received 9876 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 581 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=453, total=2503).
#TF#SERVER: Sent 530788 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10008 bytes expected).
#TF#SERVER: Received 10008 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 242 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=878, total=2928).
#TF#SERVER: Sent 553495 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11708 bytes expected).
#TF#SERVER: Received 11708 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 661 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=407, total=2457).
#TF#SERVER: Sent 609581 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9824 bytes expected).
#TF#SERVER: Received 9824 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 574 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=483, total=2533).
#TF#SERVER: Sent 637079 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10128 bytes expected).
#TF#SERVER: Received 10128 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 769 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=554, total=2604).
#TF#SERVER: Sent 667879 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10412 bytes expected).
#TF#SERVER: Received 10412 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 447 clauses
#TF#SERVER: Sending data to the TF server (conj=2, context=2048, query=707, total=2757).
#TF#SERVER: Sent 667158 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11024 bytes expected).
#TF#SERVER: Received 11024 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t128_xxreal_1, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>(r1_xxreal_0(X1,X2)=>k1_xxreal_1(X1,X2)=k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t128_xxreal_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', commutativity_k2_xboole_0)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', d2_tarski)).
fof(redefinition_k2_rcomp_1, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k2_rcomp_1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', d3_xboole_0)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', commutativity_k2_tarski)).
fof(redefinition_k1_rcomp_1, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>k1_rcomp_1(X1,X2)=k1_xxreal_1(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k1_rcomp_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t2_subset)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t7_boole)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc4_xreal_0)).
fof(fraenkel_a_2_1_rcomp_1, axiom, ![X1, X2, X3]:((v1_xxreal_0(X2)&v1_xxreal_0(X3))=>(r2_hidden(X1,a_2_1_rcomp_1(X2,X3))<=>?[X4]:(((m1_subset_1(X4,k1_numbers)&X1=X4)&~(r1_xxreal_0(X4,X2)))&~(r1_xxreal_0(X3,X4))))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fraenkel_a_2_1_rcomp_1)).
fof(d2_rcomp_1, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>k2_rcomp_1(X1,X2)=a_2_1_rcomp_1(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', d2_rcomp_1)).
fof(fraenkel_a_2_0_rcomp_1, axiom, ![X1, X2, X3]:((v1_xreal_0(X2)&v1_xreal_0(X3))=>(r2_hidden(X1,a_2_0_rcomp_1(X2,X3))<=>?[X4]:(((m1_subset_1(X4,k1_numbers)&X1=X4)&r1_xxreal_0(X2,X4))&r1_xxreal_0(X4,X3)))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fraenkel_a_2_0_rcomp_1)).
fof(d1_rcomp_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>k1_rcomp_1(X1,X2)=a_2_0_rcomp_1(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', d1_rcomp_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t1_subset)).
fof(reflexivity_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>r1_xxreal_0(X1,X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', reflexivity_r1_xxreal_0)).
fof(rc3_xxreal_0, axiom, ?[X1]:(v1_xxreal_0(X1)&v3_xxreal_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', rc3_xxreal_0)).
fof(t24_xreal_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>(r1_xxreal_0(X1,X2)<=>r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t24_xreal_1)).
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k1_real_1)).
fof(involutiveness_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(k1_real_1(X1))=X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', involutiveness_k1_real_1)).
fof(fc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_xreal_0(k4_xcmplx_0(X1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fc3_xreal_0)).
fof(dt_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k1_real_1(X1),k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k1_real_1)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc14_membered)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fc3_membered)).
fof(dt_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k8_real_1(X1,X2),k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k8_real_1)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k8_real_1)).
fof(t3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t3_arithm)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc3_xreal_0)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', spc1_numerals)).
fof(dt_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k10_real_1(X1,X2),k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k10_real_1)).
fof(redefinition_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k10_real_1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc1_xreal_0)).
fof(involutiveness_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k5_xcmplx_0(k5_xcmplx_0(X1))=X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', involutiveness_k5_xcmplx_0)).
fof(spc3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', spc3_arithm)).
fof(dt_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k5_xcmplx_0(X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k5_xcmplx_0)).
fof(fc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k5_xcmplx_0(X1))&v1_xreal_0(k5_xcmplx_0(X1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fc4_xreal_0)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fc3_xboole_0)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc13_membered)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc4_membered)).
fof(t86_sincos10, conjecture, ![X1]:(r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))=>r2_hidden(k1_seq_1(k2_sincos10,X1),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t86_sincos10)).
fof(rqRealMult__k3_xcmplx_0__rnm1d2_r2_rm1, axiom, k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k4_xcmplx_0(np__1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', rqRealMult__k3_xcmplx_0__rnm1d2_r2_rm1)).
fof(rqRealDiv__k7_xcmplx_0__r1_rm2_rnm1d2, axiom, k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', rqRealDiv__k7_xcmplx_0__r1_rm2_rnm1d2)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', spc2_numerals)).
fof(t20_rfunct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))<=>![X3]:(m1_subset_1(X3,k1_numbers)=>![X4]:(m1_subset_1(X4,k1_numbers)=>~((((r2_hidden(X3,k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2)))&r2_hidden(X4,k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2))))&~(r1_xxreal_0(X4,X3)))&r1_xxreal_0(k1_seq_1(X2,X4),k1_seq_1(X2,X3)))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t20_rfunct_2)).
fof(commutativity_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k9_subset_1(X1,X3,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', commutativity_k9_subset_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t28_xboole_1)).
fof(t46_sincos10, axiom, r1_tarski(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)),k1_relset_1(k1_numbers,k2_sincos10)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t46_sincos10)).
fof(redefinition_k2_partfun1, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k2_partfun1(X1,X2,X3,X4)=k5_relat_1(X3,X4)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k2_partfun1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t4_subset)).
fof(dt_k1_rcomp_1, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k1_rcomp_1)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k9_subset_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', commutativity_k3_xboole_0)).
fof(dt_k2_sincos10, axiom, (v1_funct_1(k2_sincos10)&m1_subset_1(k2_sincos10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k2_sincos10)).
fof(fc8_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k7_xcmplx_0(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fc8_xreal_0)).
fof(t2_xxreal_1, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>(r2_hidden(X1,k3_xxreal_1(X2,X3))<=>(~(r1_xxreal_0(X1,X2))&r1_xxreal_0(X1,X3)))))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t2_xxreal_1)).
fof(redefinition_k4_rcomp_1, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xreal_0(X2))=>k4_rcomp_1(X1,X2)=k3_xxreal_1(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k4_rcomp_1)).
fof(rqRealDiv__k7_xcmplx_0__r1_rm1_rm1, axiom, k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', rqRealDiv__k7_xcmplx_0__r1_rm1_rm1)).
fof(t82_sincos10, axiom, v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t82_sincos10)).
fof(dt_k4_rcomp_1, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xreal_0(X2))=>m1_subset_1(k4_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k4_rcomp_1)).
fof(l18_sincos10, axiom, (r2_hidden(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))&r2_hidden(k32_sin_cos,k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', l18_sincos10)).
fof(dt_k32_sin_cos, axiom, m1_subset_1(k32_sin_cos,k1_numbers), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k32_sin_cos)).
fof(dt_k31_sin_cos, axiom, v1_xreal_0(k31_sin_cos), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k31_sin_cos)).
fof(redefinition_k32_sin_cos, axiom, k32_sin_cos=k31_sin_cos, file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', redefinition_k32_sin_cos)).
fof(dt_k7_square_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k7_square_1(X1),k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k7_square_1)).
fof(t74_sincos10, axiom, (k1_seq_1(k2_sincos10,k1_real_1(k7_square_1(np__2)))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)&k1_seq_1(k2_sincos10,k1_real_1(np__1))=k32_sin_cos), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t74_sincos10)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', connectedness_r1_xxreal_0)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc12_membered)).
fof(rqLessOrEqual__r1_xxreal_0__rm1_rm1, axiom, r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', rqLessOrEqual__r1_xxreal_0__rm1_rm1)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', spc3_numerals)).
fof(spc4_numerals, axiom, (((v2_xxreal_0(np__4)&m2_subset_1(np__4,k1_numbers,k5_numbers))&m1_subset_1(np__4,k5_numbers))&m1_subset_1(np__4,k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', spc4_numerals)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', involutiveness_k4_xcmplx_0)).
fof(cc1_xcmplx_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc1_xcmplx_0)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc5_membered)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(fc5_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k2_xcmplx_0(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', fc5_xreal_0)).
fof(spc2_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', spc2_arithm)).
fof(t19_square_1, axiom, ~(r1_xxreal_0(k7_square_1(np__2),np__1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', t19_square_1)).
fof(rqRealAdd__k2_xcmplx_0__rnm1d4_r1_rn3d4, axiom, k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),np__1)=k7_xcmplx_0(np__3,np__4), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', rqRealAdd__k2_xcmplx_0__rnm1d4_r1_rn3d4)).
fof(spc4_arithm, axiom, ![X1, X2, X3]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', spc4_arithm)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc18_valued_0)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', cc1_relset_1)).
fof(commutativity_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k8_real_1(X2,X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', commutativity_k8_real_1)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t86_sincos10', dt_k1_seq_1)).
fof(c_0_83, plain, ![X638, X639]:(~v1_xxreal_0(X638)|(~v1_xxreal_0(X639)|(~r1_xxreal_0(X638,X639)|k1_xxreal_1(X638,X639)=k2_xboole_0(k4_xxreal_1(X638,X639),k2_tarski(X638,X639))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t128_xxreal_1])])])).
fof(c_0_84, plain, ![X163, X164]:k2_xboole_0(X163,X164)=k2_xboole_0(X164,X163), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_85, plain, ![X182, X183, X184, X185, X186, X187, X188, X189]:(((~r2_hidden(X185,X184)|(X185=X182|X185=X183)|X184!=k2_tarski(X182,X183))&((X186!=X182|r2_hidden(X186,X184)|X184!=k2_tarski(X182,X183))&(X186!=X183|r2_hidden(X186,X184)|X184!=k2_tarski(X182,X183))))&(((esk2_3(X187,X188,X189)!=X187|~r2_hidden(esk2_3(X187,X188,X189),X189)|X189=k2_tarski(X187,X188))&(esk2_3(X187,X188,X189)!=X188|~r2_hidden(esk2_3(X187,X188,X189),X189)|X189=k2_tarski(X187,X188)))&(r2_hidden(esk2_3(X187,X188,X189),X189)|(esk2_3(X187,X188,X189)=X187|esk2_3(X187,X188,X189)=X188)|X189=k2_tarski(X187,X188)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_86, plain, (k1_xxreal_1(X1,X2)=k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_87, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_88, plain, ![X593, X594]:(~v1_xxreal_0(X593)|~v1_xxreal_0(X594)|k2_rcomp_1(X593,X594)=k4_xxreal_1(X593,X594)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_rcomp_1])])).
fof(c_0_89, plain, ![X191, X192, X193, X194, X195, X196, X197, X198]:(((~r2_hidden(X194,X193)|(r2_hidden(X194,X191)|r2_hidden(X194,X192))|X193!=k2_xboole_0(X191,X192))&((~r2_hidden(X195,X191)|r2_hidden(X195,X193)|X193!=k2_xboole_0(X191,X192))&(~r2_hidden(X195,X192)|r2_hidden(X195,X193)|X193!=k2_xboole_0(X191,X192))))&(((~r2_hidden(esk3_3(X196,X197,X198),X196)|~r2_hidden(esk3_3(X196,X197,X198),X198)|X198=k2_xboole_0(X196,X197))&(~r2_hidden(esk3_3(X196,X197,X198),X197)|~r2_hidden(esk3_3(X196,X197,X198),X198)|X198=k2_xboole_0(X196,X197)))&(r2_hidden(esk3_3(X196,X197,X198),X198)|(r2_hidden(esk3_3(X196,X197,X198),X196)|r2_hidden(esk3_3(X196,X197,X198),X197))|X198=k2_xboole_0(X196,X197)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_90, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, plain, (k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_92, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_93, plain, (k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_94, plain, ![X161, X162]:k2_tarski(X161,X162)=k2_tarski(X162,X161), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_95, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_96, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_97, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_98, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_99, plain, (k2_xboole_0(k2_tarski(X1,X2),k4_xxreal_1(X1,X2))=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(rw,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_100, plain, (k4_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_101, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_102, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_103, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_96])])).
cnf(c_0_104, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k2_xboole_0(X3,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_97])).
fof(c_0_105, plain, ![X579, X580]:(~v1_xreal_0(X579)|~v1_xreal_0(X580)|k1_rcomp_1(X579,X580)=k1_xxreal_1(X579,X580)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_rcomp_1])])).
cnf(c_0_106, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X4,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_107, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_rcomp_1(X1,X2))=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_108, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(evalgc,[status(thm)],[c_0_101])).
fof(c_0_109, plain, ![X660, X661]:(~m1_subset_1(X660,X661)|(v1_xboole_0(X661)|r2_hidden(X660,X661))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_110, plain, ![X688, X689]:(~r2_hidden(X688,X689)|~v1_xboole_0(X689)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_111, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r2_hidden(X1,k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_99])])).
cnf(c_0_112, plain, (r2_hidden(X1,k2_tarski(X2,X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_103])])).
cnf(c_0_113, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_104])])).
cnf(c_0_114, plain, (k1_rcomp_1(X1,X2)=k1_xxreal_1(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_105])).
fof(c_0_115, plain, ![X127]:(~v1_xreal_0(X127)|v1_xxreal_0(X127)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
cnf(c_0_116, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_106])])).
cnf(c_0_117, plain, (k2_xboole_0(k2_tarski(X2,X1),k2_rcomp_1(X1,X2))=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108])])).
cnf(c_0_118, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_119, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_120, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X1)|~r1_xxreal_0(X3,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_121, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_107])])).
cnf(c_0_122, plain, (k1_xxreal_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_123, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_124, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_107])])).
cnf(c_0_125, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_117])])).
cnf(c_0_126, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_127, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_128, plain, (r2_hidden(X1,k1_xxreal_1(X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_120])])).
cnf(c_0_129, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_102])])).
cnf(c_0_130, plain, (k2_xboole_0(k2_tarski(X2,X1),k4_xxreal_1(X1,X2))=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_108])])).
cnf(c_0_131, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122])])).
cnf(c_0_132, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_123])).
cnf(c_0_133, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_122])])).
cnf(c_0_134, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|v1_xboole_0(k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_135, plain, (~v1_xboole_0(k1_xxreal_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_136, plain, (r2_hidden(X1,k1_xxreal_1(X3,X2))|~r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_137, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r2_hidden(X1,k2_rcomp_1(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107])])).
cnf(c_0_138, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_131, c_0_132]), c_0_132])).
cnf(c_0_139, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_133, c_0_132]), c_0_132])).
cnf(c_0_140, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(csr,[status(thm)],[c_0_134, c_0_135])).
cnf(c_0_141, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_122])])).
cnf(c_0_142, plain, (r2_hidden(X1,X2)|r2_hidden(X1,k2_tarski(X3,X4))|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138])])).
cnf(c_0_143, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140])])).
cnf(c_0_144, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_141, c_0_132]), c_0_132])).
cnf(c_0_145, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_142, c_0_132]), c_0_132])).
cnf(c_0_146, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_143, c_0_132]), c_0_132]), c_0_144])).
fof(c_0_147, plain, ![X512, X513, X514, X516]:(((((m1_subset_1(esk7_3(X512,X513,X514),k1_numbers)|~r2_hidden(X512,a_2_1_rcomp_1(X513,X514))|(~v1_xxreal_0(X513)|~v1_xxreal_0(X514)))&(X512=esk7_3(X512,X513,X514)|~r2_hidden(X512,a_2_1_rcomp_1(X513,X514))|(~v1_xxreal_0(X513)|~v1_xxreal_0(X514))))&(~r1_xxreal_0(esk7_3(X512,X513,X514),X513)|~r2_hidden(X512,a_2_1_rcomp_1(X513,X514))|(~v1_xxreal_0(X513)|~v1_xxreal_0(X514))))&(~r1_xxreal_0(X514,esk7_3(X512,X513,X514))|~r2_hidden(X512,a_2_1_rcomp_1(X513,X514))|(~v1_xxreal_0(X513)|~v1_xxreal_0(X514))))&(~m1_subset_1(X516,k1_numbers)|X512!=X516|r1_xxreal_0(X516,X513)|r1_xxreal_0(X514,X516)|r2_hidden(X512,a_2_1_rcomp_1(X513,X514))|(~v1_xxreal_0(X513)|~v1_xxreal_0(X514)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_rcomp_1])])])])])])).
fof(c_0_148, plain, ![X180, X181]:(~v1_xxreal_0(X180)|(~v1_xxreal_0(X181)|k2_rcomp_1(X180,X181)=a_2_1_rcomp_1(X180,X181))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_rcomp_1])])])).
cnf(c_0_149, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_150, plain, (X1=esk7_3(X1,X2,X3)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_151, plain, (k2_rcomp_1(X1,X2)=a_2_1_rcomp_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_152, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_149])])).
cnf(c_0_153, plain, (esk7_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_154, plain, (a_2_1_rcomp_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_155, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_152])])).
cnf(c_0_156, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_157, plain, (esk7_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_158, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_155, c_0_132]), c_0_132])).
cnf(c_0_159, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_160, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158])])).
cnf(c_0_161, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_99])])).
cnf(c_0_162, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_159])])).
cnf(c_0_163, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_160, c_0_132]), c_0_132])).
cnf(c_0_164, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_122])])).
cnf(c_0_165, plain, (~r1_xxreal_0(X1,esk7_3(X2,X3,X1))|~r2_hidden(X2,a_2_1_rcomp_1(X3,X1))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_166, plain, (X1=X2|X1=X3|esk7_3(X1,X3,X2)=X1|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
cnf(c_0_167, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_122])])).
cnf(c_0_168, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_164, c_0_132]), c_0_132])).
cnf(c_0_169, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_170, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_2_1_rcomp_1(X3,X1))|~r1_xxreal_0(X1,esk7_3(X2,X3,X1))), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_171, plain, (esk7_3(X1,X2,X3)=X1|X1=X2|X1=X3|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_122])])).
cnf(c_0_172, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[c_0_167, c_0_168])).
cnf(c_0_173, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_174, plain, (X2=X1|X2=X3|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_2_1_rcomp_1(X3,X1))|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_rcomp_1(X3,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171])])).
cnf(c_0_175, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_172])])).
fof(c_0_176, plain, ![X507, X508, X509, X511]:(((((m1_subset_1(esk6_3(X507,X508,X509),k1_numbers)|~r2_hidden(X507,a_2_0_rcomp_1(X508,X509))|(~v1_xreal_0(X508)|~v1_xreal_0(X509)))&(X507=esk6_3(X507,X508,X509)|~r2_hidden(X507,a_2_0_rcomp_1(X508,X509))|(~v1_xreal_0(X508)|~v1_xreal_0(X509))))&(r1_xxreal_0(X508,esk6_3(X507,X508,X509))|~r2_hidden(X507,a_2_0_rcomp_1(X508,X509))|(~v1_xreal_0(X508)|~v1_xreal_0(X509))))&(r1_xxreal_0(esk6_3(X507,X508,X509),X509)|~r2_hidden(X507,a_2_0_rcomp_1(X508,X509))|(~v1_xreal_0(X508)|~v1_xreal_0(X509))))&(~m1_subset_1(X511,k1_numbers)|X507!=X511|~r1_xxreal_0(X508,X511)|~r1_xxreal_0(X511,X509)|r2_hidden(X507,a_2_0_rcomp_1(X508,X509))|(~v1_xreal_0(X508)|~v1_xreal_0(X509)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_2_0_rcomp_1])])])])])).
cnf(c_0_177, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_173])])).
cnf(c_0_178, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_174, c_0_132]), c_0_132])).
cnf(c_0_179, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_175])])).
cnf(c_0_180, plain, (r2_hidden(X2,a_2_0_rcomp_1(X3,X4))|~m1_subset_1(X1,k1_numbers)|X2!=X1|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~v1_xreal_0(X3)|~v1_xreal_0(X4)), inference(split_conjunct,[status(thm)],[c_0_176])).
fof(c_0_181, plain, ![X178, X179]:(~v1_xreal_0(X178)|(~v1_xreal_0(X179)|k1_rcomp_1(X178,X179)=a_2_0_rcomp_1(X178,X179))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_rcomp_1])])])).
cnf(c_0_182, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_177])])).
cnf(c_0_183, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_154])])).
cnf(c_0_184, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_179])])).
fof(c_0_185, plain, ![X643, X644]:(~r2_hidden(X643,X644)|m1_subset_1(X643,X644)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_186, plain, (r2_hidden(X2,a_2_0_rcomp_1(X3,X4))|X1!=X2|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_187, plain, (k1_rcomp_1(X1,X2)=a_2_0_rcomp_1(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_188, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X1,X3)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_182])])).
cnf(c_0_189, plain, (X1=X2|X1=X3|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2))|~r2_hidden(X1,k2_rcomp_1(X3,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_183, c_0_132]), c_0_132])).
cnf(c_0_190, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_184, c_0_132]), c_0_132])).
cnf(c_0_191, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_192, plain, (r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_186])])).
cnf(c_0_193, plain, (a_2_0_rcomp_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_187])).
cnf(c_0_194, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_195, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X1,X3)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_122])])).
cnf(c_0_196, plain, (X1=esk6_3(X1,X2,X3)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_197, plain, (X1=X2|X1=X3|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190])])).
cnf(c_0_198, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_199, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193])])).
cnf(c_0_200, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_201, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X1,X3)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_195, c_0_132]), c_0_132])).
cnf(c_0_202, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_196])).
fof(c_0_203, plain, ![X610, X611]:(~v1_xxreal_0(X610)|~v1_xxreal_0(X611)|r1_xxreal_0(X610,X610)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_xxreal_0])])).
fof(c_0_204, plain, (v1_xxreal_0(esk35_0)&v3_xxreal_0(esk35_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc3_xxreal_0])])).
cnf(c_0_205, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[c_0_197, c_0_162])).
cnf(c_0_206, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199])])).
cnf(c_0_207, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_193])])).
cnf(c_0_208, plain, (r2_hidden(X1,k1_rcomp_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_201])])).
cnf(c_0_209, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_193])])).
cnf(c_0_210, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_203])).
cnf(c_0_211, plain, (v1_xxreal_0(esk35_0)), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_212, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~r2_hidden(X2,a_2_0_rcomp_1(X1,X3))|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_213, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206])])).
cnf(c_0_214, plain, (m1_subset_1(esk6_3(X1,X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208])])).
cnf(c_0_215, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_168])])).
cnf(c_0_216, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_217, plain, (v1_xxreal_0(esk35_0)), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_218, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,a_2_0_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_219, plain, (esk6_3(X1,X1,X2)=X3|esk6_3(X1,X1,X2)=X4|~r1_xxreal_0(X3,esk6_3(X1,X1,X2))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,esk6_3(X1,X1,X2))|~r1_xxreal_0(esk6_3(X1,X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214])])).
cnf(c_0_220, plain, (esk6_3(X1,X2,X1)=X1|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_112])])).
cnf(c_0_221, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217])])).
cnf(c_0_222, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k1_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_193])])).
cnf(c_0_223, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220])])).
cnf(c_0_224, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_132])])).
cnf(c_0_225, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X2,k1_xxreal_1(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_146])])).
cnf(c_0_226, plain, (X1=X2|X1=X3|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_223, c_0_224])).
cnf(c_0_227, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_rcomp_1(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_122])])).
cnf(c_0_228, plain, (esk6_3(X1,X2,X3)=X2|esk6_3(X1,X2,X3)=X4|~r1_xxreal_0(X4,esk6_3(X1,X2,X3))|~r1_xxreal_0(esk6_3(X1,X2,X3),X2)|~r1_xxreal_0(X4,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(esk6_3(X1,X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227])])).
cnf(c_0_229, plain, (esk6_3(X1,X2,X3)=X2|~r1_xxreal_0(esk6_3(X1,X2,X3),X2)|~r1_xxreal_0(X2,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(esk6_3(X1,X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_227])])).
cnf(c_0_230, plain, (esk6_3(X1,X2,X3)=X2|~r1_xxreal_0(esk6_3(X1,X2,X3),X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(esk6_3(X1,X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[c_0_229, c_0_224])).
cnf(c_0_231, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_168])])).
fof(c_0_232, plain, ![X651, X652]:((~r1_xxreal_0(X651,X652)|r1_xxreal_0(k4_xcmplx_0(X652),k4_xcmplx_0(X651))|~v1_xreal_0(X652)|~v1_xreal_0(X651))&(~r1_xxreal_0(k4_xcmplx_0(X652),k4_xcmplx_0(X651))|r1_xxreal_0(X651,X652)|~v1_xreal_0(X652)|~v1_xreal_0(X651))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_xreal_1])])])])).
fof(c_0_233, plain, ![X581]:(~m1_subset_1(X581,k1_numbers)|k1_real_1(X581)=k4_xcmplx_0(X581)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
fof(c_0_234, plain, ![X527]:(~m1_subset_1(X527,k1_numbers)|k1_real_1(k1_real_1(X527))=X527), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k1_real_1])])).
cnf(c_0_235, plain, (X1=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_220])])).
cnf(c_0_236, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_112])])).
cnf(c_0_237, plain, (r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_232])).
fof(c_0_238, plain, ![X406]:((v1_xcmplx_0(k4_xcmplx_0(X406))|~v1_xreal_0(X406))&(v1_xreal_0(k4_xcmplx_0(X406))|~v1_xreal_0(X406))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xreal_0])])])).
cnf(c_0_239, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_233])).
cnf(c_0_240, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_234])).
fof(c_0_241, plain, ![X210]:(~m1_subset_1(X210,k1_numbers)|m1_subset_1(k1_real_1(X210),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_real_1])])).
cnf(c_0_242, plain, (~r1_xxreal_0(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_243, plain, (X1=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_235, c_0_236])).
cnf(c_0_244, plain, (r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_237])).
cnf(c_0_245, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_238])).
cnf(c_0_246, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_247, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_248, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_241])).
fof(c_0_249, plain, ![X28, X29]:(~v3_membered(X28)|(~m1_subset_1(X29,X28)|v1_xreal_0(X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_250, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~r1_xxreal_0(esk7_3(X1,X2,X3),X2)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_251, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244])])).
cnf(c_0_252, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_245])).
cnf(c_0_253, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247])])).
cnf(c_0_254, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_255, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_249])).
cnf(c_0_256, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_257, plain, ![X237, X238]:(~m1_subset_1(X237,k1_numbers)|~v1_xreal_0(X238)|m1_subset_1(k8_real_1(X237,X238),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_real_1])])).
fof(c_0_258, plain, ![X601, X602]:(~m1_subset_1(X601,k1_numbers)|~v1_xreal_0(X602)|k8_real_1(X601,X602)=k3_xcmplx_0(X601,X602)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
fof(c_0_259, plain, ![X669]:(~v1_xcmplx_0(X669)|k3_xcmplx_0(np__1,X669)=X669), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
fof(c_0_260, plain, ![X116]:(~v1_xreal_0(X116)|v1_xcmplx_0(X116)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_261, plain, (X1=X3|X1=X2|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_171])])).
cnf(c_0_262, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_251, c_0_252]), c_0_252])).
cnf(c_0_263, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_253, c_0_254])).
cnf(c_0_264, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_255])).
cnf(c_0_265, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_256])).
cnf(c_0_266, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_257])).
cnf(c_0_267, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_268, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_259])).
cnf(c_0_269, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_260])).
cnf(c_0_270, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_261, c_0_132]), c_0_132])).
cnf(c_0_271, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_263])])).
cnf(c_0_272, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_254]), c_0_265])])])).
cnf(c_0_273, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_266])).
cnf(c_0_274, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_275, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_276, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_277, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_278, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_154])])).
fof(c_0_279, plain, ![X206, X207]:(~m1_subset_1(X206,k1_numbers)|~v1_xreal_0(X207)|m1_subset_1(k10_real_1(X206,X207),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_real_1])])).
fof(c_0_280, plain, ![X577, X578]:(~m1_subset_1(X577,k1_numbers)|~v1_xreal_0(X578)|k10_real_1(X577,X578)=k7_xcmplx_0(X577,X578)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k10_real_1])])).
cnf(c_0_281, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_271, c_0_272])).
cnf(c_0_282, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_274])])).
cnf(c_0_283, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_276])])).
cnf(c_0_284, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_285, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2))|~r2_hidden(X1,k2_rcomp_1(X3,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_278, c_0_132]), c_0_132])).
cnf(c_0_286, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_287, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_280])).
fof(c_0_288, plain, ![X69]:(~m1_subset_1(X69,k1_numbers)|v1_xreal_0(X69)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
fof(c_0_289, plain, ![X529]:(~v1_xcmplx_0(X529)|k5_xcmplx_0(k5_xcmplx_0(X529))=X529), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k5_xcmplx_0])])).
fof(c_0_290, plain, ![X621]:(~v1_xcmplx_0(X621)|k7_xcmplx_0(np__1,X621)=k5_xcmplx_0(X621)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc3_arithm])])).
cnf(c_0_291, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_246])])).
cnf(c_0_292, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_283]), c_0_284])])])).
cnf(c_0_293, plain, (X1=X2|X1=X3|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_190])])).
cnf(c_0_294, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_286])).
cnf(c_0_295, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_287])).
cnf(c_0_296, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_288])).
cnf(c_0_297, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_289])).
cnf(c_0_298, plain, (k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_290])).
fof(c_0_299, plain, ![X231]:(~v1_xcmplx_0(X231)|v1_xcmplx_0(k5_xcmplx_0(X231))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_xcmplx_0])])).
cnf(c_0_300, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_244])])).
cnf(c_0_301, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_252])])).
cnf(c_0_302, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))), inference(csr,[status(thm)],[c_0_293, c_0_162])).
cnf(c_0_303, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_295])])).
cnf(c_0_304, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_305, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_306, plain, (k5_xcmplx_0(X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_298])).
cnf(c_0_307, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_299])).
fof(c_0_308, plain, ![X438]:((v1_xcmplx_0(k5_xcmplx_0(X438))|~v1_xreal_0(X438))&(v1_xreal_0(k5_xcmplx_0(X438))|~v1_xreal_0(X438))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_xreal_0])])])).
cnf(c_0_309, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_300, c_0_301])).
cnf(c_0_310, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_206])])).
cnf(c_0_311, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304])])).
cnf(c_0_312, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306])])).
cnf(c_0_313, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_314, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_308])).
cnf(c_0_315, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_316, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_263])])).
cnf(c_0_317, plain, (k7_xcmplx_0(X1,X2)=X3|k7_xcmplx_0(X1,X2)=X4|~r1_xxreal_0(k7_xcmplx_0(X1,X2),X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k7_xcmplx_0(X1,X2))|~r1_xxreal_0(k7_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311])])).
cnf(c_0_318, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_312, c_0_313])).
cnf(c_0_319, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_314])).
cnf(c_0_320, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[c_0_315])).
fof(c_0_321, plain, ![X402, X403]:~v1_xboole_0(k2_tarski(X402,X403)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
cnf(c_0_322, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_316, c_0_272])).
cnf(c_0_323, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_318]), c_0_284])])])).
cnf(c_0_324, plain, (m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_319])])).
cnf(c_0_325, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_193])])).
cnf(c_0_326, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_321])).
cnf(c_0_327, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_322, c_0_247])])).
cnf(c_0_328, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_263])])).
cnf(c_0_329, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324])])).
cnf(c_0_330, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_168])])).
cnf(c_0_331, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[c_0_326])).
cnf(c_0_332, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_327, c_0_254])).
cnf(c_0_333, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X3)|~r2_hidden(X2,k2_tarski(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_168])])).
cnf(c_0_334, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_328, c_0_272])).
cnf(c_0_335, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_329, c_0_276])).
cnf(c_0_336, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_126]), c_0_331])])).
cnf(c_0_337, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_263])])).
cnf(c_0_338, plain, (r1_xxreal_0(X1,esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_182])])).
cnf(c_0_339, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_334, c_0_246])])).
cnf(c_0_340, plain, (esk6_3(X1,X2,X3)=X3|esk6_3(X1,X2,X3)=X4|~r1_xxreal_0(X3,esk6_3(X1,X2,X3))|~r1_xxreal_0(esk6_3(X1,X2,X3),X4)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(esk6_3(X1,X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_336])])).
cnf(c_0_341, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_112])])).
cnf(c_0_342, plain, (m1_subset_1(X1,k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_182])])).
cnf(c_0_343, plain, (k1_real_1(esk6_3(X1,X1,X2))=k4_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(X1,X1,X2),k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_338])])).
cnf(c_0_344, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_244])])).
cnf(c_0_345, plain, (r1_xxreal_0(esk6_3(X1,X1,X2),X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_182])])).
cnf(c_0_346, plain, (esk6_3(X1,X1,X1)=X2|esk6_3(X1,X1,X1)=X1|~r1_xxreal_0(esk6_3(X1,X1,X1),X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_341]), c_0_342])])])).
cnf(c_0_347, plain, (k1_real_1(esk6_3(X1,X1,X2))=k4_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_343, c_0_214])).
cnf(c_0_348, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_344, c_0_292]), c_0_252])).
cnf(c_0_349, plain, (k4_xcmplx_0(esk6_3(X2,X2,X1))=k1_real_1(X1)|~r1_xxreal_0(X1,esk6_3(X2,X2,X1))|~v1_xreal_0(esk6_3(X2,X2,X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_345])])).
cnf(c_0_350, plain, (v1_xreal_0(esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_214]), c_0_265])])])).
cnf(c_0_351, plain, (esk6_3(X1,X1,X1)=X1|esk6_3(X1,X1,X1)=X2|~r1_xxreal_0(esk6_3(X1,X1,X1),X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_346, c_0_224])).
cnf(c_0_352, plain, (k1_real_1(esk6_3(X1,X1,X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_347, c_0_336]), c_0_342])])])).
cnf(c_0_353, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_263])])).
cnf(c_0_354, plain, (k4_xcmplx_0(esk6_3(X1,X1,X2))=k1_real_1(X2)|~r1_xxreal_0(X2,esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_349, c_0_292]), c_0_350])).
cnf(c_0_355, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k2_tarski(X1,X3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_126]), c_0_331])])).
cnf(c_0_356, plain, (esk6_3(X1,X1,X1)=X1|~r1_xxreal_0(X1,X1)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_336]), c_0_342])])])).
cnf(c_0_357, plain, (k1_real_1(esk6_3(X1,X1,X1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_352, c_0_224])).
cnf(c_0_358, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_353, c_0_272])).
cnf(c_0_359, plain, (k4_xcmplx_0(esk6_3(X1,X1,X1))=k1_real_1(X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_355]), c_0_342])])])).
cnf(c_0_360, plain, (esk6_3(X1,X1,X1)=X1|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_356, c_0_224])).
cnf(c_0_361, plain, (esk6_3(X1,X1,X1)=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_357])])).
cnf(c_0_362, plain, (k4_xcmplx_0(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_247])])).
cnf(c_0_363, plain, (k4_xcmplx_0(esk6_3(X1,X1,X1))=k1_real_1(X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_359, c_0_224])).
cnf(c_0_364, plain, (esk6_3(X1,X1,X1)=X1|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_360, c_0_350])])).
cnf(c_0_365, plain, (esk6_3(X1,X1,X1)=k1_real_1(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_214])])).
cnf(c_0_366, plain, (k4_xcmplx_0(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_362, c_0_254])).
cnf(c_0_367, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X2,X2,X2))|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_363])])).
cnf(c_0_368, plain, (esk6_3(X1,X1,X1)=X1|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_364, c_0_224])).
cnf(c_0_369, plain, (esk6_3(X1,X1,X1)=k1_real_1(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_365, c_0_224])).
cnf(c_0_370, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_244])])).
cnf(c_0_371, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_368])])).
cnf(c_0_372, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_369])])).
cnf(c_0_373, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_370, c_0_272]), c_0_252])).
cnf(c_0_374, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_372])])).
cnf(c_0_375, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_247])])).
cnf(c_0_376, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_374, c_0_252])).
cnf(c_0_377, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_375, c_0_254])).
cnf(c_0_378, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X3,X1)|~r1_xxreal_0(X3,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_122])])).
cnf(c_0_379, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_263])])).
cnf(c_0_380, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_263])])).
cnf(c_0_381, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X3,X1)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_378, c_0_132]), c_0_132])).
cnf(c_0_382, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_379, c_0_272])).
cnf(c_0_383, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_380, c_0_272])).
cnf(c_0_384, plain, (r2_hidden(X1,k1_rcomp_1(X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_381])])).
cnf(c_0_385, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_263])])).
cnf(c_0_386, plain, (esk6_3(X2,k1_real_1(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,k1_real_1(esk6_3(X2,k1_real_1(X1),X2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(X2,k1_real_1(X1),X2),k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_341])])).
cnf(c_0_387, plain, (m1_subset_1(esk6_3(X1,X2,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_384])])).
cnf(c_0_388, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_385, c_0_272])).
cnf(c_0_389, plain, (esk6_3(X1,k1_real_1(X2),X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k1_real_1(esk6_3(X1,k1_real_1(X2),X1)))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_386, c_0_387]), c_0_272])).
cnf(c_0_390, plain, (esk6_3(X1,X2,X3)=X1|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_126]), c_0_331])])).
cnf(c_0_391, plain, (m1_subset_1(X1,k2_tarski(X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_112])])).
cnf(c_0_392, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_247])])).
cnf(c_0_393, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_390]), c_0_391])])])).
cnf(c_0_394, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_392, c_0_254])).
cnf(c_0_395, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_393, c_0_272])).
cnf(c_0_396, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_357])])).
cnf(c_0_397, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_263])])).
cnf(c_0_398, plain, (X1=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_247])])).
cnf(c_0_399, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_220])])).
cnf(c_0_400, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_214])])).
cnf(c_0_401, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_397, c_0_272])).
cnf(c_0_402, plain, (X1=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_398, c_0_254])).
cnf(c_0_403, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_399, c_0_224])).
cnf(c_0_404, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_400, c_0_224])).
cnf(c_0_405, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_246])])).
cnf(c_0_406, plain, (X1=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_403])])).
cnf(c_0_407, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_368])])).
cnf(c_0_408, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_405, c_0_247])])).
cnf(c_0_409, plain, (X1=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407])])).
cnf(c_0_410, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_247])])).
cnf(c_0_411, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_408, c_0_272])).
fof(c_0_412, plain, ![X23, X24]:(~v2_membered(X23)|(~m1_subset_1(X24,X23)|v1_xxreal_0(X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
fof(c_0_413, plain, ![X120]:(~v3_membered(X120)|v2_membered(X120)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_414, plain, (esk6_3(X1,X1,X2)=k4_xcmplx_0(k1_real_1(X1))|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~v1_xreal_0(esk6_3(X1,X1,X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_338])])).
cnf(c_0_415, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_410, c_0_254])).
cnf(c_0_416, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_411, c_0_263])])).
cnf(c_0_417, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_412])).
cnf(c_0_418, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_413])).
cnf(c_0_419, plain, (esk6_3(X1,X1,X2)=k4_xcmplx_0(k1_real_1(X1))|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_414, c_0_292]), c_0_350])).
cnf(c_0_420, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_263])])).
cnf(c_0_421, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_416, c_0_254])).
cnf(c_0_422, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_417])).
cnf(c_0_423, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_418])).
cnf(c_0_424, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_220])])).
cnf(c_0_425, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_363])])).
cnf(c_0_426, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_420, c_0_272])).
cnf(c_0_427, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_246])])).
cnf(c_0_428, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_246])])).
cnf(c_0_429, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_254])])).
cnf(c_0_430, plain, (v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_265])])).
cnf(c_0_431, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_424, c_0_224])).
cnf(c_0_432, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_425, c_0_368])])).
cnf(c_0_433, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_427])])).
cnf(c_0_434, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_428])])).
cnf(c_0_435, plain, (v1_xxreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_429, c_0_430])])).
cnf(c_0_436, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_431])])).
cnf(c_0_437, plain, (v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_432]), c_0_265])])])).
cnf(c_0_438, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_434])])).
cnf(c_0_439, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_246])])).
cnf(c_0_440, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_247])])).
cnf(c_0_441, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_436, c_0_437])).
cnf(c_0_442, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)=k1_real_1(X1)|~r1_xxreal_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_338])])).
cnf(c_0_443, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_439]), c_0_265])])])).
cnf(c_0_444, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_440, c_0_254])).
fof(c_0_445, negated_conjecture, ~(![X1]:(r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))=>r2_hidden(k1_seq_1(k2_sincos10,X1),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))), inference(assume_negation,[status(cth)],[t86_sincos10])).
cnf(c_0_446, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_403])])).
cnf(c_0_447, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)=k1_real_1(X1)|~r1_xxreal_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_442, c_0_214]), c_0_443])).
cnf(c_0_448, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_444])])).
fof(c_0_449, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))&~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_445])])])).
cnf(c_0_450, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k4_xcmplx_0(np__1)), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__rnm1d2_r2_rm1])).
cnf(c_0_451, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)), inference(split_conjunct,[status(thm)],[rqRealDiv__k7_xcmplx_0__r1_rm2_rnm1d2])).
cnf(c_0_452, plain, (m1_subset_1(np__2,k1_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_453, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_446, c_0_247])])).
cnf(c_0_454, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),k4_xcmplx_0(X1))=k1_real_1(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_447, c_0_336]), c_0_342])])])).
cnf(c_0_455, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_448, c_0_439])])).
fof(c_0_456, plain, ![X645, X646, X647, X648]:((~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X646,X645))|(~m1_subset_1(X647,k1_numbers)|(~m1_subset_1(X648,k1_numbers)|(~r2_hidden(X647,k9_subset_1(k1_numbers,X645,k1_relset_1(k1_numbers,X646)))|~r2_hidden(X648,k9_subset_1(k1_numbers,X645,k1_relset_1(k1_numbers,X646)))|r1_xxreal_0(X648,X647)|~r1_xxreal_0(k1_seq_1(X646,X648),k1_seq_1(X646,X647)))))|(~v1_funct_1(X646)|~m1_subset_1(X646,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((m1_subset_1(esk50_2(X645,X646),k1_numbers)|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X646,X645))|(~v1_funct_1(X646)|~m1_subset_1(X646,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((m1_subset_1(esk51_2(X645,X646),k1_numbers)|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X646,X645))|(~v1_funct_1(X646)|~m1_subset_1(X646,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((((r2_hidden(esk50_2(X645,X646),k9_subset_1(k1_numbers,X645,k1_relset_1(k1_numbers,X646)))|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X646,X645))|(~v1_funct_1(X646)|~m1_subset_1(X646,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(r2_hidden(esk51_2(X645,X646),k9_subset_1(k1_numbers,X645,k1_relset_1(k1_numbers,X646)))|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X646,X645))|(~v1_funct_1(X646)|~m1_subset_1(X646,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&(~r1_xxreal_0(esk51_2(X645,X646),esk50_2(X645,X646))|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X646,X645))|(~v1_funct_1(X646)|~m1_subset_1(X646,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&(r1_xxreal_0(k1_seq_1(X646,esk51_2(X645,X646)),k1_seq_1(X646,esk50_2(X645,X646)))|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X646,X645))|(~v1_funct_1(X646)|~m1_subset_1(X646,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_rfunct_2])])])])])])).
fof(c_0_457, plain, ![X173, X174, X175]:(~m1_subset_1(X175,k1_zfmisc_1(X173))|k9_subset_1(X173,X174,X175)=k9_subset_1(X173,X175,X174)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k9_subset_1])])).
fof(c_0_458, plain, ![X655, X656]:(~r1_tarski(X655,X656)|k3_xboole_0(X655,X656)=X655), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_459, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)),k1_relset_1(k1_numbers,k2_sincos10))), inference(split_conjunct,[status(thm)],[t46_sincos10])).
fof(c_0_460, plain, ![X586, X587, X588, X589]:(~v1_funct_1(X588)|~m1_subset_1(X588,k1_zfmisc_1(k2_zfmisc_1(X586,X587)))|k2_partfun1(X586,X587,X588,X589)=k5_relat_1(X588,X589)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_partfun1])])).
cnf(c_0_461, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))), inference(split_conjunct,[status(thm)],[c_0_449])).
cnf(c_0_462, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[c_0_450])).
cnf(c_0_463, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))), inference(evalgc,[status(thm)],[c_0_451])).
cnf(c_0_464, plain, (m1_subset_1(np__2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_452])).
fof(c_0_465, plain, ![X676, X677, X678]:(~r2_hidden(X676,X677)|~m1_subset_1(X677,k1_zfmisc_1(X678))|m1_subset_1(X676,X678)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_466, plain, ![X208, X209]:(~v1_xreal_0(X208)|~v1_xreal_0(X209)|m1_subset_1(k1_rcomp_1(X208,X209),k1_zfmisc_1(k1_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_rcomp_1])])).
cnf(c_0_467, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_453, c_0_272])).
cnf(c_0_468, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),k4_xcmplx_0(X1))=k1_real_1(X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_454, c_0_443]), c_0_455])).
cnf(c_0_469, plain, (r1_xxreal_0(X4,X3)|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X4,k1_numbers)|~r2_hidden(X3,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r2_hidden(X4,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r1_xxreal_0(k1_seq_1(X1,X4),k1_seq_1(X1,X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_456])).
cnf(c_0_470, plain, (k9_subset_1(X2,X3,X1)=k9_subset_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_457])).
fof(c_0_471, plain, ![X603, X604, X605]:(~m1_subset_1(X605,k1_zfmisc_1(X603))|k9_subset_1(X603,X604,X605)=k3_xboole_0(X604,X605)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
cnf(c_0_472, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_458])).
cnf(c_0_473, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)),k1_relset_1(k1_numbers,k2_sincos10))), inference(evalgc,[status(thm)],[c_0_459])).
fof(c_0_474, plain, ![X167, X168]:k3_xboole_0(X167,X168)=k3_xboole_0(X168,X167), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_475, plain, (k2_partfun1(X2,X3,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_460])).
cnf(c_0_476, plain, (m1_subset_1(k2_sincos10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k2_sincos10])).
cnf(c_0_477, plain, (v1_funct_1(k2_sincos10)), inference(split_conjunct,[status(thm)],[dt_k2_sincos10])).
fof(c_0_478, plain, ![X493, X494]:(~v1_xreal_0(X493)|~v1_xreal_0(X494)|v1_xreal_0(k7_xcmplx_0(X493,X494))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_xreal_0])])).
fof(c_0_479, plain, ![X666, X667, X668]:(((~r1_xxreal_0(X666,X667)|~r2_hidden(X666,k3_xxreal_1(X667,X668))|~v1_xxreal_0(X668)|~v1_xxreal_0(X667)|~v1_xxreal_0(X666))&(r1_xxreal_0(X666,X668)|~r2_hidden(X666,k3_xxreal_1(X667,X668))|~v1_xxreal_0(X668)|~v1_xxreal_0(X667)|~v1_xxreal_0(X666)))&(r1_xxreal_0(X666,X667)|~r1_xxreal_0(X666,X668)|r2_hidden(X666,k3_xxreal_1(X667,X668))|~v1_xxreal_0(X668)|~v1_xxreal_0(X667)|~v1_xxreal_0(X666))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t2_xxreal_1])])])])])).
fof(c_0_480, plain, ![X595, X596]:(~v1_xxreal_0(X595)|~v1_xreal_0(X596)|k4_rcomp_1(X595,X596)=k3_xxreal_1(X595,X596)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_rcomp_1])])).
cnf(c_0_481, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))), inference(evalgc,[status(thm)],[c_0_461])).
cnf(c_0_482, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1)), inference(split_conjunct,[status(thm)],[rqRealDiv__k7_xcmplx_0__r1_rm1_rm1])).
cnf(c_0_483, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),np__2)=k4_xcmplx_0(np__1)), inference(rw,[status(thm)],[c_0_462, c_0_463])).
cnf(c_0_484, plain, (v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_464]), c_0_265])])])).
cnf(c_0_485, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_465])).
cnf(c_0_486, plain, (m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_466])).
cnf(c_0_487, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_467, c_0_431])])).
cnf(c_0_488, plain, (esk6_3(X1,X1,X1)=k1_real_1(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_263])])).
cnf(c_0_489, plain, (r1_xxreal_0(X4,X3)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k1_seq_1(X1,X4),k1_seq_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(X4,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r2_hidden(X3,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))), inference(evalgc,[status(thm)],[c_0_469])).
cnf(c_0_490, plain, (k9_subset_1(X2,X1,X3)=k9_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_470])).
cnf(c_0_491, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_471])).
cnf(c_0_492, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_472])).
cnf(c_0_493, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_relset_1(k1_numbers,k2_sincos10))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_246]), c_0_284])])])).
cnf(c_0_494, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_474])).
cnf(c_0_495, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))), inference(split_conjunct,[status(thm)],[t82_sincos10])).
cnf(c_0_496, plain, (k2_partfun1(X2,X3,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_475])).
cnf(c_0_497, plain, (m1_subset_1(k2_sincos10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_476])).
cnf(c_0_498, plain, (v1_funct_1(k2_sincos10)), inference(evalgc,[status(thm)],[c_0_477])).
cnf(c_0_499, plain, (r1_xxreal_0(X1,X3)|r1_xxreal_0(X4,X1)|r2_hidden(X2,a_2_1_rcomp_1(X3,X4))|~m1_subset_1(X1,k1_numbers)|X2!=X1|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)), inference(split_conjunct,[status(thm)],[c_0_147])).
fof(c_0_500, plain, ![X225, X226]:(~v1_xxreal_0(X225)|~v1_xreal_0(X226)|m1_subset_1(k4_rcomp_1(X225,X226),k1_zfmisc_1(k1_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_rcomp_1])])).
cnf(c_0_501, plain, (r2_hidden(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))), inference(split_conjunct,[status(thm)],[l18_sincos10])).
cnf(c_0_502, plain, (m1_subset_1(k32_sin_cos,k1_numbers)), inference(split_conjunct,[status(thm)],[dt_k32_sin_cos])).
cnf(c_0_503, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_478])).
cnf(c_0_504, plain, (r1_xxreal_0(X1,X2)|~r2_hidden(X1,k3_xxreal_1(X3,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_479])).
cnf(c_0_505, plain, (k4_rcomp_1(X1,X2)=k3_xxreal_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_480])).
cnf(c_0_506, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_246])])).
cnf(c_0_507, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_311]), c_0_265])])])).
cnf(c_0_508, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[c_0_482])).
cnf(c_0_509, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_483]), c_0_484])])])).
cnf(c_0_510, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_443])])).
cnf(c_0_511, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_485])).
cnf(c_0_512, plain, (m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_486])).
cnf(c_0_513, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_487, c_0_437])).
cnf(c_0_514, plain, (esk6_3(X1,X1,X1)=k1_real_1(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_488, c_0_254])).
cnf(c_0_515, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_490])])).
cnf(c_0_516, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_491])).
cnf(c_0_517, plain, (k3_xboole_0(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_relset_1(k1_numbers,k2_sincos10))=k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_493])])).
cnf(c_0_518, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_494])).
cnf(c_0_519, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))), inference(evalgc,[status(thm)],[c_0_495])).
cnf(c_0_520, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_sincos10,X1)=k5_relat_1(k2_sincos10,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_497]), c_0_498])])])).
cnf(c_0_521, plain, (r1_xxreal_0(X4,X1)|r1_xxreal_0(X1,X3)|r2_hidden(X2,a_2_1_rcomp_1(X3,X4))|X1!=X2|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_499])).
cnf(c_0_522, plain, (m1_subset_1(k4_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_500])).
cnf(c_0_523, plain, (r2_hidden(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))), inference(evalgc,[status(thm)],[c_0_501])).
cnf(c_0_524, plain, (m1_subset_1(k32_sin_cos,k1_numbers)), inference(evalgc,[status(thm)],[c_0_502])).
cnf(c_0_525, plain, (v1_xreal_0(k31_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k31_sin_cos])).
cnf(c_0_526, plain, (k32_sin_cos=k31_sin_cos), inference(split_conjunct,[status(thm)],[redefinition_k32_sin_cos])).
cnf(c_0_527, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_503])).
cnf(c_0_528, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X1,k3_xxreal_1(X3,X2))), inference(evalgc,[status(thm)],[c_0_504])).
cnf(c_0_529, plain, (k3_xxreal_1(X1,X2)=k4_rcomp_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[c_0_505])).
cnf(c_0_530, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_506, c_0_284])])).
cnf(c_0_531, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_508]), c_0_284])])])).
cnf(c_0_532, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_509, c_0_510]), c_0_284]), c_0_464])])])).
cnf(c_0_533, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_512])])).
cnf(c_0_534, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k1_real_1(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_481, c_0_246])])).
fof(c_0_535, plain, ![X236]:(~m1_subset_1(X236,k1_numbers)|m1_subset_1(k7_square_1(X236),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_square_1])])).
cnf(c_0_536, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_514])])).
cnf(c_0_537, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k3_xboole_0(k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k3_xboole_0(k1_relset_1(k1_numbers,X3),X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_515, c_0_516])])).
cnf(c_0_538, plain, (k3_xboole_0(k1_relset_1(k1_numbers,k2_sincos10),k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))=k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))), inference(rw,[status(thm)],[c_0_517, c_0_518])).
cnf(c_0_539, plain, (v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_519, c_0_520])])).
cnf(c_0_540, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_521])])).
cnf(c_0_541, plain, (m1_subset_1(k4_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[c_0_522])).
cnf(c_0_542, plain, (r2_hidden(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos))|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_523, c_0_295]), c_0_524])])])).
cnf(c_0_543, plain, (v1_xreal_0(k31_sin_cos)), inference(evalgc,[status(thm)],[c_0_525])).
cnf(c_0_544, plain, (k31_sin_cos=k32_sin_cos), inference(evalgc,[status(thm)],[c_0_526])).
cnf(c_0_545, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_527])])).
cnf(c_0_546, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X1,k4_rcomp_1(X3,X2))|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_528, c_0_529])])).
cnf(c_0_547, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_403])])).
cnf(c_0_548, plain, (v1_xreal_0(k4_xcmplx_0(np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_531, c_0_532])])).
cnf(c_0_549, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_533, c_0_534])])).
cnf(c_0_550, plain, (m1_subset_1(k7_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_535])).
cnf(c_0_551, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_536, c_0_254])).
cnf(c_0_552, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))|~m1_subset_1(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_538]), c_0_520]), c_0_498]), c_0_497])])])).
cnf(c_0_553, plain, (v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_539, c_0_246]), c_0_284])])])).
cnf(c_0_554, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|r2_hidden(X2,k2_rcomp_1(X3,X1))|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_540, c_0_154])])).
cnf(c_0_555, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_541])])).
cnf(c_0_556, plain, (r2_hidden(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_542, c_0_484])])).
cnf(c_0_557, plain, (v1_xreal_0(k32_sin_cos)), inference(rw,[status(thm)],[c_0_543, c_0_544])).
cnf(c_0_558, plain, (m1_subset_1(k7_xcmplx_0(X1,np__2),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_545, c_0_484])])).
cnf(c_0_559, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X1,k4_rcomp_1(X3,X2))), inference(csr,[status(thm)],[c_0_546, c_0_132])).
cnf(c_0_560, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_547]), c_0_548])])])).
cnf(c_0_561, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_547]), c_0_548])])])).
cnf(c_0_562, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)), inference(csr,[status(thm)],[c_0_549, c_0_443])).
cnf(c_0_563, plain, (m1_subset_1(k7_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_550])).
cnf(c_0_564, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_551, c_0_368])])).
cnf(c_0_565, plain, (esk6_3(X1,X1,X1)=k1_real_1(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_431])])).
cnf(c_0_566, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_552, c_0_553])]), c_0_511]), c_0_511])).
cnf(c_0_567, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))), inference(split_conjunct,[status(thm)],[c_0_449])).
cnf(c_0_568, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_554])])).
cnf(c_0_569, plain, (m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_556]), c_0_557])])])).
cnf(c_0_570, plain, (m1_subset_1(k7_xcmplx_0(k32_sin_cos,np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_558, c_0_557])])).
cnf(c_0_571, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_556]), c_0_557])])])).
cnf(c_0_572, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2))), inference(csr,[status(thm)],[c_0_560, c_0_252])).
cnf(c_0_573, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|~v1_xreal_0(k7_square_1(np__2))), inference(csr,[status(thm)],[c_0_561, c_0_252])).
cnf(c_0_574, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_562, c_0_563]), c_0_464])])])).
cnf(c_0_575, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_564, c_0_292])).
cnf(c_0_576, plain, (esk6_3(X1,X1,X1)=k1_real_1(k1_real_1(X1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_565, c_0_432])).
cnf(c_0_577, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_566, c_0_512]), c_0_548])])])).
cnf(c_0_578, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))), inference(evalgc,[status(thm)],[c_0_567])).
cnf(c_0_579, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|r2_hidden(X2,k1_rcomp_1(X3,X1))|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_568, c_0_132]), c_0_132])).
cnf(c_0_580, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_569]), c_0_265])])])).
cnf(c_0_581, plain, (v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_570]), c_0_430])])])).
cnf(c_0_582, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_571, c_0_444])])).
cnf(c_0_583, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_405])])).
cnf(c_0_584, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk1_0)|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_572, c_0_573])])).
cnf(c_0_585, plain, (v1_xreal_0(k7_square_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_563]), c_0_265])])])).
cnf(c_0_586, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_272]), c_0_284])])])).
cnf(c_0_587, plain, (r1_xxreal_0(k1_real_1(X1),esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_575, c_0_576])])).
cnf(c_0_588, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_577, c_0_530])])).
cnf(c_0_589, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_530])])).
cnf(c_0_590, plain, (k1_seq_1(k2_sincos10,k1_real_1(k7_square_1(np__2)))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)), inference(split_conjunct,[status(thm)],[t74_sincos10])).
cnf(c_0_591, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_578, c_0_579]), c_0_557])])])).
cnf(c_0_592, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_580, c_0_581])])])).
cnf(c_0_593, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))), inference(csr,[status(thm)],[c_0_582, c_0_569])).
cnf(c_0_594, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_578, c_0_168]), c_0_557])])])).
cnf(c_0_595, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_246])])).
cnf(c_0_596, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_583, c_0_292]), c_0_292])).
cnf(c_0_597, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_584, c_0_585]), c_0_464])])])).
cnf(c_0_598, negated_conjecture, (v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_586]), c_0_265])])])).
cnf(c_0_599, plain, (r1_xxreal_0(k1_real_1(X1),esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_587, c_0_437])).
cnf(c_0_600, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_577, c_0_199]), c_0_548])])])).
cnf(c_0_601, plain, (k1_seq_1(k2_sincos10,k1_real_1(np__1))=k32_sin_cos), inference(split_conjunct,[status(thm)],[t74_sincos10])).
cnf(c_0_602, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_126]), c_0_589])])).
cnf(c_0_603, plain, (m1_subset_1(X1,k1_rcomp_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_182])])).
cnf(c_0_604, plain, (k1_seq_1(k2_sincos10,k1_real_1(k7_square_1(np__2)))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)), inference(evalgc,[status(thm)],[c_0_590])).
cnf(c_0_605, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_591, c_0_592])])])).
cnf(c_0_606, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_593, c_0_581])])])).
cnf(c_0_607, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_405, c_0_372])])).
cnf(c_0_608, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))), inference(rw,[status(thm)],[c_0_594, c_0_108])).
cnf(c_0_609, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_595, c_0_596])])).
cnf(c_0_610, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_597]), c_0_598])])])).
fof(c_0_611, plain, ![X176, X177]:(~v1_xxreal_0(X176)|~v1_xxreal_0(X177)|(r1_xxreal_0(X176,X177)|r1_xxreal_0(X177,X176))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
cnf(c_0_612, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_247])])).
cnf(c_0_613, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_599, c_0_468])])).
fof(c_0_614, plain, ![X18, X19]:(~v1_membered(X18)|(~m1_subset_1(X19,X18)|v1_xcmplx_0(X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
cnf(c_0_615, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_600, c_0_530])])).
cnf(c_0_616, plain, (k1_seq_1(k2_sincos10,k1_real_1(np__1))=k32_sin_cos), inference(evalgc,[status(thm)],[c_0_601])).
cnf(c_0_617, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__rm1_rm1])).
cnf(c_0_618, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_602, c_0_603]), c_0_604]), c_0_548])])])).
cnf(c_0_619, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_605, c_0_606])])).
cnf(c_0_620, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_607, c_0_252])).
cnf(c_0_621, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_608, c_0_126]), c_0_331])])).
cnf(c_0_622, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_609, c_0_292]), c_0_292])).
cnf(c_0_623, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_610, c_0_585]), c_0_464])])])).
cnf(c_0_624, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_611])).
cnf(c_0_625, plain, (m1_subset_1(np__3,k1_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_626, plain, (m1_subset_1(np__4,k1_numbers)), inference(split_conjunct,[status(thm)],[spc4_numerals])).
cnf(c_0_627, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_523]), c_0_557])])])).
cnf(c_0_628, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_612, c_0_254])).
cnf(c_0_629, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_613, c_0_443])).
fof(c_0_630, plain, ![X528]:(~v1_xcmplx_0(X528)|k4_xcmplx_0(k4_xcmplx_0(X528))=X528), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
fof(c_0_631, plain, ![X68]:(~m1_subset_1(X68,k1_numbers)|v1_xcmplx_0(X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xcmplx_0])])).
cnf(c_0_632, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_614])).
fof(c_0_633, plain, ![X131]:(~v3_membered(X131)|v1_membered(X131)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_634, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_615, c_0_304])])).
cnf(c_0_635, plain, (k1_seq_1(k2_sincos10,k4_xcmplx_0(np__1))=k32_sin_cos), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_616, c_0_246]), c_0_284])])])).
cnf(c_0_636, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[c_0_617])).
cnf(c_0_637, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_618, c_0_619])])).
cnf(c_0_638, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_620, c_0_263])])).
cnf(c_0_639, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
cnf(c_0_640, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_621, c_0_592])])])).
cnf(c_0_641, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_622, c_0_623]), c_0_598])])])).
cnf(c_0_642, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_624])).
cnf(c_0_643, plain, (m1_subset_1(np__3,k1_numbers)), inference(evalgc,[status(thm)],[c_0_625])).
cnf(c_0_644, plain, (m1_subset_1(np__4,k1_numbers)), inference(evalgc,[status(thm)],[c_0_626])).
cnf(c_0_645, plain, (m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_555, c_0_523]), c_0_557])])])).
cnf(c_0_646, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_627, c_0_444])])).
fof(c_0_647, plain, ![X449, X450]:(~v1_xreal_0(X449)|~v1_xreal_0(X450)|v1_xreal_0(k2_xcmplx_0(X449,X450))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc5_xreal_0])])).
cnf(c_0_648, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X1)),X2)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_628, c_0_629])])).
cnf(c_0_649, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_630])).
cnf(c_0_650, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_631])).
fof(c_0_651, plain, ![X620]:(~v1_xcmplx_0(X620)|k3_xcmplx_0(X620,k4_xcmplx_0(np__1))=k4_xcmplx_0(X620)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc2_arithm])])).
cnf(c_0_652, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_632])).
cnf(c_0_653, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_633])).
cnf(c_0_654, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_634, c_0_635]), c_0_636]), c_0_532])])])).
cnf(c_0_655, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_637, c_0_437])])).
cnf(c_0_656, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_638, c_0_254])).
cnf(c_0_657, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_639])).
cnf(c_0_658, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_640, c_0_606])])).
cnf(c_0_659, negated_conjecture, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k1_seq_1(k2_sincos10,esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_604, c_0_641])])).
cnf(c_0_660, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_642, c_0_444])])).
fof(c_0_661, plain, ~r1_xxreal_0(k7_square_1(np__2),np__1), inference(fof_simplification,[status(thm)],[t19_square_1])).
cnf(c_0_662, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))|~v1_xreal_0(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_578, c_0_295]), c_0_643])])])).
cnf(c_0_663, plain, (v1_xreal_0(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_644]), c_0_265])])])).
cnf(c_0_664, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_645]), c_0_265])])])).
cnf(c_0_665, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))), inference(csr,[status(thm)],[c_0_646, c_0_645])).
cnf(c_0_666, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_647])).
cnf(c_0_667, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_648, c_0_292])).
cnf(c_0_668, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_649])).
cnf(c_0_669, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_650])).
cnf(c_0_670, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_651])).
cnf(c_0_671, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_652, c_0_586])])).
cnf(c_0_672, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_653])).
cnf(c_0_673, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k4_xcmplx_0(k7_square_1(np__2)),k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_654, c_0_246])])).
cnf(c_0_674, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~m1_subset_1(k7_square_1(np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_655, c_0_656]), c_0_657]), c_0_548])])])).
cnf(c_0_675, negated_conjecture, (~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_658, c_0_659]), c_0_391])])])).
cnf(c_0_676, plain, (r1_xxreal_0(k7_square_1(X1),X2)|r1_xxreal_0(X2,k7_square_1(X1))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_660, c_0_563])])).
cnf(c_0_677, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_284])])).
cnf(c_0_678, plain, (~r1_xxreal_0(k7_square_1(np__2),np__1)), inference(split_conjunct,[status(thm)],[c_0_661])).
cnf(c_0_679, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_662, c_0_663])])).
cnf(c_0_680, plain, (v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_664, c_0_295]), c_0_663]), c_0_643])])])).
cnf(c_0_681, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_665, c_0_295]), c_0_663]), c_0_643])])])).
cnf(c_0_682, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_666])).
cnf(c_0_683, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),np__1)=k7_xcmplx_0(np__3,np__4)), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__rnm1d4_r1_rn3d4])).
fof(c_0_684, plain, ![X622, X623, X624]:(~v1_xcmplx_0(X622)|~v1_xcmplx_0(X623)|~v1_xcmplx_0(X624)|k3_xcmplx_0(X622,k7_xcmplx_0(X623,X624))=k7_xcmplx_0(k3_xcmplx_0(X622,X623),X624)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc4_arithm])])).
cnf(c_0_685, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_667, c_0_403])])).
cnf(c_0_686, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_669])])).
cnf(c_0_687, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_670])).
cnf(c_0_688, negated_conjecture, (v1_xcmplx_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_671, c_0_672]), c_0_265])])])).
cnf(c_0_689, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers)), inference(csr,[status(thm)],[c_0_673, c_0_439])).
cnf(c_0_690, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_284]), c_0_265])])])).
cnf(c_0_691, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(csr,[status(thm)],[c_0_674, c_0_292])).
cnf(c_0_692, negated_conjecture, (~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_675, c_0_585]), c_0_464])])])).
cnf(c_0_693, plain, (r1_xxreal_0(X1,k7_square_1(np__2))|r1_xxreal_0(k7_square_1(np__2),X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_676, c_0_464])])).
cnf(c_0_694, plain, (v1_xxreal_0(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_677, c_0_430])])).
cnf(c_0_695, plain, (~r1_xxreal_0(k7_square_1(np__2),np__1)), inference(evalgc,[status(thm)],[c_0_678])).
fof(c_0_696, plain, ![X51, X52, X53]:(~v3_membered(X52)|(~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52)))|v3_valued_0(X53))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
fof(c_0_697, plain, ![X63, X64, X65]:(~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64)))|v1_relat_1(X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_698, negated_conjecture, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_679, c_0_168]), c_0_557])])])).
cnf(c_0_699, plain, (v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_680, c_0_295]), c_0_484]), c_0_524])])])).
cnf(c_0_700, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_681, c_0_132])])).
fof(c_0_701, plain, ![X171, X172]:(~m1_subset_1(X171,k1_numbers)|~v1_xreal_0(X172)|k8_real_1(X171,X172)=k8_real_1(X172,X171)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k8_real_1])])).
cnf(c_0_702, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_682])])).
cnf(c_0_703, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),np__1)=k7_xcmplx_0(np__3,np__4)), inference(evalgc,[status(thm)],[c_0_683])).
cnf(c_0_704, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_684])).
cnf(c_0_705, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_685, c_0_443])).
cnf(c_0_706, negated_conjecture, (r1_xxreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_547]), c_0_548])])])).
cnf(c_0_707, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_246])])).
cnf(c_0_708, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(esk1_0))=esk1_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_586])])).
cnf(c_0_709, negated_conjecture, (k3_xcmplx_0(esk1_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_687, c_0_688])])).
cnf(c_0_710, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_689, c_0_244]), c_0_690])])])).
cnf(c_0_711, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_691, c_0_403]), c_0_692])])).
cnf(c_0_712, plain, (r1_xxreal_0(np__1,k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_693, c_0_694]), c_0_695])])).
fof(c_0_713, plain, ![X213, X214]:(~v1_relat_1(X213)|~v1_funct_1(X213)|~v3_valued_0(X213)|m1_subset_1(k1_seq_1(X213,X214),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
cnf(c_0_714, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_696])).
cnf(c_0_715, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_697])).
cnf(c_0_716, negated_conjecture, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))), inference(rw,[status(thm)],[c_0_698, c_0_108])).
cnf(c_0_717, plain, (v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_699, c_0_581])])).
cnf(c_0_718, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k7_xcmplx_0(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_700, c_0_295]), c_0_484]), c_0_524])])])).
cnf(c_0_719, plain, (v1_xreal_0(k7_xcmplx_0(k32_sin_cos,np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_570]), c_0_265])])])).
cnf(c_0_720, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_701])).
cnf(c_0_721, plain, (m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_702, c_0_703]), c_0_690])])])).
cnf(c_0_722, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_704])).
cnf(c_0_723, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_705, c_0_371])])).
cnf(c_0_724, negated_conjecture, (r1_xxreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2))), inference(csr,[status(thm)],[c_0_706, c_0_252])).
cnf(c_0_725, negated_conjecture, (k1_real_1(esk1_0)=k4_xcmplx_0(esk1_0)|~m1_subset_1(k4_xcmplx_0(esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_708])])).
cnf(c_0_726, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_709]), c_0_548]), c_0_586])])])).
cnf(c_0_727, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))), inference(csr,[status(thm)],[c_0_710, c_0_292])).
cnf(c_0_728, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_711, c_0_712])])])).
cnf(c_0_729, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_713])).
cnf(c_0_730, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_714])).
cnf(c_0_731, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_715])).
cnf(c_0_732, negated_conjecture, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_716, c_0_717])])])).
cnf(c_0_733, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_718, c_0_719])])).
cnf(c_0_734, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_720])).
cnf(c_0_735, plain, (m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_721, c_0_527]), c_0_663]), c_0_548])])])).
cnf(c_0_736, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_722, c_0_276])])).
cnf(c_0_737, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_723, c_0_443]), c_0_437])).
cnf(c_0_738, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_246])])).
cnf(c_0_739, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_724, c_0_573])])).
cnf(c_0_740, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_657])])).
cnf(c_0_741, negated_conjecture, (k1_real_1(esk1_0)=k4_xcmplx_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_725, c_0_726])])).
cnf(c_0_742, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k7_square_1(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_727, c_0_712])])])).
cnf(c_0_743, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_728, c_0_585]), c_0_464])])])).
cnf(c_0_744, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[c_0_729])).
cnf(c_0_745, plain, (v3_valued_0(k2_sincos10)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_730, c_0_497]), c_0_265])])])).
cnf(c_0_746, plain, (v1_relat_1(k2_sincos10)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_731, c_0_497])])).
cnf(c_0_747, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_732, c_0_733])])).
cnf(c_0_748, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_734])])).
cnf(c_0_749, plain, (v1_xreal_0(k7_xcmplx_0(np__3,np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_735]), c_0_265])])])).
cnf(c_0_750, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_736, c_0_276])])).
cnf(c_0_751, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_737, c_0_738])])).
cnf(c_0_752, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_739, c_0_585]), c_0_464])])])).
cnf(c_0_753, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_740, c_0_532])])])).
cnf(c_0_754, negated_conjecture, (k1_real_1(k4_xcmplx_0(esk1_0))=esk1_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_741]), c_0_586])])])).
cnf(c_0_755, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_742, c_0_585]), c_0_464])])])).
cnf(c_0_756, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_743, c_0_744]), c_0_498]), c_0_745]), c_0_746])])])).
cnf(c_0_757, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_747, c_0_748]), c_0_749]), c_0_524])])])).
cnf(c_0_758, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_750, c_0_669])])).
cnf(c_0_759, plain, (v1_xreal_0(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_643]), c_0_265])])])).
cnf(c_0_760, negated_conjecture, (k4_xcmplx_0(np__1)=esk1_0|~r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_751, c_0_752]), c_0_753]), c_0_741]), c_0_754]), c_0_598]), c_0_284])])])).
cnf(c_0_761, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_755, c_0_756])])])).
cnf(c_0_762, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_757, c_0_758]), c_0_663]), c_0_759]), c_0_524])])])).
cnf(c_0_763, negated_conjecture, (k4_xcmplx_0(np__1)=esk1_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_760, c_0_761])])])).
cnf(c_0_764, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_762, c_0_126]), c_0_331])])).
cnf(c_0_765, plain, (k1_seq_1(k2_sincos10,esk1_0)=k32_sin_cos), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_635, c_0_763])])).
cnf(c_0_766, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_764, c_0_765]), c_0_342])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 767
# Proof object clause steps            : 619
# Proof object formula steps           : 148
# Proof object conjectures             : 82
# Proof object clause conjectures      : 79
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 97
# Proof object initial formulas used   : 83
# Proof object generating inferences   : 297
# Proof object simplifying inferences  : 362
# Training examples: 410 positive, 4934 negative
# Training: Positive examples begin
cnf(c_0_544, plain, (k31_sin_cos=k32_sin_cos)).# trainpos
cnf(c_0_498, plain, (v1_funct_1(k2_sincos10))).# trainpos
cnf(c_0_481, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).# trainpos
cnf(c_0_265, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_246, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_274, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_295, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_217, plain, (v1_xxreal_0(esk35_0))).# trainpos
cnf(c_0_126, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_543, plain, (v1_xreal_0(k31_sin_cos))).# trainpos
cnf(c_0_198, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_423, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_657, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_695, plain, (~r1_xxreal_0(k7_square_1(np__2),np__1))).# trainpos
cnf(c_0_464, plain, (m1_subset_1(np__2,k1_numbers))).# trainpos
cnf(c_0_284, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_331, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).# trainpos
cnf(c_0_643, plain, (m1_subset_1(np__3,k1_numbers))).# trainpos
cnf(c_0_644, plain, (m1_subset_1(np__4,k1_numbers))).# trainpos
cnf(c_0_524, plain, (m1_subset_1(k32_sin_cos,k1_numbers))).# trainpos
cnf(c_0_616, plain, (k1_seq_1(k2_sincos10,k1_real_1(np__1))=k32_sin_cos)).# trainpos
cnf(c_0_604, plain, (k1_seq_1(k2_sincos10,k1_real_1(k7_square_1(np__2)))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))).# trainpos
cnf(c_0_578, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))).# trainpos
cnf(c_0_127, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_672, plain, (v1_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_276, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_132, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_636, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_508, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_275, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_669, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_108, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).# trainpos
cnf(c_0_92, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).# trainpos
cnf(c_0_518, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1))).# trainpos
cnf(c_0_313, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_304, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_252, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_668, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_319, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_305, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_497, plain, (m1_subset_1(k2_sincos10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_306, plain, (k5_xcmplx_0(X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_463, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))).# trainpos
cnf(c_0_247, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_492, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_687, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_462, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_106, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X4,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_563, plain, (m1_subset_1(k7_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_254, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_102, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_731, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_216, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_703, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),np__1)=k7_xcmplx_0(np__3,np__4))).# trainpos
cnf(c_0_652, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_422, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_264, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_96, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2))).# trainpos
cnf(c_0_193, plain, (a_2_0_rcomp_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_122, plain, (k1_xxreal_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_473, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)),k1_relset_1(k1_numbers,k2_sincos10)))).# trainpos
cnf(c_0_173, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4))).# trainpos
cnf(c_0_730, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_642, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_523, plain, (r2_hidden(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).# trainpos
cnf(c_0_519, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))))).# trainpos
cnf(c_0_528, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X1,k3_xxreal_1(X3,X2)))).# trainpos
cnf(c_0_682, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_527, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_202, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_154, plain, (a_2_1_rcomp_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_529, plain, (k3_xxreal_1(X1,X2)=k4_rcomp_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_100, plain, (k4_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_159, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_153, plain, (esk7_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_516, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_734, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_512, plain, (m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_541, plain, (m1_subset_1(k4_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_294, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_273, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_170, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_2_1_rcomp_1(X3,X1))|~r1_xxreal_0(X1,esk7_3(X2,X3,X1)))).# trainpos
cnf(c_0_511, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_490, plain, (k9_subset_1(X2,X1,X3)=k9_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_250, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~r1_xxreal_0(esk7_3(X1,X2,X3),X2))).# trainpos
cnf(c_0_744, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_244, plain, (r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_104, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k2_xboole_0(X3,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_496, plain, (k2_partfun1(X2,X3,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_722, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_200, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_186, plain, (r2_hidden(X2,a_2_0_rcomp_1(X3,X4))|X1!=X2|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_218, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,a_2_0_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_320, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,a_2_0_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_91, plain, (k2_xboole_0(k4_xxreal_1(X1,X2),k2_tarski(X1,X2))=k1_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_521, plain, (r1_xxreal_0(X4,X1)|r1_xxreal_0(X1,X3)|r2_hidden(X2,a_2_1_rcomp_1(X3,X4))|X1!=X2|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_489, plain, (r1_xxreal_0(X4,X3)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k1_seq_1(X1,X4),k1_seq_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(X4,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1)))|~r2_hidden(X3,k9_subset_1(k1_numbers,X2,k1_relset_1(k1_numbers,X1))))).# trainpos
cnf(c_0_430, plain, (v2_membered(k1_numbers))).# trainpos
cnf(c_0_103, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1))).# trainpos
cnf(c_0_177, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3))).# trainpos
cnf(c_0_129, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_116, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_746, plain, (v1_relat_1(k2_sincos10))).# trainpos
cnf(c_0_162, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_506, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_534, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k1_real_1(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).# trainpos
cnf(c_0_546, plain, (r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X1,k4_rcomp_1(X3,X2))|~v1_xreal_0(X2))).# trainpos
cnf(c_0_493, plain, (r1_tarski(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_relset_1(k1_numbers,k2_sincos10)))).# trainpos
cnf(c_0_663, plain, (v1_xreal_0(np__4))).# trainpos
cnf(c_0_690, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_253, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_484, plain, (v1_xreal_0(np__2))).# trainpos
cnf(c_0_635, plain, (k1_seq_1(k2_sincos10,k4_xcmplx_0(np__1))=k32_sin_cos)).# trainpos
cnf(c_0_759, plain, (v1_xreal_0(np__3))).# trainpos
cnf(c_0_221, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_439, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_542, plain, (r2_hidden(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos))|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_209, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_707, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_272, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_677, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_686, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_585, plain, (v1_xreal_0(k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_662, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))|~v1_xreal_0(np__4))).# trainpos
cnf(c_0_745, plain, (v3_valued_0(k2_sincos10))).# trainpos
cnf(c_0_429, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_555, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_748, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_533, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_303, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_157, plain, (esk7_3(X1,X2,X3)=X1|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_283, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_312, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1)))).# trainpos
cnf(c_0_282, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_517, plain, (k3_xboole_0(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_relset_1(k1_numbers,k2_sincos10))=k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_192, plain, (r2_hidden(X1,a_2_0_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_182, plain, (r2_hidden(X1,k2_tarski(X1,X2)))).# trainpos
cnf(c_0_107, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_rcomp_1(X1,X2))=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_113, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X2,X3)))).# trainpos
cnf(c_0_740, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).# trainpos
cnf(c_0_540, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_112, plain, (r2_hidden(X1,k2_tarski(X2,X1)))).# trainpos
cnf(c_0_589, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).# trainpos
cnf(c_0_520, plain, (k2_partfun1(k1_numbers,k1_numbers,k2_sincos10,X1)=k5_relat_1(k2_sincos10,X1))).# trainpos
cnf(c_0_161, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_292, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_130, plain, (k2_xboole_0(k2_tarski(X2,X1),k4_xxreal_1(X1,X2))=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_645, plain, (m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_222, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k1_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_627, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).# trainpos
cnf(c_0_111, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r2_hidden(X1,k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3))).# trainpos
cnf(c_0_325, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_569, plain, (m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_571, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).# trainpos
cnf(c_0_311, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_440, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_443, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_736, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_509, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),k1_numbers))).# trainpos
cnf(c_0_515, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_224, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_207, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_549, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).# trainpos
cnf(c_0_539, plain, (v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))))).# trainpos
cnf(c_0_188, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X1,X3)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_199, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_117, plain, (k2_xboole_0(k2_tarski(X2,X1),k2_rcomp_1(X1,X2))=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_120, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X1)|~r1_xxreal_0(X3,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3))).# trainpos
cnf(c_0_164, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_646, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers))).# trainpos
cnf(c_0_301, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_121, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_136, plain, (r2_hidden(X1,k1_xxreal_1(X3,X2))|~r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3))).# trainpos
cnf(c_0_582, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers))).# trainpos
cnf(c_0_124, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_664, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_545, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_391, plain, (m1_subset_1(X1,k2_tarski(X2,X1)))).# trainpos
cnf(c_0_537, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k3_xboole_0(k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k3_xboole_0(k1_relset_1(k1_numbers,X3),X4)))).# trainpos
cnf(c_0_580, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_554, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|r2_hidden(X2,k2_rcomp_1(X3,X1))|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_750, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_510, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_342, plain, (m1_subset_1(X1,k2_tarski(X1,X2)))).# trainpos
cnf(c_0_137, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r2_hidden(X1,k2_rcomp_1(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3))).# trainpos
cnf(c_0_702, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_324, plain, (m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_507, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_574, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(np__1)))).# trainpos
cnf(c_0_448, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_553, plain, (v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))))).# trainpos
cnf(c_0_532, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).# trainpos
cnf(c_0_128, plain, (r2_hidden(X1,k1_xxreal_1(X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_195, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X1,X3)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_531, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).# trainpos
cnf(c_0_586, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_552, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))|~m1_subset_1(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).# trainpos
cnf(c_0_141, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_333, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X3)|~r2_hidden(X2,k2_tarski(X1,X3)))).# trainpos
cnf(c_0_681, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_378, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X3,X1)|~r1_xxreal_0(X3,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_215, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_698, negated_conjecture, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).# trainpos
cnf(c_0_131, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_133, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_330, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_594, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))).# trainpos
cnf(c_0_721, plain, (m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).# trainpos
cnf(c_0_680, plain, (v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_758, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_558, plain, (m1_subset_1(k7_xcmplx_0(X1,np__2),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_231, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_125, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_598, negated_conjecture, (v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_753, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_338, plain, (r1_xxreal_0(X1,esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_208, plain, (r2_hidden(X1,k1_rcomp_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_570, plain, (m1_subset_1(k7_xcmplx_0(k32_sin_cos,np__2),k1_numbers))).# trainpos
cnf(c_0_345, plain, (r1_xxreal_0(esk6_3(X1,X1,X2),X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_341, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_384, plain, (r2_hidden(X1,k1_rcomp_1(X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_568, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_577, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).# trainpos
cnf(c_0_220, plain, (esk6_3(X1,X2,X1)=X1|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_671, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_735, plain, (m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers))).# trainpos
cnf(c_0_708, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(esk1_0))=esk1_0)).# trainpos
cnf(c_0_142, plain, (r2_hidden(X1,X2)|r2_hidden(X1,k2_tarski(X3,X4))|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X4)))).# trainpos
cnf(c_0_700, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_603, plain, (m1_subset_1(X1,k1_rcomp_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_719, plain, (v1_xreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_688, negated_conjecture, (v1_xcmplx_0(esk1_0))).# trainpos
cnf(c_0_588, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).# trainpos
cnf(c_0_581, plain, (v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_600, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_591, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).# trainpos
cnf(c_0_699, plain, (v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_336, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_718, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).# trainpos
cnf(c_0_214, plain, (m1_subset_1(esk6_3(X1,X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_621, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))))).# trainpos
cnf(c_0_749, plain, (v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).# trainpos
cnf(c_0_355, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k2_tarski(X1,X3)))).# trainpos
cnf(c_0_135, plain, (~v1_xboole_0(k1_xxreal_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_390, plain, (esk6_3(X1,X2,X3)=X1|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_592, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).# trainpos
cnf(c_0_606, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))).# trainpos
cnf(c_0_602, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).# trainpos
cnf(c_0_236, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_615, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_709, negated_conjecture, (k3_xcmplx_0(esk1_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk1_0))).# trainpos
cnf(c_0_387, plain, (m1_subset_1(esk6_3(X1,X2,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1))).# trainpos
cnf(c_0_732, negated_conjecture, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))))).# trainpos
cnf(c_0_605, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).# trainpos
cnf(c_0_726, negated_conjecture, (m1_subset_1(k4_xcmplx_0(esk1_0),k1_numbers))).# trainpos
cnf(c_0_640, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))))).# trainpos
cnf(c_0_757, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))))).# trainpos
cnf(c_0_634, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).# trainpos
cnf(c_0_725, negated_conjecture, (k1_real_1(esk1_0)=k4_xcmplx_0(esk1_0)|~m1_subset_1(k4_xcmplx_0(esk1_0),k1_numbers))).# trainpos
cnf(c_0_134, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|v1_xboole_0(k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).# trainpos
cnf(c_0_455, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_762, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4))))).# trainpos
cnf(c_0_754, negated_conjecture, (k1_real_1(k4_xcmplx_0(esk1_0))=esk1_0)).# trainpos
cnf(c_0_143, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).# trainpos
cnf(c_0_654, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).# trainpos
cnf(c_0_225, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X2,k1_xxreal_1(X1,X3)))).# trainpos
cnf(c_0_764, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4))))).# trainpos
cnf(c_0_149, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).# trainpos
cnf(c_0_350, plain, (v1_xreal_0(esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_227, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_rcomp_1(X1,X3)))).# trainpos
cnf(c_0_152, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).# trainpos
cnf(c_0_618, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_206, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_155, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).# trainpos
cnf(c_0_167, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_637, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).# trainpos
cnf(c_0_673, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k4_xcmplx_0(k7_square_1(np__2)),k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).# trainpos
cnf(c_0_710, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_160, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).# trainpos
cnf(c_0_175, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_166, plain, (X1=X2|X1=X3|esk7_3(X1,X3,X2)=X1|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).# trainpos
cnf(c_0_179, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_184, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_171, plain, (esk7_3(X1,X2,X3)=X1|X1=X2|X1=X3|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_174, plain, (X2=X1|X2=X3|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_2_1_rcomp_1(X3,X1))|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_rcomp_1(X3,X1)))).# trainpos
cnf(c_0_261, plain, (X1=X3|X1=X2|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_183, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_278, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_197, plain, (X1=X2|X1=X3|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2)))).# trainpos
cnf(c_0_293, plain, (X1=X2|X1=X3|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2)))).# trainpos
cnf(c_0_660, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_310, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_213, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_676, plain, (r1_xxreal_0(k7_square_1(X1),X2)|r1_xxreal_0(X2,k7_square_1(X1))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_317, plain, (k7_xcmplx_0(X1,X2)=X3|k7_xcmplx_0(X1,X2)=X4|~r1_xxreal_0(k7_xcmplx_0(X1,X2),X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k7_xcmplx_0(X1,X2))|~r1_xxreal_0(k7_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_219, plain, (esk6_3(X1,X1,X2)=X3|esk6_3(X1,X1,X2)=X4|~r1_xxreal_0(X3,esk6_3(X1,X1,X2))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,esk6_3(X1,X1,X2))|~r1_xxreal_0(esk6_3(X1,X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_223, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_323, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_329, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_228, plain, (esk6_3(X1,X2,X3)=X2|esk6_3(X1,X2,X3)=X4|~r1_xxreal_0(X4,esk6_3(X1,X2,X3))|~r1_xxreal_0(esk6_3(X1,X2,X3),X2)|~r1_xxreal_0(X4,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(esk6_3(X1,X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_229, plain, (esk6_3(X1,X2,X3)=X2|~r1_xxreal_0(esk6_3(X1,X2,X3),X2)|~r1_xxreal_0(X2,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(esk6_3(X1,X2,X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).# trainpos
cnf(c_0_340, plain, (esk6_3(X1,X2,X3)=X3|esk6_3(X1,X2,X3)=X4|~r1_xxreal_0(X3,esk6_3(X1,X2,X3))|~r1_xxreal_0(esk6_3(X1,X2,X3),X4)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(esk6_3(X1,X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_346, plain, (esk6_3(X1,X1,X1)=X2|esk6_3(X1,X1,X1)=X1|~r1_xxreal_0(esk6_3(X1,X1,X1),X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_235, plain, (X1=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X1)))).# trainpos
cnf(c_0_356, plain, (esk6_3(X1,X1,X1)=X1|~r1_xxreal_0(X1,X1)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_251, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_364, plain, (esk6_3(X1,X1,X1)=X1|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1))).# trainpos
cnf(c_0_271, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_328, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_291, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_410, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_339, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_420, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_300, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_344, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_353, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_612, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_362, plain, (k4_xcmplx_0(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_595, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_370, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1)))).# trainpos
cnf(c_0_375, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_380, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_386, plain, (esk6_3(X2,k1_real_1(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,k1_real_1(esk6_3(X2,k1_real_1(X1),X2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(X2,k1_real_1(X1),X2),k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1)))).# trainpos
cnf(c_0_393, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2)))).# trainpos
cnf(c_0_316, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_398, plain, (X1=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_327, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_337, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_349, plain, (k4_xcmplx_0(esk6_3(X2,X2,X1))=k1_real_1(X1)|~r1_xxreal_0(X1,esk6_3(X2,X2,X1))|~v1_xreal_0(esk6_3(X2,X2,X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_343, plain, (k1_real_1(esk6_3(X1,X1,X2))=k4_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(X1,X1,X2),k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_399, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_359, plain, (k4_xcmplx_0(esk6_3(X1,X1,X1))=k1_real_1(X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_352, plain, (k1_real_1(esk6_3(X1,X1,X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_547, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_361, plain, (esk6_3(X1,X1,X1)=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_365, plain, (esk6_3(X1,X1,X1)=k1_real_1(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1))).# trainpos
cnf(c_0_560, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_561, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_706, negated_conjecture, (r1_xxreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_406, plain, (X1=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_367, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X2,X2,X2))|~v1_xreal_0(X2))).# trainpos
cnf(c_0_425, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_584, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_372, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_371, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_739, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_432, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_437, plain, (v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_597, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk1_0))).# trainpos
cnf(c_0_752, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_397, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_374, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_405, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_610, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_379, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_623, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_square_1(np__2)))).# trainpos
cnf(c_0_408, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_607, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_385, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_583, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_428, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_416, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_609, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_427, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_434, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2)))).# trainpos
cnf(c_0_641, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_659, negated_conjecture, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k1_seq_1(k2_sincos10,esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_433, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2))).# trainpos
cnf(c_0_675, negated_conjecture, (~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_438, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_638, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_692, negated_conjecture, (~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2))))).# trainpos
cnf(c_0_655, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_674, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).# trainpos
cnf(c_0_442, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)=k1_real_1(X1)|~r1_xxreal_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_454, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),k4_xcmplx_0(X1))=k1_real_1(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_711, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).# trainpos
cnf(c_0_392, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_396, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_400, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X2))).# trainpos
cnf(c_0_407, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_409, plain, (X1=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_414, plain, (esk6_3(X1,X1,X2)=k4_xcmplx_0(k1_real_1(X1))|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~v1_xreal_0(esk6_3(X1,X1,X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_424, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_436, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_565, plain, (esk6_3(X1,X1,X1)=k1_real_1(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_446, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_453, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_487, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_738, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_488, plain, (esk6_3(X1,X1,X1)=k1_real_1(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_536, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_564, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_587, plain, (r1_xxreal_0(k1_real_1(X1),esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X2))).# trainpos
cnf(c_0_613, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_648, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X1)),X2)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_685, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_723, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_751, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2))).# trainpos
cnf(c_0_693, plain, (r1_xxreal_0(X1,k7_square_1(np__2))|r1_xxreal_0(k7_square_1(np__2),X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_712, plain, (r1_xxreal_0(np__1,k7_square_1(np__2)))).# trainpos
cnf(c_0_728, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).# trainpos
cnf(c_0_742, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).# trainpos
cnf(c_0_755, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0)))).# trainpos
cnf(c_0_743, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).# trainpos
cnf(c_0_760, negated_conjecture, (k4_xcmplx_0(np__1)=esk1_0|~r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0))).# trainpos
cnf(c_0_756, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0)))).# trainpos
cnf(c_0_761, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0))).# trainpos
cnf(c_0_763, negated_conjecture, (k4_xcmplx_0(np__1)=esk1_0)).# trainpos
cnf(c_0_765, plain, (k1_seq_1(k2_sincos10,esk1_0)=k32_sin_cos)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_652, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_1182, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_428, plain, (v5_membered(k4_numbers))).#trainneg
cnf(i_0_541, plain, (v5_membered(esk21_0))).#trainneg
cnf(i_0_592, plain, (v5_membered(esk36_0))).#trainneg
cnf(i_0_502, plain, (v1_relat_1(esk9_0))).#trainneg
cnf(i_0_528, plain, (v1_relat_1(esk15_0))).#trainneg
cnf(i_0_532, plain, (v1_relat_1(esk16_0))).#trainneg
cnf(i_0_548, plain, (v1_relat_1(esk22_0))).#trainneg
cnf(i_0_561, plain, (v1_relat_1(esk26_0))).#trainneg
cnf(i_0_616, plain, (v1_relat_1(esk42_0))).#trainneg
cnf(i_0_1237, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_1231, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_1243, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_629, plain, (~v1_xxreal_2(esk46_0))).#trainneg
cnf(i_0_639, plain, (~v1_xxreal_2(esk49_0))).#trainneg
cnf(i_0_279, plain, (~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_634, plain, (~v2_xxreal_2(esk48_0))).#trainneg
cnf(i_0_638, plain, (~v2_xxreal_2(esk49_0))).#trainneg
cnf(i_0_523, plain, (v1_valued_0(esk15_0))).#trainneg
cnf(i_0_512, plain, (v3_membered(esk13_0))).#trainneg
cnf(i_0_543, plain, (v3_membered(esk21_0))).#trainneg
cnf(i_0_594, plain, (v3_membered(esk36_0))).#trainneg
cnf(i_0_538, plain, (v1_xxreal_2(esk21_0))).#trainneg
cnf(i_0_589, plain, (v1_xxreal_2(esk36_0))).#trainneg
cnf(i_0_621, plain, (v1_xxreal_2(esk44_0))).#trainneg
cnf(i_0_635, plain, (v1_xxreal_2(esk48_0))).#trainneg
cnf(i_0_446, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_508, plain, (v6_membered(esk11_0))).#trainneg
cnf(i_0_540, plain, (v6_membered(esk21_0))).#trainneg
cnf(i_0_1249, plain, (~v1_xboole_0(np__4))).#trainneg
cnf(i_0_278, plain, (~v4_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_447, plain, (~v1_finset_1(k4_numbers))).#trainneg
cnf(i_0_456, plain, (~v1_finset_1(k3_numbers))).#trainneg
cnf(i_0_468, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_414, plain, (~v1_xboole_0(k4_numbers))).#trainneg
cnf(i_0_389, plain, (~v1_xboole_0(k3_numbers))).#trainneg
cnf(i_0_318, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_550, plain, (v6_membered(esk24_0))).#trainneg
cnf(i_0_574, plain, (v6_membered(esk31_0))).#trainneg
cnf(i_0_576, plain, (v6_membered(esk32_0))).#trainneg
cnf(i_0_591, plain, (v6_membered(esk36_0))).#trainneg
cnf(i_0_522, plain, (v2_valued_0(esk15_0))).#trainneg
cnf(i_0_537, plain, (v2_xxreal_2(esk21_0))).#trainneg
cnf(i_0_620, plain, (v2_xxreal_2(esk44_0))).#trainneg
cnf(i_0_628, plain, (v2_xxreal_2(esk46_0))).#trainneg
cnf(i_0_514, plain, (v1_membered(esk13_0))).#trainneg
cnf(i_0_545, plain, (v1_membered(esk21_0))).#trainneg
cnf(i_0_596, plain, (v1_membered(esk36_0))).#trainneg
cnf(i_0_450, plain, (v2_funct_1(k2_sincos10))).#trainneg
cnf(i_0_505, plain, (v1_xcmplx_0(esk10_0))).#trainneg
cnf(i_0_449, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_1214, plain, (~v1_xboole_0(np__11))).#trainneg
cnf(i_0_1220, plain, (~v1_xboole_0(np__13))).#trainneg
cnf(i_0_1255, plain, (~v1_xboole_0(np__5))).#trainneg
cnf(i_0_1262, plain, (~v1_xboole_0(np__7))).#trainneg
cnf(i_0_1269, plain, (~v1_xboole_0(np__9))).#trainneg
cnf(i_0_1225, plain, (~v1_xboole_0(np__15))).#trainneg
cnf(i_0_509, plain, (~v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_534, plain, (v1_xcmplx_0(esk18_0))).#trainneg
cnf(i_0_564, plain, (v1_xcmplx_0(esk28_0))).#trainneg
cnf(i_0_569, plain, (v1_xcmplx_0(esk29_0))).#trainneg
cnf(i_0_585, plain, (v1_xcmplx_0(esk34_0))).#trainneg
cnf(i_0_607, plain, (v1_xcmplx_0(esk39_0))).#trainneg
cnf(i_0_521, plain, (v3_valued_0(esk15_0))).#trainneg
cnf(i_0_588, plain, (v5_xxreal_2(esk36_0))).#trainneg
cnf(i_0_572, plain, (v1_finset_1(esk31_0))).#trainneg
cnf(i_0_513, plain, (v2_membered(esk13_0))).#trainneg
cnf(i_0_544, plain, (v2_membered(esk21_0))).#trainneg
cnf(i_0_595, plain, (v2_membered(esk36_0))).#trainneg
cnf(i_0_1241, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_613, plain, (v2_membered(esk41_0))).#trainneg
cnf(i_0_539, plain, (~v1_xboole_0(esk21_0))).#trainneg
cnf(i_0_551, plain, (~v1_xboole_0(esk24_0))).#trainneg
cnf(i_0_555, plain, (~v1_xboole_0(esk25_0))).#trainneg
cnf(i_0_562, plain, (~v1_xboole_0(esk26_0))).#trainneg
cnf(i_0_563, plain, (~v1_xboole_0(esk27_0))).#trainneg
cnf(i_0_565, plain, (~v1_xboole_0(esk28_0))).#trainneg
cnf(i_0_573, plain, (~v1_xboole_0(esk31_0))).#trainneg
cnf(i_0_577, plain, (~v1_xboole_0(esk32_0))).#trainneg
cnf(i_0_622, plain, (v2_membered(esk44_0))).#trainneg
cnf(i_0_630, plain, (v2_membered(esk46_0))).#trainneg
cnf(i_0_636, plain, (v2_membered(esk48_0))).#trainneg
cnf(i_0_641, plain, (v2_membered(esk49_0))).#trainneg
cnf(i_0_506, plain, (v1_xxreal_0(esk10_0))).#trainneg
cnf(i_0_536, plain, (v1_xxreal_0(esk20_0))).#trainneg
cnf(i_0_568, plain, (v1_xxreal_0(esk29_0))).#trainneg
cnf(i_0_571, plain, (v1_xxreal_0(esk30_0))).#trainneg
cnf(i_0_584, plain, (v1_xxreal_0(esk34_0))).#trainneg
cnf(i_0_606, plain, (v1_xxreal_0(esk39_0))).#trainneg
cnf(i_0_1235, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_609, plain, (v1_xxreal_0(esk40_0))).#trainneg
cnf(i_0_590, plain, (~v1_xboole_0(esk36_0))).#trainneg
cnf(i_0_612, plain, (~v1_xboole_0(esk41_0))).#trainneg
cnf(i_0_618, plain, (~v1_xboole_0(esk43_0))).#trainneg
cnf(i_0_632, plain, (~v1_xboole_0(esk47_0))).#trainneg
cnf(i_0_640, plain, (~v1_xboole_0(esk49_0))).#trainneg
cnf(i_0_614, plain, (~v3_funct_1(esk42_0))).#trainneg
cnf(i_0_713, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_712, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_504, plain, (v1_xreal_0(esk10_0))).#trainneg
cnf(i_0_535, plain, (v1_xreal_0(esk19_0))).#trainneg
cnf(i_0_566, plain, (v1_xreal_0(esk29_0))).#trainneg
cnf(i_0_582, plain, (v1_xreal_0(esk34_0))).#trainneg
cnf(i_0_605, plain, (v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_413, plain, (v4_membered(k3_numbers))).#trainneg
cnf(i_0_542, plain, (v4_membered(esk21_0))).#trainneg
cnf(i_0_593, plain, (v4_membered(esk36_0))).#trainneg
cnf(i_0_530, plain, (v4_valued_0(esk16_0))).#trainneg
cnf(i_0_1247, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_556, plain, (v4_valued_0(esk26_0))).#trainneg
cnf(i_0_732, plain, (~r1_xxreal_0(np__3,np__2))).#trainneg
cnf(i_0_729, plain, (~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_747, plain, (~r1_xxreal_0(np__4,np__2))).#trainneg
cnf(i_0_696, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_728, plain, (~r1_xxreal_0(np__3,np__0))).#trainneg
cnf(i_0_744, plain, (~r1_xxreal_0(np__4,np__1))).#trainneg
cnf(i_0_748, plain, (~r1_xxreal_0(np__4,np__3))).#trainneg
cnf(i_0_743, plain, (~r1_xxreal_0(np__4,np__0))).#trainneg
cnf(i_0_326, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_1208, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_533, plain, (v1_xboole_0(esk17_0))).#trainneg
cnf(i_0_608, plain, (v1_xboole_0(esk39_0))).#trainneg
cnf(i_0_610, plain, (v1_xboole_0(esk40_0))).#trainneg
cnf(i_0_307, plain, (v6_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_611, plain, (v6_xxreal_2(esk41_0))).#trainneg
cnf(i_0_619, plain, (v6_xxreal_2(esk44_0))).#trainneg
cnf(i_0_627, plain, (v6_xxreal_2(esk46_0))).#trainneg
cnf(i_0_633, plain, (v6_xxreal_2(esk48_0))).#trainneg
cnf(i_0_637, plain, (v6_xxreal_2(esk49_0))).#trainneg
cnf(i_0_1253, plain, (v2_xxreal_0(np__4))).#trainneg
cnf(i_0_1309, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_503, plain, (v1_int_1(esk10_0))).#trainneg
cnf(i_0_680, plain, (~r1_xxreal_0(np__11,np__2))).#trainneg
cnf(i_0_677, plain, (~r1_xxreal_0(np__11,np__1))).#trainneg
cnf(i_0_681, plain, (~r1_xxreal_0(np__11,np__3))).#trainneg
cnf(i_0_682, plain, (~r1_xxreal_0(np__11,np__4))).#trainneg
cnf(i_0_676, plain, (~r1_xxreal_0(np__11,np__0))).#trainneg
cnf(i_0_683, plain, (~r1_xxreal_0(np__11,np__5))).#trainneg
cnf(i_0_684, plain, (~r1_xxreal_0(np__11,np__7))).#trainneg
cnf(i_0_685, plain, (~r1_xxreal_0(np__11,np__9))).#trainneg
cnf(i_0_549, plain, (v1_int_1(esk23_0))).#trainneg
cnf(i_0_597, plain, (v7_ordinal1(esk37_0))).#trainneg
cnf(i_0_617, plain, (v7_ordinal1(esk43_0))).#trainneg
cnf(i_0_218, plain, (v1_funct_1(k1_fdiff_9))).#trainneg
cnf(i_0_216, plain, (v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_501, plain, (v1_funct_1(esk9_0))).#trainneg
cnf(i_0_525, plain, (v1_funct_1(esk15_0))).#trainneg
cnf(i_0_531, plain, (v1_funct_1(esk16_0))).#trainneg
cnf(i_0_547, plain, (v1_funct_1(esk22_0))).#trainneg
cnf(i_0_558, plain, (v1_funct_1(esk26_0))).#trainneg
cnf(i_0_615, plain, (v1_funct_1(esk42_0))).#trainneg
cnf(i_0_716, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_101, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_448, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_690, plain, (~r1_xxreal_0(np__13,np__2))).#trainneg
cnf(i_0_687, plain, (~r1_xxreal_0(np__13,np__1))).#trainneg
cnf(i_0_691, plain, (~r1_xxreal_0(np__13,np__3))).#trainneg
cnf(i_0_692, plain, (~r1_xxreal_0(np__13,np__4))).#trainneg
cnf(i_0_686, plain, (~r1_xxreal_0(np__13,np__0))).#trainneg
cnf(i_0_688, plain, (~r1_xxreal_0(np__13,np__11))).#trainneg
cnf(i_0_693, plain, (~r1_xxreal_0(np__13,np__5))).#trainneg
cnf(i_0_694, plain, (~r1_xxreal_0(np__13,np__7))).#trainneg
cnf(i_0_510, plain, (v3_ordinal1(esk12_0))).#trainneg
cnf(i_0_552, plain, (v3_ordinal1(esk25_0))).#trainneg
cnf(i_0_554, plain, (v1_ordinal1(esk25_0))).#trainneg
cnf(i_0_553, plain, (v2_ordinal1(esk25_0))).#trainneg
cnf(i_0_524, plain, (v1_zfmisc_1(esk15_0))).#trainneg
cnf(i_0_426, plain, (v7_membered(k4_numbers))).#trainneg
cnf(i_0_425, plain, (v7_membered(k3_numbers))).#trainneg
cnf(i_0_424, plain, (v7_membered(k1_numbers))).#trainneg
cnf(i_0_427, plain, (v7_membered(k4_ordinal1))).#trainneg
cnf(i_0_575, plain, (v7_membered(esk32_0))).#trainneg
cnf(i_0_546, plain, (v2_funct_1(esk22_0))).#trainneg
cnf(i_0_717, plain, (r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_88, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_323, plain, (v2_xxreal_0(k31_sin_cos))).#trainneg
cnf(i_0_762, plain, (~r1_xxreal_0(np__5,np__2))).#trainneg
cnf(i_0_759, plain, (~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_763, plain, (~r1_xxreal_0(np__5,np__3))).#trainneg
cnf(i_0_695, plain, (~r1_xxreal_0(np__13,np__9))).#trainneg
cnf(i_0_764, plain, (~r1_xxreal_0(np__5,np__4))).#trainneg
cnf(i_0_758, plain, (~r1_xxreal_0(np__5,np__0))).#trainneg
cnf(i_0_772, plain, (~r1_xxreal_0(np__7,np__2))).#trainneg
cnf(i_0_769, plain, (~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_1218, plain, (v2_xxreal_0(np__11))).#trainneg
cnf(i_0_1224, plain, (v2_xxreal_0(np__13))).#trainneg
cnf(i_0_1259, plain, (v2_xxreal_0(np__5))).#trainneg
cnf(i_0_1266, plain, (v2_xxreal_0(np__7))).#trainneg
cnf(i_0_1273, plain, (v2_xxreal_0(np__9))).#trainneg
cnf(i_0_1229, plain, (v2_xxreal_0(np__15))).#trainneg
cnf(i_0_567, plain, (v2_xxreal_0(esk29_0))).#trainneg
cnf(i_0_570, plain, (v2_xxreal_0(esk30_0))).#trainneg
cnf(i_0_583, plain, (v3_xxreal_0(esk34_0))).#trainneg
cnf(i_0_586, plain, (v3_xxreal_0(esk35_0))).#trainneg
cnf(i_0_557, plain, (v3_funct_1(esk26_0))).#trainneg
cnf(i_0_718, plain, (r1_xxreal_0(np__2,np__4))).#trainneg
cnf(i_0_631, plain, (v4_funct_1(esk47_0))).#trainneg
cnf(i_0_773, plain, (~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_774, plain, (~r1_xxreal_0(np__7,np__4))).#trainneg
cnf(i_0_782, plain, (~r1_xxreal_0(np__9,np__2))).#trainneg
cnf(i_0_768, plain, (~r1_xxreal_0(np__7,np__0))).#trainneg
cnf(i_0_775, plain, (~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_779, plain, (~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_783, plain, (~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_784, plain, (~r1_xxreal_0(np__9,np__4))).#trainneg
cnf(i_0_511, plain, (v3_rcomp_1(esk13_0))).#trainneg
cnf(i_0_1190, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_1191, plain, (k4_xcmplx_0(k4_xcmplx_0(np__3))=np__3)).#trainneg
cnf(i_0_1192, plain, (k4_xcmplx_0(k4_xcmplx_0(np__4))=np__4)).#trainneg
cnf(i_0_1193, plain, (k4_xcmplx_0(k4_xcmplx_0(np__5))=np__5)).#trainneg
cnf(i_0_581, plain, (v1_relat_1(esk33_1(X1)))).#trainneg
cnf(i_0_603, plain, (v1_relat_1(esk38_1(X1)))).#trainneg
cnf(i_0_579, plain, (v1_funct_1(esk33_1(X1)))).#trainneg
cnf(i_0_600, plain, (v1_funct_1(esk38_1(X1)))).#trainneg
cnf(i_0_599, plain, (v3_funct_1(esk38_1(X1)))).#trainneg
cnf(i_0_700, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_90, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_578, plain, (v5_ordinal1(esk33_1(X1)))).#trainneg
cnf(i_0_778, plain, (~r1_xxreal_0(np__9,np__0))).#trainneg
cnf(i_0_785, plain, (~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_786, plain, (~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_722, plain, (~r1_xxreal_0(np__2,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_706, plain, (~r1_xxreal_0(np__1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_738, plain, (~r1_xxreal_0(np__3,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_753, plain, (~r1_xxreal_0(np__4,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_671, plain, (~r1_xxreal_0(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_860, plain, (k2_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_859, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).#trainneg
cnf(i_0_858, plain, (k2_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_851, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).#trainneg
cnf(i_0_850, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_852, plain, (k2_xcmplx_0(np__1,np__3)=np__4)).#trainneg
cnf(i_0_853, plain, (k2_xcmplx_0(np__1,np__4)=np__5)).#trainneg
cnf(i_0_849, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_864, plain, (k2_xcmplx_0(np__3,np__1)=np__4)).#trainneg
cnf(i_0_865, plain, (k2_xcmplx_0(np__3,np__4)=np__7)).#trainneg
cnf(i_0_867, plain, (k2_xcmplx_0(np__4,np__1)=np__5)).#trainneg
cnf(i_0_697, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_100, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_847, plain, (k2_xcmplx_0(np__0,np__2)=np__2)).#trainneg
cnf(i_0_1112, plain, (k3_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_1111, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_723, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_724, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_725, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_726, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_707, plain, (~r1_xxreal_0(np__1,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_846, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_845, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_1110, plain, (k3_xcmplx_0(np__2,np__0)=np__0)).#trainneg
cnf(i_0_1101, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_1100, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_1102, plain, (k3_xcmplx_0(np__1,np__3)=np__3)).#trainneg
cnf(i_0_1103, plain, (k3_xcmplx_0(np__1,np__4)=np__4)).#trainneg
cnf(i_0_1099, plain, (k3_xcmplx_0(np__1,np__0)=np__0)).#trainneg
cnf(i_0_1104, plain, (k3_xcmplx_0(np__1,np__5)=np__5)).#trainneg
cnf(i_0_1119, plain, (k3_xcmplx_0(np__3,np__1)=np__3)).#trainneg
cnf(i_0_1123, plain, (k3_xcmplx_0(np__4,np__1)=np__4)).#trainneg
cnf(i_0_701, plain, (r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_89, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1097, plain, (k3_xcmplx_0(np__0,np__2)=np__0)).#trainneg
cnf(i_0_934, plain, (k6_xcmplx_0(np__2,np__1)=np__1)).#trainneg
cnf(i_0_950, plain, (k6_xcmplx_0(np__3,np__2)=np__1)).#trainneg
cnf(i_0_949, plain, (k6_xcmplx_0(np__3,np__1)=np__2)).#trainneg
cnf(i_0_708, plain, (~r1_xxreal_0(np__1,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_710, plain, (~r1_xxreal_0(np__1,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_739, plain, (~r1_xxreal_0(np__3,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_740, plain, (~r1_xxreal_0(np__3,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_741, plain, (~r1_xxreal_0(np__3,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_1096, plain, (k3_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_1095, plain, (k3_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_935, plain, (k6_xcmplx_0(np__2,np__2)=np__0)).#trainneg
cnf(i_0_933, plain, (k6_xcmplx_0(np__2,np__0)=np__2)).#trainneg
cnf(i_0_915, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_914, plain, (k6_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_951, plain, (k6_xcmplx_0(np__3,np__3)=np__0)).#trainneg
cnf(i_0_948, plain, (k6_xcmplx_0(np__3,np__0)=np__3)).#trainneg
cnf(i_0_960, plain, (k6_xcmplx_0(np__4,np__2)=np__2)).#trainneg
cnf(i_0_959, plain, (k6_xcmplx_0(np__4,np__1)=np__3)).#trainneg
cnf(i_0_961, plain, (k6_xcmplx_0(np__4,np__3)=np__1)).#trainneg
cnf(i_0_702, plain, (r1_xxreal_0(np__1,np__4))).#trainneg
cnf(i_0_79, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_962, plain, (k6_xcmplx_0(np__4,np__4)=np__0)).#trainneg
cnf(i_0_1092, plain, (k7_xcmplx_0(np__2,np__2)=np__1)).#trainneg
cnf(i_0_1080, plain, (k7_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_1093, plain, (k7_xcmplx_0(np__4,np__4)=np__1)).#trainneg
cnf(i_0_742, plain, (~r1_xxreal_0(np__3,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_754, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_755, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_756, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_757, plain, (~r1_xxreal_0(np__4,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_958, plain, (k6_xcmplx_0(np__4,np__0)=np__4)).#trainneg
cnf(i_0_968, plain, (k6_xcmplx_0(np__5,np__2)=np__3)).#trainneg
cnf(i_0_967, plain, (k6_xcmplx_0(np__5,np__1)=np__4)).#trainneg
cnf(i_0_969, plain, (k6_xcmplx_0(np__5,np__4)=np__1)).#trainneg
cnf(i_0_971, plain, (k6_xcmplx_0(np__7,np__2)=np__5)).#trainneg
cnf(i_0_972, plain, (k6_xcmplx_0(np__7,np__3)=np__4)).#trainneg
cnf(i_0_973, plain, (k6_xcmplx_0(np__7,np__4)=np__3)).#trainneg
cnf(i_0_1239, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_1233, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_733, plain, (r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_150, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1245, plain, (m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_734, plain, (r1_xxreal_0(np__3,np__4))).#trainneg
cnf(i_0_749, plain, (r1_xxreal_0(np__4,np__4))).#trainneg
cnf(i_0_799, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__0))).#trainneg
cnf(i_0_810, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__4),np__0))).#trainneg
cnf(i_0_820, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),np__1))).#trainneg
cnf(i_0_819, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),np__0))).#trainneg
cnf(i_0_829, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__4),np__0))).#trainneg
cnf(i_0_1251, plain, (m1_subset_1(np__4,k5_numbers))).#trainneg
cnf(i_0_1277, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_1210, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_1209, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_1216, plain, (m1_subset_1(np__11,k5_numbers))).#trainneg
cnf(i_0_1215, plain, (m1_subset_1(np__11,k1_numbers))).#trainneg
cnf(i_0_1222, plain, (m1_subset_1(np__13,k5_numbers))).#trainneg
cnf(i_0_1221, plain, (m1_subset_1(np__13,k1_numbers))).#trainneg
cnf(i_0_24, plain, (v3_membered(X1)|~v5_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1257, plain, (m1_subset_1(np__5,k5_numbers))).#trainneg
cnf(i_0_714, plain, (r1_xxreal_0(np__2,np__11))).#trainneg
cnf(i_0_715, plain, (r1_xxreal_0(np__2,np__13))).#trainneg
cnf(i_0_719, plain, (r1_xxreal_0(np__2,np__5))).#trainneg
cnf(i_0_804, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_815, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__4),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_824, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_834, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__4),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_727, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_1256, plain, (m1_subset_1(np__5,k1_numbers))).#trainneg
cnf(i_0_1264, plain, (m1_subset_1(np__7,k5_numbers))).#trainneg
cnf(i_0_1263, plain, (m1_subset_1(np__7,k1_numbers))).#trainneg
cnf(i_0_1271, plain, (m1_subset_1(np__9,k5_numbers))).#trainneg
cnf(i_0_1270, plain, (m1_subset_1(np__9,k1_numbers))).#trainneg
cnf(i_0_1227, plain, (m1_subset_1(np__15,k5_numbers))).#trainneg
cnf(i_0_1226, plain, (m1_subset_1(np__15,k1_numbers))).#trainneg
cnf(i_0_507, plain, (m1_subset_1(esk10_0,k1_numbers))).#trainneg
cnf(i_0_526, plain, (v5_relat_1(esk15_0,k1_numbers))).#trainneg
cnf(i_0_559, plain, (v5_relat_1(esk26_0,k4_numbers))).#trainneg
cnf(i_0_560, plain, (v5_relat_1(esk26_0,k3_numbers))).#trainneg
cnf(i_0_171, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_527, plain, (v4_relat_1(esk15_0,k1_numbers))).#trainneg
cnf(i_0_720, plain, (r1_xxreal_0(np__2,np__7))).#trainneg
cnf(i_0_721, plain, (r1_xxreal_0(np__2,np__9))).#trainneg
cnf(i_0_698, plain, (r1_xxreal_0(np__1,np__11))).#trainneg
cnf(i_0_711, plain, (~r1_xxreal_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_842, plain, (~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_806, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_825, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_826, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_699, plain, (r1_xxreal_0(np__1,np__13))).#trainneg
cnf(i_0_703, plain, (r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_704, plain, (r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_705, plain, (r1_xxreal_0(np__1,np__9))).#trainneg
cnf(i_0_730, plain, (r1_xxreal_0(np__3,np__11))).#trainneg
cnf(i_0_731, plain, (r1_xxreal_0(np__3,np__13))).#trainneg
cnf(i_0_735, plain, (r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_736, plain, (r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_737, plain, (r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_745, plain, (r1_xxreal_0(np__4,np__11))).#trainneg
cnf(i_0_746, plain, (r1_xxreal_0(np__4,np__13))).#trainneg
cnf(i_0_44, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_15, plain, (v1_finset_1(X1)|~v5_membered(X1)|~v5_xxreal_2(X1))).#trainneg
cnf(i_0_750, plain, (r1_xxreal_0(np__4,np__5))).#trainneg
cnf(i_0_751, plain, (r1_xxreal_0(np__4,np__7))).#trainneg
cnf(i_0_752, plain, (r1_xxreal_0(np__4,np__9))).#trainneg
cnf(i_0_665, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_828, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_835, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_836, plain, (~r1_xxreal_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_809, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_662, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_666, plain, (r1_xxreal_0(np__0,np__3))).#trainneg
cnf(i_0_667, plain, (r1_xxreal_0(np__0,np__4))).#trainneg
cnf(i_0_661, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_663, plain, (r1_xxreal_0(np__0,np__11))).#trainneg
cnf(i_0_664, plain, (r1_xxreal_0(np__0,np__13))).#trainneg
cnf(i_0_668, plain, (r1_xxreal_0(np__0,np__5))).#trainneg
cnf(i_0_669, plain, (r1_xxreal_0(np__0,np__7))).#trainneg
cnf(i_0_670, plain, (r1_xxreal_0(np__0,np__9))).#trainneg
cnf(i_0_678, plain, (r1_xxreal_0(np__11,np__11))).#trainneg
cnf(i_0_679, plain, (r1_xxreal_0(np__11,np__13))).#trainneg
cnf(i_0_180, plain, (v6_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_20, plain, (v1_finset_1(X1)|~v6_membered(X1)|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_689, plain, (r1_xxreal_0(np__13,np__13))).#trainneg
cnf(i_0_936, plain, (k6_xcmplx_0(np__2,np__3)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_937, plain, (k6_xcmplx_0(np__2,np__4)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_760, plain, (r1_xxreal_0(np__5,np__11))).#trainneg
cnf(i_0_761, plain, (r1_xxreal_0(np__5,np__13))).#trainneg
cnf(i_0_765, plain, (r1_xxreal_0(np__5,np__5))).#trainneg
cnf(i_0_766, plain, (r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_767, plain, (r1_xxreal_0(np__5,np__9))).#trainneg
cnf(i_0_770, plain, (r1_xxreal_0(np__7,np__11))).#trainneg
cnf(i_0_771, plain, (r1_xxreal_0(np__7,np__13))).#trainneg
cnf(i_0_776, plain, (r1_xxreal_0(np__7,np__7))).#trainneg
cnf(i_0_777, plain, (r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_780, plain, (r1_xxreal_0(np__9,np__11))).#trainneg
cnf(i_0_781, plain, (r1_xxreal_0(np__9,np__13))).#trainneg
cnf(i_0_787, plain, (r1_xxreal_0(np__9,np__9))).#trainneg
cnf(i_0_1291, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_1276, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).#trainneg
cnf(i_0_659, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_493, plain, (k2_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_183, plain, (v3_xxreal_2(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_81, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_38, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_916, plain, (k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_917, plain, (k6_xcmplx_0(np__1,np__3)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_918, plain, (k6_xcmplx_0(np__1,np__4)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_254, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_515, plain, (m1_subset_1(esk13_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_790, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_789, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_791, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__3))).#trainneg
cnf(i_0_494, plain, (k3_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_938, plain, (k6_xcmplx_0(np__2,np__5)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_939, plain, (k6_xcmplx_0(np__2,np__7)=k4_xcmplx_0(np__5))).#trainneg
cnf(i_0_919, plain, (k6_xcmplx_0(np__1,np__5)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_952, plain, (k6_xcmplx_0(np__3,np__4)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_953, plain, (k6_xcmplx_0(np__3,np__5)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_954, plain, (k6_xcmplx_0(np__3,np__7)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_903, plain, (k6_xcmplx_0(np__0,np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_902, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_904, plain, (k6_xcmplx_0(np__0,np__3)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_905, plain, (k6_xcmplx_0(np__0,np__4)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_145, plain, (v5_xxreal_2(X1)|~v3_xxreal_2(X1)|~v4_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_80, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_158, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_792, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_788, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__0))).#trainneg
cnf(i_0_215, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_1240, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1234, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_906, plain, (k6_xcmplx_0(np__0,np__5)=k4_xcmplx_0(np__5))).#trainneg
cnf(i_0_907, plain, (k6_xcmplx_0(np__0,np__7)=k4_xcmplx_0(np__7))).#trainneg
cnf(i_0_970, plain, (k6_xcmplx_0(np__5,np__7)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_862, plain, (k2_xcmplx_0(np__2,k4_xcmplx_0(np__2))=np__0)).#trainneg
cnf(i_0_861, plain, (k2_xcmplx_0(np__2,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_854, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_866, plain, (k2_xcmplx_0(np__3,k4_xcmplx_0(np__2))=np__1)).#trainneg
cnf(i_0_868, plain, (k2_xcmplx_0(np__4,k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_941, plain, (k6_xcmplx_0(np__2,k4_xcmplx_0(np__2))=np__4)).#trainneg
cnf(i_0_940, plain, (k6_xcmplx_0(np__2,k4_xcmplx_0(np__1))=np__3)).#trainneg
cnf(i_0_921, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__2))=np__3)).#trainneg
cnf(i_0_42, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_134, plain, (v3_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_102, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_46, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_45, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_75, plain, (v7_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1246, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1252, plain, (m2_subset_1(np__4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1211, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1217, plain, (m2_subset_1(np__11,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1223, plain, (m2_subset_1(np__13,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_920, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__2)).#trainneg
cnf(i_0_922, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__3))=np__4)).#trainneg
cnf(i_0_923, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(np__4))=np__5)).#trainneg
cnf(i_0_955, plain, (k6_xcmplx_0(np__3,k4_xcmplx_0(np__1))=np__4)).#trainneg
cnf(i_0_963, plain, (k6_xcmplx_0(np__4,k4_xcmplx_0(np__1))=np__5)).#trainneg
cnf(i_0_908, plain, (k6_xcmplx_0(np__0,k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_876, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),np__2)=np__0)).#trainneg
cnf(i_0_871, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__2)=np__1)).#trainneg
cnf(i_0_870, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0)).#trainneg
cnf(i_0_602, plain, (v4_relat_1(esk38_1(X1),k5_numbers))).#trainneg
cnf(i_0_268, plain, (m1_subset_1(esk4_1(X1),X1))).#trainneg
cnf(i_0_120, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_157, plain, (v4_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_32, plain, (~v1_xxreal_2(X1)|~v2_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_36, plain, (~v2_xxreal_2(X1)|~v2_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_112, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_1258, plain, (m2_subset_1(np__5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1265, plain, (m2_subset_1(np__7,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1272, plain, (m2_subset_1(np__9,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1228, plain, (m2_subset_1(np__15,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_580, plain, (v5_relat_1(esk33_1(X1),X1))).#trainneg
cnf(i_0_601, plain, (v5_relat_1(esk38_1(X1),X1))).#trainneg
cnf(i_0_518, plain, (v1_relat_1(esk14_2(X1,X2)))).#trainneg
cnf(i_0_626, plain, (v1_relat_1(esk45_2(X1,X2)))).#trainneg
cnf(i_0_519, plain, (v1_xboole_0(esk14_2(X1,X2)))).#trainneg
cnf(i_0_623, plain, (v1_funct_1(esk45_2(X1,X2)))).#trainneg
cnf(i_0_193, plain, (k6_rfunct_1(k1_numbers,k1_numbers,k19_sin_cos)=k1_fdiff_9)).#trainneg
cnf(i_0_855, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_848, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1113, plain, (k3_xcmplx_0(np__2,k4_xcmplx_0(np__2))=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_1105, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_113, plain, (v1_xxreal_2(X1)|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_387, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1316, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_144, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_709, plain, (r1_xxreal_0(np__1,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_672, plain, (r1_xxreal_0(np__0,k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_673, plain, (r1_xxreal_0(np__0,k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_674, plain, (r1_xxreal_0(np__0,k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_675, plain, (r1_xxreal_0(np__0,k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_875, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_869, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1128, plain, (k3_xcmplx_0(k4_xcmplx_0(np__2),np__2)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_1127, plain, (k3_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1132, plain, (k3_xcmplx_0(k4_xcmplx_0(np__4),np__1)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_993, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__2)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_992, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_976, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_975, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_977, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__3)=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_151, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_454, plain, (v1_xboole_0(X1)|~v1_xcmplx_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_252, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_110, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_1308, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_801, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__2))).#trainneg
cnf(i_0_800, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__1))).#trainneg
cnf(i_0_802, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__3))).#trainneg
cnf(i_0_803, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),np__4))).#trainneg
cnf(i_0_812, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__4),np__2))).#trainneg
cnf(i_0_978, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__4)=k4_xcmplx_0(np__5))).#trainneg
cnf(i_0_974, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__0)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1129, plain, (k3_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))=np__4)).#trainneg
cnf(i_0_995, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__2))=np__0)).#trainneg
cnf(i_0_996, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__3))=np__1)).#trainneg
cnf(i_0_997, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__4))=np__2)).#trainneg
cnf(i_0_980, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2))=np__1)).#trainneg
cnf(i_0_979, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_981, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__3))=np__2)).#trainneg
cnf(i_0_982, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__4))=np__3)).#trainneg
cnf(i_0_1003, plain, (k6_xcmplx_0(k4_xcmplx_0(np__3),k4_xcmplx_0(np__3))=np__0)).#trainneg
cnf(i_0_7, plain, (v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xxreal_2(X1))).#trainneg
cnf(i_0_459, plain, (v1_xboole_0(X1)|~v1_xcmplx_0(X1)|~v1_xboole_0(k5_xcmplx_0(X1)))).#trainneg
cnf(i_0_396, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_74, plain, (v6_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_811, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__4),np__1))).#trainneg
cnf(i_0_813, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__4),np__3))).#trainneg
cnf(i_0_814, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__4),np__4))).#trainneg
cnf(i_0_821, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__2),np__2))).#trainneg
cnf(i_0_822, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__2),np__3))).#trainneg
cnf(i_0_1004, plain, (k6_xcmplx_0(k4_xcmplx_0(np__3),k4_xcmplx_0(np__4))=np__1)).#trainneg
cnf(i_0_1006, plain, (k6_xcmplx_0(k4_xcmplx_0(np__4),k4_xcmplx_0(np__4))=np__0)).#trainneg
cnf(i_0_1114, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__1,np__2))=np__1)).#trainneg
cnf(i_0_1121, plain, (k3_xcmplx_0(np__3,k7_xcmplx_0(np__1,np__3))=np__1)).#trainneg
cnf(i_0_1124, plain, (k3_xcmplx_0(np__4,k7_xcmplx_0(np__1,np__2))=np__2)).#trainneg
cnf(i_0_1125, plain, (k3_xcmplx_0(np__4,k7_xcmplx_0(np__1,np__4))=np__1)).#trainneg
cnf(i_0_1098, plain, (k3_xcmplx_0(np__0,k7_xcmplx_0(np__1,np__2))=np__0)).#trainneg
cnf(i_0_1086, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=np__2)).#trainneg
cnf(i_0_130, plain, (v3_xxreal_2(X1)|~v5_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_356, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k5_xcmplx_0(X1)))).#trainneg
cnf(i_0_43, plain, (v1_int_1(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_170, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_823, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__2),np__4))).#trainneg
cnf(i_0_831, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__4),np__2))).#trainneg
cnf(i_0_830, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__4),np__1))).#trainneg
cnf(i_0_832, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__4),np__3))).#trainneg
cnf(i_0_833, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__4),np__4))).#trainneg
cnf(i_0_1087, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__4))=np__4)).#trainneg
cnf(i_0_1137, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__2)=np__1)).#trainneg
cnf(i_0_1139, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__4)=np__2)).#trainneg
cnf(i_0_1150, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__4)=np__1)).#trainneg
cnf(i_0_1155, plain, (k3_xcmplx_0(k7_xcmplx_0(np__3,np__2),np__2)=np__3)).#trainneg
cnf(i_0_516, plain, (v5_relat_1(esk14_2(X1,X2),X2))).#trainneg
cnf(i_0_624, plain, (v5_relat_1(esk45_2(X1,X2),X2))).#trainneg
cnf(i_0_517, plain, (v4_relat_1(esk14_2(X1,X2),X1))).#trainneg
cnf(i_0_625, plain, (v4_relat_1(esk45_2(X1,X2),X1))).#trainneg
cnf(i_0_598, plain, (v1_funct_2(esk38_1(X1),k5_numbers,X1))).#trainneg
cnf(i_0_423, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_140, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_294, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_4, plain, (v5_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_numbers)))).#trainneg
cnf(i_0_794, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_795, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_796, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_797, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_841, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2))).#trainneg
cnf(i_0_259, plain, (v1_xreal_0(k6_square_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_386, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_872, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_994, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1002, plain, (k6_xcmplx_0(k4_xcmplx_0(np__3),k4_xcmplx_0(np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1001, plain, (k6_xcmplx_0(k4_xcmplx_0(np__3),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1005, plain, (k6_xcmplx_0(k4_xcmplx_0(np__4),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_1008, plain, (k6_xcmplx_0(k4_xcmplx_0(np__5),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_840, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__1))).#trainneg
cnf(i_0_11, plain, (v4_xxreal_2(X1)|~v3_membered(X1)|~v2_xxreal_2(X1))).#trainneg
cnf(i_0_297, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_129, plain, (v4_xxreal_2(X1)|~v5_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_169, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9, plain, (v6_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_179, plain, (v4_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_839, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__0))).#trainneg
cnf(i_0_217, plain, (m1_subset_1(k1_fdiff_9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_214, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_529, plain, (m1_subset_1(esk15_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_654, plain, (k6_square_1(X1)=k7_square_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_397, plain, (v5_membered(k3_xboole_0(X2,X1))|~v5_membered(X1))).#trainneg
cnf(i_0_385, plain, (v5_membered(k3_xboole_0(X1,X2))|~v5_membered(X1))).#trainneg
cnf(i_0_255, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_461, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_382, plain, (v3_membered(k3_xboole_0(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_381, plain, (v3_membered(k3_xboole_0(X1,X2))|~v3_membered(X1))).#trainneg
cnf(i_0_399, plain, (v6_membered(k3_xboole_0(X2,X1))|~v6_membered(X1))).#trainneg
cnf(i_0_155, plain, (v5_xxreal_2(X1)|~v3_membered(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_371, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k5_xcmplx_0(X1)))).#trainneg
cnf(i_0_181, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_28, plain, (v6_xxreal_2(X1)|~v2_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_434, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_420, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_798, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_805, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_807, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_808, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_816, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_398, plain, (v6_membered(k3_xboole_0(X1,X2))|~v6_membered(X1))).#trainneg
cnf(i_0_462, plain, (v2_valued_0(k2_zfmisc_1(X2,X1))|~v2_membered(X1))).#trainneg
cnf(i_0_375, plain, (v1_membered(k3_xboole_0(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_373, plain, (v1_membered(k3_xboole_0(X1,X2))|~v1_membered(X1))).#trainneg
cnf(i_0_463, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_379, plain, (v2_membered(k3_xboole_0(X2,X1))|~v2_membered(X1))).#trainneg
cnf(i_0_377, plain, (v2_membered(k3_xboole_0(X1,X2))|~v2_membered(X1))).#trainneg
cnf(i_0_384, plain, (v4_membered(k3_xboole_0(X2,X1))|~v4_membered(X1))).#trainneg
cnf(i_0_383, plain, (v4_membered(k3_xboole_0(X1,X2))|~v4_membered(X1))).#trainneg
cnf(i_0_473, plain, (v4_valued_0(k2_zfmisc_1(X2,X1))|~v6_membered(X1))).#trainneg
cnf(i_0_127, plain, (v1_xboole_0(X1)|v2_xxreal_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_60, plain, (v5_valued_0(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_funct_1(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_50, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_159, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_817, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_818, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_827, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_837, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_838, plain, (r1_xxreal_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_1197, plain, (k4_xcmplx_0(k7_xcmplx_0(np__2,np__3))=k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))).#trainneg
cnf(i_0_1194, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1195, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__3))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))).#trainneg
cnf(i_0_1196, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_1198, plain, (k4_xcmplx_0(k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_1199, plain, (k4_xcmplx_0(k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_1200, plain, (k4_xcmplx_0(k7_xcmplx_0(np__4,np__3))=k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))).#trainneg
cnf(i_0_1130, plain, (k3_xcmplx_0(k4_xcmplx_0(np__2),k7_xcmplx_0(np__1,np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1133, plain, (k3_xcmplx_0(k4_xcmplx_0(np__4),k7_xcmplx_0(np__1,np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1140, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1156, plain, (k3_xcmplx_0(k7_xcmplx_0(np__3,np__2),k4_xcmplx_0(np__2))=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_70, plain, (v8_valued_0(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_funct_1(X1)|~v6_valued_0(X1))).#trainneg
cnf(i_0_107, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_124, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_95, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_499, plain, (r2_hidden(k32_sin_cos,k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_843, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_304, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_338, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_303, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_863, plain, (k2_xcmplx_0(np__2,k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(np__7,np__2))).#trainneg
cnf(i_0_856, plain, (k2_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1115, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__1,np__3))=k7_xcmplx_0(np__2,np__3))).#trainneg
cnf(i_0_1116, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1106, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1107, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1108, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1109, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1120, plain, (k3_xcmplx_0(np__3,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1122, plain, (k3_xcmplx_0(np__3,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_942, plain, (k6_xcmplx_0(np__2,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_76, plain, (v2_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_308, plain, (v1_zfmisc_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_zfmisc_1(k9_xtuple_0(X1)))).#trainneg
cnf(i_0_97, plain, (v4_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_316, plain, (~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xxreal_2(k3_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_335, plain, (~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xxreal_2(k4_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_337, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_367, plain, (v1_relat_1(k5_relat_1(k1_fdiff_9,k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))))).#trainneg
cnf(i_0_313, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))))).#trainneg
cnf(i_0_344, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))))).#trainneg
cnf(i_0_366, plain, (v2_funct_1(k5_relat_1(k1_fdiff_9,k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))))).#trainneg
cnf(i_0_943, plain, (k6_xcmplx_0(np__2,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__7,np__4))).#trainneg
cnf(i_0_944, plain, (k6_xcmplx_0(np__2,k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_945, plain, (k6_xcmplx_0(np__2,k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(np__5,np__4))).#trainneg
cnf(i_0_927, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__2,np__3))=k7_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_924, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_925, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__3))=k7_xcmplx_0(np__2,np__3))).#trainneg
cnf(i_0_926, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_929, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_956, plain, (k6_xcmplx_0(np__3,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__5,np__2))).#trainneg
cnf(i_0_957, plain, (k6_xcmplx_0(np__3,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__11,np__4))).#trainneg
cnf(i_0_964, plain, (k6_xcmplx_0(np__4,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__7,np__2))).#trainneg
cnf(i_0_116, plain, (v3_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_283, plain, (v3_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_zfmisc_1(k10_xtuple_0(X1)))).#trainneg
cnf(i_0_334, plain, (~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v2_xxreal_2(k4_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_1301, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_1302, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_312, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))))).#trainneg
cnf(i_0_343, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))))).#trainneg
cnf(i_0_965, plain, (k6_xcmplx_0(np__4,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__15,np__4))).#trainneg
cnf(i_0_966, plain, (k6_xcmplx_0(np__4,k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(np__13,np__4))).#trainneg
cnf(i_0_1088, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(np__4,np__3))).#trainneg
cnf(i_0_879, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__2)=k7_xcmplx_0(np__5,np__2))).#trainneg
cnf(i_0_878, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__1)=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1153, plain, (k3_xcmplx_0(k7_xcmplx_0(np__2,np__3),np__1)=k7_xcmplx_0(np__2,np__3))).#trainneg
cnf(i_0_1136, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__1)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1138, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__3)=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1146, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__3),np__2)=k7_xcmplx_0(np__2,np__3))).#trainneg
cnf(i_0_1145, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__3),np__1)=k7_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_1148, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__2)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1147, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__1)=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1149, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__3)=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1154, plain, (k3_xcmplx_0(k7_xcmplx_0(np__3,np__2),np__1)=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_175, plain, (v1_xxreal_2(X1)|v1_xboole_0(X1)|~v5_membered(X1)|~v3_xxreal_2(X1))).#trainneg
cnf(i_0_99, plain, (v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1157, plain, (k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),np__1)=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1033, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),np__1)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_229, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1009, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__0)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1023, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__0)=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1041, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),np__0)=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1204, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))=k7_xcmplx_0(np__2,np__3))).#trainneg
cnf(i_0_1201, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1202, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))=k7_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_1203, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1205, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1206, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1207, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=k7_xcmplx_0(np__4,np__3))).#trainneg
cnf(i_0_1117, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1126, plain, (k3_xcmplx_0(np__4,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1089, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1090, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k4_xcmplx_0(np__4))).#trainneg
cnf(i_0_1162, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__4)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_92, plain, (v1_xxreal_2(X1)|v1_xboole_0(X1)|~v1_finset_1(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_122, plain, (v3_funct_1(X1)|~v1_funct_1(X1)|~v7_valued_0(X1)|~v8_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_392, plain, (v3_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_452, plain, (v6_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_325, plain, (v1_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_1166, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),np__4)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1170, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),np__2)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_1173, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4),np__4)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_1179, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__2),np__2)=k4_xcmplx_0(np__5))).#trainneg
cnf(i_0_1181, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__4),np__4)=k4_xcmplx_0(np__5))).#trainneg
cnf(i_0_1131, plain, (k3_xcmplx_0(k4_xcmplx_0(np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__1)).#trainneg
cnf(i_0_1134, plain, (k3_xcmplx_0(k4_xcmplx_0(np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__2)).#trainneg
cnf(i_0_1135, plain, (k3_xcmplx_0(k4_xcmplx_0(np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=np__1)).#trainneg
cnf(i_0_1163, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k4_xcmplx_0(np__2))=np__1)).#trainneg
cnf(i_0_1171, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),k4_xcmplx_0(np__2))=np__3)).#trainneg
cnf(i_0_888, plain, (k2_xcmplx_0(k7_xcmplx_0(np__2,np__3),k7_xcmplx_0(np__1,np__3))=np__1)).#trainneg
cnf(i_0_881, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__2))=np__1)).#trainneg
cnf(i_0_883, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__3,np__2))=np__2)).#trainneg
cnf(i_0_890, plain, (k2_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,np__4))=np__1)).#trainneg
cnf(i_0_1159, plain, (k3_xcmplx_0(k7_xcmplx_0(np__4,np__3),k7_xcmplx_0(np__3,np__4))=np__1)).#trainneg
cnf(i_0_1016, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__2))=np__0)).#trainneg
cnf(i_0_164, plain, (v2_xxreal_2(X1)|v1_xboole_0(X1)|~v5_membered(X1)|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_368, plain, (v2_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_1029, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__1,np__4))=np__0)).#trainneg
cnf(i_0_1038, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(np__1,np__2))=np__1)).#trainneg
cnf(i_0_1039, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(np__3,np__2))=np__0)).#trainneg
cnf(i_0_1049, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__3,np__4))=np__0)).#trainneg
cnf(i_0_844, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)))).#trainneg
cnf(i_0_604, plain, (m1_subset_1(esk38_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_465, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k4_numbers)|~v5_membered(X1))).#trainneg
cnf(i_0_464, plain, (v5_relat_1(k2_zfmisc_1(X2,X1),k3_numbers)|~v4_membered(X1))).#trainneg
cnf(i_0_928, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_909, plain, (k6_xcmplx_0(np__0,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_910, plain, (k6_xcmplx_0(np__0,k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_911, plain, (k6_xcmplx_0(np__0,k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_912, plain, (k6_xcmplx_0(np__0,k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_1011, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__2)=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_91, plain, (v2_xxreal_2(X1)|v1_xboole_0(X1)|~v1_finset_1(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_228, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_353, plain, (v2_xxreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_359, plain, (v3_xxreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_1010, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1012, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__3)=k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))).#trainneg
cnf(i_0_1013, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__4)=k7_xcmplx_0(k4_xcmplx_0(np__7),np__2))).#trainneg
cnf(i_0_1014, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),np__5)=k7_xcmplx_0(k4_xcmplx_0(np__9),np__2))).#trainneg
cnf(i_0_1025, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__2)=k7_xcmplx_0(k4_xcmplx_0(np__7),np__4))).#trainneg
cnf(i_0_1024, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_1026, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__3)=k7_xcmplx_0(k4_xcmplx_0(np__11),np__4))).#trainneg
cnf(i_0_1027, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),np__4)=k7_xcmplx_0(k4_xcmplx_0(np__15),np__4))).#trainneg
cnf(i_0_1034, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),np__2)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1035, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),np__3)=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_1043, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),np__2)=k7_xcmplx_0(k4_xcmplx_0(np__5),np__4))).#trainneg
cnf(i_0_1042, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_1044, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),np__3)=k7_xcmplx_0(k4_xcmplx_0(np__9),np__4))).#trainneg
cnf(i_0_1045, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),np__4)=k7_xcmplx_0(k4_xcmplx_0(np__13),np__4))).#trainneg
cnf(i_0_874, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_891, plain, (k2_xcmplx_0(k7_xcmplx_0(np__5,np__2),k4_xcmplx_0(np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1015, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__1))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_65, plain, (v7_valued_0(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_funct_1(X1)|~v5_valued_0(X1))).#trainneg
cnf(i_0_433, plain, (v5_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_419, plain, (v4_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_86, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_1037, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),k4_xcmplx_0(np__2))=k7_xcmplx_0(np__7,np__2))).#trainneg
cnf(i_0_1036, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),k4_xcmplx_0(np__1))=k7_xcmplx_0(np__5,np__2))).#trainneg
cnf(i_0_1046, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),k4_xcmplx_0(np__1))=k7_xcmplx_0(np__7,np__4))).#trainneg
cnf(i_0_857, plain, (k2_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_946, plain, (k6_xcmplx_0(np__2,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__5,np__2))).#trainneg
cnf(i_0_947, plain, (k6_xcmplx_0(np__2,k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(np__9,np__4))).#trainneg
cnf(i_0_930, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_931, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(np__5,np__4))).#trainneg
cnf(i_0_932, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=k7_xcmplx_0(np__5,np__2))).#trainneg
cnf(i_0_913, plain, (k6_xcmplx_0(np__0,k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_893, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_892, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__1)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_900, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),np__2)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1018, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__3,np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_884, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__0)).#trainneg
cnf(i_0_59, plain, (v6_valued_0(X1)|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_funct_1(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_31, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_1307, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1294, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk52_2(X3,X2)))).#trainneg
cnf(i_0_168, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_1020, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__1)).#trainneg
cnf(i_0_1022, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=np__2)).#trainneg
cnf(i_0_1040, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__2)).#trainneg
cnf(i_0_1050, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=np__1)).#trainneg
cnf(i_0_895, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__1,np__2))=np__0)).#trainneg
cnf(i_0_896, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__3,np__2))=np__1)).#trainneg
cnf(i_0_520, plain, (m1_subset_1(esk14_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26, plain, (v1_rat_1(X2)|~v4_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_30, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_34, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_167, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_136, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1314, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_35, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_251, plain, (v1_relat_1(k4_rfunct_1(X1))|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1295, plain, (r2_hidden(esk52_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_328, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))))).#trainneg
cnf(i_0_348, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))))).#trainneg
cnf(i_0_327, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))))).#trainneg
cnf(i_0_347, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))))).#trainneg
cnf(i_0_877, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_873, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_998, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))).#trainneg
cnf(i_0_983, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_984, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__5),np__4))).#trainneg
cnf(i_0_985, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))).#trainneg
cnf(i_0_986, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__7),np__4))).#trainneg
cnf(i_0_880, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_889, plain, (k2_xcmplx_0(k7_xcmplx_0(np__3,np__2),k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1151, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_302, plain, (v5_membered(k2_tarski(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_358, plain, (v5_membered(k2_xboole_0(X1,X2))|~v5_membered(X2)|~v5_membered(X1))).#trainneg
cnf(i_0_273, plain, (v1_relat_1(k5_relat_1(k1_fdiff_9,k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))))).#trainneg
cnf(i_0_272, plain, (v2_funct_1(k5_relat_1(k1_fdiff_9,k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))))).#trainneg
cnf(i_0_1118, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1091, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))=k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))).#trainneg
cnf(i_0_901, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__2),np__2)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1160, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1165, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_1169, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_1172, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_1176, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))).#trainneg
cnf(i_0_1178, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__2),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))).#trainneg
cnf(i_0_1180, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__4),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__5),np__4))).#trainneg
cnf(i_0_1052, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))).#trainneg
cnf(i_0_1051, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_1063, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),np__1)=k7_xcmplx_0(k4_xcmplx_0(np__5),np__4))).#trainneg
cnf(i_0_989, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1306, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_57, plain, (v4_valued_0(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_250, plain, (v1_funct_1(k4_rfunct_1(X1))|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_400, plain, (v1_valued_0(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_991, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=k7_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_1167, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k4_xcmplx_0(np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1174, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4),k4_xcmplx_0(np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1054, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k4_xcmplx_0(np__2))=k7_xcmplx_0(np__3,np__2))).#trainneg
cnf(i_0_1053, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k4_xcmplx_0(np__1))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1064, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k4_xcmplx_0(np__1))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1072, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),k4_xcmplx_0(np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_882, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_886, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_887, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_1141, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1142, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__3,np__2))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1017, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1047, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1048, plain, (k6_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_885, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1310, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_121, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_415, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_292, plain, (v3_membered(k2_tarski(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_350, plain, (v3_membered(k2_xboole_0(X1,X2))|~v3_membered(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_309, plain, (v6_membered(k2_tarski(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1158, plain, (k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1055, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__1,np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1057, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__3,np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1067, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k7_xcmplx_0(np__3,np__4))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_1073, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),k7_xcmplx_0(np__1,np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1175, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4),k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=np__1)).#trainneg
cnf(i_0_1177, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3),k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))=np__1)).#trainneg
cnf(i_0_1059, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=np__0)).#trainneg
cnf(i_0_1061, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=np__1)).#trainneg
cnf(i_0_1069, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=np__0)).#trainneg
cnf(i_0_1075, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=np__0)).#trainneg
cnf(i_0_1077, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4),k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))=np__0)).#trainneg
cnf(i_0_1078, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__2),k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))=np__0)).#trainneg
cnf(i_0_1079, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__4),k7_xcmplx_0(k4_xcmplx_0(np__5),np__4))=np__0)).#trainneg
cnf(i_0_495, plain, (k9_subset_1(X2,X3,X3)=X3|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_64, plain, (v5_membered(X2)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_103, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_280, plain, (v1_zfmisc_1(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_funct_1(X1))).#trainneg
cnf(i_0_370, plain, (v6_membered(k2_xboole_0(X1,X2))|~v6_membered(X2)|~v6_membered(X1))).#trainneg
cnf(i_0_402, plain, (v2_valued_0(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_284, plain, (v1_membered(k2_tarski(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_999, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_1000, plain, (k6_xcmplx_0(k4_xcmplx_0(np__2),k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_987, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_5, plain, (v6_ordinal1(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_56, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_69, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_40, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_54, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_58, plain, (v4_membered(X2)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_178, plain, (v4_funct_1(X2)|~v4_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_342, plain, (v1_membered(k2_xboole_0(X1,X2))|~v1_membered(X2)|~v1_membered(X1))).#trainneg
cnf(i_0_369, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1278, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_443, plain, (v2_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_394, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_421, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_435, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_988, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_990, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_1007, plain, (k6_xcmplx_0(k4_xcmplx_0(np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(k4_xcmplx_0(np__7),np__2))).#trainneg
cnf(i_0_894, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))).#trainneg
cnf(i_0_1071, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),k4_xcmplx_0(np__1))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_404, plain, (v3_valued_0(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_288, plain, (v2_membered(k2_tarski(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_346, plain, (v2_membered(k2_xboole_0(X1,X2))|~v2_membered(X2)|~v2_membered(X1))).#trainneg
cnf(i_0_455, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_460, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_296, plain, (v4_membered(k2_tarski(X1,X2))|~v1_rat_1(X2)|~v1_rat_1(X1))).#trainneg
cnf(i_0_355, plain, (v4_membered(k2_xboole_0(X1,X2))|~v4_membered(X2)|~v4_membered(X1))).#trainneg
cnf(i_0_410, plain, (v4_valued_0(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_287, plain, (v6_xxreal_2(k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_138, plain, (v7_valued_0(X1)|~v1_funct_1(X1)|~v3_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_187, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_290, plain, (v6_xxreal_2(k3_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_299, plain, (v6_xxreal_2(k4_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_317, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1019, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_1028, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_1030, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_1021, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_1031, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__3,np__4))).#trainneg
cnf(i_0_361, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_412, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_466, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_319, plain, (v3_rcomp_1(k4_xxreal_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_362, plain, (v2_rcomp_1(k1_xxreal_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1032, plain, (k6_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_897, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_898, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_1074, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_82, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_83, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1292, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_189, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_209, plain, (a_2_3_rcomp_1(X1,X2)=k4_rcomp_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_331, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_352, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_330, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_351, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k8_real_1(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_1143, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_645, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_1144, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))).#trainneg
cnf(i_0_1152, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,np__4),k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))).#trainneg
cnf(i_0_1056, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_1058, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(np__3,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__5),np__4))).#trainneg
cnf(i_0_1065, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_1066, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k7_xcmplx_0(np__1,np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_210, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_48, plain, (v1_rfunct_2(X1)|~v1_funct_1(X1)|~v7_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_84, plain, (v1_rfunct_2(X1)|~v1_funct_1(X1)|~v8_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1164, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=k7_xcmplx_0(np__2,np__3))).#trainneg
cnf(i_0_1168, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))=k7_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_1062, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1068, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_1070, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_211, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_234, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_408, plain, (v5_relat_1(k5_relat_1(X1,X2),k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_657, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_491, plain, (esk8_3(X1,X2,X3)=X1|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~r2_hidden(X1,a_2_3_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_646, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_137, plain, (v8_valued_0(X1)|~v1_funct_1(X1)|~v3_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1060, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))).#trainneg
cnf(i_0_1076, plain, (k6_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))).#trainneg
cnf(i_0_196, plain, (k2_partfun2(k1_numbers,k1_numbers,k2_partfun1(k1_numbers,k1_numbers,k1_fdiff_9,k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))=k2_sincos10)).#trainneg
cnf(i_0_1298, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,k3_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_406, plain, (v5_relat_1(k5_relat_1(X1,X2),k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_264, plain, (m1_subset_1(k9_subset_1(X2,X3,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_474, plain, (v1_relat_1(k10_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k2_zfmisc_1(X3,X4)))))).#trainneg
cnf(i_0_469, plain, (v1_relat_1(k9_xtuple_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))))).#trainneg
cnf(i_0_6, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10, plain, (v2_valued_0(X2)|~v1_relat_1(X1)|~v2_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_14, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_27, plain, (v4_valued_0(X2)|~v1_relat_1(X1)|~v4_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_285, plain, (v1_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_658, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1287, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_293, plain, (v2_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_437, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_305, plain, (v3_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_438, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v2_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_490, plain, (~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~r2_hidden(X1,a_2_3_rcomp_1(X2,X3))|~r1_xxreal_0(esk8_3(X1,X2,X3),X2))).#trainneg
cnf(i_0_55, plain, (v5_relat_1(X2,k4_numbers)|~v5_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_41, plain, (v5_relat_1(X2,k3_numbers)|~v4_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_23, plain, (v5_relat_1(X2,k4_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_439, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_339, plain, (v4_valued_0(k3_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v4_valued_0(X1))).#trainneg
cnf(i_0_442, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v4_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_321, plain, (v2_funct_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1))).#trainneg
cnf(i_0_1268, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1213, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),k5_xcmplx_0(X2))=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_223, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_1289, plain, (r1_tarski(k4_xxreal_1(X1,X2),k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_440, plain, (v1_rat_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_441, plain, (v1_int_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_267, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_471, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_315, plain, (v2_xxreal_0(X2)|v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_275, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_345, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_380, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k7_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_18, plain, (v5_relat_1(X2,k3_numbers)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_418, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v7_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_365, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v5_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_391, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v6_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_417, plain, (v7_valued_0(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v7_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1267, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1212, plain, (k3_xcmplx_0(k5_xcmplx_0(X1),k5_xcmplx_0(X2))=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_332, plain, (v5_relat_1(k3_xboole_0(X1,X2),k4_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,k4_numbers))).#trainneg
cnf(i_0_314, plain, (v5_relat_1(k3_xboole_0(X1,X2),k3_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,k3_numbers))).#trainneg
cnf(i_0_432, plain, (v5_relat_1(k3_xboole_0(X1,X3),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_363, plain, (v4_relat_1(k3_xboole_0(X1,X3),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_198, plain, (X3=k2_tarski(X1,X2)|esk2_3(X1,X2,X3)!=X2|~r2_hidden(esk2_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_199, plain, (X3=k2_tarski(X1,X2)|esk2_3(X1,X2,X3)!=X1|~r2_hidden(esk2_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_289, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_286, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_310, plain, (v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k6_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_329, plain, (v2_xxreal_0(X2)|v3_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_306, plain, (v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k6_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_341, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k3_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_340, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_364, plain, (v5_valued_0(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v5_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_390, plain, (v6_valued_0(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v6_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_457, plain, (v1_valued_0(k2_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_470, plain, (v2_valued_0(k2_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v2_valued_0(X2)|~v2_valued_0(X1))).#trainneg
cnf(i_0_282, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_277, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_336, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_333, plain, (v3_xxreal_0(X2)|v3_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_475, plain, (v3_valued_0(k2_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X2)|~v3_valued_0(X1))).#trainneg
cnf(i_0_281, plain, (v4_valued_0(k2_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v4_valued_0(X2)|~v4_valued_0(X1))).#trainneg
cnf(i_0_311, plain, (v6_xxreal_2(k3_xboole_0(X1,X2))|~v2_membered(X2)|~v2_membered(X1)|~v6_xxreal_2(X2)|~v6_xxreal_2(X1))).#trainneg
cnf(i_0_429, plain, (v5_ordinal1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_ordinal1(X2)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_291, plain, (v4_funct_1(k2_tarski(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_476, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(k7_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_376, plain, (v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k7_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_416, plain, (v5_relat_1(k2_xboole_0(X1,X3),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v5_relat_1(X3,X2)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_477, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_374, plain, (v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(k7_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_349, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_378, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k7_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_271, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k5_relat_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k9_xtuple_0(X1))))).#trainneg
cnf(i_0_301, plain, (v1_relat_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v4_funct_1(X3)|~v5_relat_1(X1,X3))).#trainneg
cnf(i_0_300, plain, (v1_funct_1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v4_funct_1(X3)|~v5_relat_1(X1,X3))).#trainneg
cnf(i_0_1281, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|~v1_funct_1(X2)|~r1_xxreal_0(esk51_2(X1,X2),esk50_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_231, plain, (v1_funct_1(k2_partfun1(X1,X2,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_269, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk5_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_204, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(esk3_3(X1,X2,X3),X2))).#trainneg
cnf(i_0_205, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(esk3_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_430, plain, (v5_relat_1(k5_relat_1(X1,X2),k10_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_ordinal1(X2)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_1260, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1261, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_653, plain, (k6_rfunct_1(X3,X1,X2)=k4_rfunct_1(X2)|~v3_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_648, plain, (k2_partfun2(X2,X3,X1)=k2_funct_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_320, plain, (v4_relat_1(k2_xboole_0(X1,X3),X2)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v4_relat_1(X3,X2)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_487, plain, (m1_subset_1(esk7_3(X1,X2,X3),k1_numbers)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_492, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_numbers)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~r2_hidden(X1,a_2_3_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_258, plain, (v1_funct_1(k6_rfunct_1(X1,X2,X3))|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_233, plain, (v1_funct_1(k2_partfun2(X1,X2,X3))|~v1_funct_1(X3)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1285, plain, (m1_subset_1(esk50_2(X1,X2),k1_numbers)|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1284, plain, (m1_subset_1(esk51_2(X1,X2),k1_numbers)|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_489, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~r2_hidden(X1,a_2_3_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_1254, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_197, plain, (X3=k2_tarski(X1,X2)|esk2_3(X1,X2,X3)=X2|esk2_3(X1,X2,X3)=X1|r2_hidden(esk2_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_230, plain, (m1_subset_1(k2_partfun1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1296, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X1,k3_xxreal_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3))).#trainneg
cnf(i_0_257, plain, (m1_subset_1(k6_rfunct_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v3_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_232, plain, (m1_subset_1(k2_partfun2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v2_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_203, plain, (X3=k2_xboole_0(X1,X2)|r2_hidden(esk3_3(X1,X2,X3),X3)|r2_hidden(esk3_3(X1,X2,X3),X2)|r2_hidden(esk3_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_488, plain, (r1_xxreal_0(X1,X3)|r2_hidden(X2,a_2_3_rcomp_1(X3,X4))|X1!=X2|~v1_xxreal_0(X3)|~v1_xreal_0(X4)|~r1_xxreal_0(X1,X4)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1280, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|r1_xxreal_0(k1_seq_1(X1,esk51_2(X2,X1)),k1_seq_1(X1,esk50_2(X2,X1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1283, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|r2_hidden(esk50_2(X1,X2),k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1282, plain, (v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X2,X1))|r2_hidden(esk51_2(X1,X2),k9_subset_1(k1_numbers,X1,k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3372, plain, (esk40_0=k1_xboole_0)).#trainneg
cnf(i_0_3406, plain, (~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3416, plain, (~v1_xboole_0(esk48_0))).#trainneg
cnf(i_0_3417, plain, (~v1_xboole_0(esk44_0))).#trainneg
cnf(i_0_4213, plain, (k3_xboole_0(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)),k1_relset_1(k1_numbers,k2_sincos10))=k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))).#trainneg
cnf(i_0_4212, plain, (m1_subset_1(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10))))).#trainneg
cnf(i_0_3804, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_809, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,np__2),k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_3373, plain, (esk39_0=k1_xboole_0)).#trainneg
cnf(i_0_3374, plain, (esk17_0=k1_xboole_0)).#trainneg
cnf(i_0_3375, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_3420, plain, (~v1_xboole_0(esk46_0))).#trainneg
cnf(i_0_3423, plain, (~v1_xboole_0(esk30_0))).#trainneg
cnf(i_0_3424, plain, (~v1_xboole_0(esk29_0))).#trainneg
cnf(i_0_3737, plain, (k3_xboole_0(X1,X2)=X1|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_5985, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_6103, plain, (~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(X1,k4_rcomp_1(X2,X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3404, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).#trainneg
cnf(i_0_3446, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_842, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3437, plain, (~v6_membered(esk49_0))).#trainneg
cnf(i_0_3448, plain, (~v1_xboole_0(esk35_0))).#trainneg
cnf(i_0_3449, plain, (~v1_xboole_0(esk34_0))).#trainneg
cnf(i_0_3462, plain, (~v2_xxreal_0(esk35_0))).#trainneg
cnf(i_0_3463, plain, (~v2_xxreal_0(esk34_0))).#trainneg
cnf(i_0_3472, plain, (~v1_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_3552, plain, (~v2_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_3734, plain, (~v1_xboole_0(esk42_0))).#trainneg
cnf(i_0_4215, plain, (r1_tarski(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k1_real_1(np__1)),k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_4527, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_3403, negated_conjecture, (v1_xboole_0(k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_5984, plain, (r2_hidden(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3735, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3402, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_4113, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_3408, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_711, plain, (~r1_xxreal_0(np__1,k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_4108, plain, (v7_ordinal1(np__4))).#trainneg
cnf(i_0_4110, plain, (v7_ordinal1(np__3))).#trainneg
cnf(i_0_4117, plain, (v7_ordinal1(np__2))).#trainneg
cnf(i_0_727, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_3719, plain, (~v1_xboole_0(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_3717, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_4503, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_5998, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_4464, plain, (k9_subset_1(X1,X2,X2)=X2)).#trainneg
cnf(i_0_3370, negated_conjecture, (m1_subset_1(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).#trainneg
cnf(i_0_6398, plain, (r1_tarski(k2_rcomp_1(X1,X2),k1_xxreal_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_6397, plain, (r1_tarski(k4_xxreal_1(X1,X2),k1_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4180, plain, (v3_ordinal1(np__4))).#trainneg
cnf(i_0_3411, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_841, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),np__2))).#trainneg
cnf(i_0_4182, plain, (v3_ordinal1(np__3))).#trainneg
cnf(i_0_4185, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_4189, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_6092, plain, (esk8_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k4_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_6395, plain, (m1_subset_1(k4_xxreal_1(X1,X2),k1_zfmisc_1(k1_xxreal_1(X1,X2)))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_5761, plain, (v4_relat_1(k2_sincos10,k1_numbers))).#trainneg
cnf(i_0_3718, plain, (m1_subset_1(k32_sin_cos,k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_3438, plain, (v1_xboole_0(esk46_0)|~v6_membered(esk46_0))).#trainneg
cnf(i_0_840, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),np__1))).#trainneg
cnf(i_0_3431, plain, (~v1_xboole_0(k32_sin_cos)|~v1_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_3501, plain, (~v2_membered(k1_numbers)|~v5_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_4504, plain, (m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_6105, plain, (v1_xboole_0(k3_xxreal_1(X2,X3))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X1,k3_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_3972, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_6396, plain, (k3_xboole_0(k4_xxreal_1(X1,X2),k1_xxreal_1(X1,X2))=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4486, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3962, plain, (v1_xxreal_0(k32_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3720, plain, (r2_hidden(k32_sin_cos,k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_6090, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4628, plain, (v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_3850, plain, (v1_valued_0(k19_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_839, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),np__0))).#trainneg
cnf(i_0_3378, plain, (v3_membered(k3_numbers))).#trainneg
cnf(i_0_3379, plain, (v1_xcmplx_0(esk43_0))).#trainneg
cnf(i_0_3380, plain, (v1_xcmplx_0(esk37_0))).#trainneg
cnf(i_0_3385, plain, (v1_xxreal_0(esk43_0))).#trainneg
cnf(i_0_3386, plain, (v1_xxreal_0(esk37_0))).#trainneg
cnf(i_0_3387, plain, (v1_xreal_0(esk23_0))).#trainneg
cnf(i_0_3389, plain, (v1_xreal_0(esk43_0))).#trainneg
cnf(i_0_3536, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5880, plain, (~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_2(k4_rcomp_1(X2,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6115, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_3992, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_3971, plain, (v1_xxreal_0(k7_square_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5850, plain, (v3_xxreal_0(np__1)|~v3_xxreal_0(np__2)|~v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_4554, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3988, plain, (v1_xreal_0(np__13))).#trainneg
cnf(i_0_4635, plain, (v1_xcmplx_0(np__4)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3851, plain, (v1_valued_0(k1_fdiff_9)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3390, plain, (v1_xreal_0(esk37_0))).#trainneg
cnf(i_0_3391, plain, (v1_int_1(esk43_0))).#trainneg
cnf(i_0_3392, plain, (v1_int_1(esk37_0))).#trainneg
cnf(i_0_3393, plain, (v3_ordinal1(esk43_0))).#trainneg
cnf(i_0_3394, plain, (v3_ordinal1(esk37_0))).#trainneg
cnf(i_0_3398, plain, (v3_valued_0(esk26_0))).#trainneg
cnf(i_0_3399, plain, (v3_valued_0(esk16_0))).#trainneg
cnf(i_0_3400, plain, (v1_finset_1(esk36_0))).#trainneg
cnf(i_0_3414, plain, (v3_funct_1(esk15_0))).#trainneg
cnf(i_0_6075, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3956, plain, (v1_xxreal_0(np__0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3412, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3509, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5848, plain, (v3_xxreal_0(np__1)|~v3_xxreal_0(np__3)|~v1_xreal_0(np__3)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_3564, plain, (v2_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3852, plain, (v1_valued_0(k2_sincos10)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3397, plain, (k1_seq_1(k2_sincos10,k4_xcmplx_0(k7_square_1(np__2)))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_3647, plain, (v5_valued_0(esk15_0))).#trainneg
cnf(i_0_3833, plain, (v1_relat_1(k19_sin_cos))).#trainneg
cnf(i_0_3834, plain, (v1_relat_1(k1_fdiff_9))).#trainneg
cnf(i_0_3980, plain, (v1_xreal_0(np__15))).#trainneg
cnf(i_0_3982, plain, (v1_xreal_0(np__9))).#trainneg
cnf(i_0_3984, plain, (v1_xreal_0(np__7))).#trainneg
cnf(i_0_3986, plain, (v1_xreal_0(np__5))).#trainneg
cnf(i_0_3990, plain, (v1_xreal_0(np__11))).#trainneg
cnf(i_0_4101, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_3965, plain, (v1_xxreal_0(np__3)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3507, plain, (v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3589, plain, (v2_membered(k3_xboole_0(X1,X2))|~v3_membered(X1))).#trainneg
cnf(i_0_3865, plain, (v1_relat_1(esk10_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4089, plain, (v7_ordinal1(np__15))).#trainneg
cnf(i_0_4091, plain, (v7_ordinal1(np__9))).#trainneg
cnf(i_0_4093, plain, (v7_ordinal1(np__7))).#trainneg
cnf(i_0_4095, plain, (v7_ordinal1(np__5))).#trainneg
cnf(i_0_4097, plain, (v7_ordinal1(np__13))).#trainneg
cnf(i_0_4099, plain, (v7_ordinal1(np__11))).#trainneg
cnf(i_0_4161, plain, (v3_ordinal1(np__15))).#trainneg
cnf(i_0_4163, plain, (v3_ordinal1(np__9))).#trainneg
cnf(i_0_3658, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5857, plain, (v3_xxreal_0(np__2)|~v3_xxreal_0(np__3)|~v1_xreal_0(np__3)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_6118, plain, (k9_subset_1(X1,X3,X2)=k3_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3867, plain, (v1_relat_1(np__15)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3410, plain, (v1_int_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4165, plain, (v3_ordinal1(np__7))).#trainneg
cnf(i_0_4167, plain, (v3_ordinal1(np__5))).#trainneg
cnf(i_0_4169, plain, (v3_ordinal1(np__13))).#trainneg
cnf(i_0_4171, plain, (v3_ordinal1(np__11))).#trainneg
cnf(i_0_4173, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_4229, plain, (v3_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_4230, plain, (v3_valued_0(k1_fdiff_9))).#trainneg
cnf(i_0_3474, plain, (v1_xcmplx_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_5899, plain, (~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v2_xxreal_2(k2_rcomp_1(X2,X1)))).#trainneg
cnf(i_0_3963, plain, (v1_xxreal_0(np__4)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3572, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3584, plain, (v2_membered(k3_xboole_0(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_3868, plain, (v1_relat_1(np__15)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3413, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4988, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_4997, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3799, plain, (v1_xxreal_2(esk31_0)|~v2_membered(esk31_0))).#trainneg
cnf(i_0_3813, plain, (v2_xxreal_2(esk31_0)|~v2_membered(esk31_0))).#trainneg
cnf(i_0_3869, plain, (v1_relat_1(np__9)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3870, plain, (v1_relat_1(np__9)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3871, plain, (v1_relat_1(np__7)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3508, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_9395, plain, (k3_xboole_0(X2,X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3804, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_3431, plain, (~v1_xboole_0(k32_sin_cos))).#trainneg
cnf(i_0_3374, plain, (esk17_0=np__0)).#trainneg
cnf(i_0_3372, plain, (esk40_0=np__0)).#trainneg
cnf(i_0_9522, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_3373, plain, (esk39_0=np__0)).#trainneg
cnf(i_0_9525, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_9300, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9453, plain, (~m1_subset_1(esk49_0,k5_numbers))).#trainneg
cnf(i_0_9770, plain, (~m1_subset_1(esk46_0,k5_numbers))).#trainneg
cnf(i_0_9302, plain, (~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_9455, plain, (~m1_subset_1(esk49_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_9447, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_9446, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_9126, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X1,k4_rcomp_1(X2,X3))|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_488, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X1,a_2_3_rcomp_1(X2,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9556, plain, (v1_int_1(np__4))).#trainneg
cnf(i_0_9559, plain, (v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_9772, plain, (~m1_subset_1(esk46_0,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_9570, plain, (~v1_xboole_0(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_9676, plain, (esk8_3(k32_sin_cos,k10_real_1(k32_sin_cos,np__2),k32_sin_cos)=k32_sin_cos|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_9369, plain, (m1_subset_1(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10))))).#trainneg
cnf(i_0_9378, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_9561, plain, (v1_int_1(np__3))).#trainneg
cnf(i_0_9564, plain, (v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_9571, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_9573, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_9577, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_9844, plain, (r1_xxreal_0(k32_sin_cos,k32_sin_cos))).#trainneg
cnf(i_0_9648, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_9850, plain, (~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_8841, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k4_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_9566, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_9569, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_10018, plain, (v1_valued_0(k2_sincos10))).#trainneg
cnf(i_0_8858, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X2,X3))|X4!=k1_xxreal_1(X2,X3)|~r2_hidden(X1,X4)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_8859, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r2_hidden(X1,k4_xxreal_1(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3))).#trainneg
cnf(i_0_9418, plain, (X1=X2|X1=X3|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_9593, plain, (v1_xxreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5340, plain, (v7_valued_0(k2_sincos10)|~v3_funct_1(k2_sincos10))).#trainneg
cnf(i_0_9681, plain, (k1_real_1(np__0)=np__0)).#trainneg
cnf(i_0_9552, plain, (v3_membered(k4_numbers))).#trainneg
cnf(i_0_9882, plain, (v1_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_9887, plain, (v2_membered(k3_numbers))).#trainneg
cnf(i_0_9997, plain, (v1_valued_0(k1_fdiff_9))).#trainneg
cnf(i_0_9756, plain, (k4_xcmplx_0(k4_xcmplx_0(k32_sin_cos))=k32_sin_cos)).#trainneg
cnf(i_0_10057, plain, (m1_subset_1(k7_xcmplx_0(X1,k32_sin_cos),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1291, plain, (k3_xboole_0(X1,np__0)=np__0)).#trainneg
cnf(i_0_10005, plain, (v2_xxreal_2(esk36_0))).#trainneg
cnf(i_0_10067, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_10083, plain, (v1_int_1(np__15))).#trainneg
cnf(i_0_10085, plain, (v1_xxreal_0(np__15))).#trainneg
cnf(i_0_10086, plain, (v1_xcmplx_0(np__15))).#trainneg
cnf(i_0_10189, plain, (v2_membered(X1)|~v3_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9849, plain, (esk8_3(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)=k32_sin_cos|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_9422, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_9420, negated_conjecture, (m1_subset_1(esk1_0,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_9821, plain, (k3_xboole_0(k4_xxreal_1(X1,X2),k1_rcomp_1(X1,X2))=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9437, plain, (r1_tarski(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_3872, plain, (v1_relat_1(np__7)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1276, plain, (k2_xboole_0(X1,np__0)=X1)).#trainneg
cnf(i_0_10088, plain, (v1_int_1(np__9))).#trainneg
cnf(i_0_10090, plain, (v1_xxreal_0(np__9))).#trainneg
cnf(i_0_10091, plain, (v1_xcmplx_0(np__9))).#trainneg
cnf(i_0_10093, plain, (v1_int_1(np__7))).#trainneg
cnf(i_0_10095, plain, (v1_xxreal_0(np__7))).#trainneg
cnf(i_0_10096, plain, (v1_xcmplx_0(np__7))).#trainneg
cnf(i_0_9639, plain, (k3_xboole_0(k2_rcomp_1(X1,X2),k1_xxreal_1(X1,X2))=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9368, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).#trainneg
cnf(i_0_9461, negated_conjecture, (v1_xboole_0(k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_10319, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_9334, plain, (~r2_hidden(k2_xboole_0(X2,X3),X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_9622, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,k1_relset_1(k1_numbers,X3)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,X3))|~r2_hidden(X2,k1_relset_1(k1_numbers,X3)))).#trainneg
cnf(i_0_9744, plain, (m1_subset_1(k32_sin_cos,k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_3873, plain, (v1_relat_1(np__5)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3440, plain, (v1_xcmplx_0(k4_xcmplx_0(esk23_0)))).#trainneg
cnf(i_0_10098, plain, (v1_int_1(np__5))).#trainneg
cnf(i_0_10100, plain, (v1_xxreal_0(np__5))).#trainneg
cnf(i_0_10101, plain, (v1_xcmplx_0(np__5))).#trainneg
cnf(i_0_10103, plain, (v1_int_1(np__13))).#trainneg
cnf(i_0_10105, plain, (v1_xxreal_0(np__13))).#trainneg
cnf(i_0_10106, plain, (v1_xcmplx_0(np__13))).#trainneg
cnf(i_0_9768, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_10184, plain, (k5_xcmplx_0(k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_9376, plain, (m1_subset_1(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_10064, plain, (m1_subset_1(k7_xcmplx_0(X1,np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9122, plain, (m1_subset_1(X1,k3_xxreal_1(X2,X3))|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9370, plain, (m1_subset_1(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k1_real_1(np__1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_9660, negated_conjecture, (esk6_3(esk1_0,k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))=esk1_0|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_9640, plain, (r1_tarski(k2_rcomp_1(X1,X2),k1_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3874, plain, (v1_relat_1(np__5)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3441, plain, (v1_xcmplx_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_10108, plain, (v1_int_1(np__11))).#trainneg
cnf(i_0_10110, plain, (v1_xxreal_0(np__11))).#trainneg
cnf(i_0_10111, plain, (v1_xcmplx_0(np__11))).#trainneg
cnf(i_0_9425, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k1_real_1(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_3875, plain, (v1_relat_1(np__13)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3876, plain, (v1_relat_1(np__13)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_9375, plain, (~r2_hidden(k2_xboole_0(X2,X3),X1)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_10185, plain, (k5_xcmplx_0(np__1)=np__1|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_10063, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_square_1(X2)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10059, plain, (m1_subset_1(k7_xcmplx_0(X1,np__4),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10136, plain, (m1_subset_1(k3_xboole_0(X3,X2),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3877, plain, (v1_relat_1(np__11)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_8815, plain, (k4_rfunct_1(k19_sin_cos)=k1_fdiff_9)).#trainneg
cnf(i_0_3885, plain, (v1_relat_1(k32_sin_cos)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3886, plain, (v1_relat_1(np__4)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3878, plain, (v1_relat_1(np__11)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3879, plain, (v1_relat_1(np__0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3880, plain, (v1_relat_1(np__0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3887, plain, (v1_relat_1(np__4)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3888, plain, (v1_relat_1(np__3)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_9617, negated_conjecture, (~m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers)|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_9335, plain, (m1_subset_1(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_10225, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10297, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_9658, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k1_real_1(np__1))=esk1_0|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_3889, plain, (v1_relat_1(np__3)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3891, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3892, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3895, plain, (v1_relat_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3896, plain, (v1_relat_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_10163, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9659, negated_conjecture, (esk6_3(esk1_0,k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|~v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_9976, plain, (r1_xxreal_0(k7_square_1(X1),k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9649, plain, (r2_hidden(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_8827, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k4_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_8640, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_10010, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9351, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).#trainneg
cnf(i_0_3926, plain, (v1_xcmplx_0(k32_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_9984, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k1_real_1(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_9377, plain, (~v1_xboole_0(k2_xboole_0(X1,X2))|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_9939, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_9726, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|~r2_hidden(X1,k4_xxreal_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_9793, plain, (v1_xboole_0(k4_rcomp_1(X1,X2))|~r1_xxreal_0(X3,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~m1_subset_1(X3,k4_rcomp_1(X1,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10262, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),k1_numbers))).#trainneg
cnf(i_0_9519, negated_conjecture, (~m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers)|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_8618, plain, (k2_partfun1(k1_numbers,k1_numbers,k1_fdiff_9,X1)=k5_relat_1(k1_fdiff_9,X1))).#trainneg
cnf(i_0_9983, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_6892, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,X2))=k6_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10173, plain, (v1_xreal_0(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9784, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v3_membered(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_9275, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k9_subset_1(k1_numbers,k1_relset_1(k1_numbers,X3),X4))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6939, plain, (k5_xcmplx_0(k3_xcmplx_0(X1,X2))=k7_xcmplx_0(k5_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11576, plain, (v1_xcmplx_0(k32_sin_cos))).#trainneg
cnf(i_0_10783, plain, (v2_membered(k4_numbers))).#trainneg
cnf(i_0_11574, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_10604, plain, (~m1_subset_1(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_10609, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_196, plain, (k2_partfun2(k1_numbers,k1_numbers,k5_relat_1(k1_fdiff_9,k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))=k2_sincos10)).#trainneg
cnf(i_0_10510, plain, (k7_xcmplx_0(np__4,np__1)=np__4)).#trainneg
cnf(i_0_10601, plain, (k7_xcmplx_0(np__3,np__1)=np__3)).#trainneg
cnf(i_0_10705, plain, (k7_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_10473, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10475, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2))|~v1_xboole_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10684, plain, (~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10515, plain, (m1_subset_1(k5_relat_1(k2_sincos10,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10819, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_11085, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X2,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_11876, plain, (k5_xcmplx_0(np__4)=k7_xcmplx_0(np__1,np__4))).#trainneg
cnf(i_0_10686, plain, (~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xboole_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10474, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k5_numbers))).#trainneg
cnf(i_0_10685, plain, (~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k7_xcmplx_0(k32_sin_cos,np__2),k5_numbers))).#trainneg
cnf(i_0_10815, plain, (~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k7_xcmplx_0(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_10514, plain, (v1_funct_1(k5_relat_1(k2_sincos10,X1)))).#trainneg
cnf(i_0_10818, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_8567, plain, (v3_xxreal_0(np__4)|~v3_xxreal_0(np__1)|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_11877, plain, (k5_xcmplx_0(np__3)=k7_xcmplx_0(np__1,np__3))).#trainneg
cnf(i_0_11879, plain, (k5_xcmplx_0(np__2)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_10513, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_11650, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k1_fdiff_9,X1),k1_seq_1(k1_fdiff_9,X2))|~v5_valued_0(k5_relat_1(k1_fdiff_9,k1_relset_1(k1_numbers,k1_fdiff_9)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_fdiff_9))|~r2_hidden(X2,k1_relset_1(k1_numbers,k1_fdiff_9)))).#trainneg
cnf(i_0_11647, plain, (m1_subset_1(k5_relat_1(k1_fdiff_9,X1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10636, plain, (r2_hidden(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_10420, plain, (r2_hidden(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_8576, plain, (v3_xxreal_0(np__2)|~v3_xxreal_0(np__1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_4142, plain, (v1_funct_1(k32_sin_cos)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4143, plain, (v1_funct_1(np__4)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4144, plain, (v1_funct_1(np__4)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_10646, plain, (v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10503, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X1,k4_rcomp_1(X2,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10392, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10389, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4145, plain, (v1_funct_1(np__3)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4146, plain, (v1_funct_1(np__3)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_4148, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4014, plain, (v1_rat_1(esk10_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4016, plain, (v1_rat_1(np__15)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_10391, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10594, plain, (r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k4_xxreal_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4122, plain, (v1_funct_1(esk10_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4149, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_4152, plain, (v1_funct_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4153, plain, (v1_funct_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_4017, plain, (v1_rat_1(np__15)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4018, plain, (v1_rat_1(np__9)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4019, plain, (v1_rat_1(np__9)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4020, plain, (v1_rat_1(np__7)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4021, plain, (v1_rat_1(np__7)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_10768, plain, (v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_11646, plain, (v1_funct_1(k5_relat_1(k1_fdiff_9,X1)))).#trainneg
cnf(i_0_10696, plain, (r1_xxreal_0(esk8_3(k32_sin_cos,k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10694, plain, (r1_xxreal_0(esk8_3(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10631, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|X4!=k1_xxreal_1(X3,X2)|~r2_hidden(X1,X4)|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3))).#trainneg
cnf(i_0_4124, plain, (v1_funct_1(np__15)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4022, plain, (v1_rat_1(np__5)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4023, plain, (v1_rat_1(np__5)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4024, plain, (v1_rat_1(np__13)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4025, plain, (v1_rat_1(np__13)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4026, plain, (v1_rat_1(np__11)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_10717, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k2_tarski(X3,X4)))).#trainneg
cnf(i_0_11057, negated_conjecture, (m1_subset_1(esk1_0,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_10477, plain, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_11770, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10491, plain, (esk8_3(X1,X2,X3)=X1|r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_11495, plain, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_10553, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_4125, plain, (v1_funct_1(np__15)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3465, plain, (k7_xcmplx_0(esk34_0,np__1)=esk34_0)).#trainneg
cnf(i_0_4027, plain, (v1_rat_1(np__11)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4028, plain, (v1_rat_1(np__0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4029, plain, (v1_rat_1(np__0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4034, plain, (v1_rat_1(k32_sin_cos)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4035, plain, (v1_rat_1(np__4)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_11509, plain, (m1_subset_1(esk8_3(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_11058, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).#trainneg
cnf(i_0_10698, plain, (r1_xxreal_0(esk8_3(X1,X2,X3),X3)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_11511, plain, (m1_subset_1(esk8_3(k32_sin_cos,k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10710, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r2_hidden(X1,k2_tarski(X4,X3)))).#trainneg
cnf(i_0_4126, plain, (v1_funct_1(np__9)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3466, plain, (k7_xcmplx_0(esk29_0,np__1)=esk29_0)).#trainneg
cnf(i_0_4036, plain, (v1_rat_1(np__4)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4037, plain, (v1_rat_1(np__3)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4038, plain, (v1_rat_1(np__3)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_4040, plain, (v1_rat_1(np__1)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4041, plain, (v1_rat_1(np__1)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_10575, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10490, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(X1,X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_11048, plain, (k3_xboole_0(k2_rcomp_1(X1,X2),k1_rcomp_1(X1,X2))=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10434, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_10805, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4127, plain, (v1_funct_1(np__9)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3467, plain, (k7_xcmplx_0(esk28_0,np__1)=esk28_0)).#trainneg
cnf(i_0_4044, plain, (v1_rat_1(np__2)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4045, plain, (v1_rat_1(np__2)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_11419, plain, (v1_xxreal_0(X1)|~v2_membered(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_10709, plain, (r2_hidden(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10413, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|X4!=k1_xxreal_1(X2,X3)|~r2_hidden(X1,X4)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_10695, plain, (r1_xxreal_0(esk8_3(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_4128, plain, (v1_funct_1(np__7)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3468, plain, (k7_xcmplx_0(esk18_0,np__1)=esk18_0)).#trainneg
cnf(i_0_5754, plain, (v5_relat_1(k2_sincos10,k1_numbers))).#trainneg
cnf(i_0_10505, plain, (v1_xcmplx_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_11399, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k3_xboole_0(X4,k1_relset_1(k1_numbers,X3)))|~r2_hidden(X2,k3_xboole_0(X4,k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_11152, plain, (v1_xxreal_0(X1)|~v2_membered(k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_11496, plain, (r1_xxreal_0(esk8_3(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10570, plain, (m1_subset_1(k6_square_1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10430, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|v1_xboole_0(k2_xboole_0(X3,X2))|~m1_subset_1(X1,k2_xboole_0(X3,X2)))).#trainneg
cnf(i_0_11773, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k3_xboole_0(X4,k1_relset_1(k1_numbers,X3)))|~r2_hidden(X2,k3_xboole_0(X4,k1_relset_1(k1_numbers,X3))))).#trainneg
cnf(i_0_11506, plain, (m1_subset_1(esk8_3(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_4129, plain, (v1_funct_1(np__7)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_3469, plain, (k7_xcmplx_0(esk10_0,np__1)=esk10_0)).#trainneg
cnf(i_0_10596, plain, (v1_xcmplx_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_10700, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_10798, plain, (v1_xxreal_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11483, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_11510, plain, (m1_subset_1(esk8_3(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_11508, plain, (m1_subset_1(esk8_3(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10650, plain, (m1_subset_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_4130, plain, (v1_funct_1(np__5)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3470, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12075, plain, (~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_12028, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_12107, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_12122, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_11947, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_11951, plain, (~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_11911, plain, (r2_hidden(X1,k1_xxreal_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_12910, plain, (v1_xxreal_0(X1)|~v2_membered(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_10184, plain, (k5_xcmplx_0(k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_11900, plain, (k3_xcmplx_0(np__1,k32_sin_cos)=k32_sin_cos)).#trainneg
cnf(i_0_11901, plain, (k7_xcmplx_0(k32_sin_cos,np__1)=k32_sin_cos)).#trainneg
cnf(i_0_12425, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k1_real_1(np__1))))).#trainneg
cnf(i_0_11997, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X2,k1_relset_1(k1_numbers,k2_sincos10))|~v5_valued_0(k2_sincos10))).#trainneg
cnf(i_0_12421, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_relset_1(k1_numbers,X3))))).#trainneg
cnf(i_0_11894, plain, (v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k1_real_1(np__1))))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_3471, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11584, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,np__0))).#trainneg
cnf(i_0_4070, plain, (v1_int_1(k32_sin_cos)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_13225, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_11980, plain, (k2_partfun1(k1_numbers,k1_numbers,k5_relat_1(k2_sincos10,X1),X2)=k5_relat_1(k5_relat_1(k2_sincos10,X1),X2)|~v1_funct_1(k5_relat_1(k2_sincos10,X1)))).#trainneg
cnf(i_0_3473, plain, (k7_xcmplx_0(k4_xcmplx_0(X1),np__1)=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12070, plain, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12209, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13057, plain, (v1_xxreal_0(X1)|~v2_membered(X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_13115, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,X4)|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(k1_relset_1(k1_numbers,X3))))).#trainneg
cnf(i_0_12862, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_12309, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12350, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13116, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,X3))|~r2_hidden(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_4131, plain, (v1_funct_1(np__5)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_12337, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1)|r1_xxreal_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(X1)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12419, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k3_xboole_0(X4,k1_relset_1(k1_numbers,X3)))|~r2_hidden(X2,k3_xboole_0(X4,k1_relset_1(k1_numbers,X3))))).#trainneg
cnf(i_0_12349, plain, (r1_xxreal_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_4132, plain, (v1_funct_1(np__13)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4178, plain, (v3_ordinal1(k32_sin_cos)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_4207, plain, (v2_valued_0(k2_sincos10)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4133, plain, (v1_funct_1(np__13)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_4134, plain, (v1_funct_1(np__11)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4135, plain, (v1_funct_1(np__11)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_4136, plain, (v1_funct_1(np__0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4137, plain, (v1_funct_1(np__0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_12413, plain, (r1_xxreal_0(esk8_3(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_5839, plain, (v3_xxreal_0(np__3)|~v3_xxreal_0(np__4)|~v1_xreal_0(np__4)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_9553, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_4158, plain, (v3_ordinal1(esk10_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_12935, plain, (r2_hidden(X1,X2)|r1_xxreal_0(X1,X3)|r1_xxreal_0(X4,X1)|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12420, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,X3))|~r2_hidden(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_12244, plain, (v1_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12184, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_12324, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_8449, plain, (v3_xxreal_0(np__1)|~v3_xxreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_4205, plain, (v2_valued_0(k19_sin_cos)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_4206, plain, (v2_valued_0(k1_fdiff_9)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_13298, plain, (m1_subset_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_numbers)|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_13084, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12008, negated_conjecture, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_12937, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_13114, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(X3,X1),k1_seq_1(X3,X2))|~v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k3_xboole_0(k1_relset_1(k1_numbers,X3),X4))|~r2_hidden(X2,k3_xboole_0(k1_relset_1(k1_numbers,X3),X4)))).#trainneg
cnf(i_0_12279, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k4_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_8456, plain, (v3_xxreal_0(np__2)|~v3_xxreal_0(np__4)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_3481, plain, (k3_xcmplx_0(np__1,esk34_0)=esk34_0)).#trainneg
cnf(i_0_12007, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_13295, plain, (m1_subset_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_11988, plain, (v2_valued_0(k5_relat_1(k2_sincos10,X1)))).#trainneg
cnf(i_0_3482, plain, (k3_xcmplx_0(np__1,esk29_0)=esk29_0)).#trainneg
cnf(i_0_4472, plain, (v5_membered(esk13_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4476, plain, (v5_membered(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_13263, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12630, plain, (m1_subset_1(k7_xcmplx_0(np__1,np__2),k1_numbers)|~v1_xreal_0(k4_xcmplx_0(np__2))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_13090, plain, (m1_subset_1(k7_xcmplx_0(np__1,np__2),k1_numbers))).#trainneg
cnf(i_0_13294, plain, (m1_subset_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12895, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_3483, plain, (k3_xcmplx_0(np__1,esk28_0)=esk28_0)).#trainneg
cnf(i_0_10855, plain, (v4_membered(np__0)|~v4_membered(X1))).#trainneg
cnf(i_0_4514, plain, (v6_ordinal1(esk13_0)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_4518, plain, (v6_ordinal1(k5_numbers)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_4545, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_13085, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12381, plain, (m1_subset_1(k7_xcmplx_0(X1,k32_sin_cos),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12922, plain, (v1_xxreal_0(X1)|~v2_membered(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_11989, plain, (v3_valued_0(k5_relat_1(k2_sincos10,X1)))).#trainneg
cnf(i_0_12126, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_12211, plain, (v1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12844, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_13264, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12132, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k1_fdiff_9,X1),k1_seq_1(k1_fdiff_9,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_fdiff_9))|~r2_hidden(X2,k1_relset_1(k1_numbers,k1_fdiff_9))|~v5_valued_0(k1_fdiff_9))).#trainneg
cnf(i_0_10857, plain, (v2_membered(np__0)|~v2_membered(X1))).#trainneg
cnf(i_0_3484, plain, (k3_xcmplx_0(np__1,esk18_0)=esk18_0)).#trainneg
cnf(i_0_10859, plain, (v1_membered(np__0)|~v1_membered(X1))).#trainneg
cnf(i_0_10861, plain, (v6_membered(np__0)|~v6_membered(X1))).#trainneg
cnf(i_0_10863, plain, (v3_membered(np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_4559, plain, (v4_membered(esk13_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4563, plain, (v4_membered(k5_numbers)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4568, plain, (v4_funct_1(esk13_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_4572, plain, (v4_funct_1(k5_numbers)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_12473, plain, (v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12383, plain, (m1_subset_1(k7_xcmplx_0(X1,np__4),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12389, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12582, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12182, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_12453, plain, (v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10865, plain, (v5_membered(np__0)|~v5_membered(X1))).#trainneg
cnf(i_0_3485, plain, (k3_xcmplx_0(np__1,esk10_0)=esk10_0)).#trainneg
cnf(i_0_12227, plain, (esk8_3(X1,X2,X3)=X1|r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12351, plain, (v1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12643, plain, (m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_3486, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5335, plain, (v7_valued_0(esk15_0)|~v3_funct_1(esk15_0))).#trainneg
cnf(i_0_5338, plain, (v7_valued_0(k19_sin_cos)|~v3_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_5339, plain, (v7_valued_0(k1_fdiff_9)|~v3_funct_1(k1_fdiff_9))).#trainneg
cnf(i_0_12519, plain, (m1_subset_1(esk8_3(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12160, plain, (k2_partfun1(k1_numbers,k1_numbers,k5_relat_1(k1_fdiff_9,X1),X2)=k5_relat_1(k5_relat_1(k1_fdiff_9,X1),X2)|~v1_funct_1(k5_relat_1(k1_fdiff_9,X1)))).#trainneg
cnf(i_0_12175, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_12325, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,k2_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_12412, plain, (r1_xxreal_0(esk8_3(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_8450, plain, (v3_xxreal_0(np__1)|v3_xxreal_0(np__2)|~v3_xxreal_0(np__3)|~v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_3487, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13563, negated_conjecture, (v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_13832, plain, (v3_membered(esk32_0))).#trainneg
cnf(i_0_13833, plain, (v3_membered(esk31_0))).#trainneg
cnf(i_0_13834, plain, (v3_membered(esk24_0))).#trainneg
cnf(i_0_13836, plain, (v3_membered(esk11_0))).#trainneg
cnf(i_0_14133, plain, (v4_membered(np__0))).#trainneg
cnf(i_0_14362, plain, (v2_membered(np__0))).#trainneg
cnf(i_0_14396, plain, (v1_membered(np__0))).#trainneg
cnf(i_0_14406, plain, (v6_membered(np__0))).#trainneg
cnf(i_0_13706, plain, (esk6_3(X1,X1,X2)=X1|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13846, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|r2_hidden(X2,k1_xxreal_1(X3,X1))|~r1_xxreal_0(X3,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_13578, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_13567, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_13558, negated_conjecture, (v1_funct_1(esk1_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_9659, negated_conjecture, (esk6_3(esk1_0,k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_14421, plain, (v3_membered(np__0))).#trainneg
cnf(i_0_14657, plain, (v5_membered(np__0))).#trainneg
cnf(i_0_14274, plain, (k7_xcmplx_0(np__4,np__2)=np__2)).#trainneg
cnf(i_0_10909, negated_conjecture, (~r2_hidden(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),esk1_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_13946, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_13384, plain, (v5_valued_0(k5_relat_1(k2_sincos10,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_13808, plain, (r1_xxreal_0(esk6_3(X1,X2,X1),X1)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_13515, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v5_valued_0(k5_relat_1(k2_sincos10,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,X3)|~r2_hidden(X2,X3))).#trainneg
cnf(i_0_3489, plain, (k3_xcmplx_0(np__1,X1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14276, plain, (k7_xcmplx_0(np__3,np__3)=np__1)).#trainneg
cnf(i_0_13989, plain, (m1_subset_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13743, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v5_valued_0(k5_relat_1(k2_sincos10,X3))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(X3))|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X2,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_13992, plain, (m1_subset_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_13844, plain, (v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13565, negated_conjecture, (v1_relat_1(esk1_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3490, plain, (k7_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13479, plain, (k5_xcmplx_0(k32_sin_cos)=k7_xcmplx_0(np__1,k32_sin_cos))).#trainneg
cnf(i_0_13712, plain, (v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_13557, negated_conjecture, (v3_ordinal1(esk1_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_13675, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13693, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v5_valued_0(k5_relat_1(k2_sincos10,X3))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,X3)|~r2_hidden(X2,X3))).#trainneg
cnf(i_0_13849, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v5_valued_0(k5_relat_1(k2_sincos10,X3))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X2,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_14847, plain, (r2_hidden(X1,X2)|r2_hidden(X1,k2_tarski(X4,X3))|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_12215, plain, (v1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3513, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_3502, plain, (k3_xcmplx_0(np__1,k5_xcmplx_0(X1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13713, plain, (v1_xxreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_13560, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_13561, negated_conjecture, (v1_rat_1(esk1_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_13732, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13784, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_13764, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1)|r1_xxreal_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(X1)|~v1_xreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13669, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_12344, plain, (v1_xreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3503, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),np__1)=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9891, plain, (k3_xcmplx_0(np__1,esk43_0)=esk43_0)).#trainneg
cnf(i_0_9892, plain, (k7_xcmplx_0(esk43_0,np__1)=esk43_0)).#trainneg
cnf(i_0_9898, plain, (k3_xcmplx_0(np__1,esk37_0)=esk37_0)).#trainneg
cnf(i_0_14821, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_14359, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_13853, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_14829, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,k4_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_12201, plain, (r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3504, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_9899, plain, (k7_xcmplx_0(esk37_0,np__1)=esk37_0)).#trainneg
cnf(i_0_10146, plain, (k7_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_10883, plain, (k3_xcmplx_0(np__1,np__15)=np__15)).#trainneg
cnf(i_0_13608, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_13509, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(np__1)),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14043, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14044, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_14358, plain, (k3_xcmplx_0(X1,X2)=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14206, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_3602, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3505, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_10884, plain, (k7_xcmplx_0(np__15,np__1)=np__15)).#trainneg
cnf(i_0_11011, plain, (k3_xcmplx_0(np__1,np__9)=np__9)).#trainneg
cnf(i_0_11012, plain, (k7_xcmplx_0(np__9,np__1)=np__9)).#trainneg
cnf(i_0_3864, plain, (v1_relat_1(esk13_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3905, plain, (v1_xcmplx_0(esk13_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3941, plain, (v1_xxreal_0(esk13_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_14467, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_14047, plain, (v1_xreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_14711, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),np__2),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_14828, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_13655, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k2_tarski(X3,X1)))).#trainneg
cnf(i_0_12194, plain, (v3_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v3_xxreal_0(k32_sin_cos)|~v1_xreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_3506, plain, (v3_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11022, plain, (k3_xcmplx_0(np__1,np__7)=np__7)).#trainneg
cnf(i_0_15253, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_9519, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_9617, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_14988, plain, (v2_membered(esk32_0))).#trainneg
cnf(i_0_14990, plain, (v2_membered(esk31_0))).#trainneg
cnf(i_0_14992, plain, (v2_membered(esk24_0))).#trainneg
cnf(i_0_14994, plain, (v2_membered(esk11_0))).#trainneg
cnf(i_0_15055, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_11058, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_14920, negated_conjecture, (k7_xcmplx_0(esk1_0,np__1)=esk1_0)).#trainneg
cnf(i_0_14929, plain, (k1_real_1(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_3977, plain, (v1_xreal_0(esk13_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4013, plain, (v1_rat_1(esk13_0)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4049, plain, (v1_int_1(esk13_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4085, plain, (v7_ordinal1(esk13_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15351, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14941, plain, (~r2_hidden(k1_rcomp_1(X1,X2),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15114, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_hidden(k4_xcmplx_0(np__1),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_15116, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_real_1(k7_square_1(np__2)),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_15084, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_15090, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_14961, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_15057, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_15198, plain, (v2_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9658, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_11023, plain, (k7_xcmplx_0(np__7,np__1)=np__7)).#trainneg
cnf(i_0_4121, plain, (v1_funct_1(esk13_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4157, plain, (v3_ordinal1(esk13_0)|~v3_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_14917, negated_conjecture, (m1_subset_1(k7_xcmplx_0(X1,esk1_0),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15094, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X1,k2_tarski(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_14916, negated_conjecture, (m1_subset_1(k7_xcmplx_0(X1,esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15410, negated_conjecture, (esk7_3(esk1_0,k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|r2_hidden(esk1_0,k2_tarski(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_16356, plain, (v3_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v3_xxreal_0(k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_9370, plain, (m1_subset_1(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_11113, plain, (k7_xcmplx_0(np__5,np__1)=np__5)).#trainneg
cnf(i_0_11123, plain, (k3_xcmplx_0(np__1,np__13)=np__13)).#trainneg
cnf(i_0_11124, plain, (k7_xcmplx_0(np__13,np__1)=np__13)).#trainneg
cnf(i_0_15337, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X3,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15899, plain, (k7_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15050, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_15574, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_10014, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11234, plain, (k3_xcmplx_0(np__1,np__11)=np__11)).#trainneg
cnf(i_0_11235, plain, (k7_xcmplx_0(np__11,np__1)=np__11)).#trainneg
cnf(i_0_15317, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(k4_xcmplx_0(np__1),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_16171, plain, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos),k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13508, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(np__1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15091, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_13523, plain, (k7_xcmplx_0(esk23_0,np__1)=esk23_0)).#trainneg
cnf(i_0_10832, plain, (v1_int_1(k32_sin_cos)|~v1_int_1(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_14550, plain, (m1_subset_1(k7_xcmplx_0(k7_square_1(X1),X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15282, negated_conjecture, (k3_xcmplx_0(np__1,esk1_0)=esk1_0|~v1_xreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_14535, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15355, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|r1_xxreal_0(X3,X1)|r1_xxreal_0(X1,X2)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13679, plain, (r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3510, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_13524, plain, (k7_xcmplx_0(esk19_0,np__1)=esk19_0)).#trainneg
cnf(i_0_14549, plain, (m1_subset_1(k7_xcmplx_0(k1_real_1(X1),X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14544, plain, (m1_subset_1(k7_xcmplx_0(k32_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12058, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|X4!=k1_xxreal_1(X3,X2)|~r2_hidden(X1,X4)|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3))).#trainneg
cnf(i_0_14925, negated_conjecture, (m1_subset_1(k7_xcmplx_0(esk1_0,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15686, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12302, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_3511, plain, (k3_xcmplx_0(np__1,k5_xcmplx_0(X1))=k5_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15277, negated_conjecture, (v1_int_1(esk1_0)|~v1_int_1(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_16264, plain, (v1_relat_1(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_13258, plain, (k3_xcmplx_0(k32_sin_cos,X1)=k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12385, plain, (m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16312, plain, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),np__2),k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13098, plain, (v1_xxreal_0(k6_square_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14200, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_13610, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_3512, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),np__1)=k5_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_16265, plain, (v1_xcmplx_0(esk13_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3555, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_4470, plain, (v5_membered(esk15_0)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4474, plain, (v5_membered(k19_sin_cos)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4475, plain, (v5_membered(k1_fdiff_9)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_14719, plain, (m1_subset_1(k7_xcmplx_0(k1_real_1(X1),np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13924, plain, (m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13854, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12249, plain, (m1_subset_1(k7_xcmplx_0(np__1,np__4),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_14232, plain, (r1_xxreal_0(X1,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_13399, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_3514, plain, (v2_xxreal_0(k4_xcmplx_0(np__5))|~v3_xxreal_0(np__5)|~v1_xreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_14253, plain, (k7_xcmplx_0(np__0,np__2)=np__0)).#trainneg
cnf(i_0_4477, plain, (v5_membered(k2_sincos10)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4512, plain, (v6_ordinal1(esk15_0)|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4516, plain, (v6_ordinal1(k19_sin_cos)|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4517, plain, (v6_ordinal1(k1_fdiff_9)|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4519, plain, (v6_ordinal1(k2_sincos10)|~v6_ordinal1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_14214, plain, (m1_subset_1(k7_xcmplx_0(k7_square_1(X1),k32_sin_cos),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13091, plain, (m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13885, plain, (m1_subset_1(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_14213, plain, (m1_subset_1(k7_xcmplx_0(k1_real_1(X1),k32_sin_cos),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14720, plain, (m1_subset_1(k7_xcmplx_0(k7_square_1(X1),np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12497, plain, (esk7_3(X1,X2,X3)=X1|r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9733, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4521, plain, (v3_membered(esk15_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4525, plain, (v3_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4526, plain, (v3_membered(k1_fdiff_9)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4528, plain, (v3_membered(k2_sincos10)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4530, plain, (v6_membered(esk15_0)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_14974, plain, (~r2_hidden(k1_rcomp_1(X3,k2_xboole_0(X1,X2)),X1)|~r1_xxreal_0(X3,k2_xboole_0(X1,X2))|~v1_xreal_0(k2_xboole_0(X1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_14062, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_12629, plain, (m1_subset_1(k7_xcmplx_0(np__1,np__4),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_13079, plain, (v1_xxreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9795, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4534, plain, (v6_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4535, plain, (v6_membered(k1_fdiff_9)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4537, plain, (v6_membered(k2_sincos10)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4539, plain, (v1_membered(esk15_0)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4543, plain, (v1_membered(k19_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_12696, plain, (m1_subset_1(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_12433, plain, (r2_hidden(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_14986, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9857, plain, (m1_subset_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_10768, plain, (v1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_17661, plain, (r1_xxreal_0(np__1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,np__1))).#trainneg
cnf(i_0_13764, plain, (r1_xxreal_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_17356, plain, (~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_3813, plain, (v2_xxreal_2(esk31_0))).#trainneg
cnf(i_0_12070, plain, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_10818, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_16485, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16781, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16486, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17788, plain, (r1_xxreal_0(np__2,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,np__2))).#trainneg
cnf(i_0_3799, plain, (v1_xxreal_2(esk31_0))).#trainneg
cnf(i_0_16412, negated_conjecture, (k3_xcmplx_0(np__1,esk1_0)=esk1_0)).#trainneg
cnf(i_0_17221, plain, (k3_xcmplx_0(k32_sin_cos,np__1)=k32_sin_cos)).#trainneg
cnf(i_0_9939, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_16375, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos)))).#trainneg
cnf(i_0_17494, plain, (~v1_xboole_0(k4_xcmplx_0(np__5))|~v1_xxreal_0(k4_xcmplx_0(np__5))|~v3_xxreal_0(np__5)|~v1_xreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_16477, plain, (v1_xreal_0(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9849, plain, (esk8_3(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)=k32_sin_cos)).#trainneg
cnf(i_0_16579, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r2_hidden(X1,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_16519, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~r2_hidden(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_17749, plain, (v1_xxreal_0(X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16883, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(k4_xcmplx_0(np__1),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_17846, negated_conjecture, (r1_xxreal_0(esk1_0,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,esk1_0))).#trainneg
cnf(i_0_14467, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_16478, plain, (v1_rat_1(esk13_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16480, plain, (v1_int_1(esk13_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16483, plain, (v7_ordinal1(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16590, plain, (v1_xxreal_0(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16606, plain, (v1_funct_1(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16448, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X1,k2_tarski(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_16433, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16998, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16782, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_rcomp_1(X3,X2)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17847, negated_conjecture, (r1_xxreal_0(esk1_0,np__4)|r1_xxreal_0(np__4,esk1_0))).#trainneg
cnf(i_0_6753, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_13732, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_3661, plain, (v5_relat_1(esk16_0,k4_numbers))).#trainneg
cnf(i_0_16607, plain, (v3_ordinal1(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16952, plain, (v1_int_1(k32_sin_cos)|~v1_xboole_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_17185, negated_conjecture, (v1_int_1(esk1_0)|~v1_xboole_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_17076, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17722, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k2_tarski(X1,X3)))).#trainneg
cnf(i_0_11496, plain, (r1_xxreal_0(esk8_3(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_16648, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k1_real_1(X1),esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17848, negated_conjecture, (r1_xxreal_0(esk1_0,np__3)|r1_xxreal_0(np__3,esk1_0))).#trainneg
cnf(i_0_6109, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk13_0))).#trainneg
cnf(i_0_12412, plain, (r1_xxreal_0(esk8_3(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_3665, plain, (v5_relat_1(esk16_0,k3_numbers))).#trainneg
cnf(i_0_5752, plain, (v5_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_5753, plain, (v5_relat_1(k1_fdiff_9,k1_numbers))).#trainneg
cnf(i_0_17355, plain, (v1_xcmplx_0(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4544, plain, (v1_membered(k1_fdiff_9)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4546, plain, (v1_membered(k2_sincos10)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4548, plain, (v2_membered(esk15_0)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4552, plain, (v2_membered(k19_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17148, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_11506, plain, (m1_subset_1(esk8_3(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_16647, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk1_0),k1_numbers))).#trainneg
cnf(i_0_17302, plain, (r1_xxreal_0(esk6_3(X1,X1,X2),esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11508, plain, (m1_subset_1(esk8_3(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17849, negated_conjecture, (r1_xxreal_0(esk1_0,np__1)|r1_xxreal_0(np__1,esk1_0))).#trainneg
cnf(i_0_11158, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_12351, plain, (v1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_5759, plain, (v4_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_5760, plain, (v4_relat_1(k1_fdiff_9,k1_numbers))).#trainneg
cnf(i_0_9579, plain, (r1_xxreal_0(esk35_0,esk35_0))).#trainneg
cnf(i_0_4553, plain, (v2_membered(k1_fdiff_9)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4555, plain, (v2_membered(k2_sincos10)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4557, plain, (v4_membered(esk15_0)|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4561, plain, (v4_membered(k19_sin_cos)|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4562, plain, (v4_membered(k1_fdiff_9)|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_16649, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k7_square_1(X1),esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16632, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k1_numbers))).#trainneg
cnf(i_0_14537, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12321, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_17852, negated_conjecture, (r1_xxreal_0(esk1_0,np__2)|r1_xxreal_0(np__2,esk1_0))).#trainneg
cnf(i_0_9473, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12350, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_9580, plain, (r1_xxreal_0(esk34_0,esk34_0))).#trainneg
cnf(i_0_9581, plain, (r1_xxreal_0(esk30_0,esk30_0))).#trainneg
cnf(i_0_9582, plain, (r1_xxreal_0(esk29_0,esk29_0))).#trainneg
cnf(i_0_4564, plain, (v4_membered(k2_sincos10)|~v4_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4566, plain, (v4_funct_1(esk15_0)|~v4_funct_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4570, plain, (v4_funct_1(k19_sin_cos)|~v4_funct_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4571, plain, (v4_funct_1(k1_fdiff_9)|~v4_funct_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4573, plain, (v4_funct_1(k2_sincos10)|~v4_funct_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_12673, plain, (m1_subset_1(k4_xcmplx_0(np__4),k1_numbers)|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_16673, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(X1,esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14705, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),np__2),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12921, plain, (v1_xreal_0(X1)|~v3_membered(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_14201, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),k32_sin_cos),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17659, plain, (r1_xxreal_0(np__1,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),np__1)|~v1_xreal_0(k10_real_1(np__3,np__4))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_9696, plain, (k1_real_1(np__5)=k4_xcmplx_0(np__5)|~m1_subset_1(k4_xcmplx_0(np__5),k1_numbers))).#trainneg
cnf(i_0_9583, plain, (r1_xxreal_0(esk20_0,esk20_0))).#trainneg
cnf(i_0_9584, plain, (r1_xxreal_0(esk10_0,esk10_0))).#trainneg
cnf(i_0_9903, plain, (r1_xxreal_0(esk43_0,esk43_0))).#trainneg
cnf(i_0_16537, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_real_1(k7_square_1(np__2)),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_12261, plain, (m1_subset_1(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_14976, plain, (~r2_hidden(k1_rcomp_1(X2,X1),X1)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_12259, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_15775, plain, (esk6_3(X1,X2,X3)=X1|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_12387, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_square_1(X2)),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_16674, negated_conjecture, (v1_xxreal_0(k7_xcmplx_0(X1,esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17786, plain, (r1_xxreal_0(np__2,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),np__2)|~v1_xreal_0(k10_real_1(np__3,np__4))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_9697, plain, (k1_real_1(np__4)=k4_xcmplx_0(np__4)|~m1_subset_1(k4_xcmplx_0(np__4),k1_numbers))).#trainneg
cnf(i_0_9905, plain, (r1_xxreal_0(esk37_0,esk37_0))).#trainneg
cnf(i_0_10002, plain, (v1_xcmplx_0(k4_xcmplx_0(esk43_0)))).#trainneg
cnf(i_0_10004, plain, (v1_xcmplx_0(k4_xcmplx_0(esk37_0)))).#trainneg
cnf(i_0_12376, plain, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(X2)),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14212, plain, (m1_subset_1(k7_xcmplx_0(np__1,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_13396, plain, (~r2_hidden(k1_xxreal_1(X3,k2_xboole_0(X1,X2)),X1)|~r1_xxreal_0(X3,k2_xboole_0(X1,X2))|~v1_xxreal_0(k2_xboole_0(X1,X2))|~v1_xxreal_0(X3))).#trainneg
cnf(i_0_9835, plain, (r1_xxreal_0(k32_sin_cos,X1)|r1_xxreal_0(X1,k32_sin_cos)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9698, plain, (k1_real_1(np__3)=k4_xcmplx_0(np__3)|~m1_subset_1(k4_xcmplx_0(np__3),k1_numbers))).#trainneg
cnf(i_0_10555, plain, (m1_subset_1(esk43_0,k1_numbers))).#trainneg
cnf(i_0_10556, plain, (m1_subset_1(esk37_0,k1_numbers))).#trainneg
cnf(i_0_10557, plain, (m1_subset_1(esk34_0,k1_numbers))).#trainneg
cnf(i_0_12610, plain, (v1_xxreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13408, plain, (m1_subset_1(X1,k1_xxreal_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_12883, plain, (m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_16457, plain, (m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_relset_1(k1_numbers,k2_sincos10))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_13824, plain, (r1_xxreal_0(esk8_3(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_14536, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9836, plain, (r1_xxreal_0(np__4,X1)|r1_xxreal_0(X1,np__4)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9700, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2)|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_12008, negated_conjecture, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_10558, plain, (m1_subset_1(esk29_0,k1_numbers))).#trainneg
cnf(i_0_17954, plain, (~r1_xxreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_17955, plain, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_17956, plain, (~r1_xxreal_0(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_9446, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_19016, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))))).#trainneg
cnf(i_0_18996, plain, (r1_xxreal_0(np__4,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,np__4))).#trainneg
cnf(i_0_19021, plain, (v1_funct_1(esk43_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_18784, plain, (k1_real_1(k4_xcmplx_0(np__4))=np__4)).#trainneg
cnf(i_0_19339, plain, (k1_real_1(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_18014, negated_conjecture, (k5_xcmplx_0(esk1_0)=k7_xcmplx_0(np__1,esk1_0))).#trainneg
cnf(i_0_18040, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k32_sin_cos)))).#trainneg
cnf(i_0_18147, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_18997, plain, (r1_xxreal_0(np__3,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,np__3))).#trainneg
cnf(i_0_19028, plain, (v1_relat_1(esk43_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_18508, plain, (v1_xreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_18509, plain, (v1_xxreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_18044, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k32_sin_cos)))).#trainneg
cnf(i_0_19015, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))))).#trainneg
cnf(i_0_19019, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))))).#trainneg
cnf(i_0_10477, plain, (~r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_18878, negated_conjecture, (v1_xreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_19067, negated_conjecture, (v1_xxreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_12007, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_19034, plain, (v1_funct_1(esk37_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_12086, negated_conjecture, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))))).#trainneg
cnf(i_0_18876, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_19024, plain, (v1_rat_1(esk43_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_19037, plain, (v1_rat_1(esk37_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_19041, plain, (v1_relat_1(esk37_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19046, plain, (v3_ordinal1(esk34_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_19047, plain, (v1_funct_1(esk34_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_18150, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_18093, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_18180, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k2_tarski(X3,X1)))).#trainneg
cnf(i_0_18731, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_19054, plain, (v1_relat_1(esk34_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19352, plain, (v1_funct_1(esk29_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19359, plain, (v1_relat_1(esk29_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19049, plain, (v1_int_1(esk34_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_19050, plain, (v1_rat_1(esk34_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_19002, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18294, negated_conjecture, (v1_xxreal_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_18991, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18458, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_18888, negated_conjecture, (k7_xcmplx_0(esk1_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_19000, plain, (r1_xxreal_0(k7_square_1(X1),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10836, plain, (k1_real_1(k32_sin_cos)=k4_xcmplx_0(k32_sin_cos)|~m1_subset_1(k4_xcmplx_0(k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_10559, plain, (m1_subset_1(esk23_0,k1_numbers))).#trainneg
cnf(i_0_10560, plain, (m1_subset_1(esk19_0,k1_numbers))).#trainneg
cnf(i_0_10877, plain, (v1_xcmplx_0(k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_19351, plain, (v3_ordinal1(esk29_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_19354, plain, (v1_int_1(esk29_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_19355, plain, (v1_rat_1(esk29_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_17961, plain, (~r1_xxreal_0(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_17972, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18999, plain, (r1_xxreal_0(k1_real_1(X1),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18984, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos))).#trainneg
cnf(i_0_18211, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k1_real_1(X1),esk1_0),k7_xcmplx_0(k1_real_1(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18408, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3526, plain, (k4_xcmplx_0(k4_xcmplx_0(esk34_0))=esk34_0)).#trainneg
cnf(i_0_10879, plain, (r1_xxreal_0(np__15,np__15))).#trainneg
cnf(i_0_11005, plain, (v1_xcmplx_0(k4_xcmplx_0(np__9)))).#trainneg
cnf(i_0_11016, plain, (v1_xcmplx_0(k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_17962, plain, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_18126, plain, (v1_rat_1(esk13_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_18129, plain, (v1_int_1(esk13_0)|~v1_xboole_0(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_18983, plain, (r1_xxreal_0(k5_xcmplx_0(X1),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_18293, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_19005, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16627, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),esk1_0),k1_numbers))).#trainneg
cnf(i_0_14548, plain, (m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14208, plain, (m1_subset_1(k7_xcmplx_0(k32_sin_cos,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_14509, plain, (m1_subset_1(k7_xcmplx_0(k1_real_1(X1),np__4),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3563, plain, (v1_finset_1(X1)|~v6_membered(X1)|~v2_membered(X1)|~v5_xxreal_2(X1))).#trainneg
cnf(i_0_3527, plain, (k4_xcmplx_0(k4_xcmplx_0(esk29_0))=esk29_0)).#trainneg
cnf(i_0_11106, plain, (v1_xcmplx_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_11117, plain, (v1_xcmplx_0(k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_11228, plain, (v1_xcmplx_0(k4_xcmplx_0(np__11)))).#trainneg
cnf(i_0_13341, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16928, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0),k1_numbers))).#trainneg
cnf(i_0_14551, plain, (m1_subset_1(k7_xcmplx_0(np__2,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18142, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|esk7_3(esk1_0,k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_12386, plain, (m1_subset_1(k7_xcmplx_0(X1,k1_real_1(X2)),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_18105, negated_conjecture, (v3_xxreal_0(esk1_0)|r1_xxreal_0(k32_sin_cos,esk1_0)|~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_3568, plain, (v1_xboole_0(X1)|~v1_xboole_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3528, plain, (k4_xcmplx_0(k4_xcmplx_0(esk28_0))=esk28_0)).#trainneg
cnf(i_0_10888, plain, (v2_membered(esk15_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_10899, plain, (v2_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_14924, negated_conjecture, (m1_subset_1(k7_xcmplx_0(esk1_0,np__2),k1_numbers))).#trainneg
cnf(i_0_12223, plain, (m1_subset_1(X1,k4_rcomp_1(X2,X3))|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18604, negated_conjecture, (m1_subset_1(k7_xcmplx_0(esk1_0,k32_sin_cos),k1_numbers)|~v1_xreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_13825, plain, (r1_xxreal_0(esk8_3(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_11916, plain, (k2_partfun2(k1_numbers,k1_numbers,k5_relat_1(k1_fdiff_9,k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))=k2_sincos10)).#trainneg
cnf(i_0_12101, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X2)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_17720, plain, (r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_16393, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(k32_sin_cos,np__2)),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5763, plain, (v3_xxreal_0(np__9)|~v3_xxreal_0(np__13)|~v1_xreal_0(np__13)|~v1_xreal_0(np__9))).#trainneg
cnf(i_0_3529, plain, (k4_xcmplx_0(k4_xcmplx_0(esk18_0))=esk18_0)).#trainneg
cnf(i_0_11592, plain, (r1_xxreal_0(esk23_0,esk23_0))).#trainneg
cnf(i_0_11593, plain, (r1_xxreal_0(esk19_0,esk19_0))).#trainneg
cnf(i_0_10900, plain, (v2_membered(k1_fdiff_9)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_10903, plain, (v2_membered(k2_sincos10)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_11741, plain, (v1_xreal_0(esk28_0)|~v1_xreal_0(k7_xcmplx_0(np__1,esk28_0)))).#trainneg
cnf(i_0_11742, plain, (v1_xreal_0(esk18_0)|~v1_xreal_0(k7_xcmplx_0(np__1,esk18_0)))).#trainneg
cnf(i_0_12026, plain, (v3_xxreal_0(np__4)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_17891, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0),k1_numbers))).#trainneg
cnf(i_0_17877, plain, (m1_subset_1(k7_xcmplx_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17952, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1)|r1_xxreal_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19515, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)))))).#trainneg
cnf(i_0_20770, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_3530, plain, (k4_xcmplx_0(k4_xcmplx_0(esk10_0))=esk10_0)).#trainneg
cnf(i_0_19768, plain, (v1_xxreal_0(esk23_0))).#trainneg
cnf(i_0_19781, plain, (v1_xxreal_0(esk19_0))).#trainneg
cnf(i_0_20466, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,esk43_0)))).#trainneg
cnf(i_0_20468, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,esk34_0)))).#trainneg
cnf(i_0_20469, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,esk29_0)))).#trainneg
cnf(i_0_20470, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,esk28_0)))).#trainneg
cnf(i_0_20474, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__15)))).#trainneg
cnf(i_0_19396, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_3531, plain, (k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))=k5_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_20475, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__9)))).#trainneg
cnf(i_0_20476, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__7)))).#trainneg
cnf(i_0_20477, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__5)))).#trainneg
cnf(i_0_20478, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__13)))).#trainneg
cnf(i_0_20479, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,np__11)))).#trainneg
cnf(i_0_19384, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))))).#trainneg
cnf(i_0_20594, plain, (r1_xxreal_0(esk8_3(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_3532, plain, (k4_xcmplx_0(k4_xcmplx_0(k5_xcmplx_0(X1)))=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19370, plain, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_20434, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~v1_xxreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(esk1_0,a_2_1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_19527, plain, (~r1_xxreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_19649, plain, (r1_xxreal_0(k32_sin_cos,k10_real_1(X1,X2))|r1_xxreal_0(k10_real_1(X1,X2),k32_sin_cos)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3533, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19763, plain, (v1_funct_1(esk23_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19769, plain, (v1_xcmplx_0(esk23_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_19770, plain, (v1_relat_1(esk23_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19374, plain, (~r1_xxreal_0(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_19521, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))))).#trainneg
cnf(i_0_19886, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(X1,esk1_0),k7_xcmplx_0(X1,esk1_0))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19514, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))))).#trainneg
cnf(i_0_20813, plain, (r1_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_19776, plain, (v1_funct_1(esk19_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3534, plain, (k4_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19782, plain, (v1_xcmplx_0(esk19_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_19783, plain, (v1_relat_1(esk19_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19364, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_19501, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k32_sin_cos)))).#trainneg
cnf(i_0_19528, plain, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_19529, plain, (~r1_xxreal_0(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_19592, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_19510, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))))).#trainneg
cnf(i_0_19930, plain, (v1_int_1(esk34_0)|~v1_int_1(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_3539, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_11869, plain, (k5_xcmplx_0(np__5)=k7_xcmplx_0(np__1,np__5))).#trainneg
cnf(i_0_11870, plain, (k5_xcmplx_0(np__0)=k7_xcmplx_0(np__1,np__0))).#trainneg
cnf(i_0_20109, plain, (r1_xxreal_0(k32_sin_cos,k8_real_1(X1,X2))|r1_xxreal_0(k8_real_1(X1,X2),k32_sin_cos)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_20592, plain, (r1_xxreal_0(esk8_3(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_19635, plain, (r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(X1,X2))|r1_xxreal_0(k7_xcmplx_0(X1,X2),k32_sin_cos)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_20212, plain, (v1_int_1(esk29_0)|~v1_int_1(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_3540, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_19762, plain, (v3_ordinal1(esk23_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_19766, plain, (v1_rat_1(esk23_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_20358, plain, (r2_hidden(X1,X2)|r2_hidden(X1,k2_tarski(X4,X3))|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X4)))).#trainneg
cnf(i_0_20004, plain, (r1_xxreal_0(k32_sin_cos,X1)|r1_xxreal_0(X1,k32_sin_cos)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19381, plain, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_20108, plain, (r1_xxreal_0(k32_sin_cos,k8_real_1(X1,X2))|r1_xxreal_0(k8_real_1(X1,X2),k32_sin_cos)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19399, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_19604, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_19841, plain, (X1=X2|X1=X3|m1_subset_1(X1,k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20527, plain, (v1_int_1(esk28_0)|~v1_int_1(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_3547, plain, (v2_xxreal_0(k5_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13975, plain, (k5_xcmplx_0(esk34_0)=k7_xcmplx_0(np__1,esk34_0))).#trainneg
cnf(i_0_19775, plain, (v3_ordinal1(esk19_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_19778, plain, (v1_int_1(esk19_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_19779, plain, (v1_rat_1(esk19_0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_19380, plain, (~r1_xxreal_0(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_19379, plain, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_19605, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_19641, plain, (r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|r1_xxreal_0(k4_xcmplx_0(X1),k32_sin_cos)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20528, plain, (v1_xreal_0(esk28_0)|~v1_xreal_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_3548, plain, (v1_xboole_0(k5_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14003, plain, (k5_xcmplx_0(esk29_0)=k7_xcmplx_0(np__1,esk29_0))).#trainneg
cnf(i_0_14129, plain, (k5_xcmplx_0(esk28_0)=k7_xcmplx_0(np__1,esk28_0))).#trainneg
cnf(i_0_20539, plain, (v1_xxreal_0(esk28_0)|~v1_xreal_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_20644, plain, (v1_int_1(esk18_0)|~v1_int_1(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_20645, plain, (v1_xreal_0(esk18_0)|~v1_xreal_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_20656, plain, (v1_xxreal_0(esk18_0)|~v1_xreal_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_19463, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_19400, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_19634, plain, (r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(X1,X2))|r1_xxreal_0(k7_xcmplx_0(X1,X2),k32_sin_cos)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19378, plain, (~r1_xxreal_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_19567, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4))))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_19934, plain, (k1_real_1(esk34_0)=k4_xcmplx_0(esk34_0)|~m1_subset_1(k4_xcmplx_0(esk34_0),k1_numbers))).#trainneg
cnf(i_0_3549, plain, (v1_xreal_0(X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14392, plain, (k5_xcmplx_0(esk18_0)=k7_xcmplx_0(np__1,esk18_0))).#trainneg
cnf(i_0_14679, plain, (k5_xcmplx_0(esk10_0)=k7_xcmplx_0(np__1,esk10_0))).#trainneg
cnf(i_0_12105, plain, (v3_xxreal_0(np__3)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_12120, plain, (v3_xxreal_0(np__2)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_12439, negated_conjecture, (v3_ordinal1(esk1_0)|~v3_ordinal1(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_12440, negated_conjecture, (v1_funct_1(esk1_0)|~v4_funct_1(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_19566, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4))))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_20612, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_20779, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k1_real_1(X1),esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k1_real_1(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19902, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_19461, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_19564, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_19837, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|m1_subset_1(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20216, plain, (k1_real_1(esk29_0)=k4_xcmplx_0(esk29_0)|~m1_subset_1(k4_xcmplx_0(esk29_0),k1_numbers))).#trainneg
cnf(i_0_3553, plain, (v5_xxreal_2(X1)|~v2_membered(X1)|~v3_xxreal_2(X1)|~v2_xxreal_2(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_15474, plain, (k5_xcmplx_0(esk43_0)=k7_xcmplx_0(np__1,esk43_0))).#trainneg
cnf(i_0_12441, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_12442, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_12443, negated_conjecture, (v1_rat_1(esk1_0)|~v4_membered(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_20333, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0),k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_16399, plain, (m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_16915, plain, (m1_subset_1(k7_xcmplx_0(X1,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12628, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_20345, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X4,X3))|~r2_hidden(k2_rcomp_1(X3,X4),X1)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k2_xboole_0(X1,X2),k1_xxreal_1(X3,X4)))).#trainneg
cnf(i_0_13262, plain, (k3_xcmplx_0(np__2,X1)=k7_xcmplx_0(np__2,k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14107, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20764, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_20531, plain, (k1_real_1(esk28_0)=k4_xcmplx_0(esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k1_numbers))).#trainneg
cnf(i_0_3554, plain, (v1_finset_1(X1)|~v6_membered(X1)|~v2_xxreal_2(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_15501, plain, (k5_xcmplx_0(esk37_0)=k7_xcmplx_0(np__1,esk37_0))).#trainneg
cnf(i_0_15763, plain, (k5_xcmplx_0(np__15)=k7_xcmplx_0(np__1,np__15))).#trainneg
cnf(i_0_12447, negated_conjecture, (v1_relat_1(esk1_0)|~v4_funct_1(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_15251, plain, (v3_xxreal_0(k32_sin_cos)|~v3_xxreal_0(k7_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_16950, plain, (v1_int_1(k32_sin_cos)|~m1_subset_1(k4_xcmplx_0(k32_sin_cos),k4_numbers))).#trainneg
cnf(i_0_16951, plain, (v1_int_1(k32_sin_cos)|~m1_subset_1(k4_xcmplx_0(k32_sin_cos),k5_numbers))).#trainneg
cnf(i_0_19885, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)))).#trainneg
cnf(i_0_14215, plain, (m1_subset_1(k7_xcmplx_0(np__2,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_11984, plain, (v1_relat_1(k5_relat_1(k2_sincos10,X1)))).#trainneg
cnf(i_0_13882, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X3,X4))|~r2_hidden(k4_xxreal_1(X3,X4),X1)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k2_xboole_0(X1,X2),k1_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_20936, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4))))).#trainneg
cnf(i_0_3556, plain, (v3_xxreal_0(k4_xcmplx_0(np__5))|~v1_xreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_21076, plain, (v1_xcmplx_0(esk23_0))).#trainneg
cnf(i_0_21161, plain, (v1_xcmplx_0(esk19_0))).#trainneg
cnf(i_0_21245, plain, (~v3_xxreal_0(k7_xcmplx_0(np__1,np__0)))).#trainneg
cnf(i_0_21246, plain, (~v2_xxreal_0(k7_xcmplx_0(np__1,np__0)))).#trainneg
cnf(i_0_21578, plain, (~v2_xxreal_0(k7_xcmplx_0(np__1,esk34_0)))).#trainneg
cnf(i_0_21283, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_21284, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_21308, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_20950, plain, (r1_xxreal_0(esk8_3(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_3557, plain, (v3_xxreal_0(k4_xcmplx_0(np__4))|~v1_xreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_22390, plain, (~v2_xxreal_2(k3_numbers)|~v6_membered(k3_numbers))).#trainneg
cnf(i_0_22391, plain, (~v2_xxreal_2(k4_numbers)|~v6_membered(k4_numbers))).#trainneg
cnf(i_0_20826, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)))))).#trainneg
cnf(i_0_21306, plain, (esk6_3(X1,X2,X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_20949, plain, (r1_xxreal_0(esk8_3(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_21299, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_xxreal_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_20648, plain, (k1_real_1(esk18_0)=k4_xcmplx_0(esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k1_numbers))).#trainneg
cnf(i_0_3558, plain, (v3_xxreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_21209, plain, (v1_xreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_16243, plain, (k5_xcmplx_0(np__9)=k7_xcmplx_0(np__1,np__9))).#trainneg
cnf(i_0_16367, plain, (k5_xcmplx_0(np__7)=k7_xcmplx_0(np__1,np__7))).#trainneg
cnf(i_0_21078, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_21012, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(esk1_0,k2_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xxreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_21992, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))))))).#trainneg
cnf(i_0_21303, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_xxreal_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_22246, plain, (k2_xboole_0(X1,X2)=X3|k2_xboole_0(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k2_xboole_0(X1,X2),k1_xxreal_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),X1))).#trainneg
cnf(i_0_21093, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))))).#trainneg
cnf(i_0_20854, plain, (k1_real_1(esk10_0)=k4_xcmplx_0(esk10_0)|~m1_subset_1(k4_xcmplx_0(esk10_0),k1_numbers))).#trainneg
cnf(i_0_3559, plain, (v3_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_16760, plain, (k5_xcmplx_0(np__13)=k7_xcmplx_0(np__1,np__13))).#trainneg
cnf(i_0_16862, plain, (k5_xcmplx_0(np__11)=k7_xcmplx_0(np__1,np__11))).#trainneg
cnf(i_0_21772, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_21554, plain, (X1=X2|X1=X3|m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_21609, plain, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))),k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_22491, plain, (r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~r2_hidden(k4_xxreal_1(X2,X3),X1))).#trainneg
cnf(i_0_21304, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2))))|~m1_subset_1(X2,k1_xxreal_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_5764, plain, (v3_xxreal_0(np__9)|~v3_xxreal_0(np__11)|~v1_xreal_0(np__11)|~v1_xreal_0(np__9))).#trainneg
cnf(i_0_3560, plain, (v3_xxreal_0(k4_xcmplx_0(np__2))|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_21330, plain, (r1_xxreal_0(esk8_3(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_20836, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_20837, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_21889, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))))).#trainneg
cnf(i_0_5765, plain, (v3_xxreal_0(np__7)|~v3_xxreal_0(np__9)|~v1_xreal_0(np__9)|~v1_xreal_0(np__7))).#trainneg
cnf(i_0_3565, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_numbers)))).#trainneg
cnf(i_0_21176, plain, (~r1_xxreal_0(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_22121, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_22097, negated_conjecture, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0))|r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_21315, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X3,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_21178, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_21285, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos)|~v1_xreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_21293, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_5767, plain, (v3_xxreal_0(np__7)|~v3_xxreal_0(np__13)|~v1_xreal_0(np__13)|~v1_xreal_0(np__7))).#trainneg
cnf(i_0_3567, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21185, plain, (v1_xreal_0(k4_xcmplx_0(esk43_0)))).#trainneg
cnf(i_0_21186, plain, (v1_xreal_0(k4_xcmplx_0(esk37_0)))).#trainneg
cnf(i_0_21189, plain, (v1_xreal_0(k4_xcmplx_0(esk23_0)))).#trainneg
cnf(i_0_21172, plain, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_22242, plain, (esk6_3(k2_xboole_0(X1,X2),X3,X4)=k2_xboole_0(X1,X2)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k2_xboole_0(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),X1))).#trainneg
cnf(i_0_22105, negated_conjecture, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))))).#trainneg
cnf(i_0_21307, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_21995, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))))).#trainneg
cnf(i_0_21770, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_21771, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_21467, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X4))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_5768, plain, (v3_xxreal_0(np__7)|~v3_xxreal_0(np__11)|~v1_xreal_0(np__11)|~v1_xreal_0(np__7))).#trainneg
cnf(i_0_3569, plain, (v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21191, plain, (v1_xreal_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_21194, plain, (v1_xreal_0(k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_21195, plain, (v1_xreal_0(k4_xcmplx_0(np__9)))).#trainneg
cnf(i_0_21183, plain, (v1_int_1(esk34_0)|~v1_xboole_0(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_21435, plain, (v1_int_1(esk29_0)|~v1_xboole_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_21543, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_21314, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X2)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_21882, plain, (~r1_xxreal_0(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_16650, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__2,esk1_0),k1_numbers))).#trainneg
cnf(i_0_5769, plain, (v3_xxreal_0(np__5)|~v3_xxreal_0(np__9)|~v1_xreal_0(np__9)|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_3570, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21196, plain, (v1_xreal_0(k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_21197, plain, (v1_xreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_21198, plain, (v1_xreal_0(k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_21557, plain, (v1_int_1(esk28_0)|~v1_xboole_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_14838, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X4,X3))|~r2_hidden(k4_xxreal_1(X3,X4),X1)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k2_xboole_0(X1,X2),k1_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_16643, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k32_sin_cos,esk1_0),k1_numbers))).#trainneg
cnf(i_0_14706, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_22102, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_20600, plain, (m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_5770, plain, (v3_xxreal_0(np__5)|~v3_xxreal_0(np__7)|~v1_xreal_0(np__7)|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_3571, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21199, plain, (v1_xreal_0(k4_xcmplx_0(np__11)))).#trainneg
cnf(i_0_10394, plain, (~r2_hidden(k2_tarski(X1,X2),X1))).#trainneg
cnf(i_0_10467, plain, (~r2_hidden(k2_tarski(X2,X1),X1))).#trainneg
cnf(i_0_21751, plain, (v1_int_1(esk18_0)|~v1_xboole_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_17183, negated_conjecture, (v1_int_1(esk1_0)|~m1_subset_1(k4_xcmplx_0(esk1_0),k4_numbers))).#trainneg
cnf(i_0_17184, negated_conjecture, (v1_int_1(esk1_0)|~m1_subset_1(k4_xcmplx_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_16313, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),np__2),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_12098, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_19565, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4))))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_13246, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15341, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X3)|X4!=k1_xxreal_1(X3,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18266, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_14808, plain, (m1_subset_1(esk8_3(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_5772, plain, (v3_xxreal_0(np__5)|~v3_xxreal_0(np__13)|~v1_xreal_0(np__13)|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_11245, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,np__0))).#trainneg
cnf(i_0_17359, plain, (v5_membered(esk15_0)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17361, plain, (v5_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17363, plain, (v5_membered(k1_fdiff_9)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17510, plain, (v5_membered(k2_sincos10)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17512, plain, (v6_ordinal1(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_19561, negated_conjecture, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_14496, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),np__4),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15338, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X3,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_14699, plain, (m1_subset_1(k7_xcmplx_0(k5_xcmplx_0(X1),np__2),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23381, plain, (v1_xreal_0(esk28_0)|~v1_xboole_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_3558, plain, (v3_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_22594, plain, (~v2_xxreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_22595, plain, (~v1_xboole_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_22765, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_22766, plain, (~v1_xboole_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_22895, plain, (~v2_xxreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_22896, plain, (~v1_xboole_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_22568, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_22571, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_22562, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_22564, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_23127, plain, (r2_hidden(X1,X2)|X2!=k1_rcomp_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23613, plain, (v1_xreal_0(esk18_0)|~v1_xboole_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_3514, plain, (v2_xxreal_0(k4_xcmplx_0(np__5))|~v3_xxreal_0(np__5))).#trainneg
cnf(i_0_22515, plain, (k3_xcmplx_0(np__1,esk23_0)=esk23_0)).#trainneg
cnf(i_0_22567, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_22569, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_22948, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_22503, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4))))).#trainneg
cnf(i_0_23126, negated_conjecture, (r2_hidden(esk1_0,X1)|X1!=k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_3556, plain, (v3_xxreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_22540, plain, (k3_xcmplx_0(np__1,esk19_0)=esk19_0)).#trainneg
cnf(i_0_22563, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_22574, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_22560, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_22846, plain, (r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3))|~r2_hidden(k2_rcomp_1(X2,X3),X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_5773, plain, (v3_xxreal_0(np__5)|~v3_xxreal_0(np__11)|~v1_xreal_0(np__11)|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_23066, plain, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_23068, plain, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k10_real_1(k32_sin_cos,np__2))|~v1_xboole_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_17513, plain, (v6_ordinal1(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_23702, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|r2_hidden(X1,k2_tarski(X3,X2))|r2_hidden(X1,X4)|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_22597, plain, (esk6_3(X1,X2,X3)=X1|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_22561, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_22745, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2))|~r2_hidden(k2_rcomp_1(X3,X2),X1))).#trainneg
cnf(i_0_22741, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_5775, plain, (v3_xxreal_0(np__11)|~v3_xxreal_0(np__13)|~v1_xreal_0(np__13)|~v1_xreal_0(np__11))).#trainneg
cnf(i_0_9452, plain, (esk14_2(X1,X2)=np__0)).#trainneg
cnf(i_0_9747, plain, (k4_xcmplx_0(k4_xcmplx_0(np__15))=np__15)).#trainneg
cnf(i_0_17514, plain, (v6_ordinal1(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17515, plain, (v6_ordinal1(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17673, plain, (v3_membered(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17675, plain, (v3_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17677, plain, (v3_membered(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_22977, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_22881, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_rcomp_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k4_xxreal_1(X3,X4),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_23551, plain, (~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_tarski(k1_xxreal_1(X1,X2),X3),k4_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_23242, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_18095, plain, (v3_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_5777, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(np__9)|~v1_xreal_0(np__9)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_9748, plain, (k4_xcmplx_0(k4_xcmplx_0(np__9))=np__9)).#trainneg
cnf(i_0_9749, plain, (k4_xcmplx_0(k4_xcmplx_0(np__7))=np__7)).#trainneg
cnf(i_0_9751, plain, (k4_xcmplx_0(k4_xcmplx_0(np__13))=np__13)).#trainneg
cnf(i_0_17679, plain, (v3_membered(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17681, plain, (v6_membered(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17800, plain, (v6_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17803, plain, (v6_membered(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17806, plain, (v6_membered(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_23587, plain, (~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k2_tarski(X1,k1_xxreal_1(X2,X3)),k4_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_22903, plain, (r1_xxreal_0(esk8_3(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_22855, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k4_xxreal_1(X3,X4),k4_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_22984, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_22591, plain, (r1_xxreal_0(esk8_3(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_20464, negated_conjecture, (r1_xxreal_0(k32_sin_cos,esk1_0)|~v2_xxreal_0(esk1_0)|~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_5778, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(np__7)|~v1_xreal_0(np__7)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_9752, plain, (k4_xcmplx_0(k4_xcmplx_0(np__11))=np__11)).#trainneg
cnf(i_0_9893, plain, (k4_xcmplx_0(k4_xcmplx_0(esk43_0))=esk43_0)).#trainneg
cnf(i_0_9900, plain, (k4_xcmplx_0(k4_xcmplx_0(esk37_0))=esk37_0)).#trainneg
cnf(i_0_17808, plain, (v1_membered(esk15_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17809, plain, (v1_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_17959, plain, (~r1_xxreal_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_17960, plain, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_18248, plain, (v1_membered(k1_fdiff_9)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_22883, plain, (r2_hidden(k1_rcomp_1(X1,k4_xxreal_1(X2,X3)),k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(k1_rcomp_1(X1,k4_xxreal_1(X2,X3)),k1_rcomp_1(X2,X3))|~r1_xxreal_0(X1,k4_xxreal_1(X2,X3))|~v1_xreal_0(k4_xxreal_1(X2,X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22740, plain, (k2_xboole_0(X1,X2)=X3|k2_xboole_0(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k2_xboole_0(X1,X2),k1_rcomp_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_23584, plain, (r2_hidden(k2_tarski(X1,k4_xxreal_1(X2,X3)),k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k2_tarski(X1,k4_xxreal_1(X2,X3)),k2_tarski(X2,X3)))).#trainneg
cnf(i_0_20465, negated_conjecture, (r1_xxreal_0(k32_sin_cos,esk1_0)|~v1_xboole_0(esk1_0)|~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_5779, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(np__5)|~v1_xreal_0(np__5)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_18249, plain, (v1_membered(k2_sincos10)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18251, plain, (v2_membered(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18253, plain, (v2_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18362, plain, (v2_membered(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18364, plain, (v2_membered(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_22861, plain, (r2_hidden(k1_xxreal_1(X1,k4_xxreal_1(X2,X3)),k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(k1_xxreal_1(X1,k4_xxreal_1(X2,X3)),k1_rcomp_1(X2,X3))|~r1_xxreal_0(X1,k4_xxreal_1(X2,X3))|~v1_xxreal_0(k4_xxreal_1(X2,X3))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_23071, plain, (esk6_3(k1_xxreal_1(X1,X2),X3,X4)=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_22879, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_rcomp_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k4_xxreal_1(X3,X4),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_23646, negated_conjecture, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))))).#trainneg
cnf(i_0_23735, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_17896, plain, (v1_xreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_5780, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(np__13)|~v1_xreal_0(np__13)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_9913, plain, (k4_xcmplx_0(k4_xcmplx_0(esk23_0))=esk23_0)).#trainneg
cnf(i_0_9914, plain, (k4_xcmplx_0(k4_xcmplx_0(esk19_0))=esk19_0)).#trainneg
cnf(i_0_18365, plain, (v4_membered(esk15_0)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18367, plain, (v4_membered(k19_sin_cos)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18369, plain, (v4_membered(k1_fdiff_9)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18497, plain, (v4_membered(k2_sincos10)|~v5_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18499, plain, (v4_funct_1(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_22918, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_23626, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),np__2),k1_numbers)|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_22854, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k4_xxreal_1(X3,X4),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_22731, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))))))).#trainneg
cnf(i_0_22907, plain, (r1_xxreal_0(esk8_3(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_17907, plain, (v1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_5781, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(np__11)|~v1_xreal_0(np__11)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_18500, plain, (v4_funct_1(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18501, plain, (v4_funct_1(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_18502, plain, (v4_funct_1(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_19422, negated_conjecture, (v3_xxreal_0(esk1_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_19423, negated_conjecture, (v2_xxreal_0(esk1_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_23734, plain, (m1_subset_1(esk8_3(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_23389, plain, (r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X2))|~r2_hidden(k4_xxreal_1(X3,X2),X1))).#trainneg
cnf(i_0_15324, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10))|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_23763, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_3600, plain, (v3_xxreal_0(k5_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_19424, negated_conjecture, (v1_xboole_0(esk1_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_20467, plain, (v1_xboole_0(esk37_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_20471, plain, (v1_xboole_0(esk18_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk18_0)))).#trainneg
cnf(i_0_20472, plain, (v1_xboole_0(esk10_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_18456, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23579, plain, (r2_hidden(k2_tarski(X1,k4_xxreal_1(X2,X3)),k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k2_tarski(X1,k4_xxreal_1(X2,X3)),k2_tarski(X3,X2)))).#trainneg
cnf(i_0_16662, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2))))|~r2_hidden(X2,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2)))))).#trainneg
cnf(i_0_22742, plain, (k1_xxreal_1(X2,X1)=X3|k1_xxreal_1(X2,X1)=X4|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_xxreal_1(X2,X1),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2))|~r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_17943, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),np__4)|r1_xxreal_0(np__4,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_3601, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10854, plain, (k3_xboole_0(np__0,X1)=np__0)).#trainneg
cnf(i_0_20536, plain, (v1_xreal_0(esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k1_numbers))).#trainneg
cnf(i_0_20537, plain, (v1_xxreal_0(esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k1_numbers))).#trainneg
cnf(i_0_20653, plain, (v1_xreal_0(esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k1_numbers))).#trainneg
cnf(i_0_23356, plain, (m1_subset_1(k4_xcmplx_0(np__5),k1_numbers))).#trainneg
cnf(i_0_12032, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_16924, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),np__2),k1_numbers))).#trainneg
cnf(i_0_17944, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),np__3)|r1_xxreal_0(np__3,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_10988, plain, (k2_xboole_0(np__0,X1)=X1)).#trainneg
cnf(i_0_20654, plain, (v1_xxreal_0(esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k1_numbers))).#trainneg
cnf(i_0_21181, plain, (v1_int_1(esk34_0)|~m1_subset_1(k4_xcmplx_0(esk34_0),k4_numbers))).#trainneg
cnf(i_0_21182, plain, (v1_int_1(esk34_0)|~m1_subset_1(k4_xcmplx_0(esk34_0),k5_numbers))).#trainneg
cnf(i_0_19348, negated_conjecture, (~r1_xxreal_0(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_15789, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k4_xcmplx_0(np__1)),k7_xcmplx_0(X1,k4_xcmplx_0(np__1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_20765, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_14510, plain, (m1_subset_1(k7_xcmplx_0(k7_square_1(X1),np__4),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15019, plain, (m1_subset_1(X1,k1_rcomp_1(X2,X3))|r1_xxreal_0(X3,X1)|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23270, plain, (m1_subset_1(k7_xcmplx_0(esk6_3(X1,X1,X2),k32_sin_cos),k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_17945, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),np__1)|r1_xxreal_0(np__1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_21218, plain, (v3_xxreal_0(np__5)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__5)))).#trainneg
cnf(i_0_21433, plain, (v1_int_1(esk29_0)|~m1_subset_1(k4_xcmplx_0(esk29_0),k4_numbers))).#trainneg
cnf(i_0_21434, plain, (v1_int_1(esk29_0)|~m1_subset_1(k4_xcmplx_0(esk29_0),k5_numbers))).#trainneg
cnf(i_0_21555, plain, (v1_int_1(esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k4_numbers))).#trainneg
cnf(i_0_21556, plain, (v1_int_1(esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k5_numbers))).#trainneg
cnf(i_0_20700, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0),k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_18382, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k7_square_1(X1),esk1_0),k7_xcmplx_0(k7_square_1(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12435, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k2_tarski(X4,X3)))).#trainneg
cnf(i_0_13323, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k7_square_1(X2)),k7_xcmplx_0(X1,k7_square_1(X2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_19591, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_17948, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),np__2)|r1_xxreal_0(np__2,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_11674, plain, (k6_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_21658, plain, (v1_xreal_0(esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k5_numbers))).#trainneg
cnf(i_0_21659, plain, (v1_xreal_0(esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k4_numbers))).#trainneg
cnf(i_0_21685, plain, (v3_xxreal_0(esk29_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk29_0)))).#trainneg
cnf(i_0_17555, plain, (m1_subset_1(k7_xcmplx_0(np__1,k4_xcmplx_0(X1)),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13407, plain, (~r2_hidden(k1_xxreal_1(X1,X2),X1)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_14975, plain, (~r2_hidden(k1_rcomp_1(X3,k2_xboole_0(X1,X2)),X2)|~r1_xxreal_0(X3,k2_xboole_0(X1,X2))|~v1_xreal_0(k2_xboole_0(X1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_11184, plain, (r1_xxreal_0(X1,X2)|m1_subset_1(X1,k4_rcomp_1(X2,X3))|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_14202, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),k32_sin_cos),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_520, plain, (m1_subset_1(np__0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_518, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_9696, plain, (k1_real_1(np__5)=k4_xcmplx_0(np__5))).#trainneg
cnf(i_0_24041, plain, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_25142, plain, (r2_hidden(X1,X2)|r1_xxreal_0(X1,X3)|r1_xxreal_0(X4,X1)|X2!=k1_rcomp_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23971, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_24030, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_rcomp_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_20742, negated_conjecture, (r1_xxreal_0(esk1_0,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0))).#trainneg
cnf(i_0_516, plain, (v5_relat_1(np__0,X2))).#trainneg
cnf(i_0_23877, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))).#trainneg
cnf(i_0_23878, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_23879, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_24098, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_24050, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_rcomp_1(X1,X3)))).#trainneg
cnf(i_0_24931, plain, (k2_xboole_0(X1,X2)=X3|k2_xboole_0(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k2_rcomp_1(X3,X4),X1)|~r1_xxreal_0(X3,k2_xboole_0(X1,X2))|~r1_xxreal_0(k2_xboole_0(X1,X2),X4)|~m1_subset_1(k2_xboole_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_517, plain, (v4_relat_1(np__0,X1))).#trainneg
cnf(i_0_23908, plain, (k5_xcmplx_0(esk23_0)=k7_xcmplx_0(np__1,esk23_0))).#trainneg
cnf(i_0_23951, plain, (k5_xcmplx_0(esk19_0)=k7_xcmplx_0(np__1,esk19_0))).#trainneg
cnf(i_0_24115, plain, (v3_xxreal_0(k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_23923, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k32_sin_cos)))).#trainneg
cnf(i_0_25256, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(X1)),k7_xcmplx_0(np__1,k4_xcmplx_0(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24094, plain, (k2_tarski(X1,k2_rcomp_1(X2,X3))=X2|k2_tarski(X1,k2_rcomp_1(X2,X3))=X3|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(k2_tarski(X1,k2_rcomp_1(X2,X3)),k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_24772, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_12259, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k1_numbers))).#trainneg
cnf(i_0_24204, plain, (v3_xxreal_0(k4_xcmplx_0(np__9)))).#trainneg
cnf(i_0_24237, plain, (v3_xxreal_0(k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_24270, plain, (v3_xxreal_0(k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_23927, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k32_sin_cos)))).#trainneg
cnf(i_0_24163, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_24095, plain, (k2_tarski(k2_rcomp_1(X1,X2),X3)=X1|k2_tarski(k2_rcomp_1(X1,X2),X3)=X2|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k2_tarski(k2_rcomp_1(X1,X2),X3),k1_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_23983, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_24101, plain, (k1_rcomp_1(k2_rcomp_1(X1,X2),X3)=X1|k1_rcomp_1(k2_rcomp_1(X1,X2),X3)=X2|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(k2_rcomp_1(X1,X2),X3),k1_xxreal_1(X1,X2))|~r1_xxreal_0(k2_rcomp_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(k2_rcomp_1(X1,X2)))).#trainneg
cnf(i_0_24776, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_24590, plain, (v1_int_1(esk19_0)|~v1_int_1(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_24340, plain, (v3_xxreal_0(k4_xcmplx_0(np__11)))).#trainneg
cnf(i_0_24956, plain, (v1_xxreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_23917, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos)))).#trainneg
cnf(i_0_24047, plain, (r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_24079, plain, (k1_xxreal_1(k2_rcomp_1(X1,X2),X3)=X1|k1_xxreal_1(k2_rcomp_1(X1,X2),X3)=X2|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(k2_rcomp_1(X1,X2),X3),k1_xxreal_1(X1,X2))|~r1_xxreal_0(k2_rcomp_1(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(k2_rcomp_1(X1,X2)))).#trainneg
cnf(i_0_24185, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_24003, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_23884, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))),k32_sin_cos)))).#trainneg
cnf(i_0_24124, plain, (k1_real_1(np__15)=k4_xcmplx_0(np__15)|~m1_subset_1(k4_xcmplx_0(np__15),k1_numbers))).#trainneg
cnf(i_0_23870, plain, (~v2_xxreal_0(k4_xcmplx_0(np__3))|~v1_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_23871, plain, (~v1_xboole_0(k4_xcmplx_0(np__3))|~v1_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_23938, plain, (~v2_xxreal_0(k4_xcmplx_0(np__5))|~v1_xxreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_23939, plain, (~v1_xboole_0(k4_xcmplx_0(np__5))|~v1_xxreal_0(k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_24486, plain, (esk6_3(k1_rcomp_1(X1,X2),X3,X4)=k1_rcomp_1(X1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24072, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k4_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_24188, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_24456, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_18103, plain, (r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_24213, plain, (k1_real_1(np__9)=k4_xcmplx_0(np__9)|~m1_subset_1(k4_xcmplx_0(np__9),k1_numbers))).#trainneg
cnf(i_0_24189, plain, (~v2_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_24190, plain, (~v1_xboole_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_24460, plain, (X1=X2|X1=X3|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3))|~r2_hidden(k2_rcomp_1(X2,X3),X1))).#trainneg
cnf(i_0_24100, plain, (k1_rcomp_1(X1,k2_rcomp_1(X2,X3))=X2|k1_rcomp_1(X1,k2_rcomp_1(X2,X3))=X3|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(k1_rcomp_1(X1,k2_rcomp_1(X2,X3)),k1_xxreal_1(X2,X3))|~r1_xxreal_0(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(k2_rcomp_1(X2,X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24099, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k2_rcomp_1(X3,X4),k1_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_24246, plain, (k1_real_1(np__7)=k4_xcmplx_0(np__7)|~m1_subset_1(k4_xcmplx_0(np__7),k1_numbers))).#trainneg
cnf(i_0_3631, plain, (v3_membered(k3_xboole_0(X1,X2))|~v4_membered(X2))).#trainneg
cnf(i_0_21738, plain, (v1_xxreal_0(esk28_0)|~v1_xreal_0(k7_xcmplx_0(np__1,esk28_0)))).#trainneg
cnf(i_0_21749, plain, (v1_int_1(esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k4_numbers))).#trainneg
cnf(i_0_21750, plain, (v1_int_1(esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k5_numbers))).#trainneg
cnf(i_0_24746, plain, (r2_hidden(k1_rcomp_1(X1,k4_xxreal_1(X2,X3)),k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(k1_rcomp_1(X1,k4_xxreal_1(X2,X3)),k1_rcomp_1(X2,X3))|~r1_xxreal_0(X1,k4_xxreal_1(X2,X3))|~v1_xreal_0(k4_xxreal_1(X2,X3))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24323, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_24028, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_rcomp_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_24708, plain, (r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~r2_hidden(X1,k1_rcomp_1(X3,X2))|~r2_hidden(k2_rcomp_1(X3,X2),X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3))).#trainneg
cnf(i_0_24279, plain, (k1_real_1(np__13)=k4_xcmplx_0(np__13)|~m1_subset_1(k4_xcmplx_0(np__13),k1_numbers))).#trainneg
cnf(i_0_13386, plain, (v1_funct_1(k4_xcmplx_0(np__1))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_13393, plain, (v1_relat_1(k4_xcmplx_0(np__1))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_21754, plain, (v1_xreal_0(esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k5_numbers))).#trainneg
cnf(i_0_21755, plain, (v1_xreal_0(esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k4_numbers))).#trainneg
cnf(i_0_24096, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_24004, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k4_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_24073, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_rcomp_1(X1,X2)))).#trainneg
cnf(i_0_24005, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_rcomp_1(X1,X2)))).#trainneg
cnf(i_0_24301, plain, (r2_hidden(k2_tarski(X3,k1_xxreal_1(X1,X2)),k2_tarski(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_tarski(X3,k1_xxreal_1(X1,X2)),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_24349, plain, (k1_real_1(np__11)=k4_xcmplx_0(np__11)|~m1_subset_1(k4_xcmplx_0(np__11),k1_numbers))).#trainneg
cnf(i_0_13575, plain, (v1_xcmplx_0(k4_xcmplx_0(k32_sin_cos))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_13708, plain, (v1_funct_1(k4_xcmplx_0(np__2))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_13715, plain, (v1_relat_1(k4_xcmplx_0(np__2))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_21953, plain, (v1_xxreal_0(esk18_0)|~v1_xreal_0(k7_xcmplx_0(np__1,esk18_0)))).#trainneg
cnf(i_0_21958, plain, (v3_xxreal_0(esk10_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_21959, plain, (v2_xxreal_0(esk10_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_16635, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23218, plain, (m1_subset_1(k4_xcmplx_0(np__9),k1_numbers))).#trainneg
cnf(i_0_24024, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k2_xboole_0(X1,X2),k1_rcomp_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),X1))).#trainneg
cnf(i_0_20976, plain, (m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16401, plain, (m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),np__2),k1_numbers))).#trainneg
cnf(i_0_19832, plain, (m1_subset_1(X1,k2_tarski(X2,X3))|m1_subset_1(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23536, plain, (m1_subset_1(k4_xcmplx_0(np__11),k1_numbers))).#trainneg
cnf(i_0_24382, plain, (k1_real_1(esk43_0)=k4_xcmplx_0(esk43_0)|~m1_subset_1(k4_xcmplx_0(esk43_0),k1_numbers))).#trainneg
cnf(i_0_3636, plain, (v3_membered(k3_xboole_0(X1,X2))|~v4_membered(X1))).#trainneg
cnf(i_0_14285, plain, (k7_xcmplx_0(np__2,np__4)=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_18504, plain, (v1_funct_1(k4_xcmplx_0(np__4))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_22142, plain, (v3_xxreal_0(esk43_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk43_0)))).#trainneg
cnf(i_0_22143, plain, (v2_xxreal_0(esk43_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk43_0)))).#trainneg
cnf(i_0_16925, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20130, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),esk1_0),k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),esk1_0)))).#trainneg
cnf(i_0_14196, plain, (m1_subset_1(k7_xcmplx_0(k5_xcmplx_0(X1),k32_sin_cos),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22259, plain, (r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2))|~r2_hidden(k2_rcomp_1(X3,X2),X1))).#trainneg
cnf(i_0_12417, plain, (r1_xxreal_0(esk8_3(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_24415, plain, (k1_real_1(esk37_0)=k4_xcmplx_0(esk37_0)|~m1_subset_1(k4_xcmplx_0(esk37_0),k1_numbers))).#trainneg
cnf(i_0_18511, plain, (v1_relat_1(k4_xcmplx_0(np__4))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_19449, negated_conjecture, (v1_funct_1(k4_xcmplx_0(esk1_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_22172, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_22174, negated_conjecture, (v1_rat_1(esk1_0)|~v5_membered(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_22397, plain, (v3_xxreal_0(esk37_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_22398, plain, (v2_xxreal_0(esk37_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_23390, plain, (r2_hidden(k2_xboole_0(X2,X1),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r2_hidden(k2_xboole_0(X2,X1),k1_rcomp_1(X4,X3))|~r2_hidden(k4_xxreal_1(X4,X3),X1))).#trainneg
cnf(i_0_22844, plain, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_22979, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_11986, plain, (v4_relat_1(k5_relat_1(k2_sincos10,X1),k1_numbers))).#trainneg
cnf(i_0_24946, plain, (m1_subset_1(X1,k1_relset_1(k1_numbers,k2_sincos10))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15350, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_24561, plain, (k1_real_1(esk23_0)=k4_xcmplx_0(esk23_0)|~m1_subset_1(k4_xcmplx_0(esk23_0),k1_numbers))).#trainneg
cnf(i_0_19456, negated_conjecture, (v1_relat_1(k4_xcmplx_0(esk1_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_21206, plain, (v1_xreal_0(k4_xcmplx_0(k32_sin_cos))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_10196, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_22427, plain, (v3_xxreal_0(np__15)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__15)))).#trainneg
cnf(i_0_22519, plain, (v1_xboole_0(esk23_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_22544, plain, (v1_xboole_0(esk19_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_22661, plain, (v3_xxreal_0(np__9)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__9)))).#trainneg
cnf(i_0_24536, plain, (m1_subset_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k1_numbers)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_12384, plain, (m1_subset_1(k7_xcmplx_0(X1,np__3),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25572, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(k2_rcomp_1(X3,X2),X1))).#trainneg
cnf(i_0_24213, plain, (k1_real_1(np__9)=k4_xcmplx_0(np__9))).#trainneg
cnf(i_0_24349, plain, (k1_real_1(np__11)=k4_xcmplx_0(np__11))).#trainneg
cnf(i_0_25877, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_numbers))|~m1_subset_1(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_25876, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_26415, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25488, plain, (v1_valued_0(np__0)|~v1_membered(X1))).#trainneg
cnf(i_0_3644, plain, (v2_xxreal_0(X1)|v2_xxreal_0(k4_xcmplx_0(X1))|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_25489, plain, (v2_valued_0(np__0)|~v2_membered(X1))).#trainneg
cnf(i_0_25490, plain, (v3_valued_0(np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_25491, plain, (v4_valued_0(np__0)|~v6_membered(X1))).#trainneg
cnf(i_0_25547, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k32_sin_cos)))).#trainneg
cnf(i_0_25551, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k32_sin_cos)))).#trainneg
cnf(i_0_25841, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_3649, plain, (k3_xcmplx_0(esk34_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk34_0))).#trainneg
cnf(i_0_25499, plain, (k1_real_1(k4_xcmplx_0(np__5))=np__5)).#trainneg
cnf(i_0_25537, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos)))).#trainneg
cnf(i_0_24880, plain, (r1_xxreal_0(np__4,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),np__4)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_24594, plain, (k1_real_1(esk19_0)=k4_xcmplx_0(esk19_0)|~m1_subset_1(k4_xcmplx_0(esk19_0),k1_numbers))).#trainneg
cnf(i_0_3650, plain, (k3_xcmplx_0(esk29_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk29_0))).#trainneg
cnf(i_0_26368, plain, (v1_xxreal_0(k4_xcmplx_0(np__9)))).#trainneg
cnf(i_0_25797, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k32_sin_cos)))).#trainneg
cnf(i_0_26720, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_26091, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_rcomp_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_25019, plain, (r1_xxreal_0(np__3,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),np__3)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3651, plain, (k3_xcmplx_0(esk28_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk28_0))).#trainneg
cnf(i_0_26494, plain, (v1_xxreal_0(k4_xcmplx_0(np__11)))).#trainneg
cnf(i_0_25640, plain, (~v2_xxreal_0(k4_xcmplx_0(np__15))|~v1_xxreal_0(k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_25641, plain, (~v1_xboole_0(k4_xcmplx_0(np__15))|~v1_xxreal_0(k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_25788, plain, (~v2_xxreal_0(k4_xcmplx_0(np__9))|~v1_xxreal_0(k4_xcmplx_0(np__9)))).#trainneg
cnf(i_0_26722, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_25846, plain, (k1_rcomp_1(k2_rcomp_1(X1,X2),X3)=X2|k1_rcomp_1(k2_rcomp_1(X1,X2),X3)=X1|~r1_xxreal_0(k2_rcomp_1(X1,X2),X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(k2_rcomp_1(X1,X2))|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(k2_rcomp_1(X1,X2),X3),k1_rcomp_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26661, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_26675, plain, (r1_xxreal_0(esk8_3(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k10_real_1(np__3,np__4))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_26207, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r2_hidden(k1_rcomp_1(X1,X2),k1_rcomp_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_26025, plain, (r1_xxreal_0(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k32_sin_cos))).#trainneg
cnf(i_0_25176, plain, (r1_xxreal_0(np__1,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),np__1)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3652, plain, (k3_xcmplx_0(esk18_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk18_0))).#trainneg
cnf(i_0_25789, plain, (~v1_xboole_0(k4_xcmplx_0(np__9))|~v1_xxreal_0(k4_xcmplx_0(np__9)))).#trainneg
cnf(i_0_25791, plain, (~v2_xxreal_0(k4_xcmplx_0(np__7))|~v1_xxreal_0(k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_25792, plain, (~v1_xboole_0(k4_xcmplx_0(np__7))|~v1_xxreal_0(k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_25794, plain, (~v2_xxreal_0(k4_xcmplx_0(np__13))|~v1_xxreal_0(k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_25795, plain, (~v1_xboole_0(k4_xcmplx_0(np__13))|~v1_xxreal_0(k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_26315, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0),k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25903, plain, (k1_xxreal_1(k2_rcomp_1(X1,X2),X3)=X2|k1_xxreal_1(k2_rcomp_1(X1,X2),X3)=X1|~r1_xxreal_0(k2_rcomp_1(X1,X2),X3)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(k2_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~m1_subset_1(k1_xxreal_1(k2_rcomp_1(X1,X2),X3),k1_rcomp_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_26098, plain, (k1_rcomp_1(X1,k2_rcomp_1(X2,X3))=X3|k1_rcomp_1(X1,k2_rcomp_1(X2,X3))=X2|~r1_xxreal_0(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(k2_rcomp_1(X2,X3))|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(k1_rcomp_1(X1,k2_rcomp_1(X2,X3)),k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_25234, plain, (r1_xxreal_0(np__2,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),np__2)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3653, plain, (k3_xcmplx_0(esk10_0,k4_xcmplx_0(np__1))=k4_xcmplx_0(esk10_0))).#trainneg
cnf(i_0_10197, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_25882, plain, (~v2_xxreal_0(k4_xcmplx_0(np__11))|~v1_xxreal_0(k4_xcmplx_0(np__11)))).#trainneg
cnf(i_0_25883, plain, (~v1_xboole_0(k4_xcmplx_0(np__11))|~v1_xxreal_0(k4_xcmplx_0(np__11)))).#trainneg
cnf(i_0_26042, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_26047, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_26045, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_26046, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_26092, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_rcomp_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k2_rcomp_1(X4,X3),k1_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_25557, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~m1_subset_1(k2_rcomp_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_26659, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_26740, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_16828, plain, (k7_xcmplx_0(k1_real_1(X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19603, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_25534, negated_conjecture, (r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),esk1_0)|r1_xxreal_0(esk1_0,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3659, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12030, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_12033, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_26050, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_26051, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_25880, plain, (v1_int_1(esk19_0)|~v1_xboole_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_26043, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_26044, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_12576, plain, (k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25733, plain, (k2_tarski(X1,k2_rcomp_1(X2,X3))=X3|k2_tarski(X1,k2_rcomp_1(X2,X3))=X2|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(k2_tarski(X1,k2_rcomp_1(X2,X3)),k1_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_16889, plain, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos),k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_22982, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_22653, plain, (m1_subset_1(k4_xcmplx_0(np__3),k1_numbers))).#trainneg
cnf(i_0_26634, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_18704, plain, (~r1_xxreal_0(X1,k1_xxreal_1(X2,X3))|~v1_xreal_0(k1_xxreal_1(X2,X3))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k1_rcomp_1(X1,k1_xxreal_1(X2,X3)),k4_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_25776, plain, (v1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_22857, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k4_xxreal_1(X3,X4),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_3663, plain, (v5_relat_1(X1,k4_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12035, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,np__4))=np__4)).#trainneg
cnf(i_0_12109, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_12111, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_26048, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_26049, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_22691, plain, (v3_xxreal_0(np__7)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__7)))).#trainneg
cnf(i_0_22771, plain, (v3_xxreal_0(np__13)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__13)))).#trainneg
cnf(i_0_22801, plain, (v3_xxreal_0(np__11)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__11)))).#trainneg
cnf(i_0_26005, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_17758, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16562, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_23236, plain, (m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_24518, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_18476, plain, (r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3667, plain, (v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12112, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_12114, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,np__3))=np__3)).#trainneg
cnf(i_0_23751, plain, (v5_membered(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_23754, plain, (v5_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_23757, plain, (v5_membered(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_23893, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_17133, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12888, plain, (m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12010, negated_conjecture, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_25573, plain, (k2_xboole_0(X2,X1)=X3|k2_xboole_0(X2,X1)=X4|~r1_xxreal_0(X4,k2_xboole_0(X2,X1))|~r1_xxreal_0(k2_xboole_0(X2,X1),X3)|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(k2_xboole_0(X2,X1),k1_numbers)|~r2_hidden(k2_rcomp_1(X4,X3),X1))).#trainneg
cnf(i_0_3670, plain, (v1_xboole_0(k7_xcmplx_0(np__2,np__3))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))|~v1_xcmplx_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_12124, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_12127, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_23760, plain, (v5_membered(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_24434, plain, (v1_membered(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_24437, plain, (v1_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_14546, plain, (m1_subset_1(k7_xcmplx_0(np__4,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25761, plain, (m1_subset_1(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_27630, plain, (k7_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9698, plain, (k1_real_1(np__3)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_3671, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))|~v1_xreal_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_27193, plain, (~v1_xboole_0(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_27194, plain, (~v3_xxreal_0(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_27704, plain, (r1_xxreal_0(esk8_3(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_27334, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_27252, plain, (v3_xxreal_0(esk29_0)|~v2_xxreal_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_3672, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))|~v1_xcmplx_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_27250, plain, (~v1_xboole_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_27303, plain, (~v1_xboole_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_26986, plain, (v1_valued_0(np__0))).#trainneg
cnf(i_0_27573, plain, (m1_subset_1(k4_xcmplx_0(k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27180, plain, (esk6_3(X1,X2,X3)=X1|esk7_3(X1,X3,X2)=X1|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X3,X2)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_27110, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k32_sin_cos)))).#trainneg
cnf(i_0_27383, plain, (v1_xboole_0(esk18_0)|~v1_xboole_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_3673, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))|~v1_int_1(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_27044, plain, (v2_valued_0(np__0))).#trainneg
cnf(i_0_27070, plain, (v3_valued_0(np__0))).#trainneg
cnf(i_0_27092, plain, (v4_valued_0(np__0))).#trainneg
cnf(i_0_28030, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__3))))).#trainneg
cnf(i_0_27502, plain, (v1_xboole_0(esk10_0)|~v1_xboole_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_3674, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))|~v1_xreal_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_27788, plain, (v1_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_26895, plain, (k1_real_1(k4_xcmplx_0(np__9))=np__9)).#trainneg
cnf(i_0_26909, plain, (k1_real_1(k4_xcmplx_0(np__11))=np__11)).#trainneg
cnf(i_0_27520, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_27521, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4))))).#trainneg
cnf(i_0_27522, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4))))).#trainneg
cnf(i_0_27524, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_26931, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,X1))|~v5_valued_0(k5_relat_1(k2_sincos10,k1_numbers))|~m1_subset_1(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_27279, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_26972, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27503, plain, (v2_xxreal_0(esk10_0)|~v3_xxreal_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_3675, plain, (v3_xxreal_0(k7_xcmplx_0(np__1,np__2))|~v2_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_27190, plain, (v1_xcmplx_0(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_27525, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4))))).#trainneg
cnf(i_0_27526, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4))))).#trainneg
cnf(i_0_28081, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_27714, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_27553, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_27705, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_27350, plain, (r1_xxreal_0(esk8_3(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_26021, plain, (v3_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v3_xxreal_0(k32_sin_cos)|~v1_xreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_27504, plain, (v3_xxreal_0(esk10_0)|~v2_xxreal_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_27192, plain, (v1_xreal_0(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_27202, plain, (v1_xxreal_0(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_27247, plain, (v1_xcmplx_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_27808, plain, (esk6_3(k1_xxreal_1(X1,X2),X3,X4)=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_28077, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_27287, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_rcomp_1(X3,X4))|~m1_subset_1(k2_rcomp_1(X3,X4),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_27890, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_9837, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3677, plain, (v1_xboole_0(k7_xcmplx_0(np__1,np__2))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_27249, plain, (v1_xreal_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_27259, plain, (v1_xxreal_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_27300, plain, (v1_xcmplx_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_27358, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k1_rcomp_1(X1,X2),k1_rcomp_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_27554, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_rcomp_1(X2,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_28316, plain, (X1=X2|X1=X3|m1_subset_1(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_27186, plain, (r2_hidden(X1,X2)|esk7_3(X1,X4,X3)=X1|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_xxreal_1(X4,X3)))).#trainneg
cnf(i_0_3678, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_27380, plain, (v1_xcmplx_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_27511, plain, (v1_xxreal_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_27783, plain, (v1_funct_1(k4_xcmplx_0(np__3))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_24448, plain, (v1_membered(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_24476, plain, (v1_membered(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_27340, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~m1_subset_1(k2_rcomp_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_28146, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_27285, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_rcomp_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_20766, plain, (r1_xxreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_27790, plain, (v1_relat_1(k4_xcmplx_0(np__3))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_12129, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,np__2))=np__2)).#trainneg
cnf(i_0_15155, plain, (v1_xxreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_24611, plain, (v4_membered(esk15_0)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_24613, plain, (v4_membered(k19_sin_cos)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_24615, plain, (v4_membered(k1_fdiff_9)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_24617, plain, (v4_membered(k2_sincos10)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_25482, plain, (v7_valued_0(np__0)|~v3_funct_1(np__0)|~v1_funct_1(np__0))).#trainneg
cnf(i_0_16536, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v5_valued_0(k5_relat_1(k2_sincos10,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_real_1(k7_square_1(np__2)),X1))).#trainneg
cnf(i_0_22746, plain, (k2_xboole_0(X2,X1)=X3|k2_xboole_0(X2,X1)=X4|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k2_xboole_0(X2,X1),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),X1))).#trainneg
cnf(i_0_27371, plain, (r1_xxreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_13256, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k3_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16415, negated_conjecture, (k3_xcmplx_0(np__1,k4_xcmplx_0(esk1_0))=k4_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_23932, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k10_real_1(k32_sin_cos,np__2))|~v1_xreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_3680, plain, (v1_int_1(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_int_1(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_15255, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_15257, plain, (v1_xreal_0(k7_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_15258, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_25579, plain, (v3_xxreal_0(esk23_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_25580, plain, (v2_xxreal_0(esk23_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_25611, plain, (v3_xxreal_0(esk19_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_25612, plain, (v2_xxreal_0(esk19_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_22874, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k2_xboole_0(X1,X2),k1_rcomp_1(X3,X4))|~r2_hidden(k4_xxreal_1(X3,X4),X2))).#trainneg
cnf(i_0_15409, negated_conjecture, (esk7_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|r2_hidden(esk1_0,k2_tarski(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_23336, plain, (m1_subset_1(k4_xcmplx_0(np__7),k1_numbers))).#trainneg
cnf(i_0_26673, plain, (r1_xxreal_0(esk8_3(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_19533, negated_conjecture, (m1_subset_1(k7_xcmplx_0(X1,k4_xcmplx_0(esk1_0)),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3681, plain, (v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_15260, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,k32_sin_cos))=k32_sin_cos)).#trainneg
cnf(i_0_25878, plain, (v1_int_1(esk19_0)|~m1_subset_1(k4_xcmplx_0(esk19_0),k4_numbers))).#trainneg
cnf(i_0_25879, plain, (v1_int_1(esk19_0)|~m1_subset_1(k4_xcmplx_0(esk19_0),k5_numbers))).#trainneg
cnf(i_0_26698, negated_conjecture, (v1_rat_1(esk1_0)|~v6_membered(k1_relset_1(k1_numbers,k2_sincos10)))).#trainneg
cnf(i_0_13336, plain, (r1_xxreal_0(k10_real_1(X1,X2),k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14807, plain, (m1_subset_1(esk8_3(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k7_xcmplx_0(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_13320, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k4_xcmplx_0(X2)),k7_xcmplx_0(X1,k4_xcmplx_0(X2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_17989, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28680, plain, (m1_subset_1(k4_xcmplx_0(k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_24246, plain, (k1_real_1(np__7)=k4_xcmplx_0(np__7))).#trainneg
cnf(i_0_23870, plain, (~v2_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_28481, plain, (esk7_3(X1,X2,X3)=X1|X1=X2|X1=X3|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29089, plain, (X1=X2|X1=X3|m1_subset_1(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3684, plain, (v1_xboole_0(k7_xcmplx_0(np__1,np__3))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_28392, plain, (k1_real_1(k4_xcmplx_0(np__3))=np__3)).#trainneg
cnf(i_0_29754, plain, (v1_xxreal_0(k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_28701, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk34_0))))).#trainneg
cnf(i_0_28908, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk29_0))))).#trainneg
cnf(i_0_29060, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk28_0))))).#trainneg
cnf(i_0_29691, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_28530, plain, (m1_subset_1(k4_xcmplx_0(k4_xcmplx_0(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29851, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_28440, plain, (r1_xxreal_0(esk8_3(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_3685, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))|~v1_xreal_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_28618, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3))))).#trainneg
cnf(i_0_30057, plain, (r2_hidden(X1,X2)|m1_subset_1(X1,k1_xxreal_1(X4,X3))|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29893, plain, (m1_subset_1(esk8_3(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_30065, plain, (r2_hidden(X1,X2)|m1_subset_1(X1,k1_xxreal_1(X3,X4))|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17668, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28622, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4)))).#trainneg
cnf(i_0_28926, plain, (esk6_3(k1_rcomp_1(X1,X2),X3,X4)=k1_rcomp_1(X1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30046, plain, (m1_subset_1(X1,k2_tarski(X2,X3))|m1_subset_1(X1,k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28353, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(X1,esk1_0),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(X1,esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29891, plain, (m1_subset_1(esk8_3(k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17670, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3687, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))|~v1_int_1(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_19098, plain, (v1_xxreal_0(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_28720, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3))))).#trainneg
cnf(i_0_29455, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k1_real_1(k7_square_1(np__2)))|v1_xboole_0(X1)|~v5_valued_0(k5_relat_1(k2_sincos10,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),X1))).#trainneg
cnf(i_0_29085, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_28532, plain, (esk7_3(X1,X2,X3)=X1|esk6_3(X1,X3,X2)=X1|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_28752, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17795, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3688, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))|~v1_xreal_0(k7_xcmplx_0(np__1,np__3)))).#trainneg
cnf(i_0_28724, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4)))).#trainneg
cnf(i_0_12220, plain, (r1_xxreal_0(k2_xboole_0(X1,X2),X3)|~r2_hidden(k4_rcomp_1(X3,X4),X1)|~r1_xxreal_0(k2_xboole_0(X1,X2),X4)|~v1_xreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k2_xboole_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_12644, plain, (m1_subset_1(k7_xcmplx_0(np__3,np__2),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_20891, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_25913, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X1,X2))|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X3,X4)))).#trainneg
cnf(i_0_16829, plain, (k7_xcmplx_0(k7_square_1(X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20197, plain, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(X1),np__4),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17797, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3691, plain, (v1_xboole_0(k7_xcmplx_0(np__1,np__4))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_19429, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_19431, negated_conjecture, (k5_xcmplx_0(k7_xcmplx_0(np__1,esk1_0))=esk1_0)).#trainneg
cnf(i_0_28612, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_28625, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_16402, plain, (m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27116, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_10844, plain, (v1_xreal_0(k7_xcmplx_0(X1,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27548, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~m1_subset_1(k2_rcomp_1(X4,X3),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_19471, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_17853, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3692, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))|~v1_xreal_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_20161, plain, (v1_xreal_0(k7_xcmplx_0(k32_sin_cos,k32_sin_cos)))).#trainneg
cnf(i_0_20162, plain, (v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k32_sin_cos)))).#trainneg
cnf(i_0_20554, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(esk1_0,np__2)))).#trainneg
cnf(i_0_27530, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_17887, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15509, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X4,X3))|~r2_hidden(k2_rcomp_1(X3,X4),X1)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r2_hidden(k2_xboole_0(X1,X2),k1_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_26013, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X3,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_26721, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_17883, plain, (m1_subset_1(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17855, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_20555, negated_conjecture, (v1_xxreal_0(k7_xcmplx_0(esk1_0,np__2)))).#trainneg
cnf(i_0_20578, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(esk1_0,k32_sin_cos)))).#trainneg
cnf(i_0_20579, negated_conjecture, (v1_xxreal_0(k7_xcmplx_0(esk1_0,k32_sin_cos)))).#trainneg
cnf(i_0_27667, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_29394, plain, (v4_membered(esk15_0)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_29397, plain, (v4_membered(k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_29400, plain, (v4_membered(k1_fdiff_9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_15515, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X3,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_25775, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_16392, plain, (m1_subset_1(k7_xcmplx_0(X1,k7_xcmplx_0(k32_sin_cos,np__2)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19349, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_13334, plain, (r1_xxreal_0(k8_real_1(X1,X2),k8_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12573, plain, (k3_xcmplx_0(k6_xcmplx_0(X1,X2),k7_xcmplx_0(X3,X4))=k7_xcmplx_0(k3_xcmplx_0(k6_xcmplx_0(X1,X2),X3),X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_30106, plain, (X1=X3|X1=X2|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,a_2_1_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10836, plain, (k1_real_1(k32_sin_cos)=k4_xcmplx_0(k32_sin_cos))).#trainneg
cnf(i_0_25791, plain, (~v2_xxreal_0(k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_30074, plain, (v1_xreal_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_30075, plain, (v1_xxreal_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_30254, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(k4_xcmplx_0(np__1),np__3))))).#trainneg
cnf(i_0_30105, plain, (X2=X1|X2=X3|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X2,a_2_1_rcomp_1(X3,X1))|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_30949, plain, (m1_subset_1(k4_xcmplx_0(k7_square_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_24951, plain, (v1_funct_1(k4_xcmplx_0(np__5))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3694, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))|~v1_int_1(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_30089, plain, (k1_real_1(k4_xcmplx_0(np__7))=np__7)).#trainneg
cnf(i_0_30070, plain, (v1_funct_1(k4_xcmplx_0(k32_sin_cos))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_30403, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4)))).#trainneg
cnf(i_0_30859, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4)))).#trainneg
cnf(i_0_31098, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_30738, plain, (r1_xxreal_0(X1,esk6_3(X2,X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X2,k2_tarski(X3,X1)))).#trainneg
cnf(i_0_30326, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_30915, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_24958, plain, (v1_relat_1(k4_xcmplx_0(np__5))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_30076, plain, (v1_xcmplx_0(k4_xcmplx_0(k32_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_30077, plain, (v1_relat_1(k4_xcmplx_0(k32_sin_cos))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_29403, plain, (v4_membered(k2_sincos10)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_30928, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30733, plain, (r1_xxreal_0(esk6_3(X1,X2,X3),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X3)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_30302, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_30332, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_30725, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2)))))).#trainneg
cnf(i_0_26363, plain, (v1_funct_1(k4_xcmplx_0(np__9))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3698, plain, (v1_xboole_0(k7_xcmplx_0(np__3,np__2))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_22474, plain, (v1_xreal_0(k7_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_22475, plain, (v1_xxreal_0(k7_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_30116, plain, (X1=X2|X1=X3|m1_subset_1(X1,k2_rcomp_1(X3,X2))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3))).#trainneg
cnf(i_0_31521, plain, (r1_xxreal_0(k8_real_1(X2,X1),k8_real_1(X2,X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_30740, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,X4)|X4!=k1_xxreal_1(X3,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_30325, plain, (m1_subset_1(esk8_3(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_30728, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2)))))).#trainneg
cnf(i_0_30724, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r2_hidden(X2,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(X1,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2)))))).#trainneg
cnf(i_0_31363, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_26370, plain, (v1_relat_1(k4_xcmplx_0(np__9))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3699, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))|~v1_xreal_0(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_23157, plain, (v1_xreal_0(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_23248, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_3903, plain, (v1_xcmplx_0(esk15_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3939, plain, (v1_xxreal_0(esk15_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3975, plain, (v1_xreal_0(esk15_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_30489, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(esk8_3(k32_sin_cos,k10_real_1(k32_sin_cos,np__2),k32_sin_cos),esk1_0),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(esk8_3(k32_sin_cos,k10_real_1(k32_sin_cos,np__2),k32_sin_cos),esk1_0))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_31345, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r2_hidden(k1_xxreal_1(X1,X2),k1_rcomp_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_24863, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X2,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2))))|~m1_subset_1(X1,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2)))))).#trainneg
cnf(i_0_26489, plain, (v1_funct_1(k4_xcmplx_0(np__11))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3700, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_23249, negated_conjecture, (v1_xxreal_0(k7_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_23315, plain, (v1_xcmplx_0(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_23406, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_4011, plain, (v1_rat_1(esk15_0)|~v4_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4047, plain, (v1_int_1(esk15_0)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4083, plain, (v7_ordinal1(esk15_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_16691, negated_conjecture, (r2_hidden(esk1_0,X1)|esk7_3(esk1_0,k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|X1!=k1_xxreal_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xxreal_0(k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_12019, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_22309, plain, (k3_xcmplx_0(np__2,X1)=k7_xcmplx_0(np__2,k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23589, plain, (~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k2_tarski(X1,k1_xxreal_1(X2,X3)),k2_tarski(X3,X2)))).#trainneg
cnf(i_0_14530, plain, (m1_subset_1(k7_xcmplx_0(k5_xcmplx_0(X1),X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13255, plain, (k3_xcmplx_0(k6_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k6_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13960, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|v1_xboole_0(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k4_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_24724, plain, (r2_hidden(k1_xxreal_1(X1,k4_xxreal_1(X2,X3)),k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(k1_xxreal_1(X1,k4_xxreal_1(X2,X3)),k1_rcomp_1(X2,X3))|~r1_xxreal_0(X1,k4_xxreal_1(X2,X3))|~v1_xxreal_0(k4_xxreal_1(X2,X3))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_28855, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_26496, plain, (v1_relat_1(k4_xcmplx_0(np__11))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3701, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))|~v1_int_1(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_23407, negated_conjecture, (v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_27967, plain, (v1_xxreal_0(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_13397, plain, (~r2_hidden(k1_xxreal_1(X3,k2_xboole_0(X1,X2)),X2)|~r1_xxreal_0(X3,k2_xboole_0(X1,X2))|~v1_xxreal_0(k2_xboole_0(X1,X2))|~v1_xxreal_0(X3))).#trainneg
cnf(i_0_30824, plain, (r1_xxreal_0(k10_real_1(X1,X2),np__2)|r1_xxreal_0(np__2,k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27154, negated_conjecture, (esk7_3(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos,k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))=k1_seq_1(k2_sincos10,esk1_0)|~r1_xxreal_0(k32_sin_cos,k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~v1_xreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k32_sin_cos,k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))))).#trainneg
cnf(i_0_20336, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0),k1_numbers))).#trainneg
cnf(i_0_17297, plain, (v1_xxreal_0(esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23376, plain, (m1_subset_1(k4_xcmplx_0(np__13),k1_numbers))).#trainneg
cnf(i_0_10845, plain, (v1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15612, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30683, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X3,X4))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k2_rcomp_1(X3,X4),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_29749, plain, (v1_funct_1(k4_xcmplx_0(np__7))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3702, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))|~v1_xreal_0(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_30877, plain, (v1_xreal_0(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_29756, plain, (v1_relat_1(k4_xcmplx_0(np__7))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_14204, plain, (m1_subset_1(k7_xcmplx_0(k3_xcmplx_0(X1,X2),k32_sin_cos),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32090, plain, (r2_hidden(X1,X2)|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_24279, plain, (k1_real_1(np__13)=k4_xcmplx_0(np__13))).#trainneg
cnf(i_0_3699, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2)))).#trainneg
cnf(i_0_31778, plain, (k1_real_1(k4_xcmplx_0(k32_sin_cos))=k32_sin_cos)).#trainneg
cnf(i_0_19004, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3702, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2)))).#trainneg
cnf(i_0_31932, plain, (v1_xcmplx_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_33015, plain, (v1_xxreal_0(k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_31769, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_32461, negated_conjecture, (esk7_3(esk1_0,k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|r2_hidden(esk1_0,k1_xxreal_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_19334, plain, (r1_xxreal_0(X1,np__4)|r1_xxreal_0(np__4,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3705, plain, (v1_xboole_0(k7_xcmplx_0(np__3,np__4))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_33010, plain, (v1_funct_1(k4_xcmplx_0(np__13))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_33017, plain, (v1_relat_1(k4_xcmplx_0(np__13))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_31926, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),k2_tarski(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32223, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_31830, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_19336, plain, (r1_xxreal_0(X1,np__4)|r1_xxreal_0(np__4,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3706, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_13843, plain, (v3_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_18465, plain, (v3_ordinal1(X1)|~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10198, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xboole_0(X1))).#trainneg
cnf(i_0_32224, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_29012, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21222, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__5)))).#trainneg
cnf(i_0_21224, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__5)))).#trainneg
cnf(i_0_21225, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__5)))).#trainneg
cnf(i_0_32183, plain, (m1_subset_1(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_32895, plain, (~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k1_xxreal_1(X1,X2),X2))).#trainneg
cnf(i_0_31922, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_29014, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3708, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))|~v1_int_1(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_21238, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__5)))).#trainneg
cnf(i_0_21251, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__0)))).#trainneg
cnf(i_0_33203, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(esk1_0),k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_32157, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_24676, plain, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3709, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_21252, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__0)))).#trainneg
cnf(i_0_21265, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__0)))).#trainneg
cnf(i_0_32444, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,esk1_0),k7_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_32297, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(np__2,esk1_0),k7_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_32974, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_32058, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_32027, plain, (r2_hidden(X1,X2)|esk7_3(X1,X3,X4)=X1|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_12004, negated_conjecture, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~r1_xxreal_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_17266, plain, (m1_subset_1(k7_xcmplx_0(k1_real_1(X1),np__1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28785, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_3712, plain, (v1_xboole_0(k7_xcmplx_0(np__4,np__3))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))|~v1_xcmplx_0(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_21580, plain, (v3_xxreal_0(k7_xcmplx_0(np__1,esk34_0)))).#trainneg
cnf(i_0_21583, plain, (v1_xreal_0(k7_xcmplx_0(np__1,esk34_0)))).#trainneg
cnf(i_0_21584, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk34_0)))).#trainneg
cnf(i_0_10525, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_23602, plain, (r2_hidden(k2_tarski(X1,k2_rcomp_1(X2,X3)),k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k2_tarski(X1,k2_rcomp_1(X2,X3)),k2_tarski(X2,X3)))).#trainneg
cnf(i_0_17067, plain, (v1_xxreal_0(k7_xcmplx_0(k32_sin_cos,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25842, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_12580, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k3_xcmplx_0(np__2,X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25836, plain, (k2_tarski(k2_rcomp_1(X1,X2),X3)=X2|k2_tarski(k2_rcomp_1(X1,X2),X3)=X1|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k2_tarski(k2_rcomp_1(X1,X2),X3),k1_rcomp_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3713, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))|~v1_xreal_0(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_21598, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,esk34_0)))).#trainneg
cnf(i_0_21689, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,esk29_0)))).#trainneg
cnf(i_0_31113, plain, (m1_subset_1(k7_xcmplx_0(k7_xcmplx_0(X1,k32_sin_cos),np__2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17876, plain, (m1_subset_1(k7_xcmplx_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28082, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_20349, plain, (r2_hidden(k1_rcomp_1(X1,k2_rcomp_1(X2,X3)),k2_tarski(X3,X2))|~r1_xxreal_0(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(k2_rcomp_1(X2,X3))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(k1_rcomp_1(X1,k2_rcomp_1(X2,X3)),k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_14845, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_24842, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_23597, plain, (r2_hidden(k2_tarski(X1,k2_rcomp_1(X2,X3)),k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k2_tarski(X1,k2_rcomp_1(X2,X3)),k2_tarski(X3,X2)))).#trainneg
cnf(i_0_28304, plain, (m1_subset_1(X1,k2_tarski(X2,X3))|m1_subset_1(X1,k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_21691, plain, (v1_xreal_0(k7_xcmplx_0(np__1,esk29_0)))).#trainneg
cnf(i_0_21692, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk29_0)))).#trainneg
cnf(i_0_22930, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_10351, plain, (k1_xxreal_1(X1,X2)=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_zfmisc_1(k4_xxreal_1(X1,X2))))).#trainneg
cnf(i_0_16933, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~m1_subset_1(X1,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X2,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_23697, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_33331, plain, (r2_hidden(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25794, plain, (~v2_xxreal_0(k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_33334, plain, (k1_real_1(k4_xcmplx_0(np__13))=np__13)).#trainneg
cnf(i_0_33860, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__5))))).#trainneg
cnf(i_0_34350, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk34_0))))).#trainneg
cnf(i_0_34679, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk29_0))))).#trainneg
cnf(i_0_3715, plain, (v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))|~v1_int_1(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_33692, plain, (v2_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_21707, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,esk29_0)))).#trainneg
cnf(i_0_21722, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk28_0)))).#trainneg
cnf(i_0_9439, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos)))).#trainneg
cnf(i_0_9465, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_33590, plain, (k1_rcomp_1(k2_rcomp_1(X1,X2),X3)=X2|k1_rcomp_1(k2_rcomp_1(X1,X2),X3)=X1|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k2_rcomp_1(X1,X2),X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(k2_rcomp_1(X1,X2))|~m1_subset_1(k1_rcomp_1(k2_rcomp_1(X1,X2),X3),k1_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_33583, plain, (k1_rcomp_1(X1,X2)=X3|k1_rcomp_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X2,X1)))).#trainneg
cnf(i_0_3716, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))|~v1_xreal_0(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_21936, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk18_0)))).#trainneg
cnf(i_0_21964, plain, (v1_xreal_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_21965, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_9948, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_10607, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_10608, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_10614, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_34367, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_34635, plain, (m1_subset_1(X1,k4_xxreal_1(X2,X3))|m1_subset_1(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_33589, plain, (k1_rcomp_1(X1,k2_rcomp_1(X2,X3))=X3|k1_rcomp_1(X1,k2_rcomp_1(X2,X3))=X2|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,k2_rcomp_1(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(k2_rcomp_1(X2,X3))|~v1_xreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,k2_rcomp_1(X2,X3)),k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_3721, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_21981, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_22148, plain, (v1_xreal_0(k7_xcmplx_0(np__1,esk43_0)))).#trainneg
cnf(i_0_10615, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_11908, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_11909, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k32_sin_cos,k7_xcmplx_0(np__3,np__4)))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_17799, plain, (v6_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_34075, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k2_rcomp_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_34780, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,k32_sin_cos),X2)=k7_xcmplx_0(k7_xcmplx_0(X1,k32_sin_cos),k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_34785, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_34706, plain, (k1_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_zfmisc_1(k2_rcomp_1(X1,X2))))).#trainneg
cnf(i_0_33497, negated_conjecture, (esk7_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|r2_hidden(esk1_0,k1_xxreal_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_34784, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,np__2),X2)=k7_xcmplx_0(k7_xcmplx_0(X1,np__2),k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_34556, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),k7_xcmplx_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3722, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_22149, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk43_0)))).#trainneg
cnf(i_0_22165, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,esk43_0)))).#trainneg
cnf(i_0_17802, plain, (v6_membered(k1_fdiff_9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_17805, plain, (v6_membered(k2_sincos10)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_34074, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_34613, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_34707, plain, (k4_xxreal_1(X1,X2)=k1_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),k1_zfmisc_1(k4_xxreal_1(X1,X2)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_33727, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X3,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_28729, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))),k32_sin_cos)))).#trainneg
cnf(i_0_17267, plain, (m1_subset_1(k7_xcmplx_0(k7_square_1(X1),np__1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_32865, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_27606, negated_conjecture, (v3_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(esk1_0,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v3_xxreal_0(esk1_0)|~v1_xreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_3723, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_22403, plain, (v1_xreal_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_22404, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_20676, plain, (v1_xreal_0(esk28_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk28_0),k1_numbers))).#trainneg
cnf(i_0_20677, plain, (v1_xreal_0(esk28_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk28_0),k5_numbers))).#trainneg
cnf(i_0_20678, plain, (v1_xreal_0(esk28_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk28_0),k4_numbers))).#trainneg
cnf(i_0_20682, plain, (v1_xreal_0(esk18_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk18_0),k1_numbers))).#trainneg
cnf(i_0_14842, plain, (v1_xxreal_0(X1)|r2_hidden(X1,k2_tarski(X3,X2))|~v2_membered(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_20346, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X4,X3))|~r2_hidden(k2_rcomp_1(X3,X4),X2)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k2_xboole_0(X1,X2),k1_xxreal_1(X3,X4)))).#trainneg
cnf(i_0_17569, plain, (m1_subset_1(k7_xcmplx_0(np__1,k1_real_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_33047, negated_conjecture, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk1_0))=esk1_0)).#trainneg
cnf(i_0_36183, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),X1)|r1_xxreal_0(X1,k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3727, plain, (v1_xboole_0(k3_xxreal_1(X2,X1))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v6_membered(k3_xxreal_1(X2,X1)))).#trainneg
cnf(i_0_34885, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk28_0))))).#trainneg
cnf(i_0_35483, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk43_0))))).#trainneg
cnf(i_0_22420, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_35054, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3))))).#trainneg
cnf(i_0_34904, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4))),k32_sin_cos)))).#trainneg
cnf(i_0_3728, plain, (v1_xboole_0(k4_xxreal_1(X2,X1))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v6_membered(k4_xxreal_1(X2,X1)))).#trainneg
cnf(i_0_22431, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__15)))).#trainneg
cnf(i_0_22433, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__15)))).#trainneg
cnf(i_0_22434, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__15)))).#trainneg
cnf(i_0_35058, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4)))).#trainneg
cnf(i_0_34847, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|k4_xcmplx_0(np__1)=esk1_0|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_11261, plain, (k3_xcmplx_0(X1,np__1)=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3729, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k3_xcmplx_0(np__2,k32_sin_cos)))))).#trainneg
cnf(i_0_22450, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__15)))).#trainneg
cnf(i_0_22665, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__9)))).#trainneg
cnf(i_0_35064, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_36260, plain, (r1_xxreal_0(k1_real_1(X1),X2)|r1_xxreal_0(X2,k1_real_1(X1))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13229, plain, (k3_xcmplx_0(esk43_0,X1)=k7_xcmplx_0(esk43_0,k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3730, plain, (v1_relat_1(k5_relat_1(k1_fdiff_9,k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_22667, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__9)))).#trainneg
cnf(i_0_22668, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__9)))).#trainneg
cnf(i_0_35156, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3))))).#trainneg
cnf(i_0_35160, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4)))).#trainneg
cnf(i_0_35771, plain, (X1=X2|X1=X3|m1_subset_1(X1,k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X2)))).#trainneg
cnf(i_0_13230, plain, (k3_xcmplx_0(esk37_0,X1)=k7_xcmplx_0(esk37_0,k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3731, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_22684, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__9)))).#trainneg
cnf(i_0_22695, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__7)))).#trainneg
cnf(i_0_22697, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__7)))).#trainneg
cnf(i_0_35165, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_20683, plain, (v1_xreal_0(esk18_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk18_0),k5_numbers))).#trainneg
cnf(i_0_35980, plain, (v1_xxreal_0(X1)|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v2_membered(k4_xxreal_1(X3,X2))|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_35292, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(esk1_0,X1),X2)=k7_xcmplx_0(k7_xcmplx_0(esk1_0,X1),k7_xcmplx_0(np__1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_35456, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)),k7_xcmplx_0(X1,k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36007, plain, (esk7_3(k2_xboole_0(X1,X2),X3,X4)=k2_xboole_0(X1,X2)|r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k2_xboole_0(X1,X2),k1_xxreal_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),X2))).#trainneg
cnf(i_0_35353, plain, (k3_xcmplx_0(k7_xcmplx_0(k32_sin_cos,X1),X2)=k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,X1),k7_xcmplx_0(np__1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36348, plain, (r1_xxreal_0(X1,k7_xcmplx_0(k32_sin_cos,np__2))|r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36854, plain, (k3_xcmplx_0(X1,np__1)=k7_xcmplx_0(X1,np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3732, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_36782, negated_conjecture, (k3_xcmplx_0(esk1_0,np__1)=esk1_0)).#trainneg
cnf(i_0_36773, plain, (k3_xcmplx_0(esk43_0,np__1)=esk43_0)).#trainneg
cnf(i_0_36774, plain, (k3_xcmplx_0(esk37_0,np__1)=esk37_0)).#trainneg
cnf(i_0_36742, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__15))))).#trainneg
cnf(i_0_37322, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__9))))).#trainneg
cnf(i_0_36345, plain, (r1_xxreal_0(X1,k7_xcmplx_0(k32_sin_cos,np__2))|r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36682, plain, (k7_square_1(X1)=X2|k7_square_1(X1)=X3|~r1_xxreal_0(X2,k7_square_1(X1))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k7_square_1(X1))|~r1_xxreal_0(k7_square_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36523, plain, (k7_square_1(X1)=X2|k7_square_1(X1)=X3|~r1_xxreal_0(k7_square_1(X1),X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k7_square_1(X1))|~r1_xxreal_0(k7_square_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36679, plain, (k1_seq_1(X1,X2)=X3|k1_seq_1(X1,X2)=X4|~r1_xxreal_0(X3,k1_seq_1(X1,X2))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k1_seq_1(X1,X2))|~r1_xxreal_0(k1_seq_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_36666, plain, (k4_xcmplx_0(X1)=X2|k4_xcmplx_0(X1)=X3|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36520, plain, (k1_seq_1(X1,X2)=X3|k1_seq_1(X1,X2)=X4|~r1_xxreal_0(k1_seq_1(X1,X2),X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k1_seq_1(X1,X2))|~r1_xxreal_0(k1_seq_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_36507, plain, (k4_xcmplx_0(X1)=X2|k4_xcmplx_0(X1)=X3|~r1_xxreal_0(k4_xcmplx_0(X1),X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36522, plain, (k1_real_1(X1)=X2|k1_real_1(X1)=X3|~r1_xxreal_0(k1_real_1(X1),X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36661, plain, (k4_xcmplx_0(np__1)=X1|k4_xcmplx_0(np__1)=X2|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X2,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_36556, negated_conjecture, (esk1_0=X1|esk1_0=X2|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X2,esk1_0)|~r1_xxreal_0(esk1_0,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_36681, plain, (k1_real_1(X1)=X2|k1_real_1(X1)=X3|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13231, plain, (k3_xcmplx_0(esk34_0,X1)=k7_xcmplx_0(esk34_0,k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3733, plain, (v2_funct_1(k5_relat_1(k1_fdiff_9,k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_36775, plain, (k3_xcmplx_0(esk34_0,np__1)=esk34_0)).#trainneg
cnf(i_0_36776, plain, (k3_xcmplx_0(esk29_0,np__1)=esk29_0)).#trainneg
cnf(i_0_36777, plain, (k3_xcmplx_0(esk28_0,np__1)=esk28_0)).#trainneg
cnf(i_0_36762, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4)))).#trainneg
cnf(i_0_36502, plain, (k4_xcmplx_0(np__1)=X1|k4_xcmplx_0(np__1)=X2|~r1_xxreal_0(k4_xcmplx_0(np__1),X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X2,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_37406, plain, (r2_hidden(X1,X2)|X1=X3|X1=X4|X2!=k1_xxreal_1(X3,X4)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_36506, plain, (k4_xcmplx_0(X1)=X2|k4_xcmplx_0(X1)=X3|~r1_xxreal_0(k4_xcmplx_0(X1),X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36665, plain, (k4_xcmplx_0(X1)=X2|k4_xcmplx_0(X1)=X3|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36650, plain, (k7_xcmplx_0(X1,X2)=X3|k7_xcmplx_0(X1,X2)=X4|~r1_xxreal_0(X3,k7_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k7_xcmplx_0(X1,X2))|~r1_xxreal_0(k7_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_37817, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_39017, plain, (k7_square_1(X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(X1))|~r1_xxreal_0(k7_square_1(X1),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38568, negated_conjecture, (esk1_0=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,esk1_0)|~r1_xxreal_0(esk1_0,k32_sin_cos))).#trainneg
cnf(i_0_38581, negated_conjecture, (esk1_0=np__4|~r1_xxreal_0(np__4,esk1_0)|~r1_xxreal_0(esk1_0,np__4))).#trainneg
cnf(i_0_3738, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_36778, plain, (k3_xcmplx_0(esk23_0,np__1)=esk23_0)).#trainneg
cnf(i_0_36779, plain, (k3_xcmplx_0(esk19_0,np__1)=esk19_0)).#trainneg
cnf(i_0_36780, plain, (k3_xcmplx_0(esk18_0,np__1)=esk18_0)).#trainneg
cnf(i_0_36875, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_37345, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4)))).#trainneg
cnf(i_0_38477, negated_conjecture, (k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0)=esk1_0|~r1_xxreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0),esk1_0)|~r1_xxreal_0(esk1_0,k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_38476, negated_conjecture, (k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0)=esk1_0|~r1_xxreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0),esk1_0)|~r1_xxreal_0(esk1_0,k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_38261, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),esk1_0)|~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2),esk1_0)))).#trainneg
cnf(i_0_38215, plain, (esk8_3(X1,X2,X3)=k4_xcmplx_0(np__1)|k4_xcmplx_0(np__1)=X3|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,k4_xcmplx_0(np__1))|~r1_xxreal_0(esk8_3(X1,X2,X3),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X3)|~v1_xreal_0(X3)|~v1_xreal_0(esk8_3(X1,X2,X3))|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_37829, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_square_1(X1))|r1_xxreal_0(k7_square_1(X1),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38892, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0))|~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)))).#trainneg
cnf(i_0_38588, negated_conjecture, (esk1_0=np__3|~r1_xxreal_0(np__3,esk1_0)|~r1_xxreal_0(esk1_0,np__3))).#trainneg
cnf(i_0_38603, negated_conjecture, (esk1_0=np__1|~r1_xxreal_0(np__1,esk1_0)|~r1_xxreal_0(esk1_0,np__1))).#trainneg
cnf(i_0_3739, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_36781, plain, (k3_xcmplx_0(esk10_0,np__1)=esk10_0)).#trainneg
cnf(i_0_36783, plain, (k3_xcmplx_0(np__15,np__1)=np__15)).#trainneg
cnf(i_0_36784, plain, (k3_xcmplx_0(np__9,np__1)=np__9)).#trainneg
cnf(i_0_37536, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_38158, negated_conjecture, (esk1_0=X1|esk1_0=X2|~r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(esk1_0,X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_37725, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0))|r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_38009, plain, (k32_sin_cos=X1|k32_sin_cos=X2|~r1_xxreal_0(X2,k32_sin_cos)|~r1_xxreal_0(X1,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38285, negated_conjecture, (k7_xcmplx_0(k7_square_1(X1),esk1_0)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k7_xcmplx_0(k7_square_1(X1),esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k7_square_1(X1),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k7_square_1(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38385, plain, (k1_real_1(X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38616, negated_conjecture, (esk1_0=np__2|~r1_xxreal_0(np__2,esk1_0)|~r1_xxreal_0(esk1_0,np__2))).#trainneg
cnf(i_0_38402, negated_conjecture, (esk1_0=esk43_0|~r1_xxreal_0(esk43_0,esk1_0)|~r1_xxreal_0(esk1_0,esk43_0))).#trainneg
cnf(i_0_3752, plain, (v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))|~v1_xboole_0(k7_xcmplx_0(np__2,np__3))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_36785, plain, (k3_xcmplx_0(np__7,np__1)=np__7)).#trainneg
cnf(i_0_36786, plain, (k3_xcmplx_0(np__5,np__1)=np__5)).#trainneg
cnf(i_0_36787, plain, (k3_xcmplx_0(np__13,np__1)=np__13)).#trainneg
cnf(i_0_38917, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(X1,esk1_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(X1,esk1_0))|~r1_xxreal_0(k7_xcmplx_0(X1,esk1_0),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_xcmplx_0(X1,esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38284, negated_conjecture, (k7_xcmplx_0(k1_real_1(X1),esk1_0)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k7_xcmplx_0(k1_real_1(X1),esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k1_real_1(X1),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k1_real_1(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38068, plain, (k32_sin_cos=X1|k32_sin_cos=X2|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(X1,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_37812, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k4_xcmplx_0(X1))|r1_xxreal_0(k4_xcmplx_0(X1),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_39173, plain, (X1=X2|X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_37828, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38405, negated_conjecture, (esk1_0=esk34_0|~r1_xxreal_0(esk34_0,esk1_0)|~r1_xxreal_0(esk1_0,esk34_0))).#trainneg
cnf(i_0_3753, plain, (v1_xcmplx_0(k7_xcmplx_0(np__2,np__3))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_36788, plain, (k3_xcmplx_0(np__11,np__1)=np__11)).#trainneg
cnf(i_0_20684, plain, (v1_xreal_0(esk18_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk18_0),k4_numbers))).#trainneg
cnf(i_0_23753, plain, (v5_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_23756, plain, (v5_membered(k1_fdiff_9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_38271, plain, (k4_xcmplx_0(np__1)=k7_xcmplx_0(k32_sin_cos,np__2)|~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_39003, plain, (k1_seq_1(X1,esk51_2(X2,X1))=k4_xcmplx_0(np__1)|k1_seq_1(X1,esk50_2(X2,X1))=k4_xcmplx_0(np__1)|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_seq_1(X1,esk51_2(X2,X1)))|~r1_xxreal_0(k1_seq_1(X1,esk51_2(X2,X1)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_seq_1(X1,esk50_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk50_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk51_2(X2,X1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_36676, plain, (k10_real_1(X1,X2)=X3|k10_real_1(X1,X2)=X4|~r1_xxreal_0(X3,k10_real_1(X1,X2))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k10_real_1(X1,X2))|~r1_xxreal_0(k10_real_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_37764, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_xcmplx_0(k1_real_1(X1),esk1_0))|r1_xxreal_0(k7_xcmplx_0(k1_real_1(X1),esk1_0),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38474, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)=esk1_0|~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),esk1_0)|~r1_xxreal_0(esk1_0,k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)))).#trainneg
cnf(i_0_36517, plain, (k10_real_1(X1,X2)=X3|k10_real_1(X1,X2)=X4|~r1_xxreal_0(k10_real_1(X1,X2),X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k10_real_1(X1,X2))|~r1_xxreal_0(k10_real_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38478, negated_conjecture, (k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)=esk1_0|~r1_xxreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos),esk1_0)|~r1_xxreal_0(esk1_0,k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))|~v1_xreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_38407, negated_conjecture, (esk1_0=esk29_0|~r1_xxreal_0(esk29_0,esk1_0)|~r1_xxreal_0(esk1_0,esk29_0))).#trainneg
cnf(i_0_3754, plain, (v1_xcmplx_0(k7_xcmplx_0(np__2,np__3))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_23759, plain, (v5_membered(k2_sincos10)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_23934, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k10_real_1(k32_sin_cos,np__2))|~v1_xboole_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_24187, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2))|~v1_xboole_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_27678, plain, (~v3_xxreal_0(k32_sin_cos)|~v2_xxreal_0(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_36674, plain, (k8_real_1(X1,X2)=X3|k8_real_1(X1,X2)=X4|~r1_xxreal_0(X3,k8_real_1(X1,X2))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k8_real_1(X1,X2))|~r1_xxreal_0(k8_real_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_37770, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_xcmplx_0(k7_square_1(X1),esk1_0))|r1_xxreal_0(k7_xcmplx_0(k7_square_1(X1),esk1_0),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38373, plain, (k1_seq_1(X1,esk51_2(X2,X1))=k4_xcmplx_0(np__1)|k1_seq_1(X1,esk50_2(X2,X1))=k4_xcmplx_0(np__1)|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k1_seq_1(X1,esk50_2(X2,X1)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_seq_1(X1,esk51_2(X2,X1)),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_seq_1(X1,esk50_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk50_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk51_2(X2,X1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_38269, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(k32_sin_cos,esk1_0)|~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k32_sin_cos,esk1_0)))).#trainneg
cnf(i_0_36673, plain, (k8_real_1(X1,X2)=X3|k8_real_1(X1,X2)=X4|~r1_xxreal_0(X3,k8_real_1(X1,X2))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k8_real_1(X1,X2))|~r1_xxreal_0(k8_real_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36515, plain, (k8_real_1(X1,X2)=X3|k8_real_1(X1,X2)=X4|~r1_xxreal_0(k8_real_1(X1,X2),X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k8_real_1(X1,X2))|~r1_xxreal_0(k8_real_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_38411, negated_conjecture, (esk1_0=esk10_0|~r1_xxreal_0(esk10_0,esk1_0)|~r1_xxreal_0(esk1_0,esk10_0))).#trainneg
cnf(i_0_3755, plain, (v1_int_1(k7_xcmplx_0(np__2,np__3))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_27919, plain, (~v3_xxreal_0(k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_29396, plain, (v4_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_29399, plain, (v4_membered(k1_fdiff_9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_38286, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(np__2,esk1_0)|~r1_xxreal_0(k7_xcmplx_0(np__2,esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__2,esk1_0)))).#trainneg
cnf(i_0_40687, plain, (X1=X3|X1=X4|~r1_xxreal_0(X4,X1)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_40272, plain, (X1=X2|X1=X3|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_3756, plain, (v1_xreal_0(k7_xcmplx_0(np__2,np__3))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_40584, plain, (v1_xcmplx_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_29402, plain, (v4_membered(k2_sincos10)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_33678, plain, (v3_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_39878, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3757, plain, (v3_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v2_xxreal_0(k7_xcmplx_0(np__1,np__2))|~v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_39798, plain, (k32_sin_cos=np__4|~r1_xxreal_0(np__4,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,np__4))).#trainneg
cnf(i_0_39805, plain, (k32_sin_cos=np__3|~r1_xxreal_0(np__3,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,np__3))).#trainneg
cnf(i_0_33679, plain, (v3_membered(k1_fdiff_9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_33680, plain, (v3_membered(k2_sincos10)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_33682, plain, (v2_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_33683, plain, (v2_membered(k1_fdiff_9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_33684, plain, (v2_membered(k2_sincos10)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_40558, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0),k7_xcmplx_0(k32_sin_cos,np__2))|r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_39341, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0))|~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_39299, negated_conjecture, (k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0)=esk1_0|~r1_xxreal_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0),esk1_0)|~r1_xxreal_0(esk1_0,k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_39310, negated_conjecture, (k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0)=esk1_0|~r1_xxreal_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0),esk1_0)|~r1_xxreal_0(esk1_0,k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),esk1_0)))).#trainneg
cnf(i_0_39222, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_39702, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)=k32_sin_cos|~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)))).#trainneg
cnf(i_0_39820, plain, (k32_sin_cos=np__1|~r1_xxreal_0(np__1,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,np__1))).#trainneg
cnf(i_0_39832, plain, (k32_sin_cos=np__2|~r1_xxreal_0(np__2,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,np__2))).#trainneg
cnf(i_0_33685, plain, (v1_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_33687, plain, (v1_membered(k1_fdiff_9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_33688, plain, (v1_membered(k2_sincos10)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_39791, plain, (k10_real_1(X1,X2)=k32_sin_cos|~r1_xxreal_0(k10_real_1(X1,X2),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k10_real_1(X1,X2))|~v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_35198, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X3,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_38263, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36554, plain, (esk6_3(X1,X2,X1)=X3|esk6_3(X1,X2,X1)=X4|~r1_xxreal_0(X3,esk6_3(X1,X2,X1))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,esk6_3(X1,X2,X1))|~r1_xxreal_0(esk6_3(X1,X2,X1),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_36405, plain, (k6_square_1(X1)=X2|k6_square_1(X1)=X3|~r1_xxreal_0(k6_square_1(X1),X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k6_square_1(X1))|~r1_xxreal_0(k6_square_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_39969, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(esk1_0,esk1_0)|~r1_xxreal_0(k7_xcmplx_0(esk1_0,esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(esk1_0,esk1_0))|~v1_xreal_0(k7_xcmplx_0(esk1_0,esk1_0)))).#trainneg
cnf(i_0_36564, plain, (k6_square_1(X1)=X2|k6_square_1(X1)=X3|~r1_xxreal_0(X2,k6_square_1(X1))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k6_square_1(X1))|~r1_xxreal_0(k6_square_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36565, plain, (k5_xcmplx_0(X1)=X2|k5_xcmplx_0(X1)=X3|~r1_xxreal_0(X2,k5_xcmplx_0(X1))|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k5_xcmplx_0(X1))|~r1_xxreal_0(k5_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_37646, plain, (esk7_3(k2_xboole_0(X2,X1),X3,X4)=k2_xboole_0(X2,X1)|r2_hidden(k2_xboole_0(X2,X1),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~m1_subset_1(k2_xboole_0(X2,X1),k1_xxreal_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),X2))).#trainneg
cnf(i_0_38441, negated_conjecture, (esk1_0=np__15|~r1_xxreal_0(np__15,esk1_0)|~r1_xxreal_0(esk1_0,np__15))).#trainneg
cnf(i_0_26016, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_rcomp_1(X4,X3))|~m1_subset_1(k2_rcomp_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_35756, plain, (m1_subset_1(X1,k2_tarski(X2,X3))|m1_subset_1(X1,k1_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_16445, plain, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),X1)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_32350, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X4,X3),k2_tarski(X1,X2))|~r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X4,X3)))).#trainneg
cnf(i_0_30729, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2))))|~m1_subset_1(X2,k2_tarski(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2)))))).#trainneg
cnf(i_0_12181, plain, (r2_hidden(X1,k2_tarski(X2,X3))|r2_hidden(X1,k4_xxreal_1(X3,X2))|v1_xboole_0(k1_xxreal_1(X3,X2))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_34843, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|k4_xcmplx_0(np__1)=esk1_0|~r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_13259, plain, (k3_xcmplx_0(np__4,X1)=k7_xcmplx_0(np__4,k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_34786, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xcmplx_0(X1)|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_38442, negated_conjecture, (esk1_0=np__9|~r1_xxreal_0(np__9,esk1_0)|~r1_xxreal_0(esk1_0,np__9))).#trainneg
cnf(i_0_36514, plain, (k8_real_1(X1,X2)=X3|k8_real_1(X1,X2)=X4|~r1_xxreal_0(k8_real_1(X1,X2),X4)|~r1_xxreal_0(X4,X3)|~r1_xxreal_0(X4,k8_real_1(X1,X2))|~r1_xxreal_0(k8_real_1(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_41403, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_38446, negated_conjecture, (esk1_0=np__7|~r1_xxreal_0(np__7,esk1_0)|~r1_xxreal_0(esk1_0,np__7))).#trainneg
cnf(i_0_3670, plain, (v1_xboole_0(k7_xcmplx_0(np__2,np__3))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_41236, plain, (v1_xreal_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_22698, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__7)))).#trainneg
cnf(i_0_41308, plain, (X1=X2|X1=X3|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_38451, negated_conjecture, (esk1_0=np__5|~r1_xxreal_0(np__5,esk1_0)|~r1_xxreal_0(esk1_0,np__5))).#trainneg
cnf(i_0_3672, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_22714, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__7)))).#trainneg
cnf(i_0_22775, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__13)))).#trainneg
cnf(i_0_22777, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__13)))).#trainneg
cnf(i_0_41971, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xreal_0(X2)|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_41970, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(np__1,X3))|~v1_xreal_0(X2)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_41888, plain, (k3_xcmplx_0(np__4,X1)=k7_xcmplx_0(np__4,k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38454, negated_conjecture, (esk1_0=np__13|~r1_xxreal_0(np__13,esk1_0)|~r1_xxreal_0(esk1_0,np__13))).#trainneg
cnf(i_0_3762, plain, (v1_int_1(k7_xcmplx_0(np__1,np__2))|~v1_int_1(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))))).#trainneg
cnf(i_0_22778, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__13)))).#trainneg
cnf(i_0_22794, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__13)))).#trainneg
cnf(i_0_41416, plain, (k7_xcmplx_0(X1,X2)=k32_sin_cos|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_41965, plain, (k3_xcmplx_0(k7_xcmplx_0(k32_sin_cos,X1),X2)=k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,X1),k7_xcmplx_0(np__1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_41795, plain, (v1_xxreal_0(X1)|r2_hidden(X1,k2_tarski(X3,X2))|~v2_membered(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_41951, plain, (k3_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),X1),X2)=k7_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),X1),k7_xcmplx_0(np__1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_41507, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_41407, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)=k32_sin_cos|~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(k4_xcmplx_0(np__1),esk1_0)))).#trainneg
cnf(i_0_38456, negated_conjecture, (esk1_0=np__11|~r1_xxreal_0(np__11,esk1_0)|~r1_xxreal_0(esk1_0,np__11))).#trainneg
cnf(i_0_22805, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,np__11)))).#trainneg
cnf(i_0_22807, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__11)))).#trainneg
cnf(i_0_22808, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__11)))).#trainneg
cnf(i_0_41969, plain, (k3_xcmplx_0(k7_xcmplx_0(np__2,X1),X2)=k7_xcmplx_0(k7_xcmplx_0(np__2,X1),k7_xcmplx_0(np__1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12579, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(X1,X2))=k7_xcmplx_0(k3_xcmplx_0(np__1,X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13896, plain, (r2_hidden(k2_xboole_0(X1,X2),k2_tarski(X3,X4))|~r2_hidden(k4_xxreal_1(X3,X4),X1)|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r2_hidden(k2_xboole_0(X1,X2),k2_rcomp_1(X3,X4)))).#trainneg
cnf(i_0_18407, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_relset_1(k1_numbers,k2_sincos10))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_28809, plain, (m1_subset_1(esk8_3(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_numbers)|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_38462, negated_conjecture, (esk1_0=np__0|~r1_xxreal_0(np__0,esk1_0)|~r1_xxreal_0(esk1_0,np__0))).#trainneg
cnf(i_0_22824, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,np__11)))).#trainneg
cnf(i_0_25585, plain, (v1_xreal_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_35995, plain, (esk6_3(k2_xboole_0(X1,X2),X3,X4)=k2_xboole_0(X1,X2)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~r1_xxreal_0(X4,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(k2_xboole_0(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),X2))).#trainneg
cnf(i_0_26062, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_rcomp_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r2_hidden(k2_rcomp_1(X4,X3),k4_xxreal_1(X1,X2)))).#trainneg
cnf(i_0_33183, plain, (v1_xreal_0(k7_xcmplx_0(k3_xcmplx_0(X1,X2),k32_sin_cos))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43490, plain, (v1_xreal_0(k7_xcmplx_0(np__3,k32_sin_cos)))).#trainneg
cnf(i_0_43521, plain, (v1_xreal_0(k7_xcmplx_0(np__4,k32_sin_cos)))).#trainneg
cnf(i_0_42061, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__7))))).#trainneg
cnf(i_0_42846, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__13))))).#trainneg
cnf(i_0_43120, plain, (~v1_xboole_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__11))))).#trainneg
cnf(i_0_43560, plain, (v1_xreal_0(k7_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3674, plain, (v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3)))).#trainneg
cnf(i_0_43270, plain, (r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~r2_hidden(k4_xxreal_1(X2,X3),X1))).#trainneg
cnf(i_0_42111, plain, (esk6_3(X1,X2,X3)=X3|esk6_3(X1,X2,X3)=X4|~r1_xxreal_0(X3,esk6_3(X1,X2,X3))|~r1_xxreal_0(esk6_3(X1,X2,X3),X4)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(esk6_3(X1,X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_43378, plain, (k1_xxreal_1(X1,X2)=X3|k1_xxreal_1(X1,X2)=X4|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_rcomp_1(X3,X4))|~r2_hidden(k2_rcomp_1(X3,X4),k2_rcomp_1(X1,X2)))).#trainneg
cnf(i_0_3752, plain, (v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__2),np__3))|~v1_xboole_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_25586, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_25603, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_43364, plain, (esk6_3(k1_xxreal_1(X1,X2),X3,X4)=k1_xxreal_1(X1,X2)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~m1_subset_1(k1_xxreal_1(X1,X2),k1_xxreal_1(X4,X3))|~r2_hidden(k2_rcomp_1(X4,X3),k4_xxreal_1(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_42261, plain, (esk6_3(X2,X1,X3)=X1|X1=X4|~r1_xxreal_0(esk6_3(X2,X1,X3),X1)|~r1_xxreal_0(X1,X4)|~r1_xxreal_0(esk6_3(X2,X1,X3),X4)|~v1_xreal_0(X4)|~v1_xreal_0(esk6_3(X2,X1,X3))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k2_tarski(X1,X3)))).#trainneg
cnf(i_0_43266, plain, (r2_hidden(k1_xxreal_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X3,X4)|~v1_xxreal_0(X4)|~v1_xxreal_0(X3)|~r2_hidden(k1_xxreal_1(X1,X2),k2_rcomp_1(X3,X4))|~r2_hidden(k4_xxreal_1(X3,X4),k2_tarski(X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_23586, plain, (~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(k2_tarski(X1,k1_xxreal_1(X2,X3)),k2_tarski(X2,X3)))).#trainneg
cnf(i_0_41544, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(np__3,esk1_0)|~r1_xxreal_0(k7_xcmplx_0(np__3,esk1_0),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__3,esk1_0))|~v1_xreal_0(k7_xcmplx_0(np__3,esk1_0)))).#trainneg
cnf(i_0_35089, plain, (m1_subset_1(X1,k1_xxreal_1(X2,X3))|m1_subset_1(X1,k4_xxreal_1(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r1_xxreal_0(X3,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_13321, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),k7_xcmplx_0(X1,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25617, plain, (v1_xreal_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_25618, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_26250, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_40566, plain, (k7_xcmplx_0(X1,X2)=X3|k7_xcmplx_0(X1,X2)=X4|~r1_xxreal_0(k7_xcmplx_0(X1,X2),X3)|~r1_xxreal_0(X3,k7_xcmplx_0(X1,X2))|~r1_xxreal_0(k7_xcmplx_0(X1,X2),X4)|~r1_xxreal_0(X3,X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27611, negated_conjecture, (r1_xxreal_0(esk1_0,k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)))|r1_xxreal_0(k3_xcmplx_0(k32_sin_cos,k10_real_1(np__3,np__4)),esk1_0)|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_36484, plain, (k7_xcmplx_0(X1,k32_sin_cos)=X2|k7_xcmplx_0(X1,k32_sin_cos)=X3|~r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),X3)|~r1_xxreal_0(X3,X2)|~r1_xxreal_0(X3,k7_xcmplx_0(X1,k32_sin_cos))|~r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_41967, plain, (k3_xcmplx_0(k7_xcmplx_0(np__3,X1),X2)=k7_xcmplx_0(k7_xcmplx_0(np__3,X1),k7_xcmplx_0(np__1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_30815, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__2)|r1_xxreal_0(np__2,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25635, plain, (v1_xxreal_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_43422, plain, (v1_xreal_0(k7_xcmplx_0(esk43_0,k32_sin_cos)))).#trainneg
cnf(i_0_13886, plain, (v1_xxreal_0(X1)|r2_hidden(X1,k2_tarski(X2,X3))|~v2_membered(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_11418, plain, (v1_xreal_0(X1)|~v3_membered(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_18958, plain, (r1_xxreal_0(esk6_3(X1,X2,X1),esk6_3(X1,X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_17092, negated_conjecture, (v1_xxreal_0(k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3769, plain, (v1_int_1(k7_xcmplx_0(np__1,np__3))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__1),np__3)))).#trainneg
cnf(i_0_43423, plain, (v1_xreal_0(k7_xcmplx_0(esk37_0,k32_sin_cos)))).#trainneg
cnf(i_0_43425, plain, (v1_xreal_0(k7_xcmplx_0(esk34_0,k32_sin_cos)))).#trainneg
cnf(i_0_43427, plain, (v1_xreal_0(k7_xcmplx_0(esk29_0,k32_sin_cos)))).#trainneg
cnf(i_0_30868, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(k4_rcomp_1(X2,X3),X1))).#trainneg
cnf(i_0_28381, plain, (k7_xcmplx_0(k10_real_1(X1,X2),k4_xcmplx_0(np__1))=k4_xcmplx_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_24662, plain, (r1_xxreal_0(esk8_3(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_25340, plain, (v1_xboole_0(X1)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_xxreal_1(X1,X2),X1))).#trainneg
cnf(i_0_43430, plain, (v1_xreal_0(k7_xcmplx_0(esk23_0,k32_sin_cos)))).#trainneg
cnf(i_0_43431, plain, (v1_xreal_0(k7_xcmplx_0(esk19_0,k32_sin_cos)))).#trainneg
cnf(i_0_43435, plain, (v1_xreal_0(k7_xcmplx_0(esk10_0,k32_sin_cos)))).#trainneg
cnf(i_0_41887, plain, (k3_xcmplx_0(np__4,X1)=k7_xcmplx_0(np__4,k7_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_39782, plain, (k8_real_1(X1,X2)=k32_sin_cos|~r1_xxreal_0(k8_real_1(X1,X2),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k8_real_1(X1,X2))|~v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13899, plain, (m1_subset_1(X1,k4_xxreal_1(X2,X3))|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_34240, negated_conjecture, (m1_subset_1(k7_xcmplx_0(k4_xcmplx_0(esk1_0),np__1),k1_numbers))).#trainneg
cnf(i_0_24744, plain, (r2_hidden(k1_rcomp_1(X1,X2),k2_tarski(X3,X4))|~r1_xxreal_0(X4,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r2_hidden(k1_rcomp_1(X1,X2),k1_rcomp_1(X4,X3))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r2_hidden(k4_xxreal_1(X4,X3),k2_tarski(X1,X2)))).#trainneg
cnf(i_0_31453, plain, (esk7_3(X1,X2,X3)=X1|r2_hidden(X1,k2_tarski(X3,X2))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_43438, plain, (v1_xreal_0(k7_xcmplx_0(np__15,k32_sin_cos)))).#trainneg
cnf(i_0_43439, plain, (v1_xreal_0(k7_xcmplx_0(np__9,k32_sin_cos)))).#trainneg
cnf(i_0_43440, plain, (v1_xreal_0(k7_xcmplx_0(np__7,k32_sin_cos)))).#trainneg
cnf(i_0_9574, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_9575, plain, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_9735, plain, (v3_ordinal1(k32_sin_cos)|~v3_ordinal1(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_17570, plain, (m1_subset_1(k7_xcmplx_0(np__1,k7_square_1(X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25085, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_38650, plain, (np__2=X1|np__2=X2|~r1_xxreal_0(X2,np__2)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43820, plain, (~v1_xboole_0(k7_xcmplx_0(np__2,np__3)))).#trainneg
cnf(i_0_44364, plain, (v1_xreal_0(X1)|~v3_membered(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_43670, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),k7_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_44980, plain, (r1_xxreal_0(k7_xcmplx_0(np__1,k7_square_1(X1)),k7_xcmplx_0(np__1,k7_square_1(X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_43896, plain, (esk6_3(X1,X1,X1)=X1|X1=X2|~r1_xxreal_0(esk6_3(X1,X1,X1),X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43679, plain, (r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(X1,k32_sin_cos))|r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43642, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_44029, negated_conjecture, (k4_xcmplx_0(esk1_0)=X1|k4_xcmplx_0(esk1_0)=X2|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X2)|~r1_xxreal_0(X2,k4_xcmplx_0(esk1_0))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_44365, plain, (v1_xreal_0(X1)|~v3_membered(k2_xboole_0(X3,X2))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_43676, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),np__2)|r1_xxreal_0(np__2,k7_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3776, plain, (v1_int_1(k7_xcmplx_0(np__1,np__4))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_9736, plain, (v1_funct_1(k32_sin_cos)|~v4_funct_1(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_9737, plain, (v7_ordinal1(k32_sin_cos)|~v6_membered(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_9738, plain, (v1_int_1(k32_sin_cos)|~v5_membered(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_44405, negated_conjecture, (r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(esk1_0,X1))|r1_xxreal_0(k7_xcmplx_0(esk1_0,X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_44174, plain, (r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos))).#trainneg
cnf(i_0_44297, plain, (r1_xxreal_0(np__2,k4_xcmplx_0(X1))|r1_xxreal_0(k4_xcmplx_0(X1),np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_43680, plain, (r1_xxreal_0(np__4,k7_xcmplx_0(X1,k32_sin_cos))|r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),np__4)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43650, plain, (esk6_3(X1,X2,X2)=X2|~r1_xxreal_0(X2,X2)|~v1_xreal_0(esk6_3(X1,X2,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X2)))).#trainneg
cnf(i_0_45169, plain, (k8_real_1(X1,X2)=np__2|~r1_xxreal_0(k8_real_1(X1,X2),np__2)|~r1_xxreal_0(np__2,k8_real_1(X1,X2))|~v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9739, plain, (v1_rat_1(k32_sin_cos)|~v4_membered(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_9743, plain, (v1_relat_1(k32_sin_cos)|~v4_funct_1(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_45188, plain, (k10_real_1(X1,X2)=np__2|~r1_xxreal_0(k10_real_1(X1,X2),np__2)|~r1_xxreal_0(np__2,k10_real_1(X1,X2))|~v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_43646, plain, (X1=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_43947, plain, (k7_xcmplx_0(X1,k32_sin_cos)=k32_sin_cos|~r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(k7_xcmplx_0(X1,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_44600, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_43675, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),np__1)|r1_xxreal_0(np__1,k7_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45207, plain, (k1_seq_1(X1,esk50_2(X2,X1))=np__2|k1_seq_1(X1,esk51_2(X2,X1))=np__2|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k1_seq_1(X1,esk50_2(X2,X1)),np__2)|~r1_xxreal_0(k1_seq_1(X1,esk51_2(X2,X1)),np__2)|~r1_xxreal_0(np__2,k1_seq_1(X1,esk50_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk50_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk51_2(X2,X1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_44397, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(esk1_0,X1),k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3780, plain, (v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2))|~v1_xboole_0(k7_xcmplx_0(np__3,np__2))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2)))).#trainneg
cnf(i_0_10612, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_10613, plain, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_11088, plain, (v3_ordinal1(k32_sin_cos)|~v3_ordinal1(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_11089, plain, (v1_funct_1(k32_sin_cos)|~v4_funct_1(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_46165, plain, (esk6_3(X2,X1,X3)=X1|~r1_xxreal_0(esk6_3(X2,X1,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_46017, plain, (esk6_3(X1,X1,X2)=X2|~r1_xxreal_0(X2,esk6_3(X1,X1,X2))|~v1_xreal_0(esk6_3(X1,X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45437, plain, (r1_xxreal_0(k7_xcmplx_0(k7_square_1(X1),k32_sin_cos),k7_xcmplx_0(k7_square_1(X1),k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3781, plain, (v1_xcmplx_0(k7_xcmplx_0(np__3,np__2))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2)))).#trainneg
cnf(i_0_11090, plain, (v7_ordinal1(k32_sin_cos)|~v6_membered(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_11091, plain, (v1_int_1(k32_sin_cos)|~v5_membered(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_11092, plain, (v1_rat_1(k32_sin_cos)|~v4_membered(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_11096, plain, (v1_relat_1(k32_sin_cos)|~v4_funct_1(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_46013, plain, (esk6_3(X1,X2,X3)=X3|~r1_xxreal_0(X3,esk6_3(X1,X2,X3))|~v1_xreal_0(esk6_3(X1,X2,X3))|~v1_xreal_0(X3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_46144, plain, (k1_seq_1(X1,esk51_2(X2,X1))=k1_seq_1(X1,esk50_2(X2,X1))|v5_valued_0(k2_partfun1(k1_numbers,k1_numbers,X1,X2))|~r1_xxreal_0(k1_seq_1(X1,esk50_2(X2,X1)),k1_seq_1(X1,esk51_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk51_2(X2,X1)))|~v1_xreal_0(k1_seq_1(X1,esk50_2(X2,X1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_45436, plain, (r1_xxreal_0(k7_xcmplx_0(k1_real_1(X1),k32_sin_cos),k7_xcmplx_0(k1_real_1(X1),k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45400, plain, (r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(X1,k32_sin_cos),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(X1,k32_sin_cos),k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45432, plain, (r1_xxreal_0(k7_xcmplx_0(k10_real_1(X1,X2),k32_sin_cos),k7_xcmplx_0(k10_real_1(X1,X2),k32_sin_cos))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45411, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(esk1_0),k32_sin_cos),k7_xcmplx_0(k4_xcmplx_0(esk1_0),k32_sin_cos)))).#trainneg
cnf(i_0_45421, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),k32_sin_cos),k7_xcmplx_0(k4_xcmplx_0(np__1),k32_sin_cos)))).#trainneg
cnf(i_0_45424, plain, (r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(X1),k32_sin_cos),k7_xcmplx_0(k4_xcmplx_0(X1),k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_46015, plain, (esk6_3(X1,X2,X3)=X3|~r1_xxreal_0(X3,esk6_3(X1,X2,X3))|~v1_xreal_0(esk6_3(X1,X2,X3))|~v1_xreal_0(X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_47012, negated_conjecture, (k7_xcmplx_0(esk1_0,X1)=k32_sin_cos|~r1_xxreal_0(k7_xcmplx_0(esk1_0,X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(esk1_0,X1))|~v1_xreal_0(k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_46530, plain, (X1=X2|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X2))|~r1_xxreal_0(X2,X2))).#trainneg
cnf(i_0_3783, plain, (v1_int_1(k7_xcmplx_0(np__3,np__2))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2)))).#trainneg
cnf(i_0_23241, plain, (~r1_xxreal_0(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k7_xcmplx_0(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_23890, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_45396, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__2,esk1_0),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(np__2,esk1_0),k32_sin_cos)))).#trainneg
cnf(i_0_45445, plain, (k7_xcmplx_0(np__1,k7_square_1(X1))=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__1,k7_square_1(X1)))|~r1_xxreal_0(k7_xcmplx_0(np__1,k7_square_1(X1)),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_xcmplx_0(np__1,k7_square_1(X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_46751, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(esk1_0,X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(esk1_0,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_47007, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(esk1_0,X1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(esk1_0,X1))|~r1_xxreal_0(k7_xcmplx_0(esk1_0,X1),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_46406, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45431, plain, (r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k32_sin_cos),k7_xcmplx_0(k32_sin_cos,k32_sin_cos)))).#trainneg
cnf(i_0_45438, plain, (r1_xxreal_0(k7_xcmplx_0(np__2,k32_sin_cos),k7_xcmplx_0(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_46531, plain, (esk6_3(X1,X2,X2)=X2|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X2))|~m1_subset_1(esk6_3(X1,X2,X2),k1_numbers))).#trainneg
cnf(i_0_23891, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_23933, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k5_numbers))).#trainneg
cnf(i_0_23935, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,np__3)),np__4),k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k1_numbers))).#trainneg
cnf(i_0_24186, plain, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k10_real_1(k32_sin_cos,np__2))|~m1_subset_1(k10_real_1(k32_sin_cos,np__2),k5_numbers))).#trainneg
cnf(i_0_45367, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,esk1_0),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,esk1_0),k32_sin_cos)))).#trainneg
cnf(i_0_46166, plain, (esk6_3(X2,X1,X3)=X1|~r1_xxreal_0(esk6_3(X2,X1,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k2_tarski(X1,X3)))).#trainneg
cnf(i_0_46880, plain, (k7_xcmplx_0(k7_xcmplx_0(np__2,X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(np__2,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45260, plain, (v1_xreal_0(X1)|r2_hidden(X1,k2_tarski(X3,X2))|~v3_membered(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_45291, plain, (v1_xreal_0(X1)|r2_hidden(X1,k2_tarski(X3,X2))|~v3_membered(k2_rcomp_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_46872, plain, (k7_xcmplx_0(k7_xcmplx_0(k32_sin_cos,X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_46550, plain, (k7_xcmplx_0(X1,X2)=np__2|~r1_xxreal_0(k7_xcmplx_0(X1,X2),np__2)|~r1_xxreal_0(np__2,k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45579, plain, (r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(X1,k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_39680, plain, (k32_sin_cos=np__5|~r1_xxreal_0(np__5,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,np__5))).#trainneg
cnf(i_0_3787, plain, (v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4))|~v1_xboole_0(k7_xcmplx_0(np__3,np__4))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4)))).#trainneg
cnf(i_0_47015, negated_conjecture, (k7_xcmplx_0(esk1_0,X1)=np__2|~r1_xxreal_0(k7_xcmplx_0(esk1_0,X1),np__2)|~r1_xxreal_0(np__2,k7_xcmplx_0(esk1_0,X1))|~v1_xreal_0(k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_47064, plain, (esk6_3(X1,X1,X1)=X1|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers))).#trainneg
cnf(i_0_47316, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_3698, plain, (v1_xboole_0(k7_xcmplx_0(np__3,np__2))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__2)))).#trainneg
cnf(i_0_47035, plain, (~v1_xboole_0(k7_xcmplx_0(np__3,np__2)))).#trainneg
cnf(i_0_48388, plain, (~v1_xboole_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_47492, plain, (X1=X2|r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3788, plain, (v1_xcmplx_0(k7_xcmplx_0(np__3,np__4))|~v1_xreal_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4)))).#trainneg
cnf(i_0_47624, negated_conjecture, (k7_xcmplx_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(np__1))=esk1_0)).#trainneg
cnf(i_0_47058, plain, (esk6_3(X1,X2,X2)=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(esk6_3(X1,X2,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_47469, plain, (X1=X3|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_47523, plain, (k7_xcmplx_0(np__1,k7_square_1(X1))=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(np__1,k7_square_1(X1)))|~r1_xxreal_0(k7_xcmplx_0(np__1,k7_square_1(X1)),k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k7_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_47846, plain, (X1=X2|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k1_rcomp_1(X2,X2))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_47879, plain, (X1=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_39691, plain, (k32_sin_cos=np__0|~r1_xxreal_0(np__0,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,np__0))).#trainneg
cnf(i_0_3790, plain, (v1_int_1(k7_xcmplx_0(np__3,np__4))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4)))).#trainneg
cnf(i_0_47810, plain, (r1_xxreal_0(np__2,k8_real_1(X1,X2))|r1_xxreal_0(k8_real_1(X1,X2),np__2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_47852, plain, (esk6_3(X1,X2,X2)=X2|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X2))|~r1_xxreal_0(X2,X2))).#trainneg
cnf(i_0_45336, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(esk1_0,k32_sin_cos),k7_xcmplx_0(esk1_0,k32_sin_cos)))).#trainneg
cnf(i_0_9572, plain, (v1_xboole_0(k32_sin_cos)|~m1_subset_1(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_10916, negated_conjecture, (v3_ordinal1(esk1_0)|~v3_ordinal1(k1_rcomp_1(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_13385, plain, (v3_ordinal1(k4_xcmplx_0(np__1))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_13388, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_46014, plain, (esk6_3(X1,X2,X3)=X3|~r1_xxreal_0(X3,esk6_3(X1,X2,X3))|~v1_xreal_0(esk6_3(X1,X2,X3))|~v1_xreal_0(X3)|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_46168, plain, (esk6_3(X2,X1,X3)=X1|~r1_xxreal_0(esk6_3(X2,X1,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X2,X1,X3))|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k2_tarski(X3,X1)))).#trainneg
cnf(i_0_45706, negated_conjecture, (k7_xcmplx_0(k1_real_1(X1),esk1_0)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_xcmplx_0(k1_real_1(X1),esk1_0))|~r1_xxreal_0(k7_xcmplx_0(k1_real_1(X1),esk1_0),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_xcmplx_0(k1_real_1(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45710, negated_conjecture, (k4_xcmplx_0(esk1_0)=k7_xcmplx_0(X1,esk1_0)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_xcmplx_0(X1,esk1_0))|~r1_xxreal_0(k7_xcmplx_0(X1,esk1_0),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_xcmplx_0(X1,esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45583, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers))).#trainneg
cnf(i_0_46813, plain, (k7_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X2),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_46203, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(esk1_0,esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(esk1_0,esk1_0)))).#trainneg
cnf(i_0_3794, plain, (v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3))|~v1_xboole_0(k7_xcmplx_0(np__4,np__3))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3)))).#trainneg
cnf(i_0_13389, plain, (v1_rat_1(k4_xcmplx_0(np__1))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_13707, plain, (v3_ordinal1(k4_xcmplx_0(np__2))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_13710, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_13711, plain, (v1_rat_1(k4_xcmplx_0(np__2))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_46874, plain, (k7_xcmplx_0(k7_xcmplx_0(np__4,X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(np__4,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_43628, plain, (m1_subset_1(k7_xcmplx_0(np__4,k32_sin_cos),k1_numbers))).#trainneg
cnf(i_0_47650, negated_conjecture, (k4_xcmplx_0(np__1)=k7_xcmplx_0(esk1_0,esk1_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(esk1_0,esk1_0))|~r1_xxreal_0(k7_xcmplx_0(esk1_0,esk1_0),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_45262, plain, (v1_xreal_0(X1)|r2_hidden(X1,k2_tarski(X3,X2))|~v3_membered(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_45170, plain, (k8_real_1(X1,X2)=np__2|~r1_xxreal_0(k8_real_1(X1,X2),np__2)|~r1_xxreal_0(np__2,k8_real_1(X1,X2))|~v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_46871, plain, (k7_xcmplx_0(k7_xcmplx_0(k8_real_1(X1,X2),X3),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(k8_real_1(X1,X2),X3))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_47272, plain, (k4_xcmplx_0(X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_45382, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__1,esk1_0),k32_sin_cos),k7_xcmplx_0(k7_xcmplx_0(np__1,esk1_0),k32_sin_cos)))).#trainneg
cnf(i_0_38404, negated_conjecture, (esk1_0=esk35_0|~r1_xxreal_0(esk35_0,esk1_0)|~r1_xxreal_0(esk1_0,esk35_0)|~v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_18503, plain, (v3_ordinal1(k4_xcmplx_0(np__4))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_18506, plain, (v1_int_1(k4_xcmplx_0(np__4))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_45433, plain, (r1_xxreal_0(k7_xcmplx_0(np__4,k32_sin_cos),k7_xcmplx_0(np__4,k32_sin_cos)))).#trainneg
cnf(i_0_46873, plain, (k7_xcmplx_0(k7_xcmplx_0(k10_real_1(X1,X2),X3),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(k10_real_1(X1,X2),X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_47884, plain, (esk6_3(X1,X2,X2)=X2|~r1_xxreal_0(X2,X2)|~v1_xreal_0(esk6_3(X1,X2,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X2)))).#trainneg
cnf(i_0_44179, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_47774, plain, (r1_xxreal_0(np__2,k7_xcmplx_0(X1,k4_xcmplx_0(np__1)))|r1_xxreal_0(k7_xcmplx_0(X1,k4_xcmplx_0(np__1)),np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38406, negated_conjecture, (esk1_0=esk30_0|~r1_xxreal_0(esk30_0,esk1_0)|~r1_xxreal_0(esk1_0,esk30_0)|~v1_xreal_0(esk30_0))).#trainneg
cnf(i_0_18507, plain, (v1_rat_1(k4_xcmplx_0(np__4))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_19448, negated_conjecture, (v3_ordinal1(k4_xcmplx_0(esk1_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_19451, negated_conjecture, (v1_int_1(k4_xcmplx_0(esk1_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_19452, negated_conjecture, (v1_rat_1(k4_xcmplx_0(esk1_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_47642, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(esk1_0,X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(esk1_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48733, plain, (X1=X2|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48947, plain, (X2=X1|~r1_xxreal_0(X2,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_rcomp_1(X2,X2)))).#trainneg
cnf(i_0_3705, plain, (~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4)))).#trainneg
cnf(i_0_48748, plain, (~v1_xboole_0(k7_xcmplx_0(np__4,np__3)))).#trainneg
cnf(i_0_20847, plain, (v1_xboole_0(k4_xcmplx_0(esk10_0))|~v1_xboole_0(esk10_0))).#trainneg
cnf(i_0_3797, plain, (v1_int_1(k7_xcmplx_0(np__4,np__3))|~v1_int_1(k7_xcmplx_0(k4_xcmplx_0(np__4),np__3)))).#trainneg
cnf(i_0_24116, plain, (v2_xxreal_0(k4_xcmplx_0(np__15))|~v3_xxreal_0(np__15))).#trainneg
cnf(i_0_48802, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48891, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_48901, plain, (k4_xcmplx_0(k32_sin_cos)=X1|r1_xxreal_0(k32_sin_cos,k1_real_1(X1))|~r1_xxreal_0(X1,k4_xcmplx_0(k32_sin_cos))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_49658, plain, (esk6_3(X1,X2,X2)=X2|~v1_xreal_0(X2)|~m1_subset_1(X1,k2_tarski(X2,X2))|~m1_subset_1(esk6_3(X1,X2,X2),k1_numbers))).#trainneg
cnf(i_0_49579, plain, (k4_xcmplx_0(k10_real_1(X1,X2))=np__1|~r1_xxreal_0(k4_xcmplx_0(k10_real_1(X1,X2)),np__1)|~r1_xxreal_0(k10_real_1(X1,X2),k4_xcmplx_0(np__1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48886, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_48695, plain, (k7_xcmplx_0(X1,k32_sin_cos)=X2|r1_xxreal_0(X2,k7_xcmplx_0(X1,k32_sin_cos))|r1_xxreal_0(k7_xcmplx_0(X1,k32_sin_cos),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_49421, plain, (v1_xreal_0(X1)|r2_hidden(X1,k2_tarski(X2,X3))|~r1_xxreal_0(X3,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X1,k1_xxreal_1(X3,X2))|~m1_subset_1(k4_xxreal_1(X3,X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3800, plain, (v2_membered(k10_xtuple_0(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24408, plain, (v1_xboole_0(k4_xcmplx_0(esk37_0))|~v1_xboole_0(esk37_0))).#trainneg
cnf(i_0_24554, plain, (v1_xboole_0(k4_xcmplx_0(esk23_0))|~v1_xboole_0(esk23_0))).#trainneg
cnf(i_0_24950, plain, (v3_ordinal1(k4_xcmplx_0(np__5))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_48803, plain, (k4_xcmplx_0(X1)=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(X1),k32_sin_cos)|~r1_xxreal_0(X1,k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_49586, plain, (k4_xcmplx_0(X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45708, negated_conjecture, (k7_xcmplx_0(k7_square_1(X1),esk1_0)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_xcmplx_0(k7_square_1(X1),esk1_0))|~r1_xxreal_0(k7_xcmplx_0(k7_square_1(X1),esk1_0),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_xcmplx_0(k7_square_1(X1),esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_44616, plain, (v1_xboole_0(X1)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k1_rcomp_1(X1,X2),X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_46869, plain, (k7_xcmplx_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_49550, plain, (k4_xcmplx_0(k7_xcmplx_0(X1,X2))=np__1|~r1_xxreal_0(k4_xcmplx_0(k7_xcmplx_0(X1,X2)),np__1)|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45259, plain, (v1_xreal_0(X1)|r2_hidden(X1,k2_tarski(X3,X2))|~v3_membered(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~r2_hidden(X1,k2_rcomp_1(X2,X3)))).#trainneg
cnf(i_0_3821, plain, (~v1_xboole_0(k5_xcmplx_0(X1))|~v1_xxreal_0(k5_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24953, plain, (v1_int_1(k4_xcmplx_0(np__5))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_24954, plain, (v1_rat_1(k4_xcmplx_0(np__5))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_48900, plain, (k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2))=X1|r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),k1_real_1(X1))|~r1_xxreal_0(X1,k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48836, negated_conjecture, (k3_xcmplx_0(esk1_0,X1)=k7_xcmplx_0(esk1_0,k7_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_49219, plain, (k7_xcmplx_0(k7_xcmplx_0(np__1,X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_50026, plain, (X1=X2|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3823, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50510, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50084, plain, (X1=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_50267, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_50372, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3824, plain, (v2_xxreal_0(k7_xcmplx_0(np__1,X1))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50081, plain, (X1=np__1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50957, plain, (k4_xcmplx_0(X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50082, plain, (X1=np__2|~r1_xxreal_0(np__2,X1)|~r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50089, plain, (esk6_3(X1,k4_xcmplx_0(X2),X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(esk6_3(X1,k4_xcmplx_0(X2),X1)))|~v1_xreal_0(X2)|~m1_subset_1(esk6_3(X1,k4_xcmplx_0(X2),X1),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_3825, plain, (~v2_xxreal_0(k5_xcmplx_0(X1))|~v1_xxreal_0(k5_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51013, negated_conjecture, (~v2_xxreal_0(esk1_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_50187, plain, (~v1_xxreal_0(k4_xcmplx_0(np__15))|~v3_xxreal_0(np__15))).#trainneg
cnf(i_0_50030, plain, (k4_xcmplx_0(esk10_0)=np__0|~v1_xboole_0(esk10_0))).#trainneg
cnf(i_0_51005, plain, (~v2_xxreal_0(esk37_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_51009, plain, (~v2_xxreal_0(esk23_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_51010, plain, (~v2_xxreal_0(esk19_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_50078, plain, (X1=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50724, plain, (k32_sin_cos=np__0|~r1_xxreal_0(np__0,k32_sin_cos)|~r1_xxreal_0(np__0,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_3826, plain, (~v1_xboole_0(k5_xcmplx_0(X1))|~v1_xxreal_0(k5_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_50674, plain, (k4_xcmplx_0(esk37_0)=np__0|~v1_xboole_0(esk37_0))).#trainneg
cnf(i_0_50689, plain, (k4_xcmplx_0(esk23_0)=np__0|~v1_xboole_0(esk23_0))).#trainneg
cnf(i_0_51012, plain, (~v2_xxreal_0(esk10_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_51022, negated_conjecture, (~v2_xxreal_0(k7_xcmplx_0(np__1,esk1_0))|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_50861, plain, (X1=np__1|~r1_xxreal_0(X1,np__1)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(np__1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48734, plain, (X1=X2|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_46589, negated_conjecture, (k7_xcmplx_0(esk1_0,k32_sin_cos)=k32_sin_cos|~r1_xxreal_0(k7_xcmplx_0(esk1_0,k32_sin_cos),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k7_xcmplx_0(esk1_0,k32_sin_cos)))).#trainneg
cnf(i_0_49602, negated_conjecture, (esk1_0=esk35_0|~r1_xxreal_0(esk35_0,esk1_0)|~r1_xxreal_0(esk1_0,esk35_0)|~m1_subset_1(esk35_0,k1_numbers))).#trainneg
cnf(i_0_3828, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51011, plain, (~v2_xxreal_0(esk18_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk18_0))|~v1_xreal_0(esk18_0))).#trainneg
cnf(i_0_26362, plain, (v3_ordinal1(k4_xcmplx_0(np__9))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_26365, plain, (v1_int_1(k4_xcmplx_0(np__9))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_45264, plain, (v1_xreal_0(X1)|r2_hidden(X1,k2_tarski(X3,X2))|~v3_membered(k4_xxreal_1(X2,X3))|~r1_xxreal_0(X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_46875, plain, (k7_xcmplx_0(k7_xcmplx_0(np__3,X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(np__3,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_46867, plain, (k7_xcmplx_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),X1),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_xcmplx_0(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_46063, plain, (k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos)=k7_xcmplx_0(k32_sin_cos,np__2)|~r1_xxreal_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_47818, plain, (r1_xxreal_0(np__2,k7_square_1(X1))|r1_xxreal_0(k7_square_1(X1),np__2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48696, plain, (k7_xcmplx_0(X1,X2)=X3|r1_xxreal_0(X3,k7_xcmplx_0(X1,X2))|r1_xxreal_0(k7_xcmplx_0(X1,X2),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51981, plain, (~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_51403, plain, (X1=k7_square_1(X2)|~r1_xxreal_0(X1,k7_square_1(X2))|~r1_xxreal_0(k7_square_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_51274, negated_conjecture, (X1=esk1_0|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(esk1_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51382, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3829, plain, (v3_xxreal_0(k7_xcmplx_0(np__1,X1))|~v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51958, plain, (~v3_xxreal_0(esk29_0))).#trainneg
cnf(i_0_51965, plain, (~v3_xxreal_0(np__15))).#trainneg
cnf(i_0_51387, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_51402, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_51950, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_51396, plain, (X1=k32_sin_cos|~r1_xxreal_0(X1,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51684, plain, (X1=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_52061, plain, (k32_sin_cos=np__0|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),np__0)|~r1_xxreal_0(k32_sin_cos,np__0))).#trainneg
cnf(i_0_3845, plain, (v3_membered(k10_xtuple_0(X1))|~v5_relat_1(X1,k3_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_51386, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_51397, plain, (X1=k10_real_1(X2,X3)|~r1_xxreal_0(X1,k10_real_1(X2,X3))|~r1_xxreal_0(k10_real_1(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_51605, plain, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_51872, plain, (k4_xcmplx_0(X1)=np__1|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51394, plain, (X1=k8_real_1(X2,X3)|~r1_xxreal_0(X1,k8_real_1(X2,X3))|~r1_xxreal_0(k8_real_1(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_44359, plain, (v1_xreal_0(X1)|~v3_membered(X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_3846, plain, (v8_valued_0(X1)|~v1_funct_1(X1)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_52134, negated_conjecture, (~v1_xboole_0(esk1_0)|~v2_xxreal_0(esk1_0))).#trainneg
cnf(i_0_51974, plain, (~v3_xxreal_0(k7_xcmplx_0(np__1,k32_sin_cos)))).#trainneg
cnf(i_0_51433, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_52898, plain, (esk6_3(k1_real_1(X1),k1_real_1(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(k1_real_1(X1),k1_real_1(X1),X2),k4_xcmplx_0(X1))|~v1_xreal_0(esk6_3(k1_real_1(X1),k1_real_1(X1),X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_3854, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,esk15_0))).#trainneg
cnf(i_0_52897, plain, (esk6_3(X2,k1_real_1(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(X2,k1_real_1(X1),X2),k4_xcmplx_0(X1))|~v1_xreal_0(esk6_3(X2,k1_real_1(X1),X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_53197, plain, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_52652, negated_conjecture, (X1=esk1_0|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3856, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk4_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_52986, plain, (X1=X2|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3857, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_51709, plain, (~v2_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_51955, plain, (~v3_xxreal_0(esk43_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk43_0)))).#trainneg
cnf(i_0_51956, plain, (~v3_xxreal_0(esk37_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_53529, plain, (X1=k8_real_1(X3,X2)|~r1_xxreal_0(X1,k8_real_1(X3,X2))|~r1_xxreal_0(k8_real_1(X3,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_52735, plain, (k7_xcmplx_0(X1,X2)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_53179, negated_conjecture, (X1=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_51639, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_51603, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3858, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k1_fdiff_9))).#trainneg
cnf(i_0_51964, negated_conjecture, (~v3_xxreal_0(esk1_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_52111, negated_conjecture, (~v3_xxreal_0(esk1_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_51960, plain, (~v3_xxreal_0(esk23_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_51961, plain, (~v3_xxreal_0(esk19_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_51963, plain, (~v3_xxreal_0(esk10_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_51973, negated_conjecture, (~v3_xxreal_0(k7_xcmplx_0(np__1,esk1_0))|~v2_xxreal_0(esk1_0))).#trainneg
cnf(i_0_53593, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_53525, plain, (k4_xcmplx_0(X1)=np__1|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_53152, plain, (k4_xcmplx_0(np__1)=X1|~r1_xxreal_0(k1_real_1(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_53123, plain, (X1=k32_sin_cos|~r1_xxreal_0(X1,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54226, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_53754, plain, (esk6_3(k1_real_1(X1),k1_real_1(X1),k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X1))|~v1_xreal_0(esk6_3(k1_real_1(X1),k1_real_1(X1),k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_54495, negated_conjecture, (~v2_xxreal_0(esk1_0)|~v3_xxreal_0(esk1_0))).#trainneg
cnf(i_0_3860, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_sincos10))).#trainneg
cnf(i_0_52103, plain, (~v3_xxreal_0(esk37_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_52107, plain, (~v3_xxreal_0(esk23_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_52108, plain, (~v3_xxreal_0(esk19_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_53706, plain, (esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_53709, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54478, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3893, plain, (v1_relat_1(k1_real_1(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_52110, plain, (~v3_xxreal_0(esk10_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_54167, plain, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3894, plain, (v1_relat_1(k7_square_1(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_52120, negated_conjecture, (~v3_xxreal_0(k7_xcmplx_0(np__1,esk1_0))|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_51959, plain, (~v3_xxreal_0(esk28_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk28_0))|~v1_xreal_0(esk28_0))).#trainneg
cnf(i_0_51962, plain, (~v3_xxreal_0(esk18_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk18_0))|~v1_xreal_0(esk18_0))).#trainneg
cnf(i_0_54149, negated_conjecture, (X1=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54433, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3934, plain, (v1_xcmplx_0(k1_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_52109, plain, (~v3_xxreal_0(esk18_0)|~v1_xboole_0(k7_xcmplx_0(np__1,esk18_0))|~v1_xreal_0(esk18_0))).#trainneg
cnf(i_0_53831, plain, (esk6_3(X1,k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(esk6_3(X1,k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(esk6_3(X1,k4_xcmplx_0(np__1),X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_53697, plain, (esk6_3(k32_sin_cos,k32_sin_cos,X1)=k32_sin_cos|~r1_xxreal_0(esk6_3(k32_sin_cos,k32_sin_cos,X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(esk6_3(k32_sin_cos,k32_sin_cos,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54321, negated_conjecture, (X1=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54165, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3935, plain, (v1_xcmplx_0(k7_square_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54208, negated_conjecture, (k4_xcmplx_0(esk1_0)=X1|~r1_xxreal_0(X1,k4_xcmplx_0(esk1_0))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54358, plain, (k1_real_1(X1)=np__1|~r1_xxreal_0(k1_real_1(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_53822, plain, (esk6_3(X1,k32_sin_cos,X1)=k32_sin_cos|~r1_xxreal_0(esk6_3(X1,k32_sin_cos,X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(esk6_3(X1,k32_sin_cos,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54161, plain, (X1=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(X1))|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54160, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_54371, negated_conjecture, (k4_xcmplx_0(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54439, plain, (esk6_3(X2,X2,k1_real_1(X1))=k1_real_1(X1)|~r1_xxreal_0(k4_xcmplx_0(esk6_3(X2,X2,k1_real_1(X1))),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(X2,X2,k1_real_1(X1)),k1_numbers)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_54273, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55291, negated_conjecture, (~v1_xboole_0(esk1_0)|~v3_xxreal_0(esk1_0))).#trainneg
cnf(i_0_4042, plain, (v1_rat_1(k1_real_1(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_52185, plain, (v3_xxreal_0(esk29_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk29_0)))).#trainneg
cnf(i_0_54905, negated_conjecture, (esk6_3(k4_xcmplx_0(esk1_0),k4_xcmplx_0(esk1_0),k4_xcmplx_0(esk1_0))=k4_xcmplx_0(esk1_0)|~v1_xreal_0(esk6_3(k4_xcmplx_0(esk1_0),k4_xcmplx_0(esk1_0),k4_xcmplx_0(esk1_0))))).#trainneg
cnf(i_0_55815, negated_conjecture, (k7_square_1(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k7_square_1(X1),k4_xcmplx_0(esk1_0))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54880, negated_conjecture, (esk6_3(esk1_0,esk1_0,esk1_0)=esk1_0|~v1_xreal_0(esk6_3(esk1_0,esk1_0,esk1_0)))).#trainneg
cnf(i_0_4043, plain, (v1_rat_1(k7_square_1(X1))|~v4_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_52192, plain, (v3_xxreal_0(np__15)|~v3_xxreal_0(k7_xcmplx_0(np__1,np__15)))).#trainneg
cnf(i_0_26366, plain, (v1_rat_1(k4_xcmplx_0(np__9))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_55512, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_55817, negated_conjecture, (k4_xcmplx_0(esk1_0)=X1|~r1_xxreal_0(X1,k4_xcmplx_0(esk1_0))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54928, plain, (esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1)|~v1_xreal_0(esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_4078, plain, (v1_int_1(k1_real_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26488, plain, (v3_ordinal1(k4_xcmplx_0(np__11))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_26491, plain, (v1_int_1(k4_xcmplx_0(np__11))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_26492, plain, (v1_rat_1(k4_xcmplx_0(np__11))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_46111, plain, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_4079, plain, (v1_int_1(k7_square_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54931, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~v1_xreal_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),k4_xcmplx_0(X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56081, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55814, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(esk1_0))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56043, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_real_1(X2)))).#trainneg
cnf(i_0_56082, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_51987, plain, (~v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_49394, plain, (v1_xxreal_0(k7_xcmplx_0(np__4,k32_sin_cos)))).#trainneg
cnf(i_0_27782, plain, (v3_ordinal1(k4_xcmplx_0(np__3))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_54889, plain, (esk6_3(k32_sin_cos,k32_sin_cos,k32_sin_cos)=k32_sin_cos|~v1_xreal_0(esk6_3(k32_sin_cos,k32_sin_cos,k32_sin_cos)))).#trainneg
cnf(i_0_55802, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(esk1_0))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54985, plain, (esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54973, plain, (esk6_3(k1_real_1(X1),k1_real_1(X1),k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(k1_real_1(X1),k1_real_1(X1),k4_xcmplx_0(X1)),k1_numbers))).#trainneg
cnf(i_0_56041, plain, (esk6_3(X1,X1,k4_xcmplx_0(X2))=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(esk6_3(X1,X1,k4_xcmplx_0(X2))),X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_54893, plain, (esk6_3(np__2,np__2,np__2)=np__2|~v1_xreal_0(esk6_3(np__2,np__2,np__2)))).#trainneg
cnf(i_0_27785, plain, (v1_int_1(k4_xcmplx_0(np__3))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_27786, plain, (v1_rat_1(k4_xcmplx_0(np__3))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_29748, plain, (v3_ordinal1(k4_xcmplx_0(np__7))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_56005, negated_conjecture, (k4_xcmplx_0(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54976, plain, (esk6_3(k1_real_1(X1),k1_real_1(X1),k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_55783, negated_conjecture, (k4_xcmplx_0(esk1_0)=k7_xcmplx_0(X1,X2)|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k4_xcmplx_0(esk1_0))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4150, plain, (v1_funct_1(k1_real_1(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29751, plain, (v1_int_1(k4_xcmplx_0(np__7))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_29752, plain, (v1_rat_1(k4_xcmplx_0(np__7))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_30069, plain, (v3_ordinal1(k4_xcmplx_0(k32_sin_cos))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_55863, plain, (esk6_3(k32_sin_cos,k32_sin_cos,k32_sin_cos)=k32_sin_cos|~r1_xxreal_0(esk6_3(k32_sin_cos,k32_sin_cos,k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_55645, plain, (X1=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55233, negated_conjecture, (k4_xcmplx_0(esk1_0)=X1|~r1_xxreal_0(k4_xcmplx_0(X1),esk1_0)|~r1_xxreal_0(X1,k4_xcmplx_0(esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56209, negated_conjecture, (esk6_3(esk1_0,esk1_0,esk1_0)=esk1_0)).#trainneg
cnf(i_0_4151, plain, (v1_funct_1(k7_square_1(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56673, plain, (~v2_xxreal_0(k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_56129, negated_conjecture, (esk6_3(k4_xcmplx_0(esk1_0),k4_xcmplx_0(esk1_0),k4_xcmplx_0(esk1_0))=k4_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_56283, plain, (esk6_3(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_4186, plain, (v3_ordinal1(k1_real_1(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56622, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4187, plain, (v3_ordinal1(k7_square_1(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56708, plain, (esk6_3(k32_sin_cos,k32_sin_cos,k32_sin_cos)=k32_sin_cos)).#trainneg
cnf(i_0_57369, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(X1,esk1_0)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(esk1_0))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_57050, negated_conjecture, (k4_xcmplx_0(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4196, plain, (r1_xxreal_0(X1,k6_square_1(X2))|v3_xxreal_0(k6_square_1(X2))|v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_30072, plain, (v1_int_1(k4_xcmplx_0(k32_sin_cos))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_30073, plain, (v1_rat_1(k4_xcmplx_0(k32_sin_cos))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_56929, plain, (esk6_3(np__2,np__2,np__2)=np__2)).#trainneg
cnf(i_0_4197, plain, (r1_xxreal_0(X1,k5_xcmplx_0(X2))|v3_xxreal_0(k5_xcmplx_0(X2))|v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_57115, plain, (esk6_3(np__4,np__4,np__4)=np__4)).#trainneg
cnf(i_0_57116, plain, (esk6_3(np__3,np__3,np__3)=np__3)).#trainneg
cnf(i_0_57117, plain, (esk6_3(np__1,np__1,np__1)=np__1)).#trainneg
cnf(i_0_33009, plain, (v3_ordinal1(k4_xcmplx_0(np__13))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_33012, plain, (v1_int_1(k4_xcmplx_0(np__13))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_56662, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_56619, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_57156, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56615, plain, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_28006, plain, (v3_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|v2_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|v1_xboole_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_4198, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|v3_xxreal_0(k4_xcmplx_0(X2))|v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_33013, plain, (v1_rat_1(k4_xcmplx_0(np__13))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_34573, negated_conjecture, (~v1_xreal_0(k10_real_1(np__3,np__4))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos)))).#trainneg
cnf(i_0_56620, plain, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_56168, plain, (k4_xcmplx_0(X1)=k6_xcmplx_0(k4_xcmplx_0(X2),X3)|~r1_xxreal_0(k1_real_1(X1),k6_xcmplx_0(k4_xcmplx_0(X2),X3))|~r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4199, plain, (r1_xxreal_0(X1,k32_sin_cos)|v3_xxreal_0(k32_sin_cos)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_35049, plain, (~v1_xreal_0(k10_real_1(np__3,np__4))|~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4)))))).#trainneg
cnf(i_0_35069, plain, (~v1_xreal_0(k10_real_1(np__3,np__4))|~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__1,k10_real_1(np__3,np__4)))))).#trainneg
cnf(i_0_44953, plain, (v3_ordinal1(k32_sin_cos)|~m1_subset_1(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k5_numbers))).#trainneg
cnf(i_0_55949, plain, (esk6_3(X1,k4_xcmplx_0(X2),X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(esk6_3(X1,k4_xcmplx_0(X2),X1)))|~v1_xreal_0(esk6_3(X1,k4_xcmplx_0(X2),X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_57329, plain, (k1_real_1(X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(np__1,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_58950, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k32_sin_cos))).#trainneg
cnf(i_0_58914, negated_conjecture, (r1_xxreal_0(esk1_0,k32_sin_cos)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_4200, plain, (r1_xxreal_0(X1,X2)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_57585, plain, (esk6_3(k1_real_1(X1),k1_real_1(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,k1_real_1(esk6_3(k1_real_1(X1),k1_real_1(X1),X2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(k1_real_1(X1),k1_real_1(X1),X2),k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_57577, plain, (k4_xcmplx_0(k1_real_1(X1))=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9451, plain, (r1_xxreal_0(X1,np__4)|v3_xxreal_0(np__4)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4216, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__2),k32_sin_cos))))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_58908, plain, (r1_xxreal_0(esk34_0,k32_sin_cos))).#trainneg
cnf(i_0_58915, plain, (r1_xxreal_0(np__0,k32_sin_cos))).#trainneg
cnf(i_0_47048, plain, (v3_ordinal1(k32_sin_cos)|~m1_subset_1(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k5_numbers))).#trainneg
cnf(i_0_48820, negated_conjecture, (v1_xboole_0(k4_xcmplx_0(esk1_0))|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_48821, negated_conjecture, (v2_xxreal_0(k4_xcmplx_0(esk1_0))|~v3_xxreal_0(esk1_0))).#trainneg
cnf(i_0_48822, negated_conjecture, (v3_xxreal_0(k4_xcmplx_0(esk1_0))|~v2_xxreal_0(esk1_0))).#trainneg
cnf(i_0_57572, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_57035, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(X1))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9518, plain, (r1_xxreal_0(X1,np__1)|v3_xxreal_0(np__1)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4217, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k3_xcmplx_0(k10_real_1(np__3,np__2),k32_sin_cos))))|~m1_subset_1(k10_real_1(np__3,np__2),k1_numbers))).#trainneg
cnf(i_0_55605, plain, (esk6_3(k32_sin_cos,k32_sin_cos,X1)=k32_sin_cos|~r1_xxreal_0(esk6_3(k32_sin_cos,k32_sin_cos,X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(k32_sin_cos,k32_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_55156, plain, (esk6_3(X1,k4_xcmplx_0(X2),X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(esk6_3(X1,k4_xcmplx_0(X2),X1)))|~m1_subset_1(esk6_3(X1,k4_xcmplx_0(X2),X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_55860, plain, (esk6_3(X1,k32_sin_cos,X1)=k32_sin_cos|~r1_xxreal_0(esk6_3(X1,k32_sin_cos,X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(X1,k32_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_55608, plain, (esk6_3(k32_sin_cos,k32_sin_cos,X1)=k32_sin_cos|~r1_xxreal_0(esk6_3(k32_sin_cos,k32_sin_cos,X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55245, plain, (k4_xcmplx_0(np__1)=X1|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54346, negated_conjecture, (k1_real_1(X1)=esk1_0|~r1_xxreal_0(k1_real_1(X1),esk1_0)|~r1_xxreal_0(X1,k4_xcmplx_0(esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9542, plain, (r1_xxreal_0(X1,np__2)|v3_xxreal_0(np__2)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4218, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__2),k32_sin_cos))))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_58907, plain, (r1_xxreal_0(esk35_0,k32_sin_cos)|~v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_58948, plain, (r1_xxreal_0(k4_xcmplx_0(np__4),k32_sin_cos))).#trainneg
cnf(i_0_53834, plain, (esk6_3(X1,k4_xcmplx_0(X2),X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(esk6_3(X1,k4_xcmplx_0(X2),X1),k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(esk6_3(X1,k4_xcmplx_0(X2),X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_54319, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k1_real_1(X1))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_55376, negated_conjecture, (k7_xcmplx_0(X1,X2)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_xcmplx_0(X1,X2)))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k7_xcmplx_0(X1,X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_55616, negated_conjecture, (k4_xcmplx_0(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k4_xcmplx_0(X1)))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_57015, negated_conjecture, (k4_xcmplx_0(k7_xcmplx_0(X1,X2))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k7_xcmplx_0(X1,X2))|~r1_xxreal_0(k7_xcmplx_0(X1,X2),esk1_0)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9544, plain, (r1_xxreal_0(X1,np__3)|v3_xxreal_0(np__3)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4219, plain, (v1_relat_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k3_xcmplx_0(k10_real_1(np__3,np__2),k32_sin_cos))))|~m1_subset_1(k10_real_1(np__3,np__2),k1_numbers))).#trainneg
cnf(i_0_58949, plain, (r1_xxreal_0(k4_xcmplx_0(np__3),k32_sin_cos))).#trainneg
cnf(i_0_58951, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),k32_sin_cos))).#trainneg
cnf(i_0_59538, negated_conjecture, (~v1_xboole_0(k4_xcmplx_0(esk1_0))|~v3_xxreal_0(esk1_0))).#trainneg
cnf(i_0_59540, negated_conjecture, (~v2_xxreal_0(k4_xcmplx_0(esk1_0))|~v2_xxreal_0(esk1_0))).#trainneg
cnf(i_0_59541, negated_conjecture, (~v1_xboole_0(k4_xcmplx_0(esk1_0))|~v2_xxreal_0(esk1_0))).#trainneg
cnf(i_0_59440, plain, (r1_xxreal_0(esk34_0,np__4))).#trainneg
cnf(i_0_59400, plain, (k4_xcmplx_0(k1_real_1(X1))=X2|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59579, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59446, negated_conjecture, (r1_xxreal_0(esk1_0,np__4)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_39691, plain, (k32_sin_cos=np__0|~r1_xxreal_0(k32_sin_cos,np__0))).#trainneg
cnf(i_0_50724, plain, (k32_sin_cos=np__0|~r1_xxreal_0(np__0,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_59524, negated_conjecture, (k4_xcmplx_0(esk1_0)=np__0|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_59507, plain, (esk34_0=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,esk34_0))).#trainneg
cnf(i_0_59661, plain, (r1_xxreal_0(esk34_0,np__1))).#trainneg
cnf(i_0_59869, plain, (r1_xxreal_0(esk34_0,np__2))).#trainneg
cnf(i_0_60244, plain, (r1_xxreal_0(esk34_0,np__3))).#trainneg
cnf(i_0_59398, plain, (k4_xcmplx_0(k1_real_1(X1))=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_59362, plain, (esk6_3(k1_real_1(X1),k1_real_1(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(esk6_3(k1_real_1(X1),k1_real_1(X1),X2)))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(k1_real_1(X1),k1_real_1(X1),X2),k1_numbers))).#trainneg
cnf(i_0_59116, plain, (esk6_3(X1,k1_real_1(X2),X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(esk6_3(X1,k1_real_1(X2),X1)))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(esk6_3(X1,k1_real_1(X2),X1),k1_numbers))).#trainneg
cnf(i_0_59667, negated_conjecture, (r1_xxreal_0(esk1_0,np__1)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_4220, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__2),k32_sin_cos))))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_59875, negated_conjecture, (r1_xxreal_0(esk1_0,np__2)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_60250, negated_conjecture, (r1_xxreal_0(esk1_0,np__3)|~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_59520, plain, (k4_xcmplx_0(k32_sin_cos)=np__0|~r1_xxreal_0(k32_sin_cos,np__0))).#trainneg
cnf(i_0_59577, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_59433, plain, (np__1=X1|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59430, plain, (k32_sin_cos=X1|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k1_real_1(X1))|~r1_xxreal_0(k32_sin_cos,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59317, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,k1_real_1(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59107, plain, (esk6_3(X1,k4_xcmplx_0(X2),X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k1_real_1(esk6_3(X1,k4_xcmplx_0(X2),X1)))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4221, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k1_rcomp_1(k32_sin_cos,k3_xcmplx_0(k10_real_1(np__3,np__2),k32_sin_cos))))|~m1_subset_1(k10_real_1(np__3,np__2),k1_numbers))).#trainneg
cnf(i_0_59439, plain, (r1_xxreal_0(esk35_0,np__4)|~v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_59660, plain, (r1_xxreal_0(esk35_0,np__1)|~v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_59868, plain, (r1_xxreal_0(esk35_0,np__2)|~v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_59781, plain, (esk6_3(X1,k32_sin_cos,X1)=k32_sin_cos|~r1_xxreal_0(esk6_3(X1,k32_sin_cos,X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59393, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_59572, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_59089, negated_conjecture, (esk6_3(X1,esk1_0,X1)=esk1_0|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k1_real_1(esk6_3(X1,esk1_0,X1)))|~r1_xxreal_0(esk1_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_59397, plain, (k4_xcmplx_0(k1_real_1(X1))=np__1|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4222, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k8_real_1(k7_xcmplx_0(np__3,np__2),k32_sin_cos))))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_60243, plain, (r1_xxreal_0(esk35_0,np__3)|~v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_59480, plain, (r1_xxreal_0(k4_xcmplx_0(np__4),np__4))).#trainneg
cnf(i_0_59481, plain, (r1_xxreal_0(k4_xcmplx_0(np__3),np__4))).#trainneg
cnf(i_0_54783, plain, (k7_square_1(X1)=k32_sin_cos|~r1_xxreal_0(k7_square_1(X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4223, plain, (v2_funct_1(k5_relat_1(k19_sin_cos,k2_rcomp_1(k32_sin_cos,k3_xcmplx_0(k10_real_1(np__3,np__2),k32_sin_cos))))|~m1_subset_1(k10_real_1(np__3,np__2),k1_numbers))).#trainneg
cnf(i_0_59483, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__4))).#trainneg
cnf(i_0_59701, plain, (r1_xxreal_0(k4_xcmplx_0(np__4),np__1))).#trainneg
cnf(i_0_59702, plain, (r1_xxreal_0(k4_xcmplx_0(np__3),np__1))).#trainneg
cnf(i_0_53825, plain, (esk6_3(X1,np__1,X1)=np__1|~r1_xxreal_0(esk6_3(X1,np__1,X1),np__1)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(esk6_3(X1,np__1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55017, plain, (esk6_3(k5_xcmplx_0(X1),k5_xcmplx_0(X1),X2)=k5_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(k5_xcmplx_0(X1),k5_xcmplx_0(X1),X2),k5_xcmplx_0(X1))|~r1_xxreal_0(k5_xcmplx_0(X1),X2)|~v1_xreal_0(esk6_3(k5_xcmplx_0(X1),k5_xcmplx_0(X1),X2))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55018, plain, (esk6_3(k5_xcmplx_0(X1),k5_xcmplx_0(X1),X2)=k5_xcmplx_0(X1)|~r1_xxreal_0(esk6_3(k5_xcmplx_0(X1),k5_xcmplx_0(X1),X2),k5_xcmplx_0(X1))|~r1_xxreal_0(k5_xcmplx_0(X1),X2)|~v1_xreal_0(esk6_3(k5_xcmplx_0(X1),k5_xcmplx_0(X1),X2))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(k5_xcmplx_0(X1)),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_54768, plain, (k4_xcmplx_0(X1)=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_60362, plain, (esk34_0=np__4|~r1_xxreal_0(np__4,esk34_0))).#trainneg
cnf(i_0_60578, plain, (esk34_0=np__1|~r1_xxreal_0(np__1,esk34_0))).#trainneg
cnf(i_0_60588, plain, (esk34_0=np__2|~r1_xxreal_0(np__2,esk34_0))).#trainneg
cnf(i_0_60636, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60470, plain, (X1=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_60476, plain, (X1=k4_xcmplx_0(np__1)|~r1_xxreal_0(np__1,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_60467, plain, (X1=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k1_real_1(X1))|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_60441, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k1_real_1(X1),X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60479, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_61083, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60597, plain, (esk34_0=np__3|~r1_xxreal_0(np__3,esk34_0))).#trainneg
cnf(i_0_60629, plain, (k1_real_1(X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60472, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4238, plain, (v1_relat_1(k5_relat_1(k1_fdiff_9,k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_60632, plain, (k1_real_1(X1)=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|~r1_xxreal_0(k32_sin_cos,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4239, plain, (v1_relat_1(k5_relat_1(k1_fdiff_9,k1_rcomp_1(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_59704, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__1))).#trainneg
cnf(i_0_59909, plain, (r1_xxreal_0(k4_xcmplx_0(np__4),np__2))).#trainneg
cnf(i_0_59910, plain, (r1_xxreal_0(k4_xcmplx_0(np__3),np__2))).#trainneg
cnf(i_0_60471, plain, (X1=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),k1_real_1(X1))|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4240, plain, (v2_funct_1(k5_relat_1(k1_fdiff_9,k1_rcomp_1(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k32_sin_cos)))|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_59912, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__2))).#trainneg
cnf(i_0_60284, plain, (r1_xxreal_0(k4_xcmplx_0(np__4),np__3))).#trainneg
cnf(i_0_60285, plain, (r1_xxreal_0(k4_xcmplx_0(np__3),np__3))).#trainneg
cnf(i_0_61687, plain, (k4_xcmplx_0(X1)=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60683, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60626, plain, (k1_real_1(X1)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60453, negated_conjecture, (X1=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4241, plain, (v2_funct_1(k5_relat_1(k1_fdiff_9,k1_rcomp_1(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k32_sin_cos)))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_60287, plain, (r1_xxreal_0(k4_xcmplx_0(np__2),np__3))).#trainneg
cnf(i_0_60571, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4263, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(k7_xcmplx_0(np__3,np__4))|v1_xboole_0(k7_xcmplx_0(np__3,np__4))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4))|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_5443, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_4264, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(k7_xcmplx_0(np__3,np__2))|v1_xboole_0(k7_xcmplx_0(np__3,np__2))|~v1_xreal_0(k7_xcmplx_0(np__3,np__2))|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_6201, plain, (v4_valued_0(esk15_0)|~v4_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_10827, plain, (v3_xxreal_0(k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_15854, negated_conjecture, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos)))).#trainneg
cnf(i_0_62303, plain, (k1_real_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_61716, plain, (k4_xcmplx_0(X1)=np__1|~r1_xxreal_0(k1_real_1(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20534, plain, (r1_xxreal_0(esk28_0,esk28_0)|~v1_xreal_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_20538, plain, (m1_subset_1(esk28_0,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(esk28_0)))).#trainneg
cnf(i_0_62294, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20651, plain, (r1_xxreal_0(esk18_0,esk18_0)|~v1_xreal_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_20655, plain, (m1_subset_1(esk18_0,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_23083, negated_conjecture, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))))).#trainneg
cnf(i_0_61965, plain, (X1=np__0|~r1_xxreal_0(np__0,k1_real_1(X1))|~r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_61707, negated_conjecture, (k4_xcmplx_0(X1)=esk1_0|~r1_xxreal_0(k1_real_1(X1),esk1_0)|~r1_xxreal_0(X1,k4_xcmplx_0(esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61699, negated_conjecture, (k4_xcmplx_0(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(esk1_0))|~r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61970, plain, (X1=k4_xcmplx_0(np__13)|~r1_xxreal_0(np__13,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(np__13),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_61971, plain, (X1=k4_xcmplx_0(np__11)|~r1_xxreal_0(np__11,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(np__11),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_61939, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(np__1)|~r1_xxreal_0(k1_real_1(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61979, plain, (X1=np__4|~r1_xxreal_0(k4_xcmplx_0(np__4),k1_real_1(X1))|~r1_xxreal_0(np__4,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_61967, plain, (X1=k4_xcmplx_0(np__9)|~r1_xxreal_0(np__9,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(np__9),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_61791, plain, (k4_xcmplx_0(X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(np__1,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23990, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_61972, plain, (X1=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_62287, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61722, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61963, plain, (X1=np__13|~r1_xxreal_0(k4_xcmplx_0(np__13),k1_real_1(X1))|~r1_xxreal_0(np__13,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_62401, plain, (k1_real_1(X1)=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~r1_xxreal_0(np__2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61980, plain, (X1=np__3|~r1_xxreal_0(k4_xcmplx_0(np__3),k1_real_1(X1))|~r1_xxreal_0(np__3,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_61964, plain, (X1=np__11|~r1_xxreal_0(k4_xcmplx_0(np__11),k1_real_1(X1))|~r1_xxreal_0(np__11,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_36912, negated_conjecture, (r1_xxreal_0(X1,k7_square_1(esk1_0))|r1_xxreal_0(k7_square_1(esk1_0),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_27500, plain, (v1_int_1(k4_xcmplx_0(esk10_0))|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_28433, plain, (v2_xxreal_0(k4_xcmplx_0(esk34_0))|~v1_xxreal_0(k4_xcmplx_0(esk34_0)))).#trainneg
cnf(i_0_61861, plain, (k4_xcmplx_0(X1)=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61960, plain, (X1=np__9|~r1_xxreal_0(k4_xcmplx_0(np__9),k1_real_1(X1))|~r1_xxreal_0(np__9,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_62396, plain, (k4_xcmplx_0(X1)=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~r1_xxreal_0(np__2,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61976, plain, (X1=k4_xcmplx_0(np__2)|~r1_xxreal_0(np__2,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(np__2),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_64122, negated_conjecture, (r1_xxreal_0(k7_square_1(esk1_0),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_square_1(esk1_0)))).#trainneg
cnf(i_0_63275, plain, (~v2_xxreal_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_63276, plain, (~v1_xboole_0(k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_63017, plain, (k4_xcmplx_0(esk6_3(X1,X2,X1))=k1_real_1(X1)|~r1_xxreal_0(X1,esk6_3(X1,X2,X1))|~v1_xreal_0(esk6_3(X1,X2,X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_63323, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_63009, plain, (k4_xcmplx_0(esk8_3(X2,X3,X1))=k1_real_1(X1)|r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,esk8_3(X2,X3,X1))|~v1_xreal_0(esk8_3(X2,X3,X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_63260, plain, (k1_real_1(esk6_3(X1,X2,X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(esk6_3(X1,X2,X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(esk6_3(X1,X2,X1),k1_numbers)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63321, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_63859, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_64123, negated_conjecture, (r1_xxreal_0(k7_square_1(esk1_0),np__4)|r1_xxreal_0(np__4,k7_square_1(esk1_0)))).#trainneg
cnf(i_0_62972, plain, (k1_real_1(esk43_0)=k4_xcmplx_0(esk43_0))).#trainneg
cnf(i_0_62973, plain, (k1_real_1(esk37_0)=k4_xcmplx_0(esk37_0))).#trainneg
cnf(i_0_62975, plain, (k1_real_1(esk34_0)=k4_xcmplx_0(esk34_0))).#trainneg
cnf(i_0_63316, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_63964, plain, (X1=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),k1_real_1(X1))|~r1_xxreal_0(np__2,X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_63843, plain, (k4_xcmplx_0(k32_sin_cos)=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63436, plain, (k4_xcmplx_0(k5_xcmplx_0(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k5_xcmplx_0(X1))|~r1_xxreal_0(k5_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_64124, negated_conjecture, (r1_xxreal_0(k7_square_1(esk1_0),np__3)|r1_xxreal_0(np__3,k7_square_1(esk1_0)))).#trainneg
cnf(i_0_62982, plain, (k1_real_1(esk29_0)=k4_xcmplx_0(esk29_0))).#trainneg
cnf(i_0_62983, plain, (k1_real_1(esk23_0)=k4_xcmplx_0(esk23_0))).#trainneg
cnf(i_0_62985, plain, (k1_real_1(esk19_0)=k4_xcmplx_0(esk19_0))).#trainneg
cnf(i_0_64613, plain, (v1_xxreal_0(k4_xcmplx_0(esk43_0)))).#trainneg
cnf(i_0_64737, plain, (v1_xxreal_0(k4_xcmplx_0(esk37_0)))).#trainneg
cnf(i_0_65400, plain, (v1_xxreal_0(k4_xcmplx_0(esk23_0)))).#trainneg
cnf(i_0_64611, plain, (v1_relat_1(k4_xcmplx_0(esk43_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_64125, negated_conjecture, (r1_xxreal_0(k7_square_1(esk1_0),np__1)|r1_xxreal_0(np__1,k7_square_1(esk1_0)))).#trainneg
cnf(i_0_62986, plain, (k1_real_1(esk10_0)=k4_xcmplx_0(esk10_0))).#trainneg
cnf(i_0_63024, plain, (k1_real_1(np__15)=k4_xcmplx_0(np__15))).#trainneg
cnf(i_0_64615, plain, (v1_rat_1(k4_xcmplx_0(esk43_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_64616, plain, (v1_int_1(k4_xcmplx_0(esk43_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_64617, plain, (v1_funct_1(k4_xcmplx_0(esk43_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_64618, plain, (v3_ordinal1(k4_xcmplx_0(esk43_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_64735, plain, (v1_relat_1(k4_xcmplx_0(esk37_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_64268, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_64128, negated_conjecture, (r1_xxreal_0(k7_square_1(esk1_0),np__2)|r1_xxreal_0(np__2,k7_square_1(esk1_0)))).#trainneg
cnf(i_0_64739, plain, (v1_rat_1(k4_xcmplx_0(esk37_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_64740, plain, (v1_int_1(k4_xcmplx_0(esk37_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_64741, plain, (v1_funct_1(k4_xcmplx_0(esk37_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_64742, plain, (v3_ordinal1(k4_xcmplx_0(esk37_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_64859, plain, (v1_relat_1(k4_xcmplx_0(esk34_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65709, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_65525, plain, (v1_xxreal_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_65526, plain, (v1_xreal_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_66027, plain, (v1_xxreal_0(k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_64863, plain, (v1_rat_1(k4_xcmplx_0(esk34_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_65705, plain, (r1_tarski(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_relset_1(k1_numbers,k2_sincos10))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_65713, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~r2_hidden(X1,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_65722, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~r2_hidden(X1,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_64609, plain, (k1_real_1(k4_xcmplx_0(esk43_0))=esk43_0)).#trainneg
cnf(i_0_64733, plain, (k1_real_1(k4_xcmplx_0(esk37_0))=esk37_0)).#trainneg
cnf(i_0_64864, plain, (v1_int_1(k4_xcmplx_0(esk34_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_64865, plain, (v1_funct_1(k4_xcmplx_0(esk34_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_64866, plain, (v3_ordinal1(k4_xcmplx_0(esk34_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_65273, plain, (v1_relat_1(k4_xcmplx_0(esk29_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65277, plain, (v1_rat_1(k4_xcmplx_0(esk29_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_4443, plain, (v1_xboole_0(esk38_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_64857, plain, (k1_real_1(k4_xcmplx_0(esk34_0))=esk34_0)).#trainneg
cnf(i_0_65271, plain, (k1_real_1(k4_xcmplx_0(esk29_0))=esk29_0)).#trainneg
cnf(i_0_65396, plain, (k1_real_1(k4_xcmplx_0(esk23_0))=esk23_0)).#trainneg
cnf(i_0_65278, plain, (v1_int_1(k4_xcmplx_0(esk29_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_65279, plain, (v1_funct_1(k4_xcmplx_0(esk29_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65280, plain, (v3_ordinal1(k4_xcmplx_0(esk29_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_65398, plain, (v1_relat_1(k4_xcmplx_0(esk23_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65402, plain, (v1_rat_1(k4_xcmplx_0(esk23_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_66446, plain, (k4_xcmplx_0(esk6_3(X1,X1,X2))=k1_real_1(X1)|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X1,X1,X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4450, plain, (v2_membered(k2_tarski(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65521, plain, (k1_real_1(k4_xcmplx_0(esk19_0))=esk19_0)).#trainneg
cnf(i_0_65898, plain, (k1_real_1(k4_xcmplx_0(esk10_0))=esk10_0)).#trainneg
cnf(i_0_66023, plain, (k1_real_1(k4_xcmplx_0(np__15))=np__15)).#trainneg
cnf(i_0_66845, negated_conjecture, (~r2_hidden(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_4453, plain, (v2_membered(k2_xboole_0(X1,X2))|~v3_membered(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_66896, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,X1))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_66902, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,k4_xcmplx_0(k7_square_1(np__2))))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_66901, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_65829, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_66868, plain, (m1_subset_1(k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_sincos10)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_65759, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k4_xcmplx_0(k7_square_1(np__2)),k1_numbers)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_66531, plain, (esk6_3(X1,X1,X1)=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k1_real_1(X1))|~r1_xxreal_0(k32_sin_cos,esk6_3(X1,X1,X1))|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65744, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k4_xcmplx_0(k7_square_1(np__2)),k1_numbers)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_4480, plain, (r1_xxreal_0(X1,esk35_0)|r1_xxreal_0(esk35_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_66532, plain, (esk6_3(X1,X1,X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~r1_xxreal_0(np__1,esk6_3(X1,X1,X1))|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_59043, plain, (np__4=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_4481, plain, (r1_xxreal_0(X1,esk34_0)|r1_xxreal_0(esk34_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_65751, negated_conjecture, (r1_xxreal_0(esk1_0,k4_xcmplx_0(np__1))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k4_xcmplx_0(k7_square_1(np__2)),k1_numbers)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_66851, negated_conjecture, (m1_subset_1(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_numbers)|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_65746, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_65736, negated_conjecture, (esk7_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|r2_hidden(esk1_0,k1_xxreal_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_37031, plain, (r1_xxreal_0(X1,k7_square_1(k32_sin_cos))|r1_xxreal_0(k7_square_1(k32_sin_cos),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_59044, plain, (np__3=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_4482, plain, (r1_xxreal_0(X1,esk30_0)|r1_xxreal_0(esk30_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_66197, plain, (k1_real_1(esk6_3(X2,X2,X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,esk6_3(X2,X2,X1))|~v1_xreal_0(esk6_3(X2,X2,X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_66022, plain, (m1_subset_1(k4_xcmplx_0(np__15),k1_numbers))).#trainneg
cnf(i_0_65706, negated_conjecture, (esk7_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=esk1_0|r2_hidden(esk1_0,k2_tarski(k4_xcmplx_0(np__1),k4_xcmplx_0(k7_square_1(np__2))))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_66033, plain, (r1_xxreal_0(k4_xcmplx_0(np__15),k4_xcmplx_0(np__15)))).#trainneg
cnf(i_0_37033, plain, (r1_xxreal_0(X1,k7_square_1(np__4))|r1_xxreal_0(k7_square_1(np__4),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_59045, plain, (np__2=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_4483, plain, (r1_xxreal_0(X1,esk29_0)|r1_xxreal_0(esk29_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_65403, plain, (v1_int_1(k4_xcmplx_0(esk23_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_65404, plain, (v1_funct_1(k4_xcmplx_0(esk23_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65405, plain, (v3_ordinal1(k4_xcmplx_0(esk23_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_65523, plain, (v1_relat_1(k4_xcmplx_0(esk19_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65524, plain, (v1_xcmplx_0(k4_xcmplx_0(esk19_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_66142, plain, (X1=k4_xcmplx_0(np__15)|~r1_xxreal_0(np__15,k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(np__15),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65734, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X1,k2_tarski(k4_xcmplx_0(np__1),k4_xcmplx_0(k7_square_1(np__2))))|~r2_hidden(X2,k2_tarski(k4_xcmplx_0(np__1),k4_xcmplx_0(k7_square_1(np__2))))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_37034, plain, (r1_xxreal_0(X1,k7_square_1(np__3))|r1_xxreal_0(k7_square_1(np__3),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4484, plain, (r1_xxreal_0(X1,esk20_0)|r1_xxreal_0(esk20_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_65527, plain, (v1_rat_1(k4_xcmplx_0(esk19_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_65528, plain, (v1_int_1(k4_xcmplx_0(esk19_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_65529, plain, (v1_funct_1(k4_xcmplx_0(esk19_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65530, plain, (v3_ordinal1(k4_xcmplx_0(esk19_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_66533, plain, (esk6_3(X1,X1,X1)=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),k1_real_1(X1))|~r1_xxreal_0(np__2,esk6_3(X1,X1,X1))|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_66445, plain, (k4_xcmplx_0(esk6_3(X1,X2,X1))=k1_real_1(X2)|~r1_xxreal_0(esk6_3(X1,X2,X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(esk6_3(X1,X2,X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_66198, plain, (k1_real_1(esk6_3(X1,X2,X1))=k4_xcmplx_0(X1)|~r1_xxreal_0(X1,esk6_3(X1,X2,X1))|~v1_xreal_0(esk6_3(X1,X2,X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_37036, plain, (r1_xxreal_0(X1,k7_square_1(np__1))|r1_xxreal_0(k7_square_1(np__1),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4485, plain, (r1_xxreal_0(X1,esk10_0)|r1_xxreal_0(esk10_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_65900, plain, (v1_relat_1(k4_xcmplx_0(esk10_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65904, plain, (v1_rat_1(k4_xcmplx_0(esk10_0))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_65905, plain, (v1_int_1(k4_xcmplx_0(esk10_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_65906, plain, (v1_funct_1(k4_xcmplx_0(esk10_0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_65907, plain, (v3_ordinal1(k4_xcmplx_0(esk10_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_65719, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k1_relset_1(k1_numbers,k2_sincos10),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r2_hidden(X2,k1_rcomp_1(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r2_hidden(X1,k2_tarski(k4_xcmplx_0(np__1),k4_xcmplx_0(k7_square_1(np__2))))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_66261, plain, (k1_real_1(k7_xcmplx_0(k32_sin_cos,np__2))=k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_68756, plain, (r1_xxreal_0(k7_square_1(k32_sin_cos),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_square_1(k32_sin_cos)))).#trainneg
cnf(i_0_69216, plain, (v1_xcmplx_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_66025, plain, (v1_relat_1(k4_xcmplx_0(np__15))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_66029, plain, (v1_rat_1(k4_xcmplx_0(np__15))|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_66030, plain, (v1_int_1(k4_xcmplx_0(np__15))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_68757, plain, (r1_xxreal_0(k7_square_1(k32_sin_cos),np__4)|r1_xxreal_0(np__4,k7_square_1(k32_sin_cos)))).#trainneg
cnf(i_0_68073, plain, (v1_xxreal_0(X1)|~r2_hidden(X1,X2)|~v3_membered(X3)|~v3_membered(X2))).#trainneg
cnf(i_0_4494, plain, (v3_funct_1(k10_xtuple_0(X1))|~v1_funct_1(k10_xtuple_0(X1))|~v1_relat_1(k10_xtuple_0(X1))|~v3_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_66031, plain, (v1_funct_1(k4_xcmplx_0(np__15))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_66032, plain, (v3_ordinal1(k4_xcmplx_0(np__15))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_68758, plain, (r1_xxreal_0(k7_square_1(k32_sin_cos),np__3)|r1_xxreal_0(np__3,k7_square_1(k32_sin_cos)))).#trainneg
cnf(i_0_4495, plain, (v5_valued_0(k10_xtuple_0(X1))|~v1_funct_1(k10_xtuple_0(X1))|~v2_valued_0(k10_xtuple_0(X1))|~v1_relat_1(k10_xtuple_0(X1))|~v3_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_68024, plain, (r1_xxreal_0(k1_real_1(k4_xcmplx_0(X1)),X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68759, plain, (r1_xxreal_0(k7_square_1(k32_sin_cos),np__1)|r1_xxreal_0(np__1,k7_square_1(k32_sin_cos)))).#trainneg
cnf(i_0_68090, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=k4_xcmplx_0(k7_square_1(np__2))|~r1_xxreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_68762, plain, (r1_xxreal_0(k7_square_1(k32_sin_cos),np__2)|r1_xxreal_0(np__2,k7_square_1(k32_sin_cos)))).#trainneg
cnf(i_0_68039, plain, (r1_xxreal_0(X1,k1_real_1(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_69097, plain, (r1_xxreal_0(k7_square_1(np__4),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_square_1(np__4)))).#trainneg
cnf(i_0_4507, plain, (r2_hidden(k8_real_1(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_63492, plain, (np__0=np__4|~r1_xxreal_0(np__0,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_63493, plain, (np__0=np__3|~r1_xxreal_0(np__0,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_68383, plain, (esk6_3(X1,X1,X1)=k32_sin_cos|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k1_real_1(X1))|~r1_xxreal_0(k32_sin_cos,esk6_3(X1,X1,X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1))).#trainneg
cnf(i_0_68642, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=k32_sin_cos|~r1_xxreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_68507, plain, (esk6_3(X1,X1,X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~r1_xxreal_0(np__1,esk6_3(X1,X1,X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1))).#trainneg
cnf(i_0_69098, plain, (r1_xxreal_0(k7_square_1(np__4),np__4)|r1_xxreal_0(np__4,k7_square_1(np__4)))).#trainneg
cnf(i_0_62398, plain, (np__4=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_4508, plain, (r2_hidden(k3_xcmplx_0(k10_real_1(np__3,np__4),k32_sin_cos),k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_63496, plain, (np__0=np__2|~r1_xxreal_0(np__0,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_44940, plain, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_47038, plain, (~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers)|~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos))))).#trainneg
cnf(i_0_47499, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k7_xcmplx_0(k10_real_1(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)),k32_sin_cos))|~m1_subset_1(k10_real_1(np__3,np__4),k1_numbers))).#trainneg
cnf(i_0_68896, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X1))|~v1_xreal_0(esk6_3(X1,X1,X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1))).#trainneg
cnf(i_0_68149, negated_conjecture, (esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_numbers)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_69099, plain, (r1_xxreal_0(k7_square_1(np__4),np__3)|r1_xxreal_0(np__3,k7_square_1(np__4)))).#trainneg
cnf(i_0_62399, plain, (np__3=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_64251, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_51454, plain, (v2_xxreal_0(esk28_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk28_0))|~v1_xreal_0(k7_xcmplx_0(np__1,esk28_0)))).#trainneg
cnf(i_0_51457, plain, (v2_xxreal_0(esk18_0)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk18_0))|~v1_xreal_0(k7_xcmplx_0(np__1,esk18_0)))).#trainneg
cnf(i_0_52186, plain, (v3_xxreal_0(esk28_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk28_0))|~v1_xreal_0(k7_xcmplx_0(np__1,esk28_0)))).#trainneg
cnf(i_0_52189, plain, (v3_xxreal_0(esk18_0)|~v3_xxreal_0(k7_xcmplx_0(np__1,esk18_0))|~v1_xreal_0(k7_xcmplx_0(np__1,esk18_0)))).#trainneg
cnf(i_0_69814, plain, (k1_real_1(k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2)))=k7_xcmplx_0(k32_sin_cos,np__2))).#trainneg
cnf(i_0_68045, plain, (m1_subset_1(k1_real_1(k4_xcmplx_0(X1)),k1_numbers)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_68643, negated_conjecture, (k4_xcmplx_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_68150, negated_conjecture, (k4_xcmplx_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~m1_subset_1(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)),k1_numbers)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_69100, plain, (r1_xxreal_0(k7_square_1(np__4),np__1)|r1_xxreal_0(np__1,k7_square_1(np__4)))).#trainneg
cnf(i_0_64252, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),np__4))).#trainneg
cnf(i_0_64253, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),np__1))).#trainneg
cnf(i_0_64254, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),np__2))).#trainneg
cnf(i_0_59932, plain, (r1_xxreal_0(esk35_0,k32_sin_cos)|~m1_subset_1(esk35_0,k1_numbers))).#trainneg
cnf(i_0_59933, plain, (r1_xxreal_0(esk35_0,k32_sin_cos)|~m1_subset_1(esk35_0,k5_numbers))).#trainneg
cnf(i_0_59934, plain, (r1_xxreal_0(esk35_0,k32_sin_cos)|~m1_subset_1(esk35_0,k4_numbers))).#trainneg
cnf(i_0_69516, plain, (esk6_3(X1,X1,X1)=np__2|~r1_xxreal_0(k4_xcmplx_0(np__2),k1_real_1(X1))|~r1_xxreal_0(np__2,esk6_3(X1,X1,X1))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X1))).#trainneg
cnf(i_0_68240, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_70989, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_69103, plain, (r1_xxreal_0(k7_square_1(np__4),np__2)|r1_xxreal_0(np__2,k7_square_1(np__4)))).#trainneg
cnf(i_0_4578, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_70697, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_69378, plain, (r1_xxreal_0(k7_square_1(np__3),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_square_1(np__3)))).#trainneg
cnf(i_0_70322, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4579, plain, (k3_xcmplx_0(np__1,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_64255, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),np__3))).#trainneg
cnf(i_0_61030, plain, (r1_xxreal_0(esk35_0,np__4)|~m1_subset_1(esk35_0,k1_numbers))).#trainneg
cnf(i_0_61031, plain, (r1_xxreal_0(esk35_0,np__4)|~m1_subset_1(esk35_0,k5_numbers))).#trainneg
cnf(i_0_70454, plain, (v1_xxreal_0(k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_70459, plain, (r1_xxreal_0(k1_real_1(X1),X2)|r1_xxreal_0(X2,k1_real_1(X1))|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_70460, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_70279, plain, (v1_xxreal_0(X1)|v1_xboole_0(X3)|~v3_membered(X2)|~v3_membered(X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_70484, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_69379, plain, (r1_xxreal_0(k7_square_1(np__3),np__4)|r1_xxreal_0(np__4,k7_square_1(np__3)))).#trainneg
cnf(i_0_70328, plain, (k1_real_1(X1)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4580, plain, (k7_xcmplx_0(k2_xcmplx_0(X1,X2),np__1)=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_61032, plain, (r1_xxreal_0(esk35_0,np__4)|~m1_subset_1(esk35_0,k4_numbers))).#trainneg
cnf(i_0_61033, plain, (r1_xxreal_0(esk35_0,np__1)|~m1_subset_1(esk35_0,k1_numbers))).#trainneg
cnf(i_0_61034, plain, (r1_xxreal_0(esk35_0,np__1)|~m1_subset_1(esk35_0,k5_numbers))).#trainneg
cnf(i_0_61035, plain, (r1_xxreal_0(esk35_0,np__1)|~m1_subset_1(esk35_0,k4_numbers))).#trainneg
cnf(i_0_61036, plain, (r1_xxreal_0(esk35_0,np__2)|~m1_subset_1(esk35_0,k1_numbers))).#trainneg
cnf(i_0_71547, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_71830, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_71842, negated_conjecture, (v3_xxreal_0(esk1_0))).#trainneg
cnf(i_0_70333, plain, (k1_real_1(X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4581, plain, (k4_xcmplx_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_71847, negated_conjecture, (~v2_xxreal_0(esk1_0))).#trainneg
cnf(i_0_71997, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk1_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_72974, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k32_sin_cos,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k32_sin_cos)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_61037, plain, (r1_xxreal_0(esk35_0,np__2)|~m1_subset_1(esk35_0,k5_numbers))).#trainneg
cnf(i_0_61038, plain, (r1_xxreal_0(esk35_0,np__2)|~m1_subset_1(esk35_0,k4_numbers))).#trainneg
cnf(i_0_61204, plain, (r1_xxreal_0(esk35_0,np__3)|~m1_subset_1(esk35_0,k1_numbers))).#trainneg
cnf(i_0_71834, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_71833, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_71815, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_69380, plain, (r1_xxreal_0(k7_square_1(np__3),np__3)|r1_xxreal_0(np__3,k7_square_1(np__3)))).#trainneg
cnf(i_0_71989, negated_conjecture, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k1_seq_1(k2_sincos10,k4_xcmplx_0(esk1_0))|~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_61205, plain, (r1_xxreal_0(esk35_0,np__3)|~m1_subset_1(esk35_0,k5_numbers))).#trainneg
cnf(i_0_61206, plain, (r1_xxreal_0(esk35_0,np__3)|~m1_subset_1(esk35_0,k4_numbers))).#trainneg
cnf(i_0_71832, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_69381, plain, (r1_xxreal_0(k7_square_1(np__3),np__1)|r1_xxreal_0(np__1,k7_square_1(np__3)))).#trainneg
cnf(i_0_72966, plain, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k1_seq_1(k2_sincos10,k4_xcmplx_0(k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k32_sin_cos)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_71822, negated_conjecture, (v2_xxreal_0(esk1_0)|~v2_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_72493, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X1)),X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_69384, plain, (r1_xxreal_0(k7_square_1(np__3),np__2)|r1_xxreal_0(np__2,k7_square_1(np__3)))).#trainneg
cnf(i_0_50100, plain, (k4_xcmplx_0(np__1)=np__4|~r1_xxreal_0(np__1,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_71829, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(k7_square_1(np__2)))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_72895, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71757, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_71826, negated_conjecture, (k1_real_1(k4_xcmplx_0(k7_square_1(np__2)))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_69656, plain, (r1_xxreal_0(k7_square_1(np__1),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_square_1(np__1)))).#trainneg
cnf(i_0_63997, plain, (np__3=np__4|~r1_xxreal_0(k4_xcmplx_0(np__3),k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_50101, plain, (k4_xcmplx_0(np__1)=np__3|~r1_xxreal_0(np__1,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_50103, plain, (k4_xcmplx_0(np__1)=np__2|~r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_50328, plain, (k4_xcmplx_0(np__4)=np__1|~r1_xxreal_0(np__1,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_71840, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(k7_square_1(np__2)))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_71490, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~r1_xxreal_0(X1,k4_xcmplx_0(k1_real_1(X1)))|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_71001, negated_conjecture, (k4_xcmplx_0(np__1)=esk1_0|~r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_69657, plain, (r1_xxreal_0(k7_square_1(np__1),np__4)|r1_xxreal_0(np__4,k7_square_1(np__1)))).#trainneg
cnf(i_0_50329, plain, (k4_xcmplx_0(np__3)=np__1|~r1_xxreal_0(np__1,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_50331, plain, (k4_xcmplx_0(np__2)=np__1|~r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_72878, plain, (X1=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(X1,k4_xcmplx_0(k32_sin_cos))|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71824, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(k7_square_1(np__2)))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_69658, plain, (r1_xxreal_0(k7_square_1(np__1),np__3)|r1_xxreal_0(np__3,k7_square_1(np__1)))).#trainneg
cnf(i_0_73508, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_74187, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),np__1)|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_74185, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k32_sin_cos)|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_74800, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_73991, negated_conjecture, (r1_xxreal_0(esk1_0,k32_sin_cos))).#trainneg
cnf(i_0_59538, negated_conjecture, (~v1_xboole_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_55291, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_52111, negated_conjecture, (~v1_xboole_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_51964, negated_conjecture, (~v2_xxreal_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_48821, negated_conjecture, (v2_xxreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_74483, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(esk1_0)|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_73993, negated_conjecture, (r1_xxreal_0(esk1_0,np__1))).#trainneg
cnf(i_0_73499, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_73992, negated_conjecture, (r1_xxreal_0(esk1_0,np__4))).#trainneg
cnf(i_0_74732, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,esk1_0))).#trainneg
cnf(i_0_27606, negated_conjecture, (r1_xxreal_0(esk1_0,k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|v3_xxreal_0(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4)))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_73994, negated_conjecture, (r1_xxreal_0(esk1_0,np__2))).#trainneg
cnf(i_0_73995, negated_conjecture, (r1_xxreal_0(esk1_0,np__3))).#trainneg
cnf(i_0_74761, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_73469, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),k4_xcmplx_0(X1))|~v1_xreal_0(esk6_3(X2,X2,X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_70386, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_74860, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_75388, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_74923, negated_conjecture, (r1_xxreal_0(np__1,k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_38568, negated_conjecture, (esk1_0=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,esk1_0))).#trainneg
cnf(i_0_69659, plain, (r1_xxreal_0(k7_square_1(np__1),np__1)|r1_xxreal_0(np__1,k7_square_1(np__1)))).#trainneg
cnf(i_0_38603, negated_conjecture, (esk1_0=np__1|~r1_xxreal_0(np__1,esk1_0))).#trainneg
cnf(i_0_38581, negated_conjecture, (esk1_0=np__4|~r1_xxreal_0(np__4,esk1_0))).#trainneg
cnf(i_0_75737, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk1_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk1_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_38616, negated_conjecture, (esk1_0=np__2|~r1_xxreal_0(np__2,esk1_0))).#trainneg
cnf(i_0_74856, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_75197, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k32_sin_cos)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_75984, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X2))).#trainneg
cnf(i_0_75193, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_76100, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_69662, plain, (r1_xxreal_0(k7_square_1(np__1),np__2)|r1_xxreal_0(np__2,k7_square_1(np__1)))).#trainneg
cnf(i_0_38588, negated_conjecture, (esk1_0=np__3|~r1_xxreal_0(np__3,esk1_0))).#trainneg
cnf(i_0_75268, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_74864, negated_conjecture, (k4_xcmplx_0(k4_xcmplx_0(k7_square_1(np__2)))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_75603, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_74738, plain, (k1_real_1(X1)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k32_sin_cos))).#trainneg
cnf(i_0_76121, negated_conjecture, (~v3_xxreal_0(k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_76309, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(esk1_0,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_77327, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k32_sin_cos,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k32_sin_cos)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_76521, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_76530, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k32_sin_cos))).#trainneg
cnf(i_0_77165, plain, (k1_real_1(X1)=np__0|~r1_xxreal_0(np__0,X1)|~r1_xxreal_0(np__0,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_77230, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_77255, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_76675, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_77205, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_77704, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_75371, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,esk1_0))).#trainneg
cnf(i_0_78116, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(esk1_0,k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk1_0,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_78251, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(np__0,k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(np__0,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_78778, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_78185, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_74743, plain, (k1_real_1(X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_75397, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(np__1)|~r1_xxreal_0(np__1,esk1_0))).#trainneg
cnf(i_0_78598, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_77625, plain, (r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_77696, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_78166, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_77700, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_77701, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_77865, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_75671, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(np__4)|~r1_xxreal_0(np__4,esk1_0))).#trainneg
cnf(i_0_80123, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_80210, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_79872, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_78977, negated_conjecture, (~r1_xxreal_0(np__0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(np__0,k7_square_1(np__2))|~v1_xboole_0(k1_rcomp_1(np__0,k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_78863, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k7_square_1(np__2))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_79073, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k7_square_1(np__2)))).#trainneg
cnf(i_0_80202, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_80184, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_79061, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_80072, plain, (k4_xcmplx_0(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_80025, plain, (r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers))).#trainneg
cnf(i_0_79577, plain, (k1_real_1(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))).#trainneg
cnf(i_0_80175, plain, (X1=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_80110, plain, (k1_real_1(X1)=np__1|~r1_xxreal_0(np__1,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_81142, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_80414, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_80473, negated_conjecture, (k1_real_1(k7_square_1(np__2))=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_4606, plain, (v1_xcmplx_0(k7_xcmplx_0(np__5,np__2))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_81138, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_81139, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_80777, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_75593, plain, (k4_xcmplx_0(k1_real_1(X1))=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_80180, plain, (X1=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X1)),X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_71762, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_76651, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2)))).#trainneg
cnf(i_0_75134, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(X1))|~r1_xxreal_0(X1,k4_xcmplx_0(esk1_0))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_78581, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5569, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v1_int_1(np__7)|~v1_int_1(np__5))).#trainneg
cnf(i_0_5573, plain, (v1_int_1(k4_xcmplx_0(np__7))|~v1_int_1(np__7)|~v1_int_1(np__0))).#trainneg
cnf(i_0_5574, plain, (v1_int_1(k4_xcmplx_0(np__5))|~v1_int_1(np__5)|~v1_int_1(np__0))).#trainneg
cnf(i_0_5581, plain, (v1_int_1(k4_xcmplx_0(np__4))|~v1_int_1(np__4)|~v1_int_1(np__0))).#trainneg
cnf(i_0_5582, plain, (v1_int_1(k4_xcmplx_0(np__3))|~v1_int_1(np__3)|~v1_int_1(np__0))).#trainneg
cnf(i_0_75279, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_80174, plain, (X1=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_82895, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_82432, plain, (k1_real_1(X1)=k4_xcmplx_0(esk43_0)|~r1_xxreal_0(esk43_0,X1)|~r1_xxreal_0(X1,esk43_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_81985, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k32_sin_cos,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_5583, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(np__1)|~v1_int_1(np__0))).#trainneg
cnf(i_0_82112, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_82165, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_82433, plain, (k1_real_1(X1)=k4_xcmplx_0(esk37_0)|~r1_xxreal_0(esk37_0,X1)|~r1_xxreal_0(X1,esk37_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_75945, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(np__2)|~r1_xxreal_0(np__2,esk1_0))).#trainneg
cnf(i_0_82554, plain, (X1=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_82892, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X2))).#trainneg
cnf(i_0_82829, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_82843, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_82434, plain, (k1_real_1(X1)=k4_xcmplx_0(esk34_0)|~r1_xxreal_0(esk34_0,X1)|~r1_xxreal_0(X1,esk34_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_27500, plain, (v1_int_1(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_83294, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xboole_0(k1_rcomp_1(k32_sin_cos,k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_83062, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk43_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_83459, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk37_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_84398, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk34_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk34_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk34_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_83842, plain, (X1=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)|~r1_xxreal_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_82435, plain, (k1_real_1(X1)=k4_xcmplx_0(esk29_0)|~r1_xxreal_0(esk29_0,X1)|~r1_xxreal_0(X1,esk29_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_82231, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_82641, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_82436, plain, (k1_real_1(X1)=k4_xcmplx_0(esk23_0)|~r1_xxreal_0(esk23_0,X1)|~r1_xxreal_0(X1,esk23_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_84282, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_82460, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_82550, plain, (X1=X2|~r1_xxreal_0(k32_sin_cos,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(X1,k32_sin_cos))).#trainneg
cnf(i_0_82437, plain, (k1_real_1(X1)=k4_xcmplx_0(esk19_0)|~r1_xxreal_0(esk19_0,X1)|~r1_xxreal_0(X1,esk19_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4631, plain, (v1_xcmplx_0(k7_xcmplx_0(np__7,np__2))|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__2))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_84718, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk29_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk29_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk29_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_85159, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk23_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_85840, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk19_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_85443, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_85548, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_85499, plain, (k1_real_1(X1)=np__0|~r1_xxreal_0(X1,np__0)|~r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_82678, plain, (k1_real_1(X1)=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_71903, plain, (r1_xxreal_0(np__1,k1_real_1(X1))|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_77891, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_75591, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(X2)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_77160, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|r1_xxreal_0(k32_sin_cos,k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_82438, plain, (k1_real_1(X1)=k4_xcmplx_0(esk10_0)|~r1_xxreal_0(esk10_0,X1)|~r1_xxreal_0(X1,esk10_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_86453, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(np__0,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__0)|~r1_xxreal_0(np__0,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_87174, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk10_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_86172, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_86801, plain, (r1_xxreal_0(np__1,k4_xcmplx_0(X1))|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_86714, plain, (k1_real_1(X1)=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_86790, plain, (r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_79559, negated_conjecture, (k1_real_1(k1_real_1(k7_square_1(np__2)))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_77650, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X2),X1))).#trainneg
cnf(i_0_86817, plain, (k1_real_1(X1)=esk6_3(X2,X2,X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~r1_xxreal_0(esk6_3(X2,X2,X2),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_86164, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)=k1_real_1(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2)))|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),X2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_85472, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X1)),X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X2)),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_71900, plain, (r1_xxreal_0(k32_sin_cos,k1_real_1(X1))|~r1_xxreal_0(X1,k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_75549, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_76731, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k4_xcmplx_0(k1_real_1(X2)))|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_88398, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X2)),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_87701, plain, (~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_87702, plain, (~r1_xxreal_0(k7_xcmplx_0(k4_xcmplx_0(np__3),np__4),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_88416, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X2)),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_88235, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_88134, plain, (k4_xcmplx_0(X1)=X2|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_88505, plain, (~r1_xxreal_0(k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2)),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_87419, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),np__0)|~r1_xxreal_0(np__0,k7_square_1(np__2))|~v1_xboole_0(k1_rcomp_1(np__0,k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_87793, plain, (k1_real_1(X1)=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_87615, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(esk1_0,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_87828, plain, (k1_real_1(X1)=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_87607, negated_conjecture, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k1_seq_1(k2_sincos10,esk1_0)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_88222, plain, (k1_real_1(X1)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_87635, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,X1))|~r2_hidden(X1,k2_tarski(k4_xcmplx_0(np__1),esk1_0))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_88689, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_75402, plain, (r1_xxreal_0(esk6_3(X1,X1,X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14147, plain, (m1_subset_1(esk28_0,k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,esk28_0)))).#trainneg
cnf(i_0_89536, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_89309, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_90136, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_90078, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_90082, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_75958, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(np__3)|~r1_xxreal_0(np__3,esk1_0))).#trainneg
cnf(i_0_89412, plain, (~r1_xxreal_0(k4_xcmplx_0(k7_xcmplx_0(k32_sin_cos,np__2)),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_90208, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_88919, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_89341, plain, (k1_real_1(k1_real_1(X1))=X1|~r1_xxreal_0(X1,k4_xcmplx_0(k1_real_1(X1)))|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers))).#trainneg
cnf(i_0_89296, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_90230, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_90425, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_85493, plain, (k1_real_1(X1)=k4_xcmplx_0(np__15)|~r1_xxreal_0(X1,np__15)|~r1_xxreal_0(np__15,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14148, plain, (m1_subset_1(esk18_0,k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,esk18_0)))).#trainneg
cnf(i_0_19347, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(k7_xcmplx_0(np__3,np__4),k32_sin_cos),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_90530, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_85494, plain, (k1_real_1(X1)=k4_xcmplx_0(np__9)|~r1_xxreal_0(X1,np__9)|~r1_xxreal_0(np__9,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4734, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(k3_xxreal_1(X3,X2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X1,k3_xxreal_1(X3,X2)))).#trainneg
cnf(i_0_20535, plain, (m1_subset_1(esk28_0,k1_numbers)|~m1_subset_1(k4_xcmplx_0(esk28_0),k1_numbers))).#trainneg
cnf(i_0_20540, plain, (r1_xxreal_0(esk28_0,esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k1_numbers))).#trainneg
cnf(i_0_20652, plain, (m1_subset_1(esk18_0,k1_numbers)|~m1_subset_1(k4_xcmplx_0(esk18_0),k1_numbers))).#trainneg
cnf(i_0_20657, plain, (r1_xxreal_0(esk18_0,esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k1_numbers))).#trainneg
cnf(i_0_91437, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_85495, plain, (k1_real_1(X1)=k4_xcmplx_0(np__7)|~r1_xxreal_0(X1,np__7)|~r1_xxreal_0(np__7,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4735, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_91703, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__15),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__15)|~r1_xxreal_0(np__15,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_90437, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_85496, plain, (k1_real_1(X1)=k4_xcmplx_0(np__5)|~r1_xxreal_0(X1,np__5)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4736, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_92284, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__9),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__9)|~r1_xxreal_0(np__9,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_92746, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_92840, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_93396, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X2))).#trainneg
cnf(i_0_85497, plain, (k1_real_1(X1)=k4_xcmplx_0(np__13)|~r1_xxreal_0(X1,np__13)|~r1_xxreal_0(np__13,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4737, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),np__1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_92926, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__7),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__7)|~r1_xxreal_0(np__7,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_93463, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__5),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__5)|~r1_xxreal_0(np__5,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_81323, plain, (k1_real_1(k1_real_1(k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4))))=k8_real_1(k32_sin_cos,k10_real_1(np__3,np__4))|~v1_xreal_0(k10_real_1(np__3,np__4)))).#trainneg
cnf(i_0_85498, plain, (k1_real_1(X1)=k4_xcmplx_0(np__11)|~r1_xxreal_0(X1,np__11)|~r1_xxreal_0(np__11,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4738, plain, (k4_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X1,X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_94498, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__13),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__13)|~r1_xxreal_0(np__13,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_95154, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__11),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__11)|~r1_xxreal_0(np__11,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_94410, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_85525, plain, (k1_real_1(X1)=k4_xcmplx_0(np__4)|~r1_xxreal_0(X1,np__4)|~r1_xxreal_0(np__4,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_92704, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_91079, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_85526, plain, (k1_real_1(X1)=k4_xcmplx_0(np__3)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__3,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_96080, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__4),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__4)|~r1_xxreal_0(np__4,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_96665, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__3),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__3)|~r1_xxreal_0(np__3,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_95988, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_94310, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(esk6_3(X2,X2,X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_92577, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_31506, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,k7_xcmplx_0(np__4,np__3)),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_92736, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_90526, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_92836, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_85532, plain, (k1_real_1(X1)=k4_xcmplx_0(np__2)|~r1_xxreal_0(X1,np__2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_97972, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__2),k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),np__2)|~r1_xxreal_0(np__2,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_31510, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(k3_xcmplx_0(k32_sin_cos,np__3),np__4),k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_58906, plain, (r1_xxreal_0(esk37_0,k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_58909, plain, (r1_xxreal_0(esk23_0,k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_58910, plain, (r1_xxreal_0(esk19_0,k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_58912, plain, (r1_xxreal_0(esk10_0,k32_sin_cos)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_96999, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_97680, plain, (k7_square_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k7_square_1(X1),X2)|~r1_xxreal_0(X2,k7_square_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59438, plain, (r1_xxreal_0(esk37_0,np__4)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_59441, plain, (r1_xxreal_0(esk23_0,np__4)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_59442, plain, (r1_xxreal_0(esk19_0,np__4)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_59444, plain, (r1_xxreal_0(esk10_0,np__4)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_89214, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_96970, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_59659, plain, (r1_xxreal_0(esk37_0,np__1)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_59662, plain, (r1_xxreal_0(esk23_0,np__1)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_59663, plain, (r1_xxreal_0(esk19_0,np__1)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_59665, plain, (r1_xxreal_0(esk10_0,np__1)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_91089, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_89264, plain, (k1_real_1(X1)=esk6_3(X2,X2,X2)|~r1_xxreal_0(esk6_3(X2,X2,X2),k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(esk6_3(X2,X2,X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_90051, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_90424, plain, (k4_xcmplx_0(k1_real_1(X1))=X2|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_92830, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_95839, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_91126, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),X1)|~m1_subset_1(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59867, plain, (r1_xxreal_0(esk37_0,np__2)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_59870, plain, (r1_xxreal_0(esk23_0,np__2)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_59871, plain, (r1_xxreal_0(esk19_0,np__2)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_59873, plain, (r1_xxreal_0(esk10_0,np__2)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_91980, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_92689, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_99428, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_98727, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_99504, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_98804, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_99330, plain, (~r1_xxreal_0(np__1,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_99331, plain, (~r1_xxreal_0(np__1,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_99334, plain, (~r1_xxreal_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_60242, plain, (r1_xxreal_0(esk37_0,np__3)|~v1_xboole_0(k7_xcmplx_0(np__1,esk37_0)))).#trainneg
cnf(i_0_60245, plain, (r1_xxreal_0(esk23_0,np__3)|~v1_xboole_0(k7_xcmplx_0(np__1,esk23_0)))).#trainneg
cnf(i_0_60246, plain, (r1_xxreal_0(esk19_0,np__3)|~v1_xboole_0(k7_xcmplx_0(np__1,esk19_0)))).#trainneg
cnf(i_0_99138, negated_conjecture, (k1_real_1(X1)=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|~r1_xxreal_0(k7_square_1(np__2),X1)|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_97343, negated_conjecture, (X1=X2|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(esk1_0))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~r1_xxreal_0(esk1_0,X1))).#trainneg
cnf(i_0_60248, plain, (r1_xxreal_0(esk10_0,np__3)|~v1_xboole_0(k7_xcmplx_0(np__1,esk10_0)))).#trainneg
cnf(i_0_63370, plain, (r1_xxreal_0(esk28_0,esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k5_numbers))).#trainneg
cnf(i_0_63371, plain, (r1_xxreal_0(esk28_0,esk28_0)|~m1_subset_1(k4_xcmplx_0(esk28_0),k4_numbers))).#trainneg
cnf(i_0_99417, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_92092, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_90520, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63375, plain, (m1_subset_1(esk28_0,k1_numbers)|~m1_subset_1(k4_xcmplx_0(esk28_0),k5_numbers))).#trainneg
cnf(i_0_63376, plain, (m1_subset_1(esk28_0,k1_numbers)|~m1_subset_1(k4_xcmplx_0(esk28_0),k4_numbers))).#trainneg
cnf(i_0_63453, plain, (r1_xxreal_0(esk18_0,esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k5_numbers))).#trainneg
cnf(i_0_97462, negated_conjecture, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X2))|~r1_xxreal_0(esk1_0,X2))).#trainneg
cnf(i_0_97470, negated_conjecture, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k1_real_1(X1))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X2))|~r1_xxreal_0(esk1_0,X2))).#trainneg
cnf(i_0_63454, plain, (r1_xxreal_0(esk18_0,esk18_0)|~m1_subset_1(k4_xcmplx_0(esk18_0),k4_numbers))).#trainneg
cnf(i_0_63458, plain, (m1_subset_1(esk18_0,k1_numbers)|~m1_subset_1(k4_xcmplx_0(esk18_0),k5_numbers))).#trainneg
cnf(i_0_63459, plain, (m1_subset_1(esk18_0,k1_numbers)|~m1_subset_1(k4_xcmplx_0(esk18_0),k4_numbers))).#trainneg
cnf(i_0_63474, plain, (np__9=np__0|~r1_xxreal_0(np__0,k4_xcmplx_0(np__9)))).#trainneg
cnf(i_0_97679, plain, (k1_real_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63475, plain, (np__7=np__0|~r1_xxreal_0(np__0,k4_xcmplx_0(np__7)))).#trainneg
cnf(i_0_63476, plain, (np__5=np__0|~r1_xxreal_0(np__0,k4_xcmplx_0(np__5)))).#trainneg
cnf(i_0_63477, plain, (np__13=np__0|~r1_xxreal_0(np__0,k4_xcmplx_0(np__13)))).#trainneg
cnf(i_0_63478, plain, (np__11=np__0|~r1_xxreal_0(np__0,k4_xcmplx_0(np__11)))).#trainneg
cnf(i_0_68139, negated_conjecture, (~v2_xxreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(esk6_3(esk1_0,k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_89305, negated_conjecture, (k1_real_1(X1)=k7_square_1(np__2)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(X1))|~r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_90301, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(X2)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_100079, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_100684, negated_conjecture, (k4_xcmplx_0(esk1_0)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_99649, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_100487, plain, (k1_real_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_99646, plain, (k1_real_1(X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(X1),np__1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_100436, plain, (k4_xcmplx_0(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_100854, plain, (k1_real_1(X1)=k4_xcmplx_0(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_100128, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X1),X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_92135, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~m1_subset_1(esk6_3(X2,X2,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_97434, negated_conjecture, (X1=X2|~r1_xxreal_0(esk1_0,X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~r1_xxreal_0(esk1_0,X1))).#trainneg
cnf(i_0_97034, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(esk6_3(X2,X2,X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_89232, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_90231, plain, (k1_real_1(X1)=esk6_3(X2,X2,X2)|~r1_xxreal_0(esk6_3(X2,X2,X2),k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(esk6_3(X2,X2,X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_91013, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_100975, plain, (esk6_3(X1,X1,X1)=k4_xcmplx_0(k1_real_1(X1))|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_70463, plain, (r1_xxreal_0(k32_sin_cos,k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_101208, plain, (k1_real_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_90280, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k1_real_1(X1)))|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~r1_xxreal_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers))).#trainneg
cnf(i_0_70473, plain, (r1_xxreal_0(np__2,k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_89808, plain, (r1_xxreal_0(X1,k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(np__1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(X1))|~r1_xxreal_0(np__1,X1))).#trainneg
cnf(i_0_90528, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|~r1_xxreal_0(X1,k7_square_1(np__2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_71718, plain, (r1_xxreal_0(k1_real_1(X1),np__1)|r1_xxreal_0(np__1,k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_74051, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k4_xcmplx_0(esk1_0)),k7_xcmplx_0(k32_sin_cos,np__2))|~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_74141, plain, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k4_xcmplx_0(k32_sin_cos)),k7_xcmplx_0(k32_sin_cos,np__2))|~r1_xxreal_0(k32_sin_cos,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k32_sin_cos)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_89488, plain, (k1_real_1(k1_real_1(X1))=k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)|~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos),X1)|~r1_xxreal_0(X1,k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_102549, plain, (k1_real_1(k1_real_1(X1))=X1|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_102696, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_81487, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(np__1,k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__1,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_82744, negated_conjecture, (r1_xxreal_0(esk1_0,np__0)|~r1_xxreal_0(np__0,k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_88489, plain, (r1_xxreal_0(k32_sin_cos,np__0)|~r1_xxreal_0(np__0,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_102381, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_92107, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_105083, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_101280, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(np__1,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),np__1)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_105147, plain, (esk6_3(k4_xcmplx_0(X1),k4_xcmplx_0(X1),k4_xcmplx_0(X1))=k1_real_1(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_105088, plain, (k1_real_1(k1_real_1(X1))=k4_xcmplx_0(k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_104491, plain, (r1_xxreal_0(X1,k4_xcmplx_0(k1_real_1(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_106275, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_105052, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(X1)))|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_105037, plain, (r1_xxreal_0(k1_real_1(k1_real_1(X1)),X1)|~r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71720, negated_conjecture, (r1_xxreal_0(k1_real_1(X1),esk1_0)|r1_xxreal_0(esk1_0,k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_86194, negated_conjecture, (k1_real_1(k7_square_1(np__2))=k32_sin_cos|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_106014, plain, (k1_real_1(k4_xcmplx_0(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107793, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k32_sin_cos,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_106613, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107398, plain, (r1_xxreal_0(k4_xcmplx_0(k1_real_1(X1)),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_108117, plain, (k4_xcmplx_0(k1_real_1(X1))=k4_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107033, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k1_real_1(X1),np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_106609, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(k1_real_1(X1),np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71723, plain, (r1_xxreal_0(np__4,k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),np__4)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4821, plain, (k3_xcmplx_0(k6_xcmplx_0(X1,X2),k4_xcmplx_0(np__1))=k4_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_108165, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xboole_0(k1_rcomp_1(k32_sin_cos,k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_108226, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_108759, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_108764, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4822, plain, (k3_xcmplx_0(np__1,k6_xcmplx_0(X1,X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_108279, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),X1)|r1_xxreal_0(np__1,k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_106612, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_107395, plain, (r1_xxreal_0(esk6_3(X1,X1,X1),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_109600, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_81977, negated_conjecture, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k1_seq_1(k2_sincos10,k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_4823, plain, (k7_xcmplx_0(k6_xcmplx_0(X1,X2),np__1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5783, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(k7_xcmplx_0(np__3,np__4))|~v1_xreal_0(k7_xcmplx_0(np__3,np__4))|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_5784, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(k7_xcmplx_0(np__3,np__2))|~v1_xreal_0(k7_xcmplx_0(np__3,np__2))|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_110179, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107785, negated_conjecture, (k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)=k1_seq_1(k2_sincos10,k32_sin_cos)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_4824, plain, (k4_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_109647, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_109657, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_106990, plain, (r1_xxreal_0(X1,esk6_3(X1,X1,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71725, plain, (r1_xxreal_0(np__3,k1_real_1(X1))|r1_xxreal_0(k1_real_1(X1),np__3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_109726, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_84242, plain, (k1_real_1(X1)=np__0|~r1_xxreal_0(k4_xcmplx_0(X1),np__0)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_110815, negated_conjecture, (~v2_xxreal_0(k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111716, plain, (k1_real_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_111535, plain, (X1=k4_xcmplx_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_110810, negated_conjecture, (v3_xxreal_0(k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110929, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111261, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110912, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k1_seq_1(k2_sincos10,k32_sin_cos),k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110916, negated_conjecture, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_seq_1(k2_sincos10,k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110919, negated_conjecture, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_seq_1(k2_sincos10,k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110938, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110935, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k1_seq_1(k2_sincos10,k32_sin_cos)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110936, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k1_seq_1(k2_sincos10,k32_sin_cos)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110960, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k1_seq_1(k2_sincos10,k32_sin_cos),k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110962, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k1_seq_1(k2_sincos10,k32_sin_cos),k32_sin_cos))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111244, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k1_seq_1(k2_sincos10,k32_sin_cos),k32_sin_cos))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111647, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_102396, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_90859, plain, (k1_real_1(X1)=np__0|~r1_xxreal_0(np__0,X1)|~r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_112636, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v2_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112667, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_112668, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_112936, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(np__0,k7_square_1(np__2)))).#trainneg
cnf(i_0_112952, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__0)|~r1_xxreal_0(np__0,k7_square_1(np__2)))).#trainneg
cnf(i_0_4844, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__9),np__2))|~v1_xcmplx_0(np__5)|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_112965, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_112969, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k32_sin_cos),k7_xcmplx_0(k32_sin_cos,np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_112637, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk1_0))).#trainneg
cnf(i_0_113170, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_112800, plain, (k1_real_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_113099, plain, (r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_112727, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X1)),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_112647, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk34_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk34_0))).#trainneg
cnf(i_0_112926, negated_conjecture, (~v1_xboole_0(k1_real_1(k7_square_1(np__2)))|~v1_xxreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112805, plain, (k1_real_1(X1)=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_102321, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_110934, negated_conjecture, (r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k1_seq_1(k2_sincos10,k32_sin_cos))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_113152, plain, (r1_xxreal_0(X1,k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_113588, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_113592, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_114382, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xboole_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_113600, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__0)|~r1_xxreal_0(np__0,k7_square_1(np__2)))).#trainneg
cnf(i_0_113647, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk1_0))).#trainneg
cnf(i_0_114380, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk34_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk34_0))).#trainneg
cnf(i_0_113596, negated_conjecture, (~r1_xxreal_0(np__0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(np__0,k7_square_1(np__2)))).#trainneg
cnf(i_0_114215, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(X1))|~r1_xxreal_0(X1,k7_square_1(np__2))|~v1_xreal_0(X1)|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_112669, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k32_sin_cos)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_112945, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),np__0)|~r1_xxreal_0(k7_square_1(np__2),np__0))).#trainneg
cnf(i_0_113837, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_112670, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__4)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_111248, negated_conjecture, (~r2_hidden(k4_rcomp_1(k10_real_1(k32_sin_cos,np__2),k32_sin_cos),k1_seq_1(k2_sincos10,k32_sin_cos))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111251, negated_conjecture, (~r2_hidden(k4_rcomp_1(k7_xcmplx_0(k32_sin_cos,np__2),k32_sin_cos),k1_seq_1(k2_sincos10,k32_sin_cos))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111267, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k1_seq_1(k2_sincos10,k32_sin_cos)))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111268, negated_conjecture, (~r2_hidden(k1_seq_1(k2_sincos10,esk1_0),k2_tarski(k32_sin_cos,k1_seq_1(k2_sincos10,k32_sin_cos)))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_111292, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_xxreal_1(k1_seq_1(k2_sincos10,k32_sin_cos),k32_sin_cos))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112671, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__1)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_90819, plain, (k1_real_1(X1)=k4_xcmplx_0(esk43_0)|~r1_xxreal_0(esk43_0,X1)|~r1_xxreal_0(X1,esk43_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_111294, negated_conjecture, (~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_rcomp_1(k1_seq_1(k2_sincos10,k32_sin_cos),k32_sin_cos))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112316, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(np__0,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),np__0)|~r1_xxreal_0(k7_square_1(np__2),np__0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110678, plain, (r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_112715, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_112672, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__2)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_90820, plain, (k1_real_1(X1)=k4_xcmplx_0(esk37_0)|~r1_xxreal_0(esk37_0,X1)|~r1_xxreal_0(X1,esk37_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_115037, negated_conjecture, (~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),np__0)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__0))).#trainneg
cnf(i_0_115655, negated_conjecture, (~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),np__0)|~r1_xxreal_0(k7_square_1(np__2),np__0)|~v1_xboole_0(k1_rcomp_1(np__0,k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_115765, negated_conjecture, (r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers)|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_90821, plain, (k1_real_1(X1)=k4_xcmplx_0(esk34_0)|~r1_xxreal_0(esk34_0,X1)|~r1_xxreal_0(X1,esk34_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_112645, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v2_xxreal_0(k4_xcmplx_0(esk43_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0))).#trainneg
cnf(i_0_112646, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v2_xxreal_0(k4_xcmplx_0(esk37_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0))).#trainneg
cnf(i_0_102894, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_108758, plain, (k1_real_1(X1)=k4_xcmplx_0(X2)|~r1_xxreal_0(X2,k4_xcmplx_0(k4_xcmplx_0(X1)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_108546, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),k1_real_1(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_90822, plain, (k1_real_1(X1)=k4_xcmplx_0(esk29_0)|~r1_xxreal_0(esk29_0,X1)|~r1_xxreal_0(X1,esk29_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_117039, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0)|~v2_xxreal_0(k4_xcmplx_0(esk43_0)))).#trainneg
cnf(i_0_117043, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0)|~v2_xxreal_0(k4_xcmplx_0(esk37_0)))).#trainneg
cnf(i_0_112648, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v2_xxreal_0(k4_xcmplx_0(esk29_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk29_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk29_0))).#trainneg
cnf(i_0_112649, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v2_xxreal_0(k4_xcmplx_0(esk23_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0))).#trainneg
cnf(i_0_112650, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v2_xxreal_0(k4_xcmplx_0(esk19_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0))).#trainneg
cnf(i_0_116569, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_seq_1(k2_sincos10,esk1_0))|~m1_subset_1(k7_square_1(np__2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),k4_xcmplx_0(np__1))|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_90823, plain, (k1_real_1(X1)=k4_xcmplx_0(esk23_0)|~r1_xxreal_0(esk23_0,X1)|~r1_xxreal_0(X1,esk23_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4874, plain, (v1_xcmplx_0(k7_xcmplx_0(np__7,np__4))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_112651, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v2_xxreal_0(k4_xcmplx_0(esk10_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0))).#trainneg
cnf(i_0_96955, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_90129, plain, (r1_xxreal_0(k1_real_1(X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),X1)|~m1_subset_1(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_97020, plain, (X1=k4_xcmplx_0(k4_xcmplx_0(X2))|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_90824, plain, (k1_real_1(X1)=k4_xcmplx_0(esk19_0)|~r1_xxreal_0(esk19_0,X1)|~r1_xxreal_0(X1,esk19_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_117899, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk29_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk29_0)|~v2_xxreal_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_117903, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0)|~v2_xxreal_0(k4_xcmplx_0(esk23_0)))).#trainneg
cnf(i_0_117907, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0)|~v2_xxreal_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_118378, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0)|~v2_xxreal_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_118673, plain, (X1=k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))|~r1_xxreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),X1)|~r1_xxreal_0(X1,k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_118582, plain, (r1_xxreal_0(k1_real_1(X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_97350, negated_conjecture, (X1=X2|~r1_xxreal_0(k4_xcmplx_0(X2),k4_xcmplx_0(esk1_0))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(X1))|~r1_xxreal_0(esk1_0,X1))).#trainneg
cnf(i_0_102366, plain, (k1_real_1(X1)=X2|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_112673, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__3)|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~v1_xreal_0(k1_real_1(k7_square_1(np__2))))).#trainneg
cnf(i_0_90825, plain, (k1_real_1(X1)=k4_xcmplx_0(esk10_0)|~r1_xxreal_0(esk10_0,X1)|~r1_xxreal_0(X1,esk10_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_120438, plain, (esk6_3(X1,X1,X1)=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_120251, plain, (r1_xxreal_0(k1_real_1(X1),esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_114391, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xboole_0(k4_xcmplx_0(esk43_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0))).#trainneg
cnf(i_0_119279, plain, (r1_xxreal_0(k1_real_1(X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)))).#trainneg
cnf(i_0_120268, plain, (r1_xxreal_0(k4_xcmplx_0(X1),esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_114392, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xboole_0(k4_xcmplx_0(esk37_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0))).#trainneg
cnf(i_0_114395, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xboole_0(k4_xcmplx_0(esk23_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0))).#trainneg
cnf(i_0_114396, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xboole_0(k4_xcmplx_0(esk19_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0))).#trainneg
cnf(i_0_119468, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(np__1)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_119205, plain, (r1_xxreal_0(k1_real_1(X1),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|~r1_xxreal_0(k4_xcmplx_0(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos)),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_120250, plain, (r1_xxreal_0(k1_real_1(X1),k1_real_1(X2))|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_119261, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(esk1_0),k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))|r1_xxreal_0(esk1_0,k1_real_1(k8_real_1(k10_real_1(np__3,np__4),k32_sin_cos))))).#trainneg
cnf(i_0_120006, plain, (k4_xcmplx_0(esk6_3(X1,X1,X1))=k1_real_1(X1)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_121040, plain, (r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_90853, plain, (k1_real_1(X1)=k4_xcmplx_0(np__15)|~r1_xxreal_0(np__15,X1)|~r1_xxreal_0(X1,np__15)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_120781, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0)|~v1_xboole_0(k4_xcmplx_0(esk43_0)))).#trainneg
cnf(i_0_121146, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0)|~v1_xboole_0(k4_xcmplx_0(esk37_0)))).#trainneg
cnf(i_0_121150, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0)|~v1_xboole_0(k4_xcmplx_0(esk23_0)))).#trainneg
cnf(i_0_121154, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0)|~v1_xboole_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_114397, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xboole_0(k4_xcmplx_0(esk10_0))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0))).#trainneg
cnf(i_0_121593, plain, (esk6_3(X1,X1,X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~r1_xxreal_0(np__1,esk6_3(X1,X1,X1))|~m1_subset_1(esk6_3(X1,X1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_90854, plain, (k1_real_1(X1)=k4_xcmplx_0(np__9)|~r1_xxreal_0(np__9,X1)|~r1_xxreal_0(X1,np__9)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4880, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__13),np__4))|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_122334, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0)|~v1_xboole_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_122391, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(X2))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_90855, plain, (k1_real_1(X1)=k4_xcmplx_0(np__7)|~r1_xxreal_0(np__7,X1)|~r1_xxreal_0(X1,np__7)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4881, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__15),np__4))|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_90856, plain, (k1_real_1(X1)=k4_xcmplx_0(np__5)|~r1_xxreal_0(np__5,X1)|~r1_xxreal_0(X1,np__5)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4882, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__7),np__2))|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_115315, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk43_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_116201, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk37_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_116696, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk34_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk34_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk34_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_117553, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk29_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk29_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk29_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_123293, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_123446, plain, (r1_xxreal_0(k4_xcmplx_0(X1),esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_123288, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_4883, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__9),np__4))|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__4)))).#trainneg
cnf(i_0_117989, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk23_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_118829, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk19_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_119655, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(esk10_0),k4_xcmplx_0(np__1)))|~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_121977, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__15),k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__15,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__15)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_122732, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__9),k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__9,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__9)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_123438, plain, (r1_xxreal_0(X1,esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_120565, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_125516, plain, (r1_xxreal_0(X1,esk6_3(X2,X2,X2))|~r1_xxreal_0(k4_xcmplx_0(X2),k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_125412, plain, (r1_xxreal_0(X1,esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_109602, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_123566, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__7),k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__7,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__7)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_124049, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__5),k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__5,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__5)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_111270, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,k32_sin_cos),k10_real_1(k32_sin_cos,np__2))|~v1_xxreal_0(k10_real_1(k32_sin_cos,np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_125579, plain, (r1_xxreal_0(X1,esk6_3(X2,X2,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_90857, plain, (k1_real_1(X1)=k4_xcmplx_0(np__13)|~r1_xxreal_0(np__13,X1)|~r1_xxreal_0(X1,np__13)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4885, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__11),np__4))|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_125394, plain, (r1_xxreal_0(k32_sin_cos,esk6_3(X1,X1,X1))|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_90858, plain, (k1_real_1(X1)=k4_xcmplx_0(np__11)|~r1_xxreal_0(np__11,X1)|~r1_xxreal_0(X1,np__11)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4886, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__2))|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_126603, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__13),k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__13,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__13)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_14086, plain, (m1_subset_1(esk28_0,k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk28_0),k1_numbers))).#trainneg
cnf(i_0_14087, plain, (m1_subset_1(esk18_0,k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk18_0),k1_numbers))).#trainneg
cnf(i_0_126516, plain, (r1_xxreal_0(X1,k1_real_1(k4_xcmplx_0(X2)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4887, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__5),np__4))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4)))).#trainneg
cnf(i_0_125082, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X1)),X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_125141, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_128152, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_128198, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_127680, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_127149, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(k4_xcmplx_0(np__11),k4_xcmplx_0(np__1)))|~r1_xxreal_0(np__11,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),np__11)|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_128197, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_125518, plain, (r1_xxreal_0(X1,esk6_3(X2,X2,X2))|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k32_sin_cos,X2)|~r1_xxreal_0(X2,k32_sin_cos))).#trainneg
cnf(i_0_125414, plain, (r1_xxreal_0(X1,esk6_3(X2,X2,X2))|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(k32_sin_cos))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(X1,k32_sin_cos))).#trainneg
cnf(i_0_125869, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_114996, negated_conjecture, (k1_real_1(k7_square_1(np__2))=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_128962, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_128781, plain, (k1_real_1(X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_128964, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X2)),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_128844, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_128837, plain, (k1_real_1(X1)=k1_real_1(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_115091, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__4|~r1_xxreal_0(np__4,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_59058, plain, (k4_xcmplx_0(np__1)=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_20485, plain, (v1_xboole_0(k4_xcmplx_0(esk43_0))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk43_0))))).#trainneg
cnf(i_0_20486, plain, (v1_xboole_0(k4_xcmplx_0(esk37_0))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk37_0))))).#trainneg
cnf(i_0_128687, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X2)),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_129948, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(np__1,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_130011, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_130386, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_59061, plain, (k4_xcmplx_0(k32_sin_cos)=np__1|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_129955, negated_conjecture, (~v1_xboole_0(k1_rcomp_1(np__1,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_20487, plain, (v1_xboole_0(k4_xcmplx_0(esk23_0))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk23_0))))).#trainneg
cnf(i_0_20488, plain, (v1_xboole_0(k4_xcmplx_0(esk10_0))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk10_0))))).#trainneg
cnf(i_0_20489, negated_conjecture, (v1_xboole_0(k4_xcmplx_0(esk1_0))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk1_0))))).#trainneg
cnf(i_0_130640, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1))).#trainneg
cnf(i_0_129959, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k1_seq_1(k2_sincos10,X1),k1_seq_1(k2_sincos10,X2))|~r2_hidden(X1,k1_rcomp_1(np__1,k4_xcmplx_0(np__1)))|~r2_hidden(X2,k1_rcomp_1(np__1,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_129961, negated_conjecture, (esk6_3(esk1_0,np__1,k4_xcmplx_0(np__1))=esk1_0|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_130638, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_130229, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X2),np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k4_xcmplx_0(np__1),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_128955, plain, (k1_real_1(k1_real_1(X1))=k7_xcmplx_0(k32_sin_cos,np__2)|~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),X1)|~r1_xxreal_0(X1,k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_115194, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__1|~r1_xxreal_0(np__1,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_59064, plain, (k4_xcmplx_0(k32_sin_cos)=np__1|~r1_xxreal_0(np__1,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_59066, plain, (k4_xcmplx_0(np__1)=k32_sin_cos|~r1_xxreal_0(np__1,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_20490, plain, (v1_xboole_0(k4_xcmplx_0(np__15))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__15))))).#trainneg
cnf(i_0_20491, plain, (v1_xboole_0(k4_xcmplx_0(np__9))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__9))))).#trainneg
cnf(i_0_20492, plain, (v1_xboole_0(k4_xcmplx_0(np__7))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__7))))).#trainneg
cnf(i_0_20493, plain, (v1_xboole_0(k4_xcmplx_0(np__5))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__5))))).#trainneg
cnf(i_0_20494, plain, (v1_xboole_0(k4_xcmplx_0(np__13))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__13))))).#trainneg
cnf(i_0_125384, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),esk6_3(X1,X1,X1))|~r1_xxreal_0(k1_real_1(X1),k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_128519, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X2)),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_126519, plain, (r1_xxreal_0(X1,k1_real_1(k1_real_1(X2)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_131969, plain, (esk6_3(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(k32_sin_cos))=k4_xcmplx_0(k32_sin_cos))).#trainneg
cnf(i_0_130998, negated_conjecture, (~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_131113, negated_conjecture, (~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xboole_0(k1_rcomp_1(np__1,k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_130973, negated_conjecture, (~r2_hidden(k1_rcomp_1(np__1,k4_xcmplx_0(np__1)),esk1_0)|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_20495, plain, (v1_xboole_0(k4_xcmplx_0(np__11))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__11))))).#trainneg
cnf(i_0_20496, plain, (v1_xboole_0(k4_xcmplx_0(np__4))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__4))))).#trainneg
cnf(i_0_132019, plain, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|~r1_xxreal_0(k1_real_1(X1),k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116085, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__2|~r1_xxreal_0(np__2,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_20497, plain, (v1_xboole_0(k4_xcmplx_0(np__3))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__3))))).#trainneg
cnf(i_0_20641, plain, (v1_xboole_0(k4_xcmplx_0(esk18_0))|~v1_xboole_0(esk18_0)|~v1_xcmplx_0(k4_xcmplx_0(esk18_0)))).#trainneg
cnf(i_0_132243, plain, (k1_real_1(X1)=k1_real_1(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_131934, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__1|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_131889, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=np__1|~r1_xxreal_0(np__1,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_132547, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_119533, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__3|~r1_xxreal_0(np__3,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_59939, plain, (k4_xcmplx_0(np__4)=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_59942, plain, (k4_xcmplx_0(k32_sin_cos)=np__4|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_132136, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_132460, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4896, plain, (v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__7),np__4))|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__4)))).#trainneg
cnf(i_0_133065, plain, (k4_xcmplx_0(esk18_0)=np__0|~v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_133577, negated_conjecture, (~r1_xxreal_0(np__1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_133064, plain, (v1_xreal_0(k4_xcmplx_0(esk18_0))|~v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_132804, negated_conjecture, (~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(k1_rcomp_1(np__1,k4_xcmplx_0(np__1)),esk1_0))).#trainneg
cnf(i_0_134082, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_133998, plain, (k1_real_1(X1)=k1_real_1(k4_xcmplx_0(X2))|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_59945, plain, (k4_xcmplx_0(k32_sin_cos)=np__4|~r1_xxreal_0(np__4,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_59947, plain, (k4_xcmplx_0(np__4)=k32_sin_cos|~r1_xxreal_0(np__4,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_23020, plain, (m1_subset_1(k4_xcmplx_0(esk43_0),k1_numbers))).#trainneg
cnf(i_0_23040, plain, (m1_subset_1(k4_xcmplx_0(esk37_0),k1_numbers))).#trainneg
cnf(i_0_23060, plain, (m1_subset_1(k4_xcmplx_0(esk23_0),k1_numbers))).#trainneg
cnf(i_0_132623, negated_conjecture, (r1_xxreal_0(X1,k4_xcmplx_0(k32_sin_cos))|~r1_xxreal_0(X1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~v1_xreal_0(X1)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_132240, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(X1))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_132213, plain, (k1_real_1(X1)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23178, plain, (m1_subset_1(k4_xcmplx_0(esk10_0),k1_numbers))).#trainneg
cnf(i_0_132459, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_60312, plain, (k4_xcmplx_0(np__3)=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_24587, plain, (v1_xboole_0(k4_xcmplx_0(esk19_0))|~v1_xboole_0(esk19_0)|~v1_xcmplx_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_132517, plain, (k4_xcmplx_0(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,k1_real_1(X2))|~m1_subset_1(k1_real_1(X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k1_real_1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1)))).#trainneg
cnf(i_0_135193, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_134134, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~r1_xxreal_0(np__1,k7_square_1(np__2)))).#trainneg
cnf(i_0_135488, plain, (k4_xcmplx_0(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X2),X1))).#trainneg
cnf(i_0_135206, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k4_xcmplx_0(X1),k1_real_1(X2))|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_132725, plain, (r1_xxreal_0(X1,k4_xcmplx_0(k1_real_1(X2)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_135308, plain, (k4_xcmplx_0(esk19_0)=np__0|~v1_xboole_0(esk19_0))).#trainneg
cnf(i_0_134207, negated_conjecture, (~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~m1_subset_1(k1_rcomp_1(np__1,k4_xcmplx_0(np__1)),esk1_0))).#trainneg
cnf(i_0_134203, plain, (m1_subset_1(esk18_0,k1_numbers)|~v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_132543, plain, (X1=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(k1_real_1(X2)),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(X2))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_135216, plain, (k1_real_1(X1)=k1_real_1(k4_xcmplx_0(X2))|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_131843, plain, (k1_real_1(k4_xcmplx_0(X1))=k7_xcmplx_0(k32_sin_cos,np__2)|~r1_xxreal_0(k7_xcmplx_0(k32_sin_cos,np__2),X1)|~r1_xxreal_0(X1,k7_xcmplx_0(k32_sin_cos,np__2))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_135744, plain, (k4_xcmplx_0(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_136222, plain, (r1_xxreal_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_128352, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_132003, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k4_xcmplx_0(esk1_0)))).#trainneg
cnf(i_0_132239, negated_conjecture, (k1_real_1(X1)=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(X1))|~v1_xreal_0(k7_square_1(np__2))|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_64129, negated_conjecture, (r1_xxreal_0(k7_square_1(esk1_0),X1)|r1_xxreal_0(X1,k7_square_1(esk1_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_132622, negated_conjecture, (r1_xxreal_0(X1,k4_xcmplx_0(k32_sin_cos))|~r1_xxreal_0(X1,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2))|~v1_xreal_0(X1)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_128896, plain, (k1_real_1(k1_real_1(X1))=X2|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers))).#trainneg
cnf(i_0_64131, negated_conjecture, (r1_xxreal_0(k7_square_1(esk1_0),X1)|r1_xxreal_0(X1,k7_square_1(esk1_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_60315, plain, (k4_xcmplx_0(k32_sin_cos)=np__3|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_60318, plain, (k4_xcmplx_0(k32_sin_cos)=np__3|~r1_xxreal_0(np__3,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_134331, negated_conjecture, (k1_real_1(k1_real_1(k7_square_1(np__2)))=k4_xcmplx_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_132457, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(X1)|~m1_subset_1(k1_real_1(X2),k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_138241, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_137251, plain, (r1_xxreal_0(k7_square_1(np__2),k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k7_square_1(np__2)))).#trainneg
cnf(i_0_138245, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_138240, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k32_sin_cos)|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_137499, plain, (k1_real_1(k4_xcmplx_0(X1))=X2|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_138274, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_138182, negated_conjecture, (k1_real_1(k7_square_1(np__2))=X1|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(X1,k1_real_1(k7_square_1(np__2)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_137252, plain, (r1_xxreal_0(k7_square_1(np__2),np__4)|r1_xxreal_0(np__4,k7_square_1(np__2)))).#trainneg
cnf(i_0_138361, negated_conjecture, (X1=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,X1)|~r1_xxreal_0(X1,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(X1)|~v1_xreal_0(k1_real_1(k7_square_1(np__2)))|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_138280, negated_conjecture, (k1_real_1(X1)=k1_real_1(k7_square_1(np__2))|~r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k1_real_1(X1))|~r1_xxreal_0(k4_xcmplx_0(k32_sin_cos),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_real_1(k7_square_1(np__2)),k1_numbers)|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_113588, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_138591, plain, (v2_xxreal_0(k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_134134, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_113592, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_138585, plain, (~v3_xxreal_0(k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_114382, negated_conjecture, (~v1_xboole_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112926, negated_conjecture, (~v1_xboole_0(k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112636, negated_conjecture, (~v2_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_137253, plain, (r1_xxreal_0(k7_square_1(np__2),np__3)|r1_xxreal_0(np__3,k7_square_1(np__2)))).#trainneg
cnf(i_0_131889, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=np__1|~r1_xxreal_0(np__1,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110815, negated_conjecture, (~v2_xxreal_0(k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_113647, negated_conjecture, (~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk1_0))).#trainneg
cnf(i_0_113600, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),np__0)|~r1_xxreal_0(np__0,k7_square_1(np__2)))).#trainneg
cnf(i_0_114380, negated_conjecture, (~r1_xxreal_0(esk34_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk34_0))).#trainneg
cnf(i_0_113596, negated_conjecture, (~r1_xxreal_0(np__0,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(np__0,k7_square_1(np__2)))).#trainneg
cnf(i_0_137258, plain, (r1_xxreal_0(k7_square_1(np__2),np__2)|r1_xxreal_0(np__2,k7_square_1(np__2)))).#trainneg
cnf(i_0_122334, negated_conjecture, (~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0)|~v1_xboole_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_115037, negated_conjecture, (~r1_xxreal_0(k4_xcmplx_0(k7_square_1(np__2)),np__0)|~r1_xxreal_0(k7_square_1(np__2),np__0))).#trainneg
cnf(i_0_121154, negated_conjecture, (~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0)|~v1_xboole_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_121150, negated_conjecture, (~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0)|~v1_xboole_0(k4_xcmplx_0(esk23_0)))).#trainneg
cnf(i_0_118378, negated_conjecture, (~r1_xxreal_0(esk10_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk10_0)|~v2_xxreal_0(k4_xcmplx_0(esk10_0)))).#trainneg
cnf(i_0_117907, negated_conjecture, (~r1_xxreal_0(esk19_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk19_0)|~v2_xxreal_0(k4_xcmplx_0(esk19_0)))).#trainneg
cnf(i_0_139364, plain, (esk6_3(X1,X1,X2)=k1_real_1(k4_xcmplx_0(X1))|~r1_xxreal_0(esk6_3(X1,X1,X2),X1)|~v1_xreal_0(esk6_3(X1,X1,X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_119533, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__3|~r1_xxreal_0(np__3,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_117903, negated_conjecture, (~r1_xxreal_0(esk23_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk23_0)|~v2_xxreal_0(k4_xcmplx_0(esk23_0)))).#trainneg
cnf(i_0_117899, negated_conjecture, (~r1_xxreal_0(esk29_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk29_0)|~v2_xxreal_0(k4_xcmplx_0(esk29_0)))).#trainneg
cnf(i_0_117043, negated_conjecture, (~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0)|~v2_xxreal_0(k4_xcmplx_0(esk37_0)))).#trainneg
cnf(i_0_121146, negated_conjecture, (~r1_xxreal_0(esk37_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk37_0)|~v1_xboole_0(k4_xcmplx_0(esk37_0)))).#trainneg
cnf(i_0_120781, negated_conjecture, (~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0)|~v1_xboole_0(k4_xcmplx_0(esk43_0)))).#trainneg
cnf(i_0_114996, negated_conjecture, (k1_real_1(k7_square_1(np__2))=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_116085, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__2|~r1_xxreal_0(np__2,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_117039, negated_conjecture, (~r1_xxreal_0(esk43_0,k7_square_1(np__2))|~r1_xxreal_0(k7_square_1(np__2),esk43_0)|~v2_xxreal_0(k4_xcmplx_0(esk43_0)))).#trainneg
cnf(i_0_92107, negated_conjecture, (~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k7_xcmplx_0(k7_xcmplx_0(np__3,np__4),k7_xcmplx_0(np__1,k32_sin_cos)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_115194, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__1|~r1_xxreal_0(np__1,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_109600, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),esk1_0)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112669, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),k32_sin_cos)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_115091, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__4|~r1_xxreal_0(np__4,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112673, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__3)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_110810, negated_conjecture, (v3_xxreal_0(k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112672, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__2)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_109602, negated_conjecture, (k1_real_1(k7_square_1(np__2))=esk1_0|~r1_xxreal_0(esk1_0,k1_real_1(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112671, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__1)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_112670, negated_conjecture, (r1_xxreal_0(k1_real_1(k7_square_1(np__2)),np__4)|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_27199, plain, (m1_subset_1(k4_xcmplx_0(esk34_0),k1_numbers))).#trainneg
cnf(i_0_139379, plain, (k1_real_1(k4_xcmplx_0(X1))=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_139921, plain, (~v1_xboole_0(k7_square_1(np__2))|~v1_xxreal_0(k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_140996, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_141415, plain, (k1_real_1(k4_xcmplx_0(X1))=k1_real_1(X2)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k4_xcmplx_0(X2),X1))).#trainneg
cnf(i_0_4931, plain, (v1_xcmplx_0(k7_xcmplx_0(np__13,np__4))|~v1_xcmplx_0(k7_xcmplx_0(np__3,np__4))|~v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_140444, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0)|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_140047, negated_conjecture, (~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(esk1_0,k7_square_1(np__2))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_4932, plain, (v1_xcmplx_0(k7_xcmplx_0(np__15,np__4))|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__4))|~v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_140931, negated_conjecture, (k4_xcmplx_0(np__1)=esk1_0|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(k7_square_1(np__2)))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_139361, plain, (esk6_3(X2,X1,X3)=k1_real_1(k4_xcmplx_0(X1))|~r1_xxreal_0(esk6_3(X2,X1,X3),X1)|~v1_xreal_0(esk6_3(X2,X1,X3))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_rcomp_1(X1,X3)))).#trainneg
cnf(i_0_4942, plain, (v1_xcmplx_0(k7_xcmplx_0(np__11,np__4))|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__4))|~v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_140830, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_4951, plain, (v1_xcmplx_0(k7_xcmplx_0(np__5,np__4))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_60320, plain, (k4_xcmplx_0(np__3)=k32_sin_cos|~r1_xxreal_0(np__3,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_140013, negated_conjecture, (~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__1,k4_xcmplx_0(k7_square_1(np__2)))|~r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_139363, plain, (esk6_3(X2,X1,X2)=k1_real_1(k4_xcmplx_0(X1))|~r1_xxreal_0(esk6_3(X2,X1,X2),X1)|~v1_xreal_0(esk6_3(X2,X1,X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_140127, negated_conjecture, (k4_xcmplx_0(k7_square_1(np__2))=np__3|~r1_xxreal_0(np__3,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_60330, plain, (k4_xcmplx_0(np__2)=k32_sin_cos|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_60333, plain, (k4_xcmplx_0(k32_sin_cos)=np__2|~r1_xxreal_0(k32_sin_cos,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_60336, plain, (k4_xcmplx_0(k32_sin_cos)=np__2|~r1_xxreal_0(np__2,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_60338, plain, (k4_xcmplx_0(np__2)=k32_sin_cos|~r1_xxreal_0(np__2,k4_xcmplx_0(k32_sin_cos)))).#trainneg
cnf(i_0_139370, plain, (k1_real_1(k4_xcmplx_0(X1))=esk6_3(X2,X2,X2)|~r1_xxreal_0(esk6_3(X2,X2,X2),X1)|~v1_xreal_0(esk6_3(X2,X2,X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_139381, plain, (k1_real_1(k4_xcmplx_0(X1))=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_124762, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(k1_real_1(X2),X1)|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_140177, negated_conjecture, (k1_real_1(k7_square_1(np__2))=np__3|~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__3,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_125397, plain, (r1_xxreal_0(np__1,esk6_3(X1,X1,X1))|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_128284, plain, (k4_xcmplx_0(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_142603, negated_conjecture, (~v1_xreal_0(k7_square_1(np__2))|~r1_xxreal_0(np__3,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_141421, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),k7_square_1(np__2)))).#trainneg
cnf(i_0_141749, negated_conjecture, (~r1_xxreal_0(np__1,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_141699, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),k7_square_1(np__2))|~m1_subset_1(k7_square_1(np__2),k1_numbers))).#trainneg
cnf(i_0_4971, plain, (v1_xcmplx_0(k7_xcmplx_0(np__9,np__4))|~v1_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__4))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_142204, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X1),X2))).#trainneg
cnf(i_0_138592, plain, (k7_square_1(np__2)=np__2|np__2=np__1|~r1_xxreal_0(k7_square_1(np__2),np__2)|~r1_xxreal_0(np__2,k7_square_1(np__2))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_142642, plain, (esk6_3(X1,X1,X1)=np__1|~r1_xxreal_0(k4_xcmplx_0(np__1),k1_real_1(X1))|~v1_xreal_0(X1)|~r1_xxreal_0(k1_real_1(X1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_142077, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,k4_xcmplx_0(k4_xcmplx_0(X2)))|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X2),k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_142209, plain, (k1_real_1(X1)=k1_real_1(k1_real_1(X2))|~r1_xxreal_0(X1,k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(X2),X1))).#trainneg
cnf(i_0_141702, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__2),k7_square_1(np__2)))).#trainneg
cnf(i_0_27256, plain, (m1_subset_1(k4_xcmplx_0(esk29_0),k1_numbers))).#trainneg
cnf(i_0_128784, plain, (k1_real_1(k1_real_1(X1))=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_136212, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|~r1_xxreal_0(X1,k1_real_1(X2))|~v1_xreal_0(k1_real_1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_61211, plain, (k4_xcmplx_0(np__4)=np__4|~r1_xxreal_0(np__4,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_125854, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(k1_real_1(k4_xcmplx_0(X2)),X1)|~v1_xreal_0(k1_real_1(X2))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_28637, plain, (m1_subset_1(k4_xcmplx_0(esk19_0),k1_numbers))).#trainneg
cnf(i_0_142790, negated_conjecture, (k1_seq_1(k2_sincos10,esk1_0)=k32_sin_cos|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k1_seq_1(k2_sincos10,esk1_0)))).#trainneg
cnf(i_0_142786, negated_conjecture, (v2_xxreal_0(k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k1_seq_1(k2_sincos10,esk1_0)))).#trainneg
cnf(i_0_142785, negated_conjecture, (~r1_xxreal_0(np__3,k4_xcmplx_0(k7_square_1(np__2))))).#trainneg
cnf(i_0_142787, negated_conjecture, (~v3_xxreal_0(k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k1_seq_1(k2_sincos10,esk1_0)))).#trainneg
cnf(i_0_142806, negated_conjecture, (k1_seq_1(k2_sincos10,esk1_0)=X1|k1_seq_1(k2_sincos10,esk1_0)=k32_sin_cos|~r1_xxreal_0(X1,k1_seq_1(k2_sincos10,esk1_0))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~r1_xxreal_0(X1,k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143433, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(np__2)|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_143441, negated_conjecture, (k7_square_1(np__2)=k4_xcmplx_0(np__2)|~r1_xxreal_0(np__2,k4_xcmplx_0(k7_square_1(np__2)))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_143889, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k1_real_1(X2),k4_xcmplx_0(X1))|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_142799, negated_conjecture, (k1_real_1(k1_real_1(k1_seq_1(k2_sincos10,esk1_0)))=k32_sin_cos|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k1_seq_1(k2_sincos10,esk1_0)))).#trainneg
cnf(i_0_143576, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k1_real_1(X1)),X2)|~r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(k1_real_1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143530, negated_conjecture, (k1_real_1(X1)=esk1_0|~r1_xxreal_0(k4_xcmplx_0(esk1_0),X1)|~r1_xxreal_0(X1,k4_xcmplx_0(esk1_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_125495, plain, (r1_xxreal_0(X1,esk6_3(k4_xcmplx_0(X2),k4_xcmplx_0(X2),k4_xcmplx_0(X2)))|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_129948, negated_conjecture, (r2_hidden(esk1_0,k1_rcomp_1(np__1,k4_xcmplx_0(np__1)))|~r1_xxreal_0(k7_square_1(np__2),k4_xcmplx_0(np__1))|~v1_xreal_0(k7_square_1(np__2)))).#trainneg
cnf(i_0_61230, plain, (k4_xcmplx_0(np__3)=np__4|~r1_xxreal_0(np__4,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_144128, negated_conjecture, (~v1_xboole_0(k1_seq_1(k2_sincos10,esk1_0))|~v1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0))|~v1_xreal_0(k1_seq_1(k2_sincos10,esk1_0)))).#trainneg
cnf(i_0_144305, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_144125, negated_conjecture, (k1_seq_1(k2_sincos10,esk1_0)=k32_sin_cos|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~m1_subset_1(k1_seq_1(k2_sincos10,esk1_0),k1_numbers))).#trainneg
cnf(i_0_143102, plain, (k7_square_1(np__2)=np__2|np__2=np__1|~r1_xxreal_0(k7_square_1(np__2),np__2)|~r1_xxreal_0(np__2,k7_square_1(np__2)))).#trainneg
cnf(i_0_61233, plain, (k4_xcmplx_0(np__4)=np__3|~r1_xxreal_0(np__4,k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_61236, plain, (k4_xcmplx_0(np__4)=np__3|~r1_xxreal_0(np__3,k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_29063, plain, (v1_xreal_0(k4_xcmplx_0(esk28_0))|~v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk28_0))))).#trainneg
cnf(i_0_29137, plain, (v1_xboole_0(k4_xcmplx_0(esk18_0))|~v1_xboole_0(k7_xcmplx_0(np__1,k4_xcmplx_0(esk18_0))))).#trainneg
cnf(i_0_144842, plain, (k1_real_1(X1)=k1_real_1(k4_xcmplx_0(X2))|~r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_144298, plain, (k1_real_1(X1)=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X2)),X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(k4_xcmplx_0(k4_xcmplx_0(X1)),X2)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_144419, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(k1_real_1(k1_seq_1(k2_sincos10,esk1_0))),k32_sin_cos)|~v1_xreal_0(k1_real_1(k1_seq_1(k2_sincos10,esk1_0)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k1_seq_1(k2_sincos10,esk1_0)))).#trainneg
cnf(i_0_144374, negated_conjecture, (k1_real_1(k1_seq_1(k2_sincos10,esk1_0))=k4_xcmplx_0(esk1_0)|~r1_xxreal_0(k4_xcmplx_0(esk1_0),k1_real_1(k1_seq_1(k2_sincos10,esk1_0)))|~v1_xreal_0(k1_real_1(k1_seq_1(k2_sincos10,esk1_0)))|~r1_xxreal_0(k1_seq_1(k2_sincos10,esk1_0),k32_sin_cos)|~v1_xreal_0(k1_seq_1(k2_sincos10,esk1_0)))).#trainneg
cnf(i_0_144753, plain, (X1=k1_real_1(X2)|~r1_xxreal_0(k4_xcmplx_0(X1),X2)|~r1_xxreal_0(X2,k1_real_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k1_real_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_134134, negated_conjecture, (~r1_xxreal_0(k7_square_1(np__2),esk1_0))).#trainneg
cnf(i_0_9699, plain, (k1_real_1(np__1)=esk1_0)).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 1025
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 1317
# Removed in clause preprocessing      : 91
# Initial clauses in saturation        : 1226
# Processed clauses                    : 12531
# ...of these trivial                  : 1072
# ...subsumed                          : 6115
# ...remaining for further processing  : 5344
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 288
# Backward-rewritten                   : 791
# Generated clauses                    : 143631
# ...of the previous two non-trivial   : 108660
# Contextual simplify-reflections      : 605
# Paramodulations                      : 143566
# Factorizations                       : 6
# Equation resolutions                 : 59
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 4260
#    Positive orientable unit clauses  : 1279
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 293
#    Non-unit-clauses                  : 2685
# Current number of unprocessed clauses: 96234
# ...number of literals in the above   : 455454
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6423
# Proof object given clauses           : 410
# Proof search given clauses           : 5344
# Clause-clause subsumption calls (NU) : 2662417
# Rec. Clause-clause subsumption calls : 675505
# Non-unit clause-clause subsumptions  : 5067
# Unit Clause-clause subsumption calls : 113580
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 398
# BW rewrite match successes           : 133
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2707538

# -------------------------------------------------
# User time                : 59.713 s
# System time              : 1.340 s
# Total time               : 61.053 s
# Maximum resident set size: 247180 pages
59.71user 1.37system 3:22.24elapsed 30%CPU (0avgtext+0avgdata 247560maxresident)k
0inputs+1816outputs (0major+335947minor)pagefaults 0swaps
