# local_data = 0x16735a0
#TFS# Evaluating axioms..
# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
#TF#SERVER: Lgb model filter skipped 0 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=0, query=197, total=208).
#TF#SERVER: Sent 51771 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (792 bytes expected).
#TF#SERVER: Received 792 bytes from the TF server.
# Preprocessing time       : 0.824 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 4 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=194, query=288, total=493).
#TF#SERVER: Sent 119259 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1932 bytes expected).
#TF#SERVER: Received 1932 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 82 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=404, query=832, total=1247).
#TF#SERVER: Sent 303395 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4948 bytes expected).
#TF#SERVER: Received 4948 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 153 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=530, query=876, total=1417).
#TF#SERVER: Sent 607550 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5628 bytes expected).
#TF#SERVER: Received 5628 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 242 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=626, query=790, total=1427).
#TF#SERVER: Sent 612225 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5668 bytes expected).
#TF#SERVER: Received 5668 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 423 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=708, query=646, total=1365).
#TF#SERVER: Sent 629703 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5420 bytes expected).
#TF#SERVER: Received 5420 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 259 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=808, query=812, total=1631).
#TF#SERVER: Sent 854629 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6484 bytes expected).
#TF#SERVER: Received 6484 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 194 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=909, query=843, total=1763).
#TF#SERVER: Sent 1199390 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7012 bytes expected).
#TF#SERVER: Received 7012 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 180 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=995, query=854, total=1860).
#TF#SERVER: Sent 1297638 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7400 bytes expected).
#TF#SERVER: Received 7400 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 96 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1076, query=935, total=2022).
#TF#SERVER: Sent 1617183 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8048 bytes expected).
#TF#SERVER: Received 8048 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 336 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1121, query=750, total=1882).
#TF#SERVER: Sent 1259574 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7488 bytes expected).
#TF#SERVER: Received 7488 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 211 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1196, query=819, total=2026).
#TF#SERVER: Sent 1506206 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8064 bytes expected).
#TF#SERVER: Received 8064 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 201 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1261, query=925, total=2197).
#TF#SERVER: Sent 1818439 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8748 bytes expected).
#TF#SERVER: Received 8748 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 229 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1305, query=821, total=2137).
#TF#SERVER: Sent 1795199 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8508 bytes expected).
#TF#SERVER: Received 8508 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 542 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1339, query=523, total=1873).
#TF#SERVER: Sent 1215217 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7452 bytes expected).
#TF#SERVER: Received 7452 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 327 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1390, query=700, total=2101).
#TF#SERVER: Sent 1486357 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8364 bytes expected).
#TF#SERVER: Received 8364 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 237 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1438, query=811, total=2260).
#TF#SERVER: Sent 1872050 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9000 bytes expected).
#TF#SERVER: Received 9000 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 246 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1487, query=781, total=2279).
#TF#SERVER: Sent 1860002 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9076 bytes expected).
#TF#SERVER: Received 9076 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 455 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1531, query=571, total=2113).
#TF#SERVER: Sent 1611674 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8412 bytes expected).
#TF#SERVER: Received 8412 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 321 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1591, query=707, total=2309).
#TF#SERVER: Sent 1772749 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9196 bytes expected).
#TF#SERVER: Received 9196 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 149 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1645, query=892, total=2548).
#TF#SERVER: Sent 1973782 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10152 bytes expected).
#TF#SERVER: Received 10152 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 339 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1670, query=770, total=2451).
#TF#SERVER: Sent 1856694 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9764 bytes expected).
#TF#SERVER: Received 9764 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 400 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1709, query=724, total=2444).
#TF#SERVER: Sent 1835134 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9736 bytes expected).
#TF#SERVER: Received 9736 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 372 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1754, query=663, total=2428).
#TF#SERVER: Sent 1616067 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9672 bytes expected).
#TF#SERVER: Received 9672 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 177 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1805, query=847, total=2663).
#TF#SERVER: Sent 1732166 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10612 bytes expected).
#TF#SERVER: Received 10612 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 209 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1839, query=886, total=2736).
#TF#SERVER: Sent 1625740 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10904 bytes expected).
#TF#SERVER: Received 10904 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 206 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1876, query=829, total=2716).
#TF#SERVER: Sent 1513047 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10824 bytes expected).
#TF#SERVER: Received 10824 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 78 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1926, query=966, total=2903).
#TF#SERVER: Sent 2144321 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11572 bytes expected).
#TF#SERVER: Received 11572 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 103 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=1978, query=937, total=2926).
#TF#SERVER: Sent 2031152 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11664 bytes expected).
#TF#SERVER: Received 11664 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 137 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2029, query=927, total=2967).
#TF#SERVER: Sent 2007866 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11828 bytes expected).
#TF#SERVER: Received 11828 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 135 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=898, total=2957).
#TF#SERVER: Sent 1992612 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11788 bytes expected).
#TF#SERVER: Received 11788 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 98 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=971, total=3030).
#TF#SERVER: Sent 2347089 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12080 bytes expected).
#TF#SERVER: Received 12080 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 71 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1015, total=3074).
#TF#SERVER: Sent 2113807 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12256 bytes expected).
#TF#SERVER: Received 12256 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 44 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=982, total=3041).
#TF#SERVER: Sent 1992881 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12124 bytes expected).
#TF#SERVER: Received 12124 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 95 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=969, total=3028).
#TF#SERVER: Sent 2491216 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12072 bytes expected).
#TF#SERVER: Received 12072 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 125 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=899, total=2958).
#TF#SERVER: Sent 2376101 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11792 bytes expected).
#TF#SERVER: Received 11792 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 182 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=871, total=2930).
#TF#SERVER: Sent 2152499 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11680 bytes expected).
#TF#SERVER: Received 11680 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 111 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=946, total=3005).
#TF#SERVER: Sent 2447622 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11980 bytes expected).
#TF#SERVER: Received 11980 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 141 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=977, total=3036).
#TF#SERVER: Sent 2500227 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12104 bytes expected).
#TF#SERVER: Received 12104 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 54 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=999, total=3058).
#TF#SERVER: Sent 2520620 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12192 bytes expected).
#TF#SERVER: Received 12192 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 125 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=932, total=2991).
#TF#SERVER: Sent 2031696 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11924 bytes expected).
#TF#SERVER: Received 11924 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 132 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=898, total=2957).
#TF#SERVER: Sent 1970306 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11788 bytes expected).
#TF#SERVER: Received 11788 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 107 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=991, total=3050).
#TF#SERVER: Sent 2506853 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12160 bytes expected).
#TF#SERVER: Received 12160 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 97 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1000, total=3059).
#TF#SERVER: Sent 2317390 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12196 bytes expected).
#TF#SERVER: Received 12196 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 171 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=900, total=2959).
#TF#SERVER: Sent 1898979 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11796 bytes expected).
#TF#SERVER: Received 11796 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 97 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=933, total=2992).
#TF#SERVER: Sent 2009588 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11928 bytes expected).
#TF#SERVER: Received 11928 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 117 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=942, total=3001).
#TF#SERVER: Sent 1818675 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11964 bytes expected).
#TF#SERVER: Received 11964 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 105 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1016, total=3075).
#TF#SERVER: Sent 2220097 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12260 bytes expected).
#TF#SERVER: Received 12260 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 252 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=780, total=2839).
#TF#SERVER: Sent 1965773 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11316 bytes expected).
#TF#SERVER: Received 11316 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 153 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=872, total=2931).
#TF#SERVER: Sent 2089851 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11684 bytes expected).
#TF#SERVER: Received 11684 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 137 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=893, total=2952).
#TF#SERVER: Sent 1467506 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11768 bytes expected).
#TF#SERVER: Received 11768 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 155 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=946, total=3005).
#TF#SERVER: Sent 1533766 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11980 bytes expected).
#TF#SERVER: Received 11980 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 243 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=788, total=2847).
#TF#SERVER: Sent 1566960 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11348 bytes expected).
#TF#SERVER: Received 11348 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 422 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=959, total=3018).
#TF#SERVER: Sent 2058550 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12032 bytes expected).
#TF#SERVER: Received 12032 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 32 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=995, total=3054).
#TF#SERVER: Sent 1868139 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12176 bytes expected).
#TF#SERVER: Received 12176 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 58 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=974, total=3033).
#TF#SERVER: Sent 1692080 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12092 bytes expected).
#TF#SERVER: Received 12092 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 23 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1004, total=3063).
#TF#SERVER: Sent 1662074 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12212 bytes expected).
#TF#SERVER: Received 12212 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 25 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1053, total=3112).
#TF#SERVER: Sent 1653136 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12408 bytes expected).
#TF#SERVER: Received 12408 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 44 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=991, total=3050).
#TF#SERVER: Sent 1653748 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12160 bytes expected).
#TF#SERVER: Received 12160 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 94 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=950, total=3009).
#TF#SERVER: Sent 1652084 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (11996 bytes expected).
#TF#SERVER: Received 11996 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 58 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1026, total=3085).
#TF#SERVER: Sent 1847290 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12300 bytes expected).
#TF#SERVER: Received 12300 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 49 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=990, total=3049).
#TF#SERVER: Sent 1525724 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12156 bytes expected).
#TF#SERVER: Received 12156 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 43 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=986, total=3045).
#TF#SERVER: Sent 1515971 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12140 bytes expected).
#TF#SERVER: Received 12140 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 42 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=998, total=3057).
#TF#SERVER: Sent 1543592 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12188 bytes expected).
#TF#SERVER: Received 12188 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 46 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1056, total=3115).
#TF#SERVER: Sent 1627245 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12420 bytes expected).
#TF#SERVER: Received 12420 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 60 clauses
#TF#SERVER: Sending data to the TF server (conj=11, context=2048, query=1047, total=3106).
#TF#SERVER: Sent 2172294 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (12384 bytes expected).
#TF#SERVER: Received 12384 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds65new/t88_chord', redefinition_k1_nat_1)).
fof(commutativity_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k1_nat_1(X2,X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', commutativity_k1_nat_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc8_ordinal1)).
fof(fc1_abian, axiom, ![X1]:(v1_int_1(X1)=>v1_abian(k3_xcmplx_0(np__2,X1))), file('/scratch/mptp/minsubpreds65new/t88_chord', fc1_abian)).
fof(cc2_int_1, axiom, ![X1]:(v7_ordinal1(X1)=>v1_int_1(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc2_int_1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/minsubpreds65new/t88_chord', redefinition_k5_numbers)).
fof(t1_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X1))=>X1=X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', t1_xxreal_0)).
fof(t4_chord, axiom, ![X1]:((v1_int_1(X1)&~(v1_abian(X1)))=>![X2]:((v1_int_1(X2)&~(v1_abian(X2)))=>(~(r1_xxreal_0(X1,X2))=>r1_xxreal_0(k2_xcmplx_0(X2,np__2),X1)))), file('/scratch/mptp/minsubpreds65new/t88_chord', t4_chord)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/minsubpreds65new/t88_chord', fc6_membered)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', cc13_membered)).
fof(t88_chord, conjecture, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>~((((v8_glib_001(X2,X1)&v5_chord(X2,X1))&k15_glib_001(X1,X2)=np__4)&![X3]:(~(r1_glib_000(X1,k1_funct_1(X2,np__1),k1_funct_1(X2,np__5),X3))&~(r1_glib_000(X1,k1_funct_1(X2,np__3),k1_funct_1(X2,np__7),X3))))))), file('/scratch/mptp/minsubpreds65new/t88_chord', t88_chord)).
fof(t7_xreal_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>![X3]:(v1_xreal_0(X3)=>![X4]:(v1_xreal_0(X4)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X3,X4))=>r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X4))))))), file('/scratch/mptp/minsubpreds65new/t88_chord', t7_xreal_1)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', cc14_membered)).
fof(fc4_abian, axiom, ![X1]:((v1_int_1(X1)&v1_abian(X1))=>~(v1_abian(k6_xcmplx_0(X1,np__1)))), file('/scratch/mptp/minsubpreds65new/t88_chord', fc4_abian)).
fof(rqRealMult__k3_xcmplx_0__r2_r2_r4, axiom, k3_xcmplx_0(np__2,np__2)=np__4, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealMult__k3_xcmplx_0__r2_r2_r4)).
fof(spc7_numerals, axiom, (((v2_xxreal_0(np__7)&m2_subset_1(np__7,k1_numbers,k5_numbers))&m1_subset_1(np__7,k5_numbers))&m1_subset_1(np__7,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc7_numerals)).
fof(spc9_numerals, axiom, (((v2_xxreal_0(np__9)&m2_subset_1(np__9,k1_numbers,k5_numbers))&m1_subset_1(np__9,k5_numbers))&m1_subset_1(np__9,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc9_numerals)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc4_membered)).
fof(rqRealMult__k3_xcmplx_0__r2_r3_r6, axiom, k3_xcmplx_0(np__2,np__3)=np__6, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealMult__k3_xcmplx_0__r2_r3_r6)).
fof(t112_glib_001, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>k3_finseq_1(X2)=k2_nat_1(k4_nat_1(np__2,k15_glib_001(X1,X2)),np__1))), file('/scratch/mptp/minsubpreds65new/t88_chord', t112_glib_001)).
fof(rqRealAdd__k2_xcmplx_0__r1_r2_r3, axiom, k2_xcmplx_0(np__1,np__2)=np__3, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r1_r2_r3)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc2_numerals)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/minsubpreds65new/t88_chord', fc3_membered)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc1_numerals)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', cc21_membered)).
fof(dt_k4_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', dt_k4_nat_1)).
fof(redefinition_k4_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds65new/t88_chord', redefinition_k4_nat_1)).
fof(t3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1), file('/scratch/mptp/minsubpreds65new/t88_chord', t3_arithm)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc3_xreal_0)).
fof(t3_chord, axiom, ![X1]:((v1_int_1(X1)&~(v1_abian(X1)))=>![X2]:((v1_int_1(X2)&~(v1_abian(X2)))=>(~(r1_xxreal_0(X2,X1))=>r1_xxreal_0(X1,k6_xcmplx_0(X2,np__2))))), file('/scratch/mptp/minsubpreds65new/t88_chord', t3_chord)).
fof(fc8_abian, axiom, ![X1, X2]:((((v1_int_1(X1)&~(v1_abian(X1)))&v1_int_1(X2))&~(v1_abian(X2)))=>~(v1_abian(k3_xcmplx_0(X1,X2)))), file('/scratch/mptp/minsubpreds65new/t88_chord', fc8_abian)).
fof(rqRealDiff__k6_xcmplx_0__r5_r1_r4, axiom, k6_xcmplx_0(np__5,np__1)=np__4, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealDiff__k6_xcmplx_0__r5_r1_r4)).
fof(spc5_numerals, axiom, (((v2_xxreal_0(np__5)&m2_subset_1(np__5,k1_numbers,k5_numbers))&m1_subset_1(np__5,k5_numbers))&m1_subset_1(np__5,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc5_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r7_r2_r9, axiom, k2_xcmplx_0(np__7,np__2)=np__9, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r7_r2_r9)).
fof(rqRealDiff__k6_xcmplx_0__r7_r1_r6, axiom, k6_xcmplx_0(np__7,np__1)=np__6, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealDiff__k6_xcmplx_0__r7_r1_r6)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc3_numerals)).
fof(spc4_numerals, axiom, (((v2_xxreal_0(np__4)&m2_subset_1(np__4,k1_numbers,k5_numbers))&m1_subset_1(np__4,k5_numbers))&m1_subset_1(np__4,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc4_numerals)).
fof(t2_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X3))=>r1_xxreal_0(X1,X3))))), file('/scratch/mptp/minsubpreds65new/t88_chord', t2_xxreal_0)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', dt_k5_numbers)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', cc1_subset_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/minsubpreds65new/t88_chord', fc6_ordinal1)).
fof(fc2_abian, axiom, ![X1]:((v1_int_1(X1)&v1_abian(X1))=>~(v1_abian(k2_xcmplx_0(X1,np__1)))), file('/scratch/mptp/minsubpreds65new/t88_chord', fc2_abian)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc2_xxreal_0)).
fof(redefinition_k2_nat_1, axiom, ![X1, X2]:((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds65new/t88_chord', redefinition_k2_nat_1)).
fof(rqRealMult__k3_xcmplx_0__r2_r4_r8, axiom, k3_xcmplx_0(np__2,np__4)=np__8, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealMult__k3_xcmplx_0__r2_r4_r8)).
fof(t84_chord, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>~((v5_chord(X2,X1)&![X3]:((v7_ordinal1(X3)&~(v1_abian(X3)))=>![X4]:((v7_ordinal1(X4)&~(v1_abian(X4)))=>~(((((~(r1_xxreal_0(X4,k1_nat_1(X3,np__2)))&r1_xxreal_0(X4,k3_finseq_1(X2)))&k1_funct_1(X2,X3)!=k1_funct_1(X2,X4))&?[X5]:r1_glib_000(X1,k1_funct_1(X2,X3),k1_funct_1(X2,X4),X5))&(v8_glib_001(X2,X1)=>((~((X3=np__1&X4=k3_finseq_1(X2)))&~((X3=np__1&X4=k6_xcmplx_0(k3_finseq_1(X2),np__2))))&~((X3=np__3&X4=k3_finseq_1(X2))))))))))))), file('/scratch/mptp/minsubpreds65new/t88_chord', t84_chord)).
fof(dt_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', dt_k1_nat_1)).
fof(cc2_xreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc2_xreal_0)).
fof(rqRealDiff__k6_xcmplx_0__r7_r2_r5, axiom, k6_xcmplx_0(np__7,np__2)=np__5, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealDiff__k6_xcmplx_0__r7_r2_r5)).
fof(rqRealAdd__k2_xcmplx_0__r2_r1_r3, axiom, k2_xcmplx_0(np__2,np__1)=np__3, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r2_r1_r3)).
fof(rqRealAdd__k2_xcmplx_0__r8_r1_r9, axiom, k2_xcmplx_0(np__8,np__1)=np__9, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r8_r1_r9)).
fof(spc8_numerals, axiom, (((v2_xxreal_0(np__8)&m2_subset_1(np__8,k1_numbers,k5_numbers))&m1_subset_1(np__8,k5_numbers))&m1_subset_1(np__8,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc8_numerals)).
fof(cc20_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v2_membered(X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', cc20_membered)).
fof(rqRealAdd__k2_xcmplx_0__r0_r1_r1, axiom, k2_xcmplx_0(np__0,np__1)=np__1, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r0_r1_r1)).
fof(rqRealMult__k3_xcmplx_0__r2_r0_r0, axiom, k3_xcmplx_0(np__2,np__0)=np__0, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealMult__k3_xcmplx_0__r2_r0_r0)).
fof(rqRealAdd__k2_xcmplx_0__r2_r5_r7, axiom, k2_xcmplx_0(np__2,np__5)=np__7, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r2_r5_r7)).
fof(rqRealDiff__k6_xcmplx_0__r9_r2_r7, axiom, k6_xcmplx_0(np__9,np__2)=np__7, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealDiff__k6_xcmplx_0__r9_r2_r7)).
fof(rqRealDiff__k6_xcmplx_0__r9_r1_r8, axiom, k6_xcmplx_0(np__9,np__1)=np__8, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealDiff__k6_xcmplx_0__r9_r1_r8)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r2, axiom, r1_xxreal_0(np__2,np__2), file('/scratch/mptp/minsubpreds65new/t88_chord', rqLessOrEqual__r1_xxreal_0__r2_r2)).
fof(t12_abian, axiom, ![X1]:((v7_ordinal1(X1)&~(v1_abian(X1)))=>r1_xxreal_0(np__1,X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', t12_abian)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/minsubpreds65new/t88_chord', spc0_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r3_r2_r5, axiom, k2_xcmplx_0(np__3,np__2)=np__5, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r3_r2_r5)).
fof(rqRealAdd__k2_xcmplx_0__r5_r2_r7, axiom, k2_xcmplx_0(np__5,np__2)=np__7, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r5_r2_r7)).
fof(rqRealDiff__k6_xcmplx_0__r5_r2_r3, axiom, k6_xcmplx_0(np__5,np__2)=np__3, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealDiff__k6_xcmplx_0__r5_r2_r3)).
fof(rqRealAdd__k2_xcmplx_0__r2_r7_r9, axiom, k2_xcmplx_0(np__2,np__7)=np__9, file('/scratch/mptp/minsubpreds65new/t88_chord', rqRealAdd__k2_xcmplx_0__r2_r7_r9)).
fof(d24_glib_001, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>(v1_glib_001(X2,X1)<=>k3_glib_001(X1,X2)=k4_glib_001(X1,X2)))), file('/scratch/mptp/minsubpreds65new/t88_chord', d24_glib_001)).
fof(cc6_glib_001, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>(v8_glib_001(X2,X1)=>((v1_glib_001(X2,X1)&~(v3_glib_001(X2,X1)))&v5_glib_001(X2,X1))))), file('/scratch/mptp/minsubpreds65new/t88_chord', cc6_glib_001)).
fof(d7_glib_001, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>k4_glib_001(X1,X2)=k1_funct_1(X2,k3_finseq_1(X2)))), file('/scratch/mptp/minsubpreds65new/t88_chord', d7_glib_001)).
fof(d6_glib_001, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>k3_glib_001(X1,X2)=k1_funct_1(X2,np__1))), file('/scratch/mptp/minsubpreds65new/t88_chord', d6_glib_001)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc4_xreal_0)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds65new/t88_chord', cc1_xreal_0)).
fof(t14_glib_000, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2, X3, X4]:(r1_glib_000(X1,X3,X4,X2)=>r1_glib_000(X1,X4,X3,X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', t14_glib_000)).
fof(fc1_int_1, axiom, ![X1, X2]:((v1_int_1(X1)&v1_int_1(X2))=>v1_int_1(k2_xcmplx_0(X1,X2))), file('/scratch/mptp/minsubpreds65new/t88_chord', fc1_int_1)).
fof(fc11_abian, axiom, ![X1, X2]:((((v1_int_1(X1)&v1_abian(X1))&v1_int_1(X2))&~(v1_abian(X2)))=>~(v1_abian(k2_xcmplx_0(X2,X1)))), file('/scratch/mptp/minsubpreds65new/t88_chord', fc11_abian)).
fof(fc17_abian, axiom, ![X1]:((v1_int_1(X1)&~(v1_abian(X1)))=>~(v1_abian(k2_xcmplx_0(X1,np__2)))), file('/scratch/mptp/minsubpreds65new/t88_chord', fc17_abian)).
fof(c_0_77, plain, ![X85, X86]:(~v7_ordinal1(X85)|~m1_subset_1(X86,k5_numbers)|k1_nat_1(X85,X86)=k2_xcmplx_0(X85,X86)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_nat_1])])).
fof(c_0_78, plain, ![X128, X129]:(~v7_ordinal1(X128)|~m1_subset_1(X129,k5_numbers)|k1_nat_1(X128,X129)=k1_nat_1(X129,X128)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k1_nat_1])])).
fof(c_0_79, plain, ![X38]:(~m1_subset_1(X38,k4_ordinal1)|v7_ordinal1(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_80, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_81, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_78])).
fof(c_0_82, plain, ![X71]:(~v1_int_1(X71)|v1_abian(k3_xcmplx_0(np__2,X71))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_abian])])).
fof(c_0_83, plain, ![X27]:(~v7_ordinal1(X27)|v1_int_1(X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_int_1])])).
cnf(c_0_84, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_85, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_86, plain, ![X104, X105]:(~v1_xxreal_0(X104)|(~v1_xxreal_0(X105)|(~r1_xxreal_0(X104,X105)|~r1_xxreal_0(X105,X104)|X104=X105))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xxreal_0])])])).
fof(c_0_87, plain, ![X112, X113]:(~v1_int_1(X112)|v1_abian(X112)|(~v1_int_1(X113)|v1_abian(X113)|(r1_xxreal_0(X112,X113)|r1_xxreal_0(k2_xcmplx_0(X113,np__2),X112)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_chord])])])])).
fof(c_0_88, plain, ![X15, X16]:(~v6_membered(X15)|(~m1_subset_1(X16,X15)|v7_ordinal1(X16))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_89, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
fof(c_0_90, plain, ![X9, X10]:(~v2_membered(X9)|(~m1_subset_1(X10,X9)|v1_xxreal_0(X10))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
fof(c_0_91, negated_conjecture, ~(![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_finset_1(X1))&v1_glib_000(X1))=>![X2]:(m3_glib_001(X2,X1)=>~((((v8_glib_001(X2,X1)&v5_chord(X2,X1))&k15_glib_001(X1,X2)=np__4)&![X3]:(~(r1_glib_000(X1,k1_funct_1(X2,np__1),k1_funct_1(X2,np__5),X3))&~(r1_glib_000(X1,k1_funct_1(X2,np__3),k1_funct_1(X2,np__7),X3)))))))), inference(assume_negation,[status(cth)],[t88_chord])).
fof(c_0_92, plain, ![X115, X116, X117, X118]:(~v1_xreal_0(X115)|(~v1_xreal_0(X116)|(~v1_xreal_0(X117)|(~v1_xreal_0(X118)|(~r1_xxreal_0(X115,X116)|~r1_xxreal_0(X117,X118)|r1_xxreal_0(k2_xcmplx_0(X115,X117),k2_xcmplx_0(X116,X118))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_xreal_1])])])).
fof(c_0_93, plain, ![X11, X12]:(~v3_membered(X11)|(~m1_subset_1(X12,X11)|v1_xreal_0(X12))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_94, plain, (k2_xcmplx_0(X1,X2)=k1_nat_1(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_95, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_81])).
fof(c_0_96, plain, ![X78]:(~v1_int_1(X78)|~v1_abian(X78)|~v1_abian(k6_xcmplx_0(X78,np__1))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_abian])])])).
cnf(c_0_97, plain, (v1_abian(k3_xcmplx_0(np__2,X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_98, plain, (k3_xcmplx_0(np__2,np__2)=np__4), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r2_r2_r4])).
cnf(c_0_99, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_100, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_101, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_102, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_103, plain, (v1_abian(X1)|v1_abian(X2)|r1_xxreal_0(X1,X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X1)|~v1_int_1(X1)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_104, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_105, plain, (m1_subset_1(np__7,k5_numbers)), inference(split_conjunct,[status(thm)],[spc7_numerals])).
cnf(c_0_106, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_107, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_108, plain, (m1_subset_1(np__9,k1_numbers)), inference(split_conjunct,[status(thm)],[spc9_numerals])).
fof(c_0_109, plain, ![X33]:(~v3_membered(X33)|v2_membered(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_110, plain, (k3_xcmplx_0(np__2,np__3)=np__6), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r2_r3_r6])).
fof(c_0_111, plain, ![X97, X98]:(~v1_relat_1(X97)|~v4_relat_1(X97,k5_numbers)|~v1_funct_1(X97)|~v1_finset_1(X97)|~v1_glib_000(X97)|(~m3_glib_001(X98,X97)|k3_finseq_1(X98)=k2_nat_1(k4_nat_1(np__2,k15_glib_001(X97,X98)),np__1))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t112_glib_001])])])).
fof(c_0_112, negated_conjecture, ![X8]:(((((v1_relat_1(esk1_0)&v4_relat_1(esk1_0,k5_numbers))&v1_funct_1(esk1_0))&v1_finset_1(esk1_0))&v1_glib_000(esk1_0))&(m3_glib_001(esk2_0,esk1_0)&(((v8_glib_001(esk2_0,esk1_0)&v5_chord(esk2_0,esk1_0))&k15_glib_001(esk1_0,esk2_0)=np__4)&(~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__5),X8)&~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__3),k1_funct_1(esk2_0,np__7),X8))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_91])])])])])).
cnf(c_0_113, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X4))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_114, plain, (k2_xcmplx_0(np__1,np__2)=np__3), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r2_r3])).
cnf(c_0_115, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_116, plain, (m1_subset_1(np__2,k1_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_117, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_118, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_119, plain, ![X24, X25]:(~v3_membered(X24)|(~m1_subset_1(X25,k1_zfmisc_1(X24))|v3_membered(X25))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
fof(c_0_120, plain, ![X56, X57]:(~m1_subset_1(X56,k5_numbers)|~v7_ordinal1(X57)|m2_subset_1(k4_nat_1(X56,X57),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_nat_1])])).
fof(c_0_121, plain, ![X90, X91]:(~m1_subset_1(X90,k5_numbers)|~v7_ordinal1(X91)|k4_nat_1(X90,X91)=k3_xcmplx_0(X90,X91)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_nat_1])])).
fof(c_0_122, plain, ![X109]:(~v1_xcmplx_0(X109)|k3_xcmplx_0(np__1,X109)=X109), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
fof(c_0_123, plain, ![X32]:(~v1_xreal_0(X32)|v1_xcmplx_0(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
fof(c_0_124, plain, ![X110, X111]:(~v1_int_1(X110)|v1_abian(X110)|(~v1_int_1(X111)|v1_abian(X111)|(r1_xxreal_0(X111,X110)|r1_xxreal_0(X110,k6_xcmplx_0(X111,np__2))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_chord])])])])).
fof(c_0_125, plain, ![X82, X83]:(~v1_int_1(X82)|v1_abian(X82)|~v1_int_1(X83)|v1_abian(X83)|~v1_abian(k3_xcmplx_0(X82,X83))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc8_abian])])])).
cnf(c_0_126, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_127, plain, (~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k6_xcmplx_0(X1,np__1))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_128, plain, (k6_xcmplx_0(np__5,np__1)=np__4), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r5_r1_r4])).
cnf(c_0_129, plain, (v1_abian(k3_xcmplx_0(np__2,X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_130, plain, (k3_xcmplx_0(np__2,np__2)=np__4), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_131, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_132, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_133, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_134, plain, (m1_subset_1(np__5,k5_numbers)), inference(split_conjunct,[status(thm)],[spc5_numerals])).
cnf(c_0_135, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_136, plain, (v1_abian(X2)|v1_abian(X1)|r1_xxreal_0(X1,X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X1)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_137, plain, (k2_xcmplx_0(np__7,np__2)=np__9), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r7_r2_r9])).
cnf(c_0_138, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_139, plain, (m1_subset_1(np__7,k5_numbers)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_140, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_106, c_0_101])).
cnf(c_0_141, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_142, plain, (m1_subset_1(np__9,k1_numbers)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_143, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_144, plain, (k6_xcmplx_0(np__7,np__1)=np__6), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r7_r1_r6])).
cnf(c_0_145, plain, (k3_xcmplx_0(np__2,np__3)=np__6), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_146, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_147, plain, (k3_finseq_1(X2)=k2_nat_1(k4_nat_1(np__2,k15_glib_001(X1,X2)),np__1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_148, negated_conjecture, (k15_glib_001(esk1_0,esk2_0)=np__4), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_149, negated_conjecture, (m3_glib_001(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_150, negated_conjecture, (v1_glib_000(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_151, negated_conjecture, (v1_finset_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_152, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_153, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_154, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_155, plain, (m1_subset_1(np__4,k5_numbers)), inference(split_conjunct,[status(thm)],[spc4_numerals])).
fof(c_0_156, plain, ![X106, X107, X108]:(~v1_xxreal_0(X106)|(~v1_xxreal_0(X107)|(~v1_xxreal_0(X108)|(~r1_xxreal_0(X106,X107)|~r1_xxreal_0(X107,X108)|r1_xxreal_0(X106,X108))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_xxreal_0])])])).
cnf(c_0_157, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X4))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_158, plain, (k2_xcmplx_0(np__1,np__2)=np__3), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_159, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_160, plain, (m1_subset_1(np__2,k1_numbers)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_161, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_162, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_163, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_164, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
fof(c_0_165, plain, ![X94, X95, X96]:((~m2_subset_1(X96,X94,X95)|m1_subset_1(X96,X95)|(v1_xboole_0(X94)|v1_xboole_0(X95)|~m1_subset_1(X95,k1_zfmisc_1(X94))))&(~m1_subset_1(X96,X95)|m2_subset_1(X96,X94,X95)|(v1_xboole_0(X94)|v1_xboole_0(X95)|~m1_subset_1(X95,k1_zfmisc_1(X94))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_166, plain, ![X19, X20]:(~v1_xboole_0(X19)|(~m1_subset_1(X20,k1_zfmisc_1(X19))|v1_xboole_0(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_167, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
cnf(c_0_168, plain, (m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_169, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_170, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_171, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_172, plain, (v1_abian(X1)|v1_abian(X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,k6_xcmplx_0(X2,np__2))|~v1_int_1(X1)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_173, plain, (m1_subset_1(np__5,k1_numbers)), inference(split_conjunct,[status(thm)],[spc5_numerals])).
fof(c_0_174, plain, ![X76]:(~v1_int_1(X76)|~v1_abian(X76)|~v1_abian(k2_xcmplx_0(X76,np__1))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_abian])])])).
cnf(c_0_175, plain, (v1_abian(X1)|v1_abian(X2)|~v1_int_1(X1)|~v1_int_1(X2)|~v1_abian(k3_xcmplx_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_176, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_126])])).
cnf(c_0_177, plain, (m1_subset_1(np__7,k1_numbers)), inference(split_conjunct,[status(thm)],[spc7_numerals])).
cnf(c_0_178, plain, (~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k6_xcmplx_0(X1,np__1))), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_179, plain, (k6_xcmplx_0(np__5,np__1)=np__4), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_180, plain, (v1_abian(np__4)|~v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_181, plain, (v1_int_1(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132])])).
cnf(c_0_182, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_183, plain, (m1_subset_1(np__5,k5_numbers)), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_184, plain, (X1=k2_xcmplx_0(X2,np__2)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136])])).
cnf(c_0_185, plain, (k2_xcmplx_0(np__7,np__2)=np__9), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_186, plain, (v7_ordinal1(np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_140])])])).
cnf(c_0_187, plain, (v1_xxreal_0(np__9)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_188, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_143])).
fof(c_0_189, plain, ![X31]:(~v7_ordinal1(X31)|v1_xxreal_0(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
cnf(c_0_190, plain, (k6_xcmplx_0(np__7,np__1)=np__6), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_191, plain, (v1_abian(np__6)|~v1_int_1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_145])])).
cnf(c_0_192, plain, (m1_subset_1(np__3,k5_numbers)), inference(evalgc,[status(thm)],[c_0_146])).
fof(c_0_193, plain, ![X87, X88]:(~m1_subset_1(X87,k5_numbers)|~v7_ordinal1(X88)|k2_nat_1(X87,X88)=k2_xcmplx_0(X87,X88)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_1])])).
cnf(c_0_194, plain, (k3_finseq_1(X2)=k2_nat_1(k4_nat_1(np__2,k15_glib_001(X1,X2)),np__1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_195, negated_conjecture, (k15_glib_001(esk1_0,esk2_0)=np__4), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_196, negated_conjecture, (m3_glib_001(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_197, negated_conjecture, (v1_glib_000(esk1_0)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_198, negated_conjecture, (v1_finset_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_199, negated_conjecture, (v1_funct_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_200, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_201, negated_conjecture, (v1_relat_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_202, plain, (k3_xcmplx_0(np__2,np__4)=np__8), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r2_r4_r8])).
cnf(c_0_203, plain, (m1_subset_1(np__4,k5_numbers)), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_204, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_205, plain, ![X119, X120]:(((v7_ordinal1(esk4_2(X119,X120))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119)))&(~v1_abian(esk4_2(X119,X120))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119))))&(((v7_ordinal1(esk5_2(X119,X120))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119)))&(~v1_abian(esk5_2(X119,X120))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119))))&(((((~r1_xxreal_0(esk5_2(X119,X120),k1_nat_1(esk4_2(X119,X120),np__2))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119)))&(r1_xxreal_0(esk5_2(X119,X120),k3_finseq_1(X120))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119))))&(k1_funct_1(X120,esk4_2(X119,X120))!=k1_funct_1(X120,esk5_2(X119,X120))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119))))&(r1_glib_000(X119,k1_funct_1(X120,esk4_2(X119,X120)),k1_funct_1(X120,esk5_2(X119,X120)),esk6_2(X119,X120))|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119))))&(((esk4_2(X119,X120)!=np__1|esk5_2(X119,X120)!=k3_finseq_1(X120)|~v8_glib_001(X120,X119)|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119)))&(esk4_2(X119,X120)!=np__1|esk5_2(X119,X120)!=k6_xcmplx_0(k3_finseq_1(X120),np__2)|~v8_glib_001(X120,X119)|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119))))&(esk4_2(X119,X120)!=np__3|esk5_2(X119,X120)!=k3_finseq_1(X120)|~v8_glib_001(X120,X119)|~v5_chord(X120,X119)|~m3_glib_001(X120,X119)|(~v1_relat_1(X119)|~v4_relat_1(X119,k5_numbers)|~v1_funct_1(X119)|~v1_finset_1(X119)|~v1_glib_000(X119))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t84_chord])])])])])])).
cnf(c_0_206, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_156])).
cnf(c_0_207, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158])])).
cnf(c_0_208, plain, (v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_161])])])).
cnf(c_0_209, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_162]), c_0_161])])])).
cnf(c_0_210, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_211, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_212, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_213, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_214, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_167])).
fof(c_0_215, plain, ![X53, X54]:(~v7_ordinal1(X53)|~m1_subset_1(X54,k5_numbers)|m2_subset_1(k1_nat_1(X53,X54),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_nat_1])])).
cnf(c_0_216, plain, (m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_217, plain, (k3_xcmplx_0(X1,X2)=k4_nat_1(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_218, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_219, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_171])).
fof(c_0_220, plain, ![X30]:(~v7_ordinal1(X30)|v1_xreal_0(X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xreal_0])])).
cnf(c_0_221, plain, (v1_abian(X2)|v1_abian(X1)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,k6_xcmplx_0(X2,np__2))|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_222, plain, (k6_xcmplx_0(np__7,np__2)=np__5), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r7_r2_r5])).
cnf(c_0_223, plain, (m1_subset_1(np__5,k1_numbers)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_224, plain, (~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k2_xcmplx_0(X1,np__1))), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_225, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r1_r3])).
cnf(c_0_226, plain, (v1_abian(X2)|v1_abian(X1)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_abian(k3_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[c_0_175])).
cnf(c_0_227, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_176, c_0_132]), c_0_132])).
cnf(c_0_228, plain, (m1_subset_1(np__7,k1_numbers)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_229, plain, (~v1_abian(np__4)|~v1_abian(np__5)|~v1_int_1(np__5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179])])).
cnf(c_0_230, plain, (v1_abian(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_182])])])).
cnf(c_0_231, plain, (v7_ordinal1(np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_183]), c_0_140])])])).
cnf(c_0_232, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(np__7)|~v1_int_1(X1)|~v1_xxreal_0(np__9)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_185])])).
cnf(c_0_233, plain, (v1_int_1(np__7)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_186])])).
cnf(c_0_234, plain, (v1_xxreal_0(np__9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_161])])])).
cnf(c_0_235, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_189])).
cnf(c_0_236, plain, (~v1_abian(np__6)|~v1_abian(np__7)|~v1_int_1(np__7)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_190])])).
cnf(c_0_237, plain, (v1_abian(np__6)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_181]), c_0_192])])])).
cnf(c_0_238, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_239, negated_conjecture, (k2_nat_1(k4_nat_1(np__2,np__4),np__1)=k3_finseq_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_195]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_240, plain, (k3_xcmplx_0(np__2,np__4)=np__8), inference(evalgc,[status(thm)],[c_0_202])).
cnf(c_0_241, plain, (v7_ordinal1(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_203]), c_0_140])])])).
cnf(c_0_242, plain, (k2_xcmplx_0(np__8,np__1)=np__9), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r8_r1_r9])).
cnf(c_0_243, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_244, plain, (m1_subset_1(np__8,k5_numbers)), inference(split_conjunct,[status(thm)],[spc8_numerals])).
cnf(c_0_245, plain, (~r1_xxreal_0(esk5_2(X1,X2),k1_nat_1(esk4_2(X1,X2),np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_246, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_247, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_207, c_0_208]), c_0_209])])).
cnf(c_0_248, plain, (v1_xxreal_0(np__3)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_192])])).
cnf(c_0_249, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_161])])])).
cnf(c_0_250, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_251, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_252, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_253, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_215])).
cnf(c_0_254, plain, (m2_subset_1(k3_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217])])).
cnf(c_0_255, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219])])).
cnf(c_0_256, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_220])).
fof(c_0_257, plain, ![X22, X23]:(~v2_membered(X22)|(~m1_subset_1(X23,k1_zfmisc_1(X22))|v2_membered(X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc20_membered])])])).
cnf(c_0_258, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r0_r1_r1])).
cnf(c_0_259, plain, (k3_xcmplx_0(np__2,np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r2_r0_r0])).
cnf(c_0_260, plain, (k6_xcmplx_0(X1,np__2)=X2|r1_xxreal_0(X1,X2)|v1_abian(X2)|v1_abian(X1)|~r1_xxreal_0(k6_xcmplx_0(X1,np__2),X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(k6_xcmplx_0(X1,np__2))|~v1_int_1(X1)|~v1_int_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_221])])).
cnf(c_0_261, plain, (k6_xcmplx_0(np__7,np__2)=np__5), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_262, plain, (v1_xxreal_0(np__5)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_223])])).
cnf(c_0_263, plain, (~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k2_xcmplx_0(X1,np__1))), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_264, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_265, plain, (v1_abian(np__2)|~v1_abian(np__4)|~v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_130])])).
cnf(c_0_266, plain, (X1=k2_xcmplx_0(np__2,X2)|r1_xxreal_0(X1,X2)|v1_abian(X2)|v1_abian(X1)|~r1_xxreal_0(X1,k2_xcmplx_0(np__2,X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X2))|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_227]), c_0_182])])])).
cnf(c_0_267, plain, (k2_xcmplx_0(np__2,np__5)=np__7), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r5_r7])).
cnf(c_0_268, plain, (v1_xxreal_0(np__7)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_228])])).
cnf(c_0_269, plain, (~v1_abian(np__5)|~v1_int_1(np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_229, c_0_230])])])).
cnf(c_0_270, plain, (v1_int_1(np__5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_231])])).
cnf(c_0_271, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_232, c_0_233]), c_0_234])])).
cnf(c_0_272, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_273, plain, (~v1_abian(np__7)|~v1_int_1(np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_236, c_0_237])])])).
cnf(c_0_274, plain, (r1_xxreal_0(esk5_2(X1,X2),k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_275, plain, (k2_xcmplx_0(X1,X2)=k2_nat_1(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_276, negated_conjecture, (k2_nat_1(np__8,np__1)=k3_finseq_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_217]), c_0_240]), c_0_241]), c_0_182])])])).
cnf(c_0_277, plain, (k2_xcmplx_0(np__8,np__1)=np__9), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_278, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_243]), c_0_140])])])).
cnf(c_0_279, plain, (m1_subset_1(np__8,k5_numbers)), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_280, plain, (~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)|~r1_xxreal_0(esk5_2(X1,X2),k1_nat_1(esk4_2(X1,X2),np__2))), inference(evalgc,[status(thm)],[c_0_245])).
cnf(c_0_281, plain, (v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_282, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__2,X3)|~r1_xxreal_0(np__1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247])])).
cnf(c_0_283, plain, (v1_xxreal_0(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_188]), c_0_249])])])).
cnf(c_0_284, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_250, c_0_251])).
cnf(c_0_285, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_252, c_0_101])).
cnf(c_0_286, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_253])).
cnf(c_0_287, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255]), c_0_243])])])).
cnf(c_0_288, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_256])).
cnf(c_0_289, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_257])).
cnf(c_0_290, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(evalgc,[status(thm)],[c_0_258])).
cnf(c_0_291, plain, (k3_xcmplx_0(np__2,np__0)=np__0), inference(evalgc,[status(thm)],[c_0_259])).
cnf(c_0_292, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(np__7)|~v1_int_1(X1)|~v1_xxreal_0(np__5)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261])])).
cnf(c_0_293, plain, (v1_xxreal_0(np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_188]), c_0_161])])])).
cnf(c_0_294, plain, (k6_xcmplx_0(np__9,np__2)=np__7), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r9_r2_r7])).
cnf(c_0_295, plain, (k6_xcmplx_0(np__9,np__1)=np__8), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r9_r1_r8])).
cnf(c_0_296, plain, (m1_subset_1(np__9,k5_numbers)), inference(split_conjunct,[status(thm)],[spc9_numerals])).
cnf(c_0_297, plain, (~v1_abian(np__3)|~v1_abian(np__2)|~v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264])])).
cnf(c_0_298, plain, (v1_abian(np__2)|~v1_int_1(np__2)), inference(csr,[status(thm)],[c_0_265, c_0_180])).
cnf(c_0_299, plain, (X1=k2_xcmplx_0(np__2,X2)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(np__2,X2))|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X2))|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers)), inference(csr,[status(thm)],[c_0_266, c_0_181])).
cnf(c_0_300, plain, (k2_xcmplx_0(np__2,np__5)=np__7), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_301, plain, (v1_xxreal_0(np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_188]), c_0_161])])])).
cnf(c_0_302, plain, (~v1_abian(np__5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_269, c_0_270])])).
cnf(c_0_303, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272])])).
cnf(c_0_304, plain, (~v1_abian(np__7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_273, c_0_233])])).
cnf(c_0_305, plain, (r1_xxreal_0(esk5_2(X1,X2),k3_finseq_1(X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_274])).
cnf(c_0_306, negated_conjecture, (k3_finseq_1(esk2_0)=np__9), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_276]), c_0_277]), c_0_278]), c_0_279])])])).
cnf(c_0_307, plain, (~r1_xxreal_0(esk5_2(X1,X2),k2_xcmplx_0(esk4_2(X1,X2),np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_94]), c_0_182])])])).
cnf(c_0_308, plain, (v7_ordinal1(esk4_2(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_309, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__2,X3)|~r1_xxreal_0(np__1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_282, c_0_283])])).
cnf(c_0_310, plain, (r1_xxreal_0(np__2,np__2)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r2])).
cnf(c_0_311, plain, (m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_211]), c_0_285])])).
cnf(c_0_312, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_94])])).
cnf(c_0_313, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_287, c_0_288])).
cnf(c_0_314, plain, (v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_315, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_289])).
fof(c_0_316, plain, ![X99]:(~v7_ordinal1(X99)|v1_abian(X99)|r1_xxreal_0(np__1,X99)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_abian])])])).
cnf(c_0_317, plain, (~v1_abian(np__1)|~v1_abian(np__0)|~v1_int_1(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_290])])).
cnf(c_0_318, plain, (v1_abian(np__0)|~v1_int_1(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_291])])).
cnf(c_0_319, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_320, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(np__7)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_292, c_0_293])])).
cnf(c_0_321, plain, (k6_xcmplx_0(np__9,np__2)=np__7), inference(evalgc,[status(thm)],[c_0_294])).
cnf(c_0_322, plain, (k6_xcmplx_0(np__9,np__1)=np__8), inference(evalgc,[status(thm)],[c_0_295])).
cnf(c_0_323, plain, (v1_abian(np__8)|~v1_int_1(np__4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_240])])).
cnf(c_0_324, plain, (m1_subset_1(np__9,k5_numbers)), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_325, plain, (k2_xcmplx_0(np__3,np__2)=np__5), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r3_r2_r5])).
cnf(c_0_326, plain, (~v1_abian(np__3)|~v1_abian(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_181]), c_0_182])])])).
cnf(c_0_327, plain, (v1_abian(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_181]), c_0_182])])])).
cnf(c_0_328, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_300]), c_0_301]), c_0_183])]), c_0_302])])).
cnf(c_0_329, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_303, c_0_304]), c_0_131])).
cnf(c_0_330, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306])])).
cnf(c_0_331, plain, (~v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_332, plain, (~r1_xxreal_0(esk5_2(X1,X2),k2_xcmplx_0(esk4_2(X1,X2),np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_307, c_0_308])).
cnf(c_0_333, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__2,X3)|~r1_xxreal_0(np__1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_272])])).
cnf(c_0_334, plain, (r1_xxreal_0(np__2,np__2)), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_335, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_312])])).
cnf(c_0_336, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_313])])).
cnf(c_0_337, plain, (v7_ordinal1(esk5_2(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_314])).
cnf(c_0_338, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_211])])).
cnf(c_0_339, plain, (v1_abian(X1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_316])).
cnf(c_0_340, plain, (~v1_abian(np__1)|~v1_int_1(np__0)), inference(csr,[status(thm)],[c_0_317, c_0_318])).
cnf(c_0_341, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_319])).
cnf(c_0_342, plain, (k2_xcmplx_0(np__5,np__2)=np__7), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r5_r2_r7])).
cnf(c_0_343, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|v1_abian(X2)|v1_abian(X3)|~r1_xxreal_0(X1,k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(X2)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(X1)|~v1_int_1(X3)|~v1_int_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_136])])).
cnf(c_0_344, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_320, c_0_233])])])).
cnf(c_0_345, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|v1_abian(np__9)|~v1_int_1(np__9)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_321])])).
cnf(c_0_346, plain, (~v1_abian(np__8)|~v1_abian(np__9)|~v1_int_1(np__9)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_322])])).
cnf(c_0_347, plain, (v1_abian(np__8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_181]), c_0_203])])])).
cnf(c_0_348, plain, (v7_ordinal1(np__9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_324]), c_0_140])])])).
cnf(c_0_349, plain, (k6_xcmplx_0(np__5,np__2)=np__3), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r5_r2_r3])).
cnf(c_0_350, plain, (X1=k2_xcmplx_0(X2,np__2)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__2))|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_227]), c_0_182])])])).
cnf(c_0_351, plain, (k2_xcmplx_0(np__3,np__2)=np__5), inference(evalgc,[status(thm)],[c_0_325])).
cnf(c_0_352, plain, (~v1_abian(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_326, c_0_327])])])).
cnf(c_0_353, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_272])])).
cnf(c_0_354, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_330])])).
cnf(c_0_355, plain, (~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)|~v1_abian(esk5_2(X1,X2))), inference(evalgc,[status(thm)],[c_0_331])).
cnf(c_0_356, plain, (~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__3)|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,X2),np__2))|~v1_xreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_333]), c_0_334]), c_0_208])])])).
cnf(c_0_357, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_335]), c_0_140])])])).
cnf(c_0_358, plain, (m1_subset_1(esk5_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_337])])).
cnf(c_0_359, plain, (v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_188]), c_0_161])])])).
cnf(c_0_360, plain, (v1_abian(X1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_339])).
cnf(c_0_361, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_abian(X1)|v1_abian(np__1)|~v1_int_1(np__1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_158])])).
cnf(c_0_362, plain, (v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_278])])).
cnf(c_0_363, plain, (~v1_abian(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_181]), c_0_341])])])).
cnf(c_0_364, plain, (k2_xcmplx_0(np__5,np__2)=np__7), inference(evalgc,[status(thm)],[c_0_342])).
cnf(c_0_365, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|v1_abian(X2)|v1_abian(X1)|~r1_xxreal_0(X3,k2_xcmplx_0(np__2,X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_227]), c_0_182])])])).
cnf(c_0_366, plain, (k2_xcmplx_0(np__2,np__7)=np__9), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r7_r9])).
cnf(c_0_367, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(sr,[status(thm)],[c_0_344, c_0_304])).
cnf(c_0_368, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__7)|v1_abian(np__9)|v1_abian(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_181]), c_0_324])])])).
cnf(c_0_369, plain, (~v1_abian(np__9)|~v1_int_1(np__9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_346, c_0_347])])])).
cnf(c_0_370, plain, (v1_int_1(np__9)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_348])])).
cnf(c_0_371, plain, (k6_xcmplx_0(np__5,np__2)=np__3), inference(evalgc,[status(thm)],[c_0_349])).
cnf(c_0_372, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_351]), c_0_293]), c_0_192])]), c_0_352])])).
cnf(c_0_373, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_353, c_0_131])).
cnf(c_0_374, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_354, c_0_337]), c_0_355])).
cnf(c_0_375, plain, (~r1_xxreal_0(esk5_2(X1,X2),np__3)|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v1_xreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357]), c_0_182])])])).
cnf(c_0_376, plain, (v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_358]), c_0_359])])])).
cnf(c_0_377, plain, (X1=np__1|v1_abian(X1)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_360])])).
cnf(c_0_378, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_243])])).
cnf(c_0_379, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_361, c_0_362])]), c_0_363])])).
cnf(c_0_380, plain, (v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_308])])).
cnf(c_0_381, plain, (~v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_382, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|v1_abian(np__5)|~v1_int_1(np__5)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_364])])).
cnf(c_0_383, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|v1_abian(X2)|v1_abian(X3)|~r1_xxreal_0(X1,k2_xcmplx_0(np__2,X3))|~v1_int_1(X2)|~v1_xxreal_0(k2_xcmplx_0(np__2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X3,k5_numbers)), inference(csr,[status(thm)],[c_0_365, c_0_181])).
cnf(c_0_384, plain, (k2_xcmplx_0(np__2,np__7)=np__9), inference(evalgc,[status(thm)],[c_0_366])).
cnf(c_0_385, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_272])])).
cnf(c_0_386, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(np__9)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_181])])).
cnf(c_0_387, plain, (~v1_abian(np__9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_369, c_0_370])])).
cnf(c_0_388, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|v1_abian(np__5)|~v1_int_1(np__5)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_371])])).
fof(c_0_389, plain, ![X45, X46]:((~v1_glib_001(X46,X45)|k3_glib_001(X45,X46)=k4_glib_001(X45,X46)|~m3_glib_001(X46,X45)|(~v1_relat_1(X45)|~v4_relat_1(X45,k5_numbers)|~v1_funct_1(X45)|~v1_finset_1(X45)|~v1_glib_000(X45)))&(k3_glib_001(X45,X46)!=k4_glib_001(X45,X46)|v1_glib_001(X46,X45)|~m3_glib_001(X46,X45)|(~v1_relat_1(X45)|~v4_relat_1(X45,k5_numbers)|~v1_funct_1(X45)|~v1_finset_1(X45)|~v1_glib_000(X45)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d24_glib_001])])])])).
fof(c_0_390, plain, ![X36, X37]:(((v1_glib_001(X37,X36)|~v8_glib_001(X37,X36)|~m3_glib_001(X37,X36)|(~v1_relat_1(X36)|~v4_relat_1(X36,k5_numbers)|~v1_funct_1(X36)|~v1_finset_1(X36)|~v1_glib_000(X36)))&(~v3_glib_001(X37,X36)|~v8_glib_001(X37,X36)|~m3_glib_001(X37,X36)|(~v1_relat_1(X36)|~v4_relat_1(X36,k5_numbers)|~v1_funct_1(X36)|~v1_finset_1(X36)|~v1_glib_000(X36))))&(v5_glib_001(X37,X36)|~v8_glib_001(X37,X36)|~m3_glib_001(X37,X36)|(~v1_relat_1(X36)|~v4_relat_1(X36,k5_numbers)|~v1_funct_1(X36)|~v1_finset_1(X36)|~v1_glib_000(X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc6_glib_001])])])])])).
cnf(c_0_391, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_272])])).
cnf(c_0_392, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_374])])).
cnf(c_0_393, plain, (~r1_xxreal_0(esk5_2(X1,X2),np__3)|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_375, c_0_376]), c_0_288]), c_0_308])).
cnf(c_0_394, plain, (X1=np__1|v1_abian(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)|~v1_xxreal_0(np__1)), inference(csr,[status(thm)],[c_0_377, c_0_272])).
cnf(c_0_395, plain, (v1_xxreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_188]), c_0_249])])])).
cnf(c_0_396, plain, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_380])])).
cnf(c_0_397, plain, (~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)|~v1_abian(esk4_2(X1,X2))), inference(evalgc,[status(thm)],[c_0_381])).
cnf(c_0_398, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,X1)|v1_abian(np__5)|v1_abian(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_181]), c_0_183])])])).
cnf(c_0_399, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_384]), c_0_234]), c_0_139])]), c_0_304])])).
cnf(c_0_400, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_385, c_0_131])).
cnf(c_0_401, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)), inference(sr,[status(thm)],[c_0_386, c_0_387])).
cnf(c_0_402, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__3)|v1_abian(np__5)|v1_abian(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_181]), c_0_183])])])).
fof(c_0_403, plain, ![X50, X51]:(~v1_relat_1(X50)|~v4_relat_1(X50,k5_numbers)|~v1_funct_1(X50)|~v1_finset_1(X50)|~v1_glib_000(X50)|(~m3_glib_001(X51,X50)|k4_glib_001(X50,X51)=k1_funct_1(X51,k3_finseq_1(X51)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_glib_001])])])).
cnf(c_0_404, plain, (k3_glib_001(X2,X1)=k4_glib_001(X2,X1)|~v1_glib_001(X1,X2)|~m3_glib_001(X1,X2)|~v1_relat_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)), inference(split_conjunct,[status(thm)],[c_0_389])).
cnf(c_0_405, plain, (v1_glib_001(X1,X2)|~v8_glib_001(X1,X2)|~m3_glib_001(X1,X2)|~v1_relat_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_406, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_391, c_0_131])).
cnf(c_0_407, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_392, c_0_337]), c_0_355])).
cnf(c_0_408, plain, (v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_360])])).
cnf(c_0_409, plain, (X1=np__1|v1_abian(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_394, c_0_395])])])).
cnf(c_0_410, plain, (r1_xxreal_0(esk4_2(X1,X2),np__1)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_396, c_0_397])).
cnf(c_0_411, plain, (m1_subset_1(esk4_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_308])])).
cnf(c_0_412, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__5)|v1_abian(np__5)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_181])])).
cnf(c_0_413, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_272])])).
cnf(c_0_414, plain, (X1=np__7|np__5=X1|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_400]), c_0_301])])])).
cnf(c_0_415, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),np__7)|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_358])])).
cnf(c_0_416, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,X1)|v1_abian(np__5)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_181])])).
fof(c_0_417, plain, ![X47, X48]:(~v1_relat_1(X47)|~v4_relat_1(X47,k5_numbers)|~v1_funct_1(X47)|~v1_finset_1(X47)|~v1_glib_000(X47)|(~m3_glib_001(X48,X47)|k3_glib_001(X47,X48)=k1_funct_1(X48,np__1))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_glib_001])])])).
cnf(c_0_418, plain, (k4_glib_001(X1,X2)=k1_funct_1(X2,k3_finseq_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_403])).
cnf(c_0_419, plain, (k4_glib_001(X2,X1)=k3_glib_001(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v1_glib_001(X1,X2)|~v4_relat_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_404])).
cnf(c_0_420, plain, (v1_glib_001(X1,X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v8_glib_001(X1,X2)|~v4_relat_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_405])).
cnf(c_0_421, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(esk5_2(X1,esk2_0),np__3)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407])])).
cnf(c_0_422, plain, (~r1_xxreal_0(esk5_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_408, c_0_308]), c_0_397])).
cnf(c_0_423, plain, (v1_xreal_0(np__7)|~v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_139])])).
cnf(c_0_424, plain, (esk4_2(X1,X2)=np__1|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_410])])).
cnf(c_0_425, plain, (v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_411]), c_0_359])])])).
cnf(c_0_426, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)), inference(sr,[status(thm)],[c_0_412, c_0_302])).
cnf(c_0_427, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)), inference(csr,[status(thm)],[c_0_413, c_0_131])).
cnf(c_0_428, plain, (np__5=X1|X1=np__7|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_414, c_0_272])).
cnf(c_0_429, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_415, c_0_355])).
cnf(c_0_430, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)), inference(sr,[status(thm)],[c_0_416, c_0_302])).
fof(c_0_431, plain, ![X34]:(~v1_xreal_0(X34)|v1_xxreal_0(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
cnf(c_0_432, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(X3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(X3,np__9)|~v1_int_1(esk5_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_376])])).
cnf(c_0_433, plain, (v1_int_1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_337])])).
cnf(c_0_434, plain, (k3_glib_001(X1,X2)=k1_funct_1(X2,np__1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_417])).
cnf(c_0_435, plain, (k4_glib_001(X1,X2)=k1_funct_1(X2,k3_finseq_1(X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_418])).
cnf(c_0_436, plain, (k4_glib_001(X1,X2)=k3_glib_001(X1,X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_420])])).
cnf(c_0_437, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_421, c_0_337]), c_0_355]), c_0_422])).
cnf(c_0_438, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__7,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__7)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_185])])).
cnf(c_0_439, plain, (v1_xreal_0(np__7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_423, c_0_249])])).
cnf(c_0_440, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_424, c_0_308]), c_0_397])).
cnf(c_0_441, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_425])])).
cnf(c_0_442, plain, (r1_xxreal_0(esk4_2(X1,X2),np__5)|r1_xxreal_0(np__7,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_411])])).
cnf(c_0_443, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(X2,X3),X1)|v1_abian(X1)|~r1_xxreal_0(esk5_2(X2,X3),np__9)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_376])])).
cnf(c_0_444, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__9,esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_429])])).
cnf(c_0_445, plain, (r1_xxreal_0(np__5,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),np__3)|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_358])])).
cnf(c_0_446, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(X1,esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_305])])).
cnf(c_0_447, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_431])).
fof(c_0_448, plain, ![X21]:(~m1_subset_1(X21,k1_numbers)|v1_xreal_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_449, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(X3,esk5_2(X1,X2))|~r1_xxreal_0(X3,np__9)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_432, c_0_433]), c_0_355])).
cnf(c_0_450, negated_conjecture, (v5_chord(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_451, plain, (k3_glib_001(X1,X2)=k1_funct_1(X2,np__1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_434])).
cnf(c_0_452, plain, (k3_glib_001(X1,X2)=k1_funct_1(X2,k3_finseq_1(X2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_436])])).
cnf(c_0_453, negated_conjecture, (v8_glib_001(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_454, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_437])])).
cnf(c_0_455, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__7,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_438, c_0_208]), c_0_439])])).
cnf(c_0_456, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_440]), c_0_283])])])).
cnf(c_0_457, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_441, c_0_380]), c_0_397])).
cnf(c_0_458, plain, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_442, c_0_397])).
cnf(c_0_459, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_330])])).
cnf(c_0_460, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_444, c_0_337]), c_0_355])).
cnf(c_0_461, plain, (r1_xxreal_0(esk5_2(X1,X2),np__3)|r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_445, c_0_355])).
cnf(c_0_462, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(esk5_2(X3,X2))|v1_abian(X1)|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(X1)|~v1_int_1(esk5_2(X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_446, c_0_136])])).
cnf(c_0_463, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_447])).
cnf(c_0_464, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_448])).
cnf(c_0_465, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,X1),np__7)|r1_xxreal_0(X2,esk5_2(esk1_0,X1))|~r1_xxreal_0(X2,np__9)|~v1_xxreal_0(X2)|~v5_chord(X1,esk1_0)|~m3_glib_001(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_200]), c_0_197]), c_0_198]), c_0_199]), c_0_201])])])).
cnf(c_0_466, negated_conjecture, (v5_chord(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_450])).
cnf(c_0_467, plain, (r1_glib_000(X1,k1_funct_1(X2,esk4_2(X1,X2)),k1_funct_1(X2,esk5_2(X1,X2)),esk6_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_468, plain, (k1_funct_1(X2,k3_finseq_1(X2))=k1_funct_1(X2,np__1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_452])])).
cnf(c_0_469, negated_conjecture, (v8_glib_001(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_453])).
cnf(c_0_470, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_455]), c_0_334]), c_0_208])])])).
cnf(c_0_471, plain, (v1_xreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_411]), c_0_249])])])).
cnf(c_0_472, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_456, c_0_425])).
cnf(c_0_473, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_458])])).
cnf(c_0_474, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_459])])).
cnf(c_0_475, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__7|esk5_2(X2,X3)=np__5|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,esk5_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_460]), c_0_234])])])).
cnf(c_0_476, plain, (r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_461])])).
cnf(c_0_477, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_462, c_0_433]), c_0_355])).
cnf(c_0_478, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_463])])).
cnf(c_0_479, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_464])).
cnf(c_0_480, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_465, c_0_466]), c_0_196])])])).
fof(c_0_481, plain, ![X100, X101, X102, X103]:(~v1_relat_1(X100)|~v4_relat_1(X100,k5_numbers)|~v1_funct_1(X100)|~v1_finset_1(X100)|~v1_glib_000(X100)|(~r1_glib_000(X100,X102,X103,X101)|r1_glib_000(X100,X103,X102,X101))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_glib_000])])])).
cnf(c_0_482, plain, (r1_glib_000(X1,k1_funct_1(X2,esk4_2(X1,X2)),k1_funct_1(X2,esk5_2(X1,X2)),esk6_2(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_467])).
cnf(c_0_483, negated_conjecture, (k1_funct_1(esk2_0,np__9)=k1_funct_1(esk2_0,np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_469]), c_0_306]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_484, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__7,esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_470, c_0_471])).
cnf(c_0_485, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_472, c_0_473])])).
cnf(c_0_486, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_474, c_0_272]), c_0_376])).
cnf(c_0_487, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(X3,esk5_2(X1,X2))|~r1_xxreal_0(X3,np__9)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_475, c_0_376]), c_0_476])).
cnf(c_0_488, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_272])])).
cnf(c_0_489, plain, (v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_192]), c_0_140])])])).
cnf(c_0_490, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_478, c_0_479])])).
cnf(c_0_491, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|v1_abian(X2)|v1_abian(X3)|~r1_xxreal_0(X1,k2_xcmplx_0(X3,np__2))|~v1_int_1(X2)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X3,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_227]), c_0_182])])])).
cnf(c_0_492, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk5_2(X2,esk2_0)=np__9|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_374]), c_0_301])])])).
cnf(c_0_493, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(esk5_2(X1,esk2_0),esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_330])])).
cnf(c_0_494, plain, (esk4_2(X1,X2)!=np__3|esk5_2(X1,X2)!=k3_finseq_1(X2)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_495, plain, (r1_glib_000(X1,X3,X2,X4)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~r1_glib_000(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_481])).
cnf(c_0_496, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_437]), c_0_483])])).
cnf(c_0_497, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__3|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_484, c_0_485])])).
cnf(c_0_498, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__5|esk5_2(X2,esk2_0)=np__7|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_486, c_0_487])])).
cnf(c_0_499, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)), inference(csr,[status(thm)],[c_0_488, c_0_337])).
cnf(c_0_500, plain, (v1_int_1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_489])])).
cnf(c_0_501, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_490, c_0_272])])).
cnf(c_0_502, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_364]), c_0_301]), c_0_183])]), c_0_302])])).
cnf(c_0_503, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(X2,np__7)|~r1_xxreal_0(X2,esk5_2(X1,esk2_0))|~v1_xxreal_0(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_492, c_0_376])).
cnf(c_0_504, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_493, c_0_376])).
cnf(c_0_505, plain, (esk4_2(X1,X2)!=np__3|esk5_2(X1,X2)!=k3_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_494])).
cnf(c_0_506, negated_conjecture, (~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__5),X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_507, plain, (r1_glib_000(X1,X3,X2,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~v4_relat_1(X1,k5_numbers)|~r1_glib_000(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_495])).
cnf(c_0_508, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|r1_glib_000(X1,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_497])])).
cnf(c_0_509, plain, (esk4_2(X1,X2)!=np__1|esk5_2(X1,X2)!=k3_finseq_1(X2)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_510, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v1_int_1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_376])])).
cnf(c_0_511, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|X2=esk5_2(X1,esk2_0)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_498, c_0_272])).
cnf(c_0_512, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|r1_xxreal_0(esk5_2(X2,X1),np__3)|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_351]), c_0_500]), c_0_293])]), c_0_352])])).
cnf(c_0_513, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(X2,esk2_0),X1)|v1_abian(X1)|~v7_ordinal1(esk5_2(X2,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_501, c_0_330])])).
cnf(c_0_514, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_272])])).
cnf(c_0_515, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_503, c_0_504]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_516, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=k3_finseq_1(esk2_0)|esk4_2(esk1_0,esk2_0)!=np__3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_505, c_0_469]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_517, negated_conjecture, (~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__5),X1)), inference(evalgc,[status(thm)],[c_0_506])).
cnf(c_0_518, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__5),esk6_2(X1,esk2_0))|esk4_2(X1,esk2_0)=np__3|esk4_2(X1,esk2_0)=np__1|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_508])])).
cnf(c_0_519, plain, (esk4_2(X1,X2)!=np__1|esk5_2(X1,X2)!=k3_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_509])).
cnf(c_0_520, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_510, c_0_433]), c_0_355])).
cnf(c_0_521, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk5_2(X2,esk2_0)=esk5_2(X1,esk2_0)|r1_xxreal_0(esk5_2(X2,esk2_0),np__7)|v1_abian(esk5_2(X2,esk2_0))|~v7_ordinal1(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_511, c_0_330])])).
cnf(c_0_522, plain, (r1_xxreal_0(esk5_2(X1,X2),np__3)|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_512, c_0_463])])).
cnf(c_0_523, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_513, c_0_337])).
cnf(c_0_524, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)), inference(csr,[status(thm)],[c_0_514, c_0_131])).
cnf(c_0_525, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_515, c_0_376])).
cnf(c_0_526, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__9|esk4_2(esk1_0,esk2_0)!=np__3), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_516, c_0_306])])).
cnf(c_0_527, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_518]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_528, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=k3_finseq_1(esk2_0)|esk4_2(esk1_0,esk2_0)!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_519, c_0_469]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_529, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_520]), c_0_301])])])).
cnf(c_0_530, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(X2,esk2_0)|esk5_2(X2,esk2_0)=np__5|esk5_2(X2,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~v5_chord(esk2_0,X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X2)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X2)|~v1_glib_000(X1)|~v1_finset_1(X2)|~v1_finset_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_521, c_0_337]), c_0_355])).
cnf(c_0_531, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_335]), c_0_359])])])).
cnf(c_0_532, plain, (r1_xxreal_0(esk5_2(X1,X2),np__3)|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_479])])).
cnf(c_0_533, negated_conjecture, (r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__7)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_523])])).
cnf(c_0_534, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_286])])).
cnf(c_0_535, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_524, c_0_272])])).
cnf(c_0_536, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_525, c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_537, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)!=np__9), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_526, c_0_527])])).
cnf(c_0_538, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__9|esk4_2(esk1_0,esk2_0)!=np__1), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_528, c_0_306])])).
cnf(c_0_539, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_529, c_0_376]), c_0_476])).
cnf(c_0_540, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_530, c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_541, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_531]), c_0_182])])])).
cnf(c_0_542, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|~m1_subset_1(k3_finseq_1(X1),k1_numbers)|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_532, c_0_422])).
cnf(c_0_543, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_533]), c_0_301])])])).
cnf(c_0_544, plain, (v1_xxreal_0(k1_nat_1(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_534]), c_0_359])])])).
cnf(c_0_545, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__9|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_535, c_0_536])])).
cnf(c_0_546, plain, (m1_subset_1(np__3,k1_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_547, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)!=np__9), inference(csr,[status(thm)],[c_0_537, c_0_538])).
cnf(c_0_548, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_539, c_0_540])])).
cnf(c_0_549, plain, (v1_xreal_0(np__5)|~v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_183])])).
cnf(c_0_550, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)), inference(csr,[status(thm)],[c_0_541, c_0_131])).
cnf(c_0_551, negated_conjecture, (r1_xxreal_0(np__5,np__9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_542, c_0_466]), c_0_306]), c_0_306]), c_0_142]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_552, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_544]), c_0_182])])])).
cnf(c_0_553, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__5)|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_545]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_554, plain, (m1_subset_1(np__3,k1_numbers)), inference(evalgc,[status(thm)],[c_0_546])).
cnf(c_0_555, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(X1,esk2_0)!=np__9|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_547, c_0_548])])).
cnf(c_0_556, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_364])])).
cnf(c_0_557, plain, (v1_xreal_0(np__5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_549, c_0_249])])).
cnf(c_0_558, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__9)|r1_xxreal_0(esk5_2(X2,esk2_0),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_550, c_0_306]), c_0_234])])])).
cnf(c_0_559, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_551]), c_0_234]), c_0_293])])])).
cnf(c_0_560, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_552, c_0_308])).
cnf(c_0_561, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_553]), c_0_293])])])).
cnf(c_0_562, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_351])])).
cnf(c_0_563, plain, (v1_xreal_0(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_554]), c_0_161])])])).
cnf(c_0_564, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_555, c_0_437])).
cnf(c_0_565, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_556, c_0_208]), c_0_557])])).
cnf(c_0_566, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(esk5_2(X1,esk2_0),np__3)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__7,np__9)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_558]), c_0_364]), c_0_231])]), c_0_302])])).
cnf(c_0_567, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__9)|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_376])])).
cnf(c_0_568, plain, (r1_xxreal_0(k2_xcmplx_0(np__2,X1),X2)|r1_xxreal_0(X2,X1)|v1_abian(X2)|v1_abian(X1)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_227]), c_0_182])])])).
fof(c_0_569, plain, ![X74, X75]:(~v1_int_1(X74)|~v1_int_1(X75)|v1_int_1(k2_xcmplx_0(X74,X75))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_int_1])])).
cnf(c_0_570, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_560, c_0_94]), c_0_182])])])).
cnf(c_0_571, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_561, c_0_272])).
cnf(c_0_572, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_562, c_0_208]), c_0_563])])).
cnf(c_0_573, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_564, c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_574, plain, (k2_xcmplx_0(X1,np__2)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_565]), c_0_233]), c_0_301]), c_0_334]), c_0_208])]), c_0_304])])).
cnf(c_0_575, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_566, c_0_337]), c_0_355]), c_0_422])).
cnf(c_0_576, negated_conjecture, (esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_567]), c_0_234])])])).
cnf(c_0_577, plain, (r1_xxreal_0(np__7,k3_finseq_1(X1))|r1_xxreal_0(esk5_2(X2,X1),np__5)|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_364]), c_0_270]), c_0_301])]), c_0_302])])).
cnf(c_0_578, plain, (r1_xxreal_0(k2_xcmplx_0(np__2,X1),X2)|r1_xxreal_0(X2,X1)|v1_abian(X1)|v1_abian(X2)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[c_0_568, c_0_181])).
cnf(c_0_579, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X1)|~v1_int_1(X2)), inference(split_conjunct,[status(thm)],[c_0_569])).
cnf(c_0_580, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_570, c_0_308])).
fof(c_0_581, plain, ![X64, X65]:(~v1_int_1(X64)|~v1_abian(X64)|~v1_int_1(X65)|v1_abian(X65)|~v1_abian(k2_xcmplx_0(X65,X64))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc11_abian])])])).
cnf(c_0_582, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_571, c_0_572]), c_0_334]), c_0_208])])])).
cnf(c_0_583, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__3)|v1_abian(X1)|v1_abian(np__3)|~v1_int_1(np__3)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_351])])).
cnf(c_0_584, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_573]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_585, plain, (k2_xcmplx_0(X1,np__2)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_574, c_0_531]), c_0_182])])])).
cnf(c_0_586, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_575])])).
cnf(c_0_587, negated_conjecture, (esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_576, c_0_376])).
cnf(c_0_588, plain, (r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_577, c_0_272])])).
cnf(c_0_589, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_476]), c_0_293])])])).
cnf(c_0_590, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_578, c_0_300]), c_0_183])]), c_0_302])])).
cnf(c_0_591, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_579])).
cnf(c_0_592, plain, (v7_ordinal1(np__8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_279]), c_0_140])])])).
cnf(c_0_593, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_580, c_0_565]), c_0_334]), c_0_208])])])).
cnf(c_0_594, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_182]), c_0_140])])])).
cnf(c_0_595, plain, (v1_abian(X2)|~v1_int_1(X1)|~v1_abian(X1)|~v1_int_1(X2)|~v1_abian(k2_xcmplx_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_581])).
cnf(c_0_596, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_582, c_0_357]), c_0_182])])])).
cnf(c_0_597, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,X1)|v1_abian(np__3)|v1_abian(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_583, c_0_181]), c_0_192])])])).
cnf(c_0_598, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|np__7!=np__5), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_573])])).
cnf(c_0_599, negated_conjecture, (~r1_xxreal_0(np__5,np__3)|~r1_xxreal_0(np__7,np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_584]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_600, plain, (k2_xcmplx_0(X1,np__2)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_585, c_0_336]), c_0_288])).
cnf(c_0_601, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_572]), c_0_334]), c_0_208])])])).
cnf(c_0_602, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_587, c_0_588])])).
cnf(c_0_603, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_589, c_0_376])).
cnf(c_0_604, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|v1_abian(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_590, c_0_591])])).
cnf(c_0_605, plain, (v1_int_1(np__8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_592])])).
fof(c_0_606, plain, ![X70]:(~v1_int_1(X70)|v1_abian(X70)|~v1_abian(k2_xcmplx_0(X70,np__2))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc17_abian])])])).
cnf(c_0_607, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_593, c_0_471])).
cnf(c_0_608, plain, (v1_int_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_594])])).
cnf(c_0_609, plain, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__3)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_411])])).
cnf(c_0_610, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_272])])).
cnf(c_0_611, plain, (v1_abian(X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k2_xcmplx_0(X2,X1))), inference(evalgc,[status(thm)],[c_0_595])).
cnf(c_0_612, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_596, c_0_288])).
cnf(c_0_613, plain, (r1_xxreal_0(np__5,np__3)|r1_xxreal_0(np__3,np__3)|v1_abian(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_597, c_0_500])])).
cnf(c_0_614, negated_conjecture, (~r1_xxreal_0(np__5,np__3)|np__7!=np__5), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_598]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_615, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__5,np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_599, c_0_600]), c_0_351]), c_0_489])]), c_0_352])])).
cnf(c_0_616, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_601, c_0_471])).
cnf(c_0_617, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_602, c_0_603]), c_0_234])])])).
cnf(c_0_618, plain, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_604, c_0_277]), c_0_362]), c_0_605])]), c_0_387])])).
cnf(c_0_619, plain, (v1_abian(X1)|~v1_int_1(X1)|~v1_abian(k2_xcmplx_0(X1,np__2))), inference(split_conjunct,[status(thm)],[c_0_606])).
cnf(c_0_620, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_607, c_0_591]), c_0_608])])])).
cnf(c_0_621, plain, (r1_xxreal_0(esk4_2(X1,X2),np__3)|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_609, c_0_397])).
cnf(c_0_622, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_610, c_0_551]), c_0_231])])])).
cnf(c_0_623, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_611, c_0_612]), c_0_327]), c_0_608])])])).
cnf(c_0_624, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_440]), c_0_283])])])).
cnf(c_0_625, plain, (r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__5,np__3)), inference(sr,[status(thm)],[c_0_613, c_0_352])).
cnf(c_0_626, negated_conjecture, (~r1_xxreal_0(np__5,np__3)), inference(csr,[status(thm)],[c_0_614, c_0_615])).
cnf(c_0_627, negated_conjecture, (r1_xxreal_0(np__7,np__9)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_616, c_0_440])])).
cnf(c_0_628, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|r1_xxreal_0(np__7,np__9)|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_617, c_0_618])])).
cnf(c_0_629, plain, (v1_abian(X1)|~v1_int_1(X1)|~v1_abian(k2_xcmplx_0(X1,np__2))), inference(evalgc,[status(thm)],[c_0_619])).
cnf(c_0_630, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_620, c_0_380])).
cnf(c_0_631, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_472, c_0_621])])).
cnf(c_0_632, negated_conjecture, (r1_xxreal_0(np__5,np__5)|r1_xxreal_0(np__5,np__7)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_622, c_0_231]), c_0_302])])).
cnf(c_0_633, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_623, c_0_131])).
cnf(c_0_634, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_624, c_0_425])).
cnf(c_0_635, plain, (r1_xxreal_0(np__3,np__3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_625, c_0_626])])).
cnf(c_0_636, negated_conjecture, (r1_xxreal_0(np__7,np__9)|esk5_2(esk1_0,esk2_0)!=np__9), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_538, c_0_627]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_637, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_628, c_0_306]), c_0_348])])])).
cnf(c_0_638, negated_conjecture, (v1_abian(esk4_2(X1,esk2_0))|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|~v1_int_1(esk4_2(X1,esk2_0))|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_629, c_0_630])])).
cnf(c_0_639, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_631]), c_0_293])])])).
cnf(c_0_640, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__5)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_535, c_0_632]), c_0_231])])])).
cnf(c_0_641, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_633, c_0_634]), c_0_635]), c_0_283]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_642, negated_conjecture, (r1_xxreal_0(np__7,np__9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_636, c_0_637]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_643, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_638, c_0_380]), c_0_397])).
cnf(c_0_644, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|r1_xxreal_0(X3,esk4_2(X1,X2))|~r1_xxreal_0(X3,np__5)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_639, c_0_425])).
cnf(c_0_645, negated_conjecture, (r1_xxreal_0(np__5,np__5)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_640, c_0_231]), c_0_302])])).
cnf(c_0_646, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_641, c_0_308]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_647, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_642]), c_0_234]), c_0_301])])])).
cnf(c_0_648, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(np__5)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(np__5)|~v1_int_1(X1)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_371])])).
cnf(c_0_649, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_330]), c_0_234])])])).
cnf(c_0_650, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_643, c_0_644]), c_0_645]), c_0_293])])])).
cnf(c_0_651, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_646, c_0_337]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_652, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__9)|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_647, c_0_376])])).
cnf(c_0_653, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(np__5)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(np__5)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_648, c_0_283])])).
cnf(c_0_654, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_649, c_0_376])).
cnf(c_0_655, plain, (esk4_2(X1,X2)!=np__1|esk5_2(X1,X2)!=k6_xcmplx_0(k3_finseq_1(X2),np__2)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_656, negated_conjecture, (np__7=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_650, c_0_651]), c_0_223]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_657, negated_conjecture, (~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__3),k1_funct_1(esk2_0,np__7),X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_658, negated_conjecture, (esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_652]), c_0_234])])])).
cnf(c_0_659, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(np__5)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_653, c_0_270])])])).
cnf(c_0_660, negated_conjecture, (r1_xxreal_0(np__7,np__9)|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__7|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_654, c_0_600]), c_0_301])])])).
cnf(c_0_661, plain, (esk4_2(X1,X2)!=np__1|esk5_2(X1,X2)!=k6_xcmplx_0(k3_finseq_1(X2),np__2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_655])).
cnf(c_0_662, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__7),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__5), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_573]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_663, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_656]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_664, negated_conjecture, (~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__3),k1_funct_1(esk2_0,np__7),X1)), inference(evalgc,[status(thm)],[c_0_657])).
cnf(c_0_665, negated_conjecture, (esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_658, c_0_376])).
cnf(c_0_666, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)), inference(sr,[status(thm)],[c_0_659, c_0_302])).
cnf(c_0_667, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__7|r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_660, c_0_337]), c_0_476]), c_0_355])).
cnf(c_0_668, negated_conjecture, (esk5_2(X1,esk2_0)!=np__7|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_661, c_0_306]), c_0_321])])).
cnf(c_0_669, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_662, c_0_663]), c_0_664])])).
cnf(c_0_670, plain, (esk5_2(X2,X1)=k3_finseq_1(X1)|~r1_xxreal_0(k3_finseq_1(X1),esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_305])])).
cnf(c_0_671, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_665, c_0_536]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_672, plain, (r1_xxreal_0(X1,X2)|np__5=X2|v1_abian(X2)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,X2)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_400]), c_0_301])])])).
cnf(c_0_673, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_666, c_0_272])])).
cnf(c_0_674, negated_conjecture, (r1_xxreal_0(np__7,np__7)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v1_xreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_565, c_0_667]), c_0_334]), c_0_208])])])).
cnf(c_0_675, plain, (v1_xreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_358]), c_0_249])])])).
cnf(c_0_676, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__7|esk4_2(esk1_0,esk2_0)!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_668, c_0_469]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_677, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__5), inference(csr,[status(thm)],[c_0_669, c_0_598])).
cnf(c_0_678, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_670, c_0_306]), c_0_234])])])).
cnf(c_0_679, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_671, c_0_603]), c_0_234]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_680, plain, (np__5=X1|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)), inference(csr,[status(thm)],[c_0_672, c_0_272])).
cnf(c_0_681, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_673, c_0_131])).
cnf(c_0_682, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__7,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_674, c_0_675]), c_0_476])).
cnf(c_0_683, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)!=np__7), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_676, c_0_677])])).
cnf(c_0_684, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_678, c_0_272])])).
cnf(c_0_685, negated_conjecture, (~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_679]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_686, plain, (np__5=X1|r1_xxreal_0(X2,X1)|v1_abian(X1)|np__3=X1|~r1_xxreal_0(X2,np__7)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(np__3,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_680, c_0_681])])).
cnf(c_0_687, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__7,np__7)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_682, c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_688, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9), inference(csr,[status(thm)],[c_0_683, c_0_573])).
cnf(c_0_689, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_684, c_0_337])).
cnf(c_0_690, negated_conjecture, (~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_685, c_0_455]), c_0_334]), c_0_208])])])).
cnf(c_0_691, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(X3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~r1_xxreal_0(X3,np__7)|~v7_ordinal1(esk4_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_686, c_0_440])])).
cnf(c_0_692, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__7,np__7)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_610, c_0_687]), c_0_186])])])).
cnf(c_0_693, negated_conjecture, (np__9=np__5|np__9=np__7|esk5_2(esk1_0,esk2_0)=np__5), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_573, c_0_688])])).
cnf(c_0_694, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|np__9!=np__5), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_688])])).
cnf(c_0_695, plain, (r1_glib_000(X1,k1_funct_1(X2,esk5_2(X1,X2)),k1_funct_1(X2,esk4_2(X1,X2)),esk6_2(X1,X2))|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_482])])).
cnf(c_0_696, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_689, c_0_603]), c_0_234])])])).
cnf(c_0_697, negated_conjecture, (~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_690, c_0_471]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_698, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(X3,esk4_2(X1,X2))|~r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_691, c_0_308]), c_0_397])).
cnf(c_0_699, negated_conjecture, (r1_xxreal_0(np__7,np__7)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_692, c_0_186]), c_0_304])])).
cnf(c_0_700, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|np__9=np__7), inference(csr,[status(thm)],[c_0_693, c_0_694])).
cnf(c_0_701, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_695, c_0_696]), c_0_483])])).
cnf(c_0_702, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_697, c_0_698]), c_0_699]), c_0_301]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_703, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_700]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_704, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_701, c_0_702]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])]), c_0_517])])).
cnf(c_0_705, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_703, c_0_572]), c_0_334]), c_0_208])])])).
cnf(c_0_706, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_685, c_0_704]), c_0_351])])])).
cnf(c_0_707, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_705, c_0_471]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_708, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__9|~r1_xxreal_0(np__9,np__5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_538, c_0_706])])).
cnf(c_0_709, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(esk1_0,esk2_0))|np__9=np__7), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_482, c_0_700]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_710, negated_conjecture, (np__9=np__7|esk4_2(esk1_0,esk2_0)=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_707, c_0_634]), c_0_635]), c_0_283]), c_0_466]), c_0_196]), c_0_197]), c_0_198]), c_0_199]), c_0_200]), c_0_201])])])).
cnf(c_0_711, negated_conjecture, (~r1_xxreal_0(np__9,np__5)), inference(csr,[status(thm)],[c_0_708, c_0_679])).
cnf(c_0_712, negated_conjecture, (np__9=np__7), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_709, c_0_710]), c_0_517])])).
cnf(c_0_713, negated_conjecture, (~r1_xxreal_0(np__7,np__5)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_711, c_0_712])])).
cnf(c_0_714, plain, (np__7=np__5), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_321, c_0_712]), c_0_261])])).
cnf(c_0_715, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_713, c_0_714]), c_0_645])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 716
# Proof object clause steps            : 594
# Proof object formula steps           : 122
# Proof object conjectures             : 174
# Proof object clause conjectures      : 171
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 102
# Proof object initial formulas used   : 77
# Proof object generating inferences   : 277
# Proof object simplifying inferences  : 699
# Training examples: 392 positive, 2847 negative
# Training: Positive examples begin
cnf(c_0_195, negated_conjecture, (k15_glib_001(esk1_0,esk2_0)=np__4)).# trainpos
cnf(c_0_196, negated_conjecture, (m3_glib_001(esk2_0,esk1_0))).# trainpos
cnf(c_0_197, negated_conjecture, (v1_glib_000(esk1_0))).# trainpos
cnf(c_0_200, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers))).# trainpos
cnf(c_0_466, negated_conjecture, (v5_chord(esk2_0,esk1_0))).# trainpos
cnf(c_0_201, negated_conjecture, (v1_relat_1(esk1_0))).# trainpos
cnf(c_0_469, negated_conjecture, (v8_glib_001(esk2_0,esk1_0))).# trainpos
cnf(c_0_198, negated_conjecture, (v1_finset_1(esk1_0))).# trainpos
cnf(c_0_199, negated_conjecture, (v1_funct_1(esk1_0))).# trainpos
cnf(c_0_507, plain, (r1_glib_000(X1,X3,X2,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~v4_relat_1(X1,k5_numbers)|~r1_glib_000(X1,X2,X3,X4))).# trainpos
cnf(c_0_101, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_482, plain, (r1_glib_000(X1,k1_funct_1(X2,esk4_2(X1,X2)),k1_funct_1(X2,esk5_2(X1,X2)),esk6_2(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_106, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_664, negated_conjecture, (~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__3),k1_funct_1(esk2_0,np__7),X1))).# trainpos
cnf(c_0_517, negated_conjecture, (~r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__5),X1))).# trainpos
cnf(c_0_661, plain, (esk4_2(X1,X2)!=np__1|esk5_2(X1,X2)!=k6_xcmplx_0(k3_finseq_1(X2),np__2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_252, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_519, plain, (esk4_2(X1,X2)!=np__1|esk5_2(X1,X2)!=k3_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_251, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_505, plain, (esk4_2(X1,X2)!=np__3|esk5_2(X1,X2)!=k3_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v8_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_182, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_397, plain, (~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)|~v1_abian(esk4_2(X1,X2)))).# trainpos
cnf(c_0_290, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).# trainpos
cnf(c_0_451, plain, (k3_glib_001(X1,X2)=k1_funct_1(X2,np__1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_243, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_136, plain, (v1_abian(X2)|v1_abian(X1)|r1_xxreal_0(X1,X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X1)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_138, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_355, plain, (~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)|~v1_abian(esk5_2(X1,X2)))).# trainpos
cnf(c_0_419, plain, (k4_glib_001(X2,X1)=k3_glib_001(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v1_glib_001(X1,X2)|~v4_relat_1(X2,k5_numbers))).# trainpos
cnf(c_0_264, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).# trainpos
cnf(c_0_280, plain, (~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers)|~r1_xxreal_0(esk5_2(X1,X2),k1_nat_1(esk4_2(X1,X2),np__2)))).# trainpos
cnf(c_0_300, plain, (k2_xcmplx_0(np__2,np__5)=np__7)).# trainpos
cnf(c_0_192, plain, (m1_subset_1(np__3,k5_numbers))).# trainpos
cnf(c_0_183, plain, (m1_subset_1(np__5,k5_numbers))).# trainpos
cnf(c_0_221, plain, (v1_abian(X2)|v1_abian(X1)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,k6_xcmplx_0(X2,np__2))|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_139, plain, (m1_subset_1(np__7,k5_numbers))).# trainpos
cnf(c_0_334, plain, (r1_xxreal_0(np__2,np__2))).# trainpos
cnf(c_0_250, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_435, plain, (k4_glib_001(X1,X2)=k1_funct_1(X2,k3_finseq_1(X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_203, plain, (m1_subset_1(np__4,k5_numbers))).# trainpos
cnf(c_0_360, plain, (v1_abian(X1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_341, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_384, plain, (k2_xcmplx_0(np__2,np__7)=np__9)).# trainpos
cnf(c_0_158, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).# trainpos
cnf(c_0_629, plain, (v1_abian(X1)|~v1_int_1(X1)|~v1_abian(k2_xcmplx_0(X1,np__2)))).# trainpos
cnf(c_0_94, plain, (k2_xcmplx_0(X1,X2)=k1_nat_1(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_351, plain, (k2_xcmplx_0(np__3,np__2)=np__5)).# trainpos
cnf(c_0_364, plain, (k2_xcmplx_0(np__5,np__2)=np__7)).# trainpos
cnf(c_0_179, plain, (k6_xcmplx_0(np__5,np__1)=np__4)).# trainpos
cnf(c_0_130, plain, (k3_xcmplx_0(np__2,np__2)=np__4)).# trainpos
cnf(c_0_100, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_194, plain, (k3_finseq_1(X2)=k2_nat_1(k4_nat_1(np__2,k15_glib_001(X1,X2)),np__1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_211, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_371, plain, (k6_xcmplx_0(np__5,np__2)=np__3)).# trainpos
cnf(c_0_131, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_308, plain, (v7_ordinal1(esk4_2(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_291, plain, (k3_xcmplx_0(np__2,np__0)=np__0)).# trainpos
cnf(c_0_261, plain, (k6_xcmplx_0(np__7,np__2)=np__5)).# trainpos
cnf(c_0_185, plain, (k2_xcmplx_0(np__7,np__2)=np__9)).# trainpos
cnf(c_0_420, plain, (v1_glib_001(X1,X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v8_glib_001(X1,X2)|~v4_relat_1(X2,k5_numbers))).# trainpos
cnf(c_0_263, plain, (~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k2_xcmplx_0(X1,np__1)))).# trainpos
cnf(c_0_135, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_246, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_305, plain, (r1_xxreal_0(esk5_2(X1,X2),k3_finseq_1(X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_178, plain, (~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k6_xcmplx_0(X1,np__1)))).# trainpos
cnf(c_0_611, plain, (v1_abian(X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k2_xcmplx_0(X2,X1)))).# trainpos
cnf(c_0_277, plain, (k2_xcmplx_0(np__8,np__1)=np__9)).# trainpos
cnf(c_0_591, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_337, plain, (v7_ordinal1(esk5_2(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v5_chord(X2,X1)|~v4_relat_1(X1,k5_numbers))).# trainpos
cnf(c_0_321, plain, (k6_xcmplx_0(np__9,np__2)=np__7)).# trainpos
cnf(c_0_129, plain, (v1_abian(k3_xcmplx_0(np__2,X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_95, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_145, plain, (k3_xcmplx_0(np__2,np__3)=np__6)).# trainpos
cnf(c_0_279, plain, (m1_subset_1(np__8,k5_numbers))).# trainpos
cnf(c_0_288, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_161, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_160, plain, (m1_subset_1(np__2,k1_numbers))).# trainpos
cnf(c_0_190, plain, (k6_xcmplx_0(np__7,np__1)=np__6)).# trainpos
cnf(c_0_324, plain, (m1_subset_1(np__9,k5_numbers))).# trainpos
cnf(c_0_275, plain, (k2_xcmplx_0(X1,X2)=k2_nat_1(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_159, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_226, plain, (v1_abian(X2)|v1_abian(X1)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_abian(k3_xcmplx_0(X1,X2)))).# trainpos
cnf(c_0_272, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_217, plain, (k3_xcmplx_0(X1,X2)=k4_nat_1(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_162, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_479, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_240, plain, (k3_xcmplx_0(np__2,np__4)=np__8)).# trainpos
cnf(c_0_157, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X4))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_322, plain, (k6_xcmplx_0(np__9,np__1)=np__8)).# trainpos
cnf(c_0_223, plain, (m1_subset_1(np__5,k1_numbers))).# trainpos
cnf(c_0_463, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_554, plain, (m1_subset_1(np__3,k1_numbers))).# trainpos
cnf(c_0_228, plain, (m1_subset_1(np__7,k1_numbers))).# trainpos
cnf(c_0_218, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_219, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_216, plain, (m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_210, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_286, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_142, plain, (m1_subset_1(np__9,k1_numbers))).# trainpos
cnf(c_0_188, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_141, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_315, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_695, plain, (r1_glib_000(X1,k1_funct_1(X2,esk5_2(X1,X2)),k1_funct_1(X2,esk4_2(X1,X2)),esk6_2(X1,X2))|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1))).# trainpos
cnf(c_0_516, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=k3_finseq_1(esk2_0)|esk4_2(esk1_0,esk2_0)!=np__3)).# trainpos
cnf(c_0_528, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=k3_finseq_1(esk2_0)|esk4_2(esk1_0,esk2_0)!=np__1)).# trainpos
cnf(c_0_260, plain, (k6_xcmplx_0(X1,np__2)=X2|r1_xxreal_0(X1,X2)|v1_abian(X2)|v1_abian(X1)|~r1_xxreal_0(k6_xcmplx_0(X1,np__2),X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(k6_xcmplx_0(X1,np__2))|~v1_int_1(X1)|~v1_int_1(X2))).# trainpos
cnf(c_0_311, plain, (m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_231, plain, (v7_ordinal1(np__5))).# trainpos
cnf(c_0_436, plain, (k4_glib_001(X1,X2)=k3_glib_001(X1,X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X1))).# trainpos
cnf(c_0_318, plain, (v1_abian(np__0)|~v1_int_1(np__0))).# trainpos
cnf(c_0_489, plain, (v7_ordinal1(np__3))).# trainpos
cnf(c_0_241, plain, (v7_ordinal1(np__4))).# trainpos
cnf(c_0_278, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_249, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_307, plain, (~r1_xxreal_0(esk5_2(X1,X2),k2_xcmplx_0(esk4_2(X1,X2),np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).# trainpos
cnf(c_0_186, plain, (v7_ordinal1(np__7))).# trainpos
cnf(c_0_594, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_239, negated_conjecture, (k2_nat_1(k4_nat_1(np__2,np__4),np__1)=k3_finseq_1(esk2_0))).# trainpos
cnf(c_0_323, plain, (v1_abian(np__8)|~v1_int_1(np__4))).# trainpos
cnf(c_0_184, plain, (X1=k2_xcmplx_0(X2,np__2)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_361, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_abian(X1)|v1_abian(np__1)|~v1_int_1(np__1)|~v1_int_1(X1))).# trainpos
cnf(c_0_377, plain, (X1=np__1|v1_abian(X1)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_209, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_229, plain, (~v1_abian(np__4)|~v1_abian(np__5)|~v1_int_1(np__5))).# trainpos
cnf(c_0_180, plain, (v1_abian(np__4)|~v1_int_1(np__2))).# trainpos
cnf(c_0_592, plain, (v7_ordinal1(np__8))).# trainpos
cnf(c_0_388, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|v1_abian(np__5)|~v1_int_1(np__5)|~v1_int_1(X1))).# trainpos
cnf(c_0_208, plain, (v1_xreal_0(np__2))).# trainpos
cnf(c_0_382, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|v1_abian(np__5)|~v1_int_1(np__5)|~v1_int_1(X1))).# trainpos
cnf(c_0_343, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|v1_abian(X2)|v1_abian(X3)|~r1_xxreal_0(X1,k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(X2)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(X1)|~v1_int_1(X3)|~v1_int_1(X2))).# trainpos
cnf(c_0_549, plain, (v1_xreal_0(np__5)|~v3_membered(k5_numbers))).# trainpos
cnf(c_0_563, plain, (v1_xreal_0(np__3))).# trainpos
cnf(c_0_297, plain, (~v1_abian(np__3)|~v1_abian(np__2)|~v1_int_1(np__2))).# trainpos
cnf(c_0_317, plain, (~v1_abian(np__1)|~v1_abian(np__0)|~v1_int_1(np__0))).# trainpos
cnf(c_0_423, plain, (v1_xreal_0(np__7)|~v3_membered(k5_numbers))).# trainpos
cnf(c_0_191, plain, (v1_abian(np__6)|~v1_int_1(np__3))).# trainpos
cnf(c_0_255, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_583, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__3)|v1_abian(X1)|v1_abian(np__3)|~v1_int_1(np__3)|~v1_int_1(X1))).# trainpos
cnf(c_0_236, plain, (~v1_abian(np__6)|~v1_abian(np__7)|~v1_int_1(np__7))).# trainpos
cnf(c_0_265, plain, (v1_abian(np__2)|~v1_abian(np__4)|~v1_int_1(np__2))).# trainpos
cnf(c_0_126, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_670, plain, (esk5_2(X2,X1)=k3_finseq_1(X1)|~r1_xxreal_0(k3_finseq_1(X1),esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_207, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_181, plain, (v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_562, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__3))).# trainpos
cnf(c_0_346, plain, (~v1_abian(np__8)|~v1_abian(np__9)|~v1_int_1(np__9))).# trainpos
cnf(c_0_312, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_446, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(X1,esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).# trainpos
cnf(c_0_348, plain, (v7_ordinal1(np__9))).# trainpos
cnf(c_0_380, plain, (v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_556, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__5))).# trainpos
cnf(c_0_345, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|v1_abian(np__9)|~v1_int_1(np__9)|~v1_int_1(X1))).# trainpos
cnf(c_0_438, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__7,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1)|~v1_xreal_0(np__7))).# trainpos
cnf(c_0_254, plain, (m2_subset_1(k3_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_378, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_433, plain, (v1_int_1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_248, plain, (v1_xxreal_0(np__3)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_262, plain, (v1_xxreal_0(np__5)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_268, plain, (v1_xxreal_0(np__7)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_187, plain, (v1_xxreal_0(np__9)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_338, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_230, plain, (v1_abian(np__4))).# trainpos
cnf(c_0_452, plain, (k3_glib_001(X1,X2)=k1_funct_1(X2,k3_finseq_1(X2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X1))).# trainpos
cnf(c_0_363, plain, (~v1_abian(np__1))).# trainpos
cnf(c_0_326, plain, (~v1_abian(np__3)|~v1_abian(np__2))).# trainpos
cnf(c_0_327, plain, (v1_abian(np__2))).# trainpos
cnf(c_0_362, plain, (v1_int_1(np__1))).# trainpos
cnf(c_0_347, plain, (v1_abian(np__8))).# trainpos
cnf(c_0_395, plain, (v1_xxreal_0(np__1))).# trainpos
cnf(c_0_237, plain, (v1_abian(np__6))).# trainpos
cnf(c_0_597, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,X1)|v1_abian(np__3)|v1_abian(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_402, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__3)|v1_abian(np__5)|v1_abian(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_398, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,X1)|v1_abian(np__5)|v1_abian(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_283, plain, (v1_xxreal_0(np__3))).# trainpos
cnf(c_0_276, negated_conjecture, (k2_nat_1(np__8,np__1)=k3_finseq_1(esk2_0))).# trainpos
cnf(c_0_359, plain, (v2_membered(k5_numbers))).# trainpos
cnf(c_0_608, plain, (v1_int_1(np__2))).# trainpos
cnf(c_0_648, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(np__5)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(np__5)|~v1_int_1(X1)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_293, plain, (v1_xxreal_0(np__5))).# trainpos
cnf(c_0_368, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__7)|v1_abian(np__9)|v1_abian(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_292, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(np__7)|~v1_int_1(X1)|~v1_xxreal_0(np__5)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_500, plain, (v1_int_1(np__3))).# trainpos
cnf(c_0_301, plain, (v1_xxreal_0(np__7))).# trainpos
cnf(c_0_234, plain, (v1_xxreal_0(np__9))).# trainpos
cnf(c_0_270, plain, (v1_int_1(np__5))).# trainpos
cnf(c_0_233, plain, (v1_int_1(np__7))).# trainpos
cnf(c_0_287, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_176, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_335, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_605, plain, (v1_int_1(np__8))).# trainpos
cnf(c_0_370, plain, (v1_int_1(np__9))).# trainpos
cnf(c_0_534, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_462, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(esk5_2(X3,X2))|v1_abian(X1)|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(X1)|~v1_int_1(esk5_2(X3,X2)))).# trainpos
cnf(c_0_282, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__2,X3)|~r1_xxreal_0(np__1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_232, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(np__7)|~v1_int_1(X1)|~v1_xxreal_0(np__9)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_306, negated_conjecture, (k3_finseq_1(esk2_0)=np__9)).# trainpos
cnf(c_0_352, plain, (~v1_abian(np__3))).# trainpos
cnf(c_0_336, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_269, plain, (~v1_abian(np__5)|~v1_int_1(np__5))).# trainpos
cnf(c_0_273, plain, (~v1_abian(np__7)|~v1_int_1(np__7))).# trainpos
cnf(c_0_369, plain, (~v1_abian(np__9)|~v1_int_1(np__9))).# trainpos
cnf(c_0_379, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_468, plain, (k1_funct_1(X2,k3_finseq_1(X2))=k1_funct_1(X2,np__1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X1))).# trainpos
cnf(c_0_365, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|v1_abian(X2)|v1_abian(X1)|~r1_xxreal_0(X3,k2_xcmplx_0(np__2,X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_488, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,X2)))).# trainpos
cnf(c_0_659, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(np__5)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_344, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_416, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,X1)|v1_abian(np__5)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_531, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_409, plain, (X1=np__1|v1_abian(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_412, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__5)|v1_abian(np__5)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_386, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(np__9)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_266, plain, (X1=k2_xcmplx_0(np__2,X2)|r1_xxreal_0(X1,X2)|v1_abian(X2)|v1_abian(X1)|~r1_xxreal_0(X1,k2_xcmplx_0(np__2,X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X2))|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_357, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_303, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_568, plain, (r1_xxreal_0(k2_xcmplx_0(np__2,X1),X2)|r1_xxreal_0(X2,X1)|v1_abian(X2)|v1_abian(X1)|~v1_int_1(X1)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_333, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__2,X3)|~r1_xxreal_0(np__1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).# trainpos
cnf(c_0_613, plain, (r1_xxreal_0(np__5,np__3)|r1_xxreal_0(np__3,np__3)|v1_abian(np__3))).# trainpos
cnf(c_0_544, plain, (v1_xxreal_0(k1_nat_1(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_526, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__9|esk4_2(esk1_0,esk2_0)!=np__3)).# trainpos
cnf(c_0_483, negated_conjecture, (k1_funct_1(esk2_0,np__9)=k1_funct_1(esk2_0,np__1))).# trainpos
cnf(c_0_538, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__9|esk4_2(esk1_0,esk2_0)!=np__1)).# trainpos
cnf(c_0_358, plain, (m1_subset_1(esk5_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_411, plain, (m1_subset_1(esk4_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_668, negated_conjecture, (esk5_2(X1,esk2_0)!=np__7|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_590, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_330, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_541, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_512, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|r1_xxreal_0(esk5_2(X2,X1),np__3)|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_678, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_577, plain, (r1_xxreal_0(np__7,k3_finseq_1(X1))|r1_xxreal_0(esk5_2(X2,X1),np__5)|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_350, plain, (X1=k2_xcmplx_0(X2,np__2)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__2))|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_673, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_385, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_399, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_491, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X3)|v1_abian(X2)|v1_abian(X3)|~r1_xxreal_0(X1,k2_xcmplx_0(X3,np__2))|~v1_int_1(X2)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X3,k5_numbers))).# trainpos
cnf(c_0_396, plain, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_356, plain, (~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__3)|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,X2),np__2))|~v1_xreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2)))).# trainpos
cnf(c_0_328, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_676, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__7|esk4_2(esk1_0,esk2_0)!=np__1)).# trainpos
cnf(c_0_684, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).# trainpos
cnf(c_0_354, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_558, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__9)|r1_xxreal_0(esk5_2(X2,esk2_0),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_424, plain, (esk4_2(X1,X2)=np__1|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_376, plain, (v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_375, plain, (~r1_xxreal_0(esk5_2(X1,X2),np__3)|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v1_xreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).# trainpos
cnf(c_0_675, plain, (v1_xreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_425, plain, (v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_413, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_478, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_471, plain, (v1_xreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_588, plain, (r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).# trainpos
cnf(c_0_522, plain, (r1_xxreal_0(esk5_2(X1,X2),np__3)|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).# trainpos
cnf(c_0_353, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_502, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_415, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),np__7)|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_649, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_445, plain, (r1_xxreal_0(np__5,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),np__3)|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_609, plain, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__3)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_574, plain, (k2_xcmplx_0(X1,np__2)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_604, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|v1_abian(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).# trainpos
cnf(c_0_372, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v1_int_1(X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_442, plain, (r1_xxreal_0(esk4_2(X1,X2),np__5)|r1_xxreal_0(np__7,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_414, plain, (X1=np__7|np__5=X1|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_672, plain, (r1_xxreal_0(X1,X2)|np__5=X2|v1_abian(X2)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,X2)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_456, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_408, plain, (v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).# trainpos
cnf(c_0_624, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_392, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_391, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_610, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_492, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk5_2(X2,esk2_0)=np__9|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_686, plain, (np__5=X1|r1_xxreal_0(X2,X1)|v1_abian(X1)|np__3=X1|~r1_xxreal_0(X2,np__7)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(np__3,X1))).# trainpos
cnf(c_0_490, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_514, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_618, plain, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_444, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__9,esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_532, plain, (r1_xxreal_0(esk5_2(X1,X2),np__3)|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).# trainpos
cnf(c_0_443, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(X2,X3),X1)|v1_abian(X1)|~r1_xxreal_0(esk5_2(X2,X3),np__9)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_441, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_510, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v1_int_1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_585, plain, (k2_xcmplx_0(X1,np__2)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_432, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(X3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(X3,np__9)|~v1_int_1(esk5_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_551, negated_conjecture, (r1_xxreal_0(np__5,np__9))).# trainpos
cnf(c_0_476, plain, (r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_631, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_421, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(esk5_2(X1,esk2_0),np__3)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_473, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_465, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,X1),np__7)|r1_xxreal_0(X2,esk5_2(esk1_0,X1))|~r1_xxreal_0(X2,np__9)|~v1_xxreal_0(X2)|~v5_chord(X1,esk1_0)|~m3_glib_001(X1,esk1_0))).# trainpos
cnf(c_0_566, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(esk5_2(X1,esk2_0),np__3)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__7,np__9)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_459, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_660, negated_conjecture, (r1_xxreal_0(np__7,np__9)|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__7|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).# trainpos
cnf(c_0_501, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_691, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(X3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~r1_xxreal_0(X3,np__7)|~v7_ordinal1(esk4_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_529, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_475, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__7|esk5_2(X2,X3)=np__5|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,esk5_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_485, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_622, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_496, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_559, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_639, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_589, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_474, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_674, negated_conjecture, (r1_xxreal_0(np__7,np__7)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v1_xreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_687, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__7,np__7))).# trainpos
cnf(c_0_632, negated_conjecture, (r1_xxreal_0(np__5,np__5)|r1_xxreal_0(np__5,np__7))).# trainpos
cnf(c_0_696, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_567, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__9)|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_701, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_692, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__7,np__7)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_699, negated_conjecture, (r1_xxreal_0(np__7,np__7))).# trainpos
cnf(c_0_513, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(X2,esk2_0),X1)|v1_abian(X1)|~v7_ordinal1(esk5_2(X2,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_535, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_640, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__5)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_533, negated_conjecture, (r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__7)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers))).# trainpos
cnf(c_0_645, negated_conjecture, (r1_xxreal_0(np__5,np__5))).# trainpos
cnf(c_0_576, negated_conjecture, (esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_602, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).# trainpos
cnf(c_0_498, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__5|esk5_2(X2,esk2_0)=np__7|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_521, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk5_2(X2,esk2_0)=esk5_2(X1,esk2_0)|r1_xxreal_0(esk5_2(X2,esk2_0),np__7)|v1_abian(esk5_2(X2,esk2_0))|~v7_ordinal1(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).# trainpos
cnf(c_0_617, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_586, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_543, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_454, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_470, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0)))).# trainpos
cnf(c_0_540, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_601, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0)))).# trainpos
cnf(c_0_552, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).# trainpos
cnf(c_0_628, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|r1_xxreal_0(np__7,np__9)|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_497, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__3|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_548, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_627, negated_conjecture, (r1_xxreal_0(np__7,np__9)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_508, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|r1_glib_000(X1,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_637, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_636, negated_conjecture, (r1_xxreal_0(np__7,np__9)|esk5_2(esk1_0,esk2_0)!=np__9)).# trainpos
cnf(c_0_518, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__5),esk6_2(X1,esk2_0))|esk4_2(X1,esk2_0)=np__3|esk4_2(X1,esk2_0)=np__1|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1))).# trainpos
cnf(c_0_480, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_527, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3)).# trainpos
cnf(c_0_642, negated_conjecture, (r1_xxreal_0(np__7,np__9))).# trainpos
cnf(c_0_570, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).# trainpos
cnf(c_0_537, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)!=np__9)).# trainpos
cnf(c_0_647, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_593, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0)))).# trainpos
cnf(c_0_652, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__9)|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_620, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,esk2_0)))).# trainpos
cnf(c_0_555, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(X1,esk2_0)!=np__9|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_573, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5)).# trainpos
cnf(c_0_638, negated_conjecture, (v1_abian(esk4_2(X1,esk2_0))|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|~v1_int_1(esk4_2(X1,esk2_0))|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_658, negated_conjecture, (esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_584, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,np__3))).# trainpos
cnf(c_0_662, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__7),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__5)).# trainpos
cnf(c_0_650, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_599, negated_conjecture, (~r1_xxreal_0(np__5,np__3)|~r1_xxreal_0(np__7,np__3))).# trainpos
cnf(c_0_598, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|np__7!=np__5)).# trainpos
cnf(c_0_615, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__5,np__3))).# trainpos
cnf(c_0_614, negated_conjecture, (~r1_xxreal_0(np__5,np__3)|np__7!=np__5)).# trainpos
cnf(c_0_635, plain, (r1_xxreal_0(np__3,np__3))).# trainpos
cnf(c_0_493, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(esk5_2(X1,esk2_0),esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_515, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).# trainpos
cnf(c_0_536, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7))).# trainpos
cnf(c_0_671, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_545, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__9|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_679, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_553, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__5)|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).# trainpos
cnf(c_0_685, negated_conjecture, (~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_561, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_690, negated_conjecture, (~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_582, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_697, negated_conjecture, (~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_596, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_623, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_702, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_641, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_704, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_646, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_706, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_708, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__9|~r1_xxreal_0(np__9,np__5))).# trainpos
cnf(c_0_651, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_656, negated_conjecture, (np__7=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_663, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5)).# trainpos
cnf(c_0_669, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5)).# trainpos
cnf(c_0_683, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)!=np__7)).# trainpos
cnf(c_0_694, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|np__9!=np__5)).# trainpos
cnf(c_0_693, negated_conjecture, (np__9=np__5|np__9=np__7|esk5_2(esk1_0,esk2_0)=np__5)).# trainpos
cnf(c_0_703, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).# trainpos
cnf(c_0_705, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_707, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).# trainpos
cnf(c_0_709, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(esk1_0,esk2_0))|np__9=np__7)).# trainpos
cnf(c_0_710, negated_conjecture, (np__9=np__7|esk4_2(esk1_0,esk2_0)=np__1)).# trainpos
cnf(c_0_712, negated_conjecture, (np__9=np__7)).# trainpos
cnf(c_0_714, plain, (np__7=np__5)).# trainpos
cnf(c_0_713, negated_conjecture, (~r1_xxreal_0(np__7,np__5))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_43, plain, (np__2=k2_glib_000)).#trainneg
cnf(i_0_129, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_77, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_45, plain, (k7_glib_000(X1)=k1_funct_1(X1,k2_glib_000)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_40, plain, (k1_glib_000=np__1)).#trainneg
cnf(i_0_183, plain, (k1_funct_1(X1,esk5_2(X2,X1))!=k1_funct_1(X1,esk4_2(X2,X1))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v5_chord(X1,X2)|~v4_relat_1(X2,k5_numbers))).#trainneg
cnf(i_0_71, plain, (v1_abian(X1)|v1_abian(k2_xcmplx_0(X1,np__1))|~v1_int_1(X1))).#trainneg
cnf(i_0_32, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_194, plain, (k15_glib_001(X1,X2)=k3_finseq_1(k12_glib_001(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_98, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_177, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_104, plain, (k2_xcmplx_0(np__2,np__3)=np__5)).#trainneg
cnf(i_0_113, plain, (k6_xcmplx_0(np__4,np__2)=np__2)).#trainneg
cnf(i_0_93, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_108, plain, (k2_xcmplx_0(np__4,np__1)=np__5)).#trainneg
cnf(i_0_91, plain, (k1_xboole_0=k6_numbers)).#trainneg
cnf(i_0_124, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_68, plain, (~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v1_xboole_0(k1_card_1(X1))|~v4_relat_1(X2,k5_numbers))).#trainneg
cnf(i_0_66, plain, (~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v1_abian(k1_card_1(X1))|~v4_relat_1(X2,k5_numbers))).#trainneg
cnf(i_0_122, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_191, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_79, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_100, plain, (k2_xcmplx_0(np__1,np__4)=np__5)).#trainneg
cnf(i_0_57, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_110, plain, (k2_xcmplx_0(np__6,np__1)=np__7)).#trainneg
cnf(i_0_119, plain, (k6_xcmplx_0(np__8,np__4)=np__4)).#trainneg
cnf(i_0_116, plain, (k6_xcmplx_0(np__6,np__3)=np__3)).#trainneg
cnf(i_0_101, plain, (k2_xcmplx_0(np__1,np__6)=np__7)).#trainneg
cnf(i_0_199, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_39, plain, (k4_nat_1(X1,X2)=k4_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_63, plain, (v1_abian(X2)|v1_abian(X1)|v1_abian(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_62, plain, (v1_abian(X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k6_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_41, plain, (v1_glib_001(X2,X1)|k4_glib_001(X1,X2)!=k3_glib_001(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_56, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_26, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_102, plain, (k2_xcmplx_0(np__1,np__8)=np__9)).#trainneg
cnf(i_0_154, plain, (m1_subset_1(np__6,k5_numbers))).#trainneg
cnf(i_0_17, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_59, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_128, plain, (k3_xcmplx_0(np__3,np__2)=np__6)).#trainneg
cnf(i_0_75, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_94, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_55, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_37, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_22, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_60, plain, (m2_finseq_1(X2,k2_xboole_0(k6_glib_000(X1),k7_glib_000(X1)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_130, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_16, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_58, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_145, plain, (m1_subset_1(np__4,k1_numbers))).#trainneg
cnf(i_0_74, plain, (v1_card_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_54, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_161, plain, (m1_subset_1(np__8,k1_numbers))).#trainneg
cnf(i_0_47, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_153, plain, (m1_subset_1(np__6,k1_numbers))).#trainneg
cnf(i_0_88, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_38, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).#trainneg
cnf(i_0_193, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_132, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_81, plain, (~v1_abian(esk3_0))).#trainneg
cnf(i_0_14, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_139, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_85, plain, (v1_xxreal_0(esk3_0))).#trainneg
cnf(i_0_82, plain, (v1_int_1(esk3_0))).#trainneg
cnf(i_0_151, plain, (m2_subset_1(np__5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_159, plain, (m2_subset_1(np__7,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_135, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_196, plain, (m1_subset_1(k3_glib_001(X1,X2),k6_glib_000(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_147, plain, (m2_subset_1(np__4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_143, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_197, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_167, plain, (m2_subset_1(np__9,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_163, plain, (m2_subset_1(np__8,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_92, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_155, plain, (m2_subset_1(np__6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_198, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_195, plain, (m2_glib_001(k12_glib_001(X1,X2),X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_83, plain, (v1_xreal_0(esk3_0))).#trainneg
cnf(i_0_33, plain, (v5_glib_001(X1,X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v8_glib_001(X1,X2)|~v4_relat_1(X2,k5_numbers))).#trainneg
cnf(i_0_34, plain, (~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finset_1(X2)|~v1_glib_000(X2)|~m3_glib_001(X1,X2)|~v8_glib_001(X1,X2)|~v3_glib_001(X1,X2)|~v4_relat_1(X2,k5_numbers))).#trainneg
cnf(i_0_84, plain, (v1_xcmplx_0(esk3_0))).#trainneg
cnf(i_0_190, plain, (v5_glib_001(X2,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finset_1(X1)|~v1_glib_000(X1)|~m3_glib_001(X2,X1)|~v3_glib_001(X2,X1)|~v4_relat_1(X1,k5_numbers))).#trainneg
cnf(i_0_140, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_136, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_148, plain, (v2_xxreal_0(np__4))).#trainneg
cnf(i_0_164, plain, (v2_xxreal_0(np__8))).#trainneg
cnf(i_0_152, plain, (v2_xxreal_0(np__5))).#trainneg
cnf(i_0_160, plain, (v2_xxreal_0(np__7))).#trainneg
cnf(i_0_156, plain, (v2_xxreal_0(np__6))).#trainneg
cnf(i_0_144, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_168, plain, (v2_xxreal_0(np__9))).#trainneg
cnf(i_0_470, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_590, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_498, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_499, plain, (v3_ordinal1(np__4))).#trainneg
cnf(i_0_585, plain, (~r1_xxreal_0(esk5_2(X1,X2),k1_nat_1(np__2,esk4_2(X1,X2)))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_496, plain, (v3_ordinal1(np__3))).#trainneg
cnf(i_0_551, plain, (v1_abian(np__4)|v1_abian(np__5)|v1_abian(np__1)|~v1_int_1(np__1)|~v1_int_1(np__5))).#trainneg
cnf(i_0_508, plain, (~v1_abian(np__2)|~v1_abian(np__1)|~v1_int_1(np__1))).#trainneg
cnf(i_0_497, plain, (v3_ordinal1(np__5))).#trainneg
cnf(i_0_596, plain, (v7_ordinal1(np__6))).#trainneg
cnf(i_0_489, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_541, plain, (v1_abian(X1)|~v1_abian(np__1)|~v1_int_1(X1)|~v1_int_1(np__1))).#trainneg
cnf(i_0_698, plain, (v1_xreal_0(np__4))).#trainneg
cnf(i_0_475, plain, (v1_abian(np__5)|v1_abian(np__4)|~v1_int_1(np__4))).#trainneg
cnf(i_0_479, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_509, plain, (~v1_abian(np__5)|~v1_abian(np__4)|~v1_int_1(np__4))).#trainneg
cnf(i_0_495, plain, (v3_ordinal1(np__7))).#trainneg
cnf(i_0_481, plain, (v1_abian(np__1)|~v1_abian(np__3)|~v1_int_1(np__1))).#trainneg
cnf(i_0_469, plain, (v1_abian(np__2)|v1_abian(np__1)|~v1_int_1(np__1))).#trainneg
cnf(i_0_573, plain, (v1_abian(np__2)|~v1_int_1(np__1))).#trainneg
cnf(i_0_494, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_624, plain, (k2_nat_1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),np__1)=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k15_glib_001(X1,X2)))).#trainneg
cnf(i_0_565, plain, (v1_int_1(np__2)|~v1_int_1(np__1))).#trainneg
cnf(i_0_474, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__4,X1)|v1_abian(X1)|v1_abian(np__4)|~v1_int_1(np__4)|~v1_int_1(X1))).#trainneg
cnf(i_0_516, plain, (r1_xxreal_0(X1,X2)|v1_abian(X2)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X2)|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_486, plain, (v1_abian(np__5)|~v1_abian(np__7)|~v1_int_1(np__5))).#trainneg
cnf(i_0_623, plain, (v1_abian(np__2)|v1_abian(X1)|~v1_int_1(X1)|~v1_int_1(np__2))).#trainneg
cnf(i_0_502, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|v1_abian(np__7)|~v1_int_1(np__7)|~v1_int_1(X1))).#trainneg
cnf(i_0_692, plain, (m2_subset_1(k4_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_706, plain, (v3_ordinal1(k1_card_1(X1)))).#trainneg
cnf(i_0_675, plain, (v1_abian(np__8)|v1_abian(np__9)|v1_abian(np__1)|~v1_int_1(np__1)|~v1_int_1(np__9))).#trainneg
cnf(i_0_493, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_600, plain, (v1_abian(np__6)|v1_abian(np__7)|v1_abian(np__1)|~v1_int_1(np__1)|~v1_int_1(np__7))).#trainneg
cnf(i_0_543, plain, (v1_abian(np__9)|v1_abian(np__8)|~v1_int_1(np__8))).#trainneg
cnf(i_0_526, plain, (~v1_abian(np__7)|~v1_abian(np__6)|~v1_int_1(np__6))).#trainneg
cnf(i_0_668, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__4,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__4)|~v1_xreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_689, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_559, plain, (v1_int_1(np__3)|~v1_int_1(np__1)|~v1_int_1(np__2))).#trainneg
cnf(i_0_484, plain, (v1_abian(np__3)|~v1_abian(np__5)|~v1_int_1(np__3))).#trainneg
cnf(i_0_606, plain, (v1_xreal_0(np__8)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_607, plain, (v1_xreal_0(np__6)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_525, plain, (v1_abian(np__7)|v1_abian(np__6)|~v1_int_1(np__6))).#trainneg
cnf(i_0_588, plain, (v3_ordinal1(np__8))).#trainneg
cnf(i_0_595, plain, (v3_ordinal1(np__6))).#trainneg
cnf(i_0_667, plain, (r1_xxreal_0(np__2,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_544, plain, (~v1_abian(np__9)|~v1_abian(np__8)|~v1_int_1(np__8))).#trainneg
cnf(i_0_577, plain, (v1_abian(np__7)|~v1_abian(np__5)|~v1_abian(np__2)|~v1_int_1(np__7)|~v1_int_1(np__2))).#trainneg
cnf(i_0_665, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__6,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__6)|~v1_xreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_574, plain, (v1_abian(np__9)|~v1_abian(np__7)|~v1_abian(np__2)|~v1_int_1(np__9)|~v1_int_1(np__2))).#trainneg
cnf(i_0_566, plain, (v1_int_1(np__5)|~v1_int_1(np__4)|~v1_int_1(np__1))).#trainneg
cnf(i_0_673, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_722, plain, (v1_int_1(np__2)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_658, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__3,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__3)|~v1_xreal_0(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_558, plain, (v1_int_1(np__7)|~v1_int_1(np__5)|~v1_int_1(np__2))).#trainneg
cnf(i_0_604, plain, (k2_xcmplx_0(k4_nat_1(np__2,k15_glib_001(X1,X2)),np__1)=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k4_nat_1(np__2,k15_glib_001(X1,X2)),k5_numbers))).#trainneg
cnf(i_0_660, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_557, plain, (v1_int_1(np__5)|~v1_int_1(np__3)|~v1_int_1(np__2))).#trainneg
cnf(i_0_554, plain, (v1_int_1(np__7)|~v1_int_1(np__1)|~v1_int_1(np__6))).#trainneg
cnf(i_0_659, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__5,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__5)|~v1_xreal_0(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_669, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__4,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1)|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_664, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__8,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__8)|~v1_xreal_0(X1)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_730, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_655, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__6,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1)|~v1_xreal_0(np__6))).#trainneg
cnf(i_0_726, plain, (v1_int_1(np__3)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_534, plain, (v1_abian(np__7)|~v1_abian(np__9)|~v1_abian(np__2)|~v1_int_1(np__7)|~v1_int_1(np__2))).#trainneg
cnf(i_0_707, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_732, plain, (v1_int_1(np__4)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_636, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X3,X4)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_656, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__0,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_637, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,k2_xcmplx_0(X4,X5))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X4,X5))|~v1_xxreal_0(X1)|~r1_xxreal_0(X5,X3)|~r1_xxreal_0(X4,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X5)|~v1_xreal_0(X2)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_521, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m2_subset_1(X2,k2_zfmisc_1(k5_numbers,X3),X1)|~m2_finseq_1(X1,X3))).#trainneg
cnf(i_0_625, plain, (k4_nat_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_720, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_728, plain, (v1_int_1(np__5)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_657, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__7,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__7)|~v1_xreal_0(X1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_504, plain, (v1_abian(np__7)|~v1_abian(np__9)|~v1_int_1(np__7))).#trainneg
cnf(i_0_605, plain, (v1_xreal_0(np__9)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_654, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__8,X1)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1)|~v1_xreal_0(np__8))).#trainneg
cnf(i_0_476, plain, (k4_glib_001(X1,X2)=k4_glib_001(X3,X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_724, plain, (v1_int_1(np__7)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_466, plain, (v7_ordinal1(k5_numbers)|~v1_finset_1(k5_numbers))).#trainneg
cnf(i_0_652, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(X2,np__4)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(np__4)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_644, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(np__1)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_633, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_556, plain, (v1_int_1(np__9)|~v1_int_1(np__7)|~v1_int_1(np__2))).#trainneg
cnf(i_0_653, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,np__4)|~v1_xreal_0(np__1)|~v1_xreal_0(X2)|~v1_xreal_0(np__4)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_503, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|v1_abian(np__7)|~v1_int_1(np__7)|~v1_int_1(X1))).#trainneg
cnf(i_0_553, plain, (v1_int_1(np__9)|~v1_int_1(np__1)|~v1_int_1(np__8))).#trainneg
cnf(i_0_670, plain, (k2_nat_1(np__1,k4_nat_1(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k4_nat_1(np__2,k15_glib_001(X1,X2))))).#trainneg
cnf(i_0_640, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,np__0)|~v1_xreal_0(np__1)|~v1_xreal_0(X2)|~v1_xreal_0(np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_650, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(X2,np__2)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(np__2)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_651, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__2)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(np__1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_713, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_736, plain, (v3_ordinal1(np__9)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_716, plain, (v1_int_1(np__8)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_602, plain, (v3_ordinal1(np__9))).#trainneg
cnf(i_0_718, plain, (v1_int_1(np__6)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_721, plain, (v1_int_1(np__2)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_642, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(np__3)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_729, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_672, plain, (k2_nat_1(np__1,k4_nat_1(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(k4_nat_1(np__2,k15_glib_001(X1,X2)),k5_numbers))).#trainneg
cnf(i_0_725, plain, (v1_int_1(np__3)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_646, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(X2,np__2)|~r1_xxreal_0(X1,np__3)|~v1_xreal_0(np__2)|~v1_xreal_0(X2)|~v1_xreal_0(np__3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_643, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(X2,np__5)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(np__5)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_731, plain, (v1_int_1(np__4)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_719, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_757, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_714, plain, (v1_int_1(np__9)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_649, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(X2,np__6)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(np__6)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_770, plain, (v1_xxreal_0(np__4)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_639, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,np__6)|~v1_xreal_0(np__1)|~v1_xreal_0(X2)|~v1_xreal_0(np__6)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_759, plain, (v1_xxreal_0(np__2)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_727, plain, (v1_int_1(np__5)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_647, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(X2,np__2)|~r1_xxreal_0(X1,np__5)|~v1_xreal_0(np__2)|~v1_xreal_0(X2)|~v1_xreal_0(np__5)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_634, plain, (m2_subset_1(X1,k2_zfmisc_1(k5_numbers,X2),X3)|v1_xboole_0(X3)|~m1_subset_1(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_723, plain, (v1_int_1(np__7)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_60, plain, (m2_finseq_1(X1,k2_xboole_0(k7_glib_000(X2),k6_glib_000(X2)))|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_648, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(X2,np__8)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(np__8)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_591, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_641, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(X2,np__7)|~r1_xxreal_0(X1,np__2)|~v1_xreal_0(np__7)|~v1_xreal_0(X2)|~v1_xreal_0(np__2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_767, plain, (v1_xxreal_0(np__1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_766, plain, (v1_xxreal_0(np__5)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_749, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_762, plain, (v1_xxreal_0(np__7)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_638, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,np__8)|~v1_xreal_0(np__1)|~v1_xreal_0(X2)|~v1_xreal_0(np__8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_645, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(X2,np__2)|~r1_xxreal_0(X1,np__7)|~v1_xreal_0(np__2)|~v1_xreal_0(X2)|~v1_xreal_0(np__7)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_747, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_756, plain, (v1_xxreal_0(np__0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_753, plain, (v1_xxreal_0(np__8)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_715, plain, (v1_int_1(np__8)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_755, plain, (v1_xxreal_0(np__6)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_717, plain, (v1_int_1(np__6)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_763, plain, (v1_xxreal_0(np__3)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_769, plain, (v1_xxreal_0(np__4)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_758, plain, (v1_xxreal_0(np__2)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_738, plain, (v1_int_1(np__9)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_752, plain, (v1_xxreal_0(np__8)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_754, plain, (v1_xxreal_0(np__6)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_467, plain, (esk5_2(X1,k12_glib_001(X2,X3))!=k6_xcmplx_0(k15_glib_001(X2,X3),np__2)|esk4_2(X1,k12_glib_001(X2,X3))!=np__1|~v5_chord(k12_glib_001(X2,X3),X1)|~v8_glib_001(k12_glib_001(X2,X3),X1)|~m3_glib_001(k12_glib_001(X2,X3),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_751, plain, (v1_xxreal_0(np__9)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_748, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_743, plain, (m1_subset_1(k1_funct_1(X2,np__1),k6_glib_000(X1))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_745, plain, (m2_subset_1(k15_glib_001(X1,X2),k1_numbers,k5_numbers)|~v1_finseq_1(k12_glib_001(X1,X2))|~v1_funct_1(k12_glib_001(X1,X2))|~v1_relat_1(k12_glib_001(X1,X2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_477, plain, (k4_glib_001(X1,k12_glib_001(X2,X3))=k1_funct_1(k12_glib_001(X2,X3),k15_glib_001(X2,X3))|~m3_glib_001(k12_glib_001(X2,X3),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_773, plain, (k3_xcmplx_0(np__1,esk3_0)=esk3_0)).#trainneg
cnf(i_0_524, plain, (r1_xxreal_0(esk5_2(X1,k12_glib_001(X2,X3)),k15_glib_001(X2,X3))|~v5_chord(k12_glib_001(X2,X3),X1)|~m3_glib_001(k12_glib_001(X2,X3),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_774, plain, (k3_xcmplx_0(X1,esk3_0)=k3_xcmplx_0(esk3_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_520, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_696, plain, (v3_membered(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_776, plain, (v2_membered(X1)|~v2_membered(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_517, plain, (r1_xxreal_0(X1,k6_xcmplx_0(X2,np__2))|r1_xxreal_0(X2,X3)|v1_abian(X3)|v1_abian(X2)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(k6_xcmplx_0(X2,np__2))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v1_int_1(X3))).#trainneg
cnf(i_0_897, plain, (v1_abian(np__0))).#trainneg
cnf(i_0_946, plain, (v1_abian(np__6)|v1_abian(np__7))).#trainneg
cnf(i_0_923, plain, (v1_abian(np__3)|~v1_abian(np__5))).#trainneg
cnf(i_0_1042, plain, (v1_abian(np__7)|~v1_abian(np__5)|~v1_int_1(np__2)|~v1_int_1(np__5))).#trainneg
cnf(i_0_1613, plain, (v1_xxreal_0(np__4))).#trainneg
cnf(i_0_908, plain, (v1_abian(np__5)|~v1_abian(np__7))).#trainneg
cnf(i_0_970, plain, (v1_abian(np__7)|~v1_abian(np__5)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1422, plain, (v1_abian(np__9)|~v1_abian(np__7)|~v1_int_1(np__2)|~v1_int_1(np__7))).#trainneg
cnf(i_0_1702, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_799, plain, (v1_int_1(np__4))).#trainneg
cnf(i_0_991, plain, (v1_abian(np__9)|~v1_abian(np__7)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1658, plain, (v1_relat_1(X1)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_790, plain, (m1_subset_1(k4_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_901, plain, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__5)|v1_abian(np__7)|v1_abian(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_177, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1509, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1291, negated_conjecture, (k2_nat_1(k3_xcmplx_0(np__4,np__2),np__1)=k3_finseq_1(esk2_0))).#trainneg
cnf(i_0_1324, plain, (v1_abian(np__7)|~v1_abian(np__9))).#trainneg
cnf(i_0_1752, plain, (v1_xxreal_0(np__2))).#trainneg
cnf(i_0_1442, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(np__7)|v1_abian(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1656, plain, (v1_funct_1(X1)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1745, plain, (v1_xxreal_0(np__0))).#trainneg
cnf(i_0_1513, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_778, plain, (np__7=X1|r1_xxreal_0(np__9,X1)|v1_abian(np__9)|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_int_1(np__9)|~v1_int_1(X1)|~v1_xxreal_0(np__7)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1038, plain, (r1_xxreal_0(np__1,k3_finseq_1(X1))|v1_abian(esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v1_xxreal_0(np__1)|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_1349, plain, (k4_glib_001(X3,X2)=k3_glib_001(X1,X2)|~v8_glib_001(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_840, plain, (m1_subset_1(k4_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_811, plain, (X1=np__3|r1_xxreal_0(X1,np__1)|v1_abian(np__1)|v1_abian(X1)|~r1_xxreal_0(X1,np__3)|~v1_int_1(np__1)|~v1_int_1(X1)|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_994, negated_conjecture, (k2_xcmplx_0(np__1,k4_nat_1(np__2,np__4))=k3_finseq_1(esk2_0)|~v7_ordinal1(k4_nat_1(np__2,np__4)))).#trainneg
cnf(i_0_1657, plain, (v1_finseq_1(X1)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_823, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_844, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_804, negated_conjecture, (k2_xcmplx_0(k4_nat_1(np__2,np__4),np__1)=k3_finseq_1(esk2_0)|~m1_subset_1(k4_nat_1(np__2,np__4),k5_numbers))).#trainneg
cnf(i_0_1292, plain, (k2_nat_1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2),np__1)=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers))).#trainneg
cnf(i_0_1293, plain, (k2_xcmplx_0(k3_xcmplx_0(k15_glib_001(X1,X2),np__2),np__1)=k3_finseq_1(X2)|~m1_subset_1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers))).#trainneg
cnf(i_0_995, plain, (k2_xcmplx_0(np__1,k4_nat_1(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k4_nat_1(np__2,k15_glib_001(X1,X2))))).#trainneg
cnf(i_0_818, plain, (~r1_xxreal_0(esk5_2(X1,X2),k2_xcmplx_0(np__2,esk4_2(X1,X2)))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,X2),k5_numbers))).#trainneg
cnf(i_0_822, plain, (v1_int_1(np__6))).#trainneg
cnf(i_0_1757, plain, (v1_xxreal_0(np__6))).#trainneg
cnf(i_0_805, negated_conjecture, (k2_nat_1(np__1,k4_nat_1(np__2,np__4))=k3_finseq_1(esk2_0)|~v7_ordinal1(k4_nat_1(np__2,np__4)))).#trainneg
cnf(i_0_1195, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1756, plain, (v1_xxreal_0(np__8))).#trainneg
cnf(i_0_1782, negated_conjecture, (r1_xxreal_0(esk5_2(X1,k12_glib_001(esk1_0,esk2_0)),np__4)|~v5_chord(k12_glib_001(esk1_0,esk2_0),X1)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_993, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~v7_ordinal1(k15_glib_001(X1,X2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2))))).#trainneg
cnf(i_0_1659, plain, (m2_finseq_1(X1,k2_xboole_0(k1_funct_1(X2,np__2),k6_glib_000(X2)))|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1744, plain, (k3_xcmplx_0(np__1,np__1)=np__1|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_834, plain, (k2_xcmplx_0(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),np__1)=k3_finseq_1(X2)|~m1_subset_1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),k5_numbers)|~v7_ordinal1(k15_glib_001(X1,X2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1772, plain, (k1_funct_1(k12_glib_001(X2,X3),k15_glib_001(X2,X3))=k3_glib_001(X1,k12_glib_001(X2,X3))|~v8_glib_001(k12_glib_001(X2,X3),X1)|~m3_glib_001(k12_glib_001(X2,X3),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1296, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1449, plain, (k2_nat_1(np__1,k3_xcmplx_0(k15_glib_001(X1,X2),np__2))=k3_finseq_1(X2)|~v7_ordinal1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers))).#trainneg
cnf(i_0_1448, plain, (k2_nat_1(np__1,k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~v7_ordinal1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k15_glib_001(X1,X2)))).#trainneg
cnf(i_0_1549, plain, (k2_nat_1(np__1,k3_xcmplx_0(k15_glib_001(X1,X2),np__2))=k3_finseq_1(X2)|~m1_subset_1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers))).#trainneg
cnf(i_0_809, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v1_int_1(np__3)|~v1_int_1(X1)|~v1_xxreal_0(np__5)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1214, plain, (k2_xcmplx_0(X1,X2)=np__3|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(np__2,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(np__3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_803, negated_conjecture, (k2_nat_1(np__1,k4_nat_1(np__2,np__4))=k3_finseq_1(esk2_0)|~m1_subset_1(k4_nat_1(np__2,np__4),k5_numbers))).#trainneg
cnf(i_0_788, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1058, plain, (k2_xcmplx_0(X1,X2)=np__3|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~v1_xxreal_0(np__3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1237, plain, (k2_xcmplx_0(X1,X2)=np__1|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v1_xxreal_0(np__1)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__0,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_810, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_int_1(np__5)|~v1_int_1(X1)|~v1_xxreal_0(np__7)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1771, plain, (k1_funct_1(k12_glib_001(X2,X3),k3_finseq_1(k12_glib_001(X2,X3)))=k1_funct_1(k12_glib_001(X2,X3),k15_glib_001(X2,X3))|~m3_glib_001(k12_glib_001(X2,X3),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1548, plain, (k2_nat_1(np__1,k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m1_subset_1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k15_glib_001(X1,X2)))).#trainneg
cnf(i_0_1278, plain, (r1_xxreal_0(np__2,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__2)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_1275, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__9)|~r1_xxreal_0(np__8,X4)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_1290, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__0,X3))).#trainneg
cnf(i_0_1379, plain, (v1_int_1(k5_numbers)|~v1_finset_1(k5_numbers))).#trainneg
cnf(i_0_1330, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__9)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__8,X3))).#trainneg
cnf(i_0_1284, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__3)|~r1_xxreal_0(np__2,X4)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_1236, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X3,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X3))).#trainneg
cnf(i_0_1282, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__7)|~r1_xxreal_0(np__6,X4)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_1775, negated_conjecture, (k4_glib_001(X1,k12_glib_001(esk1_0,esk2_0))=k1_funct_1(k12_glib_001(esk1_0,esk2_0),np__4)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_853, plain, (k2_xcmplx_0(X1,X2)=np__5|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~v1_xxreal_0(np__5)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__4,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_949, plain, (k2_xcmplx_0(X1,X2)=np__2|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__2)|~v1_xxreal_0(np__2)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1305, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__9)|~r1_xxreal_0(np__7,X4)|~r1_xxreal_0(np__2,X3))).#trainneg
cnf(i_0_1279, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__7)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__6,X3))).#trainneg
cnf(i_0_1017, plain, (k2_xcmplx_0(X1,X2)=np__5|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~v1_xxreal_0(np__5)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__3,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1288, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__4,X3))).#trainneg
cnf(i_0_1361, plain, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__9)|~r1_xxreal_0(np__2,X4)|~r1_xxreal_0(np__7,X3))).#trainneg
cnf(i_0_1287, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__5)|~r1_xxreal_0(np__4,X4)|~r1_xxreal_0(np__1,X3))).#trainneg
cnf(i_0_1206, plain, (k2_xcmplx_0(X1,X2)=np__7|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(np__5,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(np__7)|~v1_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_1280, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__7)|~r1_xxreal_0(np__5,X4)|~r1_xxreal_0(np__2,X3))).#trainneg
cnf(i_0_1285, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__5)|~r1_xxreal_0(np__3,X4)|~r1_xxreal_0(np__2,X3))).#trainneg
cnf(i_0_1283, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__3)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__2,X3))).#trainneg
cnf(i_0_1281, plain, (r1_xxreal_0(np__7,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__7)|~r1_xxreal_0(np__2,X4)|~r1_xxreal_0(np__5,X3))).#trainneg
cnf(i_0_1107, plain, (k2_xcmplx_0(X1,X2)=np__5|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~v1_xxreal_0(np__5)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__4,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1059, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(np__3)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__2,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1204, plain, (k2_xcmplx_0(X1,X2)=np__9|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(np__7,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(np__9)|~v1_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_1128, plain, (k2_xcmplx_0(X1,X2)=np__9|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~v1_xxreal_0(np__9)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__8,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_971, plain, (k2_xcmplx_0(X1,X2)=np__7|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~v1_xxreal_0(np__7)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__6,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1238, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(np__1)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__0,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1286, plain, (r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xxreal_0(np__5)|~r1_xxreal_0(np__2,X4)|~r1_xxreal_0(np__3,X3))).#trainneg
cnf(i_0_1201, plain, (k2_xcmplx_0(X1,X2)=np__9|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__8,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(np__9)|~v1_xxreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_1150, plain, (k2_xcmplx_0(X1,X2)=np__7|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~v1_xxreal_0(np__7)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__6,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1871, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1841, negated_conjecture, (v1_relat_1(esk2_0))).#trainneg
cnf(i_0_1905, negated_conjecture, (v1_funct_1(esk2_0))).#trainneg
cnf(i_0_1995, negated_conjecture, (v1_finseq_1(esk2_0))).#trainneg
cnf(i_0_1820, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_1954, plain, (k4_glib_001(X3,X2)=k1_funct_1(X2,np__1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X1)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_2224, negated_conjecture, (k2_xcmplx_0(np__1,k4_nat_1(np__2,np__4))=k3_finseq_1(esk2_0)|~v7_ordinal1(k4_nat_1(np__2,np__4)))).#trainneg
cnf(i_0_1793, plain, (k3_glib_001(X1,X2)=k3_glib_001(X3,X2)|~v8_glib_001(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1823, plain, (k3_xcmplx_0(X1,np__1)=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1915, plain, (v1_xxreal_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2013, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_1845, plain, (v1_xxreal_0(k4_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2308, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_xreal_0(esk5_2(X3,X2)))).#trainneg
cnf(i_0_2394, negated_conjecture, (k2_nat_1(np__1,np__8)=k3_finseq_1(esk2_0))).#trainneg
cnf(i_0_1979, plain, (r1_xxreal_0(np__7,np__5)|r1_xxreal_0(np__5,np__5)|v1_abian(np__5))).#trainneg
cnf(i_0_1873, plain, (m2_subset_1(k3_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2364, plain, (v1_abian(k3_xcmplx_0(X1,np__2))|~v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1848, plain, (v7_ordinal1(k4_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_778, plain, (np__7=X1|r1_xxreal_0(np__9,X1)|v1_abian(np__9)|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_2289, negated_conjecture, (esk5_2(X1,k12_glib_001(esk1_0,esk2_0))=np__4|~r1_xxreal_0(np__4,esk5_2(X1,k12_glib_001(esk1_0,esk2_0)))|~v1_xxreal_0(esk5_2(X1,k12_glib_001(esk1_0,esk2_0)))|~v5_chord(k12_glib_001(esk1_0,esk2_0),X1)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2063, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2014, plain, (k2_xcmplx_0(np__1,k4_nat_1(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m1_subset_1(k4_nat_1(np__2,k15_glib_001(X1,X2)),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2022, plain, (v1_abian(X1)|~v1_abian(k2_xcmplx_0(np__2,X1))|~v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1918, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2002, plain, (r1_xxreal_0(np__7,np__5)|r1_xxreal_0(np__7,np__7)|v1_abian(np__7))).#trainneg
cnf(i_0_1964, plain, (v1_xxreal_0(k4_nat_1(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1991, plain, (X1=np__3|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~r1_xxreal_0(X1,np__3)|~v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1821, plain, (k3_xcmplx_0(np__1,X1)=k3_xcmplx_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2252, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2026, plain, (~v1_abian(k2_xcmplx_0(np__1,X1))|~v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2177, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(k15_glib_001(X1,X2),np__2))=k3_finseq_1(X2)|~m1_subset_1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2),k5_numbers)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2287, plain, (r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__9,np__9)|v1_abian(np__9))).#trainneg
cnf(i_0_1978, plain, (r1_xxreal_0(np__5,np__3)|r1_xxreal_0(np__5,np__5)|v1_abian(np__5))).#trainneg
cnf(i_0_1831, plain, (r1_xxreal_0(np__7,np__1)|r1_xxreal_0(np__1,np__5)|v1_abian(np__5))).#trainneg
cnf(i_0_1932, plain, (m1_subset_1(k3_xcmplx_0(X2,X1),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2373, plain, (m1_subset_1(k3_xcmplx_0(X2,X1),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2372, plain, (m2_subset_1(k3_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2172, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(k15_glib_001(X1,X2),np__2))=k3_finseq_1(X2)|~v7_ordinal1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2))|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2356, plain, (k1_funct_1(k12_glib_001(X2,X3),k15_glib_001(X2,X3))=k1_funct_1(k12_glib_001(X2,X3),np__1)|~m3_glib_001(k12_glib_001(X2,X3),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v8_glib_001(k12_glib_001(X2,X3),X1)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2175, plain, (k2_nat_1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),np__1)=k3_finseq_1(X2)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k15_glib_001(X1,X2)))).#trainneg
cnf(i_0_1816, plain, (r1_xxreal_0(np__1,np__3)|r1_xxreal_0(np__5,np__1)|v1_abian(np__5))).#trainneg
cnf(i_0_1990, plain, (X1=np__3|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~r1_xxreal_0(X1,np__3)|~v1_int_1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1902, plain, (r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__9,np__1)|v1_abian(np__9))).#trainneg
cnf(i_0_2008, plain, (v1_int_1(k1_card_1(X1))|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_2221, negated_conjecture, (k2_xcmplx_0(k3_xcmplx_0(np__4,np__2),np__1)=k3_finseq_1(esk2_0)|~m1_subset_1(k3_xcmplx_0(np__4,np__2),k5_numbers))).#trainneg
cnf(i_0_2387, negated_conjecture, (k2_nat_1(np__1,k3_xcmplx_0(np__4,np__2))=k3_finseq_1(esk2_0)|~v7_ordinal1(k3_xcmplx_0(np__4,np__2)))).#trainneg
cnf(i_0_2301, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~v7_ordinal1(k15_glib_001(X1,X2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),k5_numbers))).#trainneg
cnf(i_0_1967, plain, (v7_ordinal1(k4_nat_1(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3033, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1916, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2048, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(X2,X1))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__0,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2029, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),k2_xcmplx_0(X3,X4))|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1951, plain, (r1_xxreal_0(np__1,k15_glib_001(X1,X2))|~v1_xxreal_0(k15_glib_001(X1,X2))|~v5_chord(k12_glib_001(X1,X2),X3)|~m3_glib_001(k12_glib_001(X1,X2),X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1846, plain, (v1_xreal_0(k4_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2272, plain, (v1_xxreal_0(k1_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2007, plain, (r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__7,np__7)|v1_abian(np__7))).#trainneg
cnf(i_0_2745, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__2,X3)|~r1_xxreal_0(np__1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_2222, plain, (k2_xcmplx_0(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),np__1)=k3_finseq_1(X2)|~m1_subset_1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)),k5_numbers)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k15_glib_001(X1,X2)))).#trainneg
cnf(i_0_516, plain, (r1_xxreal_0(X1,X2)|v1_abian(X2)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3057, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__2,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_3058, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__2,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_2275, plain, (v7_ordinal1(k1_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2388, plain, (k2_nat_1(np__1,k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~v7_ordinal1(k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2229, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2397, negated_conjecture, (k2_nat_1(np__1,k3_xcmplx_0(np__4,np__2))=k3_finseq_1(esk2_0)|~m1_subset_1(k3_xcmplx_0(np__4,np__2),k5_numbers))).#trainneg
cnf(i_0_2290, negated_conjecture, (r1_xxreal_0(X1,np__4)|~r1_xxreal_0(X1,esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~v1_xxreal_0(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~v1_xxreal_0(X1)|~v5_chord(k12_glib_001(esk1_0,esk2_0),X2)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3077, plain, (k2_xcmplx_0(X1,X2)=np__9|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(np__7,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_3078, plain, (k2_xcmplx_0(X1,X2)=np__9|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(np__7,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2460, plain, (k2_xcmplx_0(X1,X2)=np__3|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2461, plain, (k2_xcmplx_0(X1,X2)=np__3|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_3097, plain, (k2_xcmplx_0(X1,X2)=np__9|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(np__8,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2064, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X4,X3))|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3254, negated_conjecture, (v1_finset_1(esk2_0))).#trainneg
cnf(i_0_4184, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),X2)|r1_xxreal_0(X2,X1)|v1_abian(X1)|v1_abian(X2)|~v1_int_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4187, plain, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3844, plain, (~v1_abian(k2_xcmplx_0(X1,np__1))|~v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4186, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_3690, plain, (v1_abian(X1)|~v1_abian(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3255, negated_conjecture, (k1_card_1(esk2_0)=np__9)).#trainneg
cnf(i_0_3821, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3420, plain, (r1_xxreal_0(np__9,k3_finseq_1(X1))|r1_xxreal_0(esk5_2(X2,X1),np__7)|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3423, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|r1_xxreal_0(esk5_2(X2,X1),np__1)|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3265, negated_conjecture, (k4_glib_001(X1,esk2_0)=k1_funct_1(esk2_0,np__1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3253, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3791, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_487, negated_conjecture, (k2_nat_1(k4_nat_1(np__2,np__4),np__1)=np__9)).#trainneg
cnf(i_0_1291, negated_conjecture, (k2_nat_1(k3_xcmplx_0(np__4,np__2),np__1)=np__9)).#trainneg
cnf(i_0_3427, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(k2_xcmplx_0(X1,np__2)))).#trainneg
cnf(i_0_3426, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_xreal_0(k2_xcmplx_0(X1,np__2)))).#trainneg
cnf(i_0_3788, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3440, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4295, plain, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_806, negated_conjecture, (k2_nat_1(np__8,np__1)=np__9)).#trainneg
cnf(i_0_3522, plain, (np__7=X1|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3521, plain, (np__7=X1|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_int_1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3454, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3394, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,np__1)|v1_abian(X2)|~r1_xxreal_0(X1,np__3)|~v1_int_1(X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3310, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__0,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4071, plain, (X1=np__3|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~r1_xxreal_0(X1,np__3)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3238, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3781, plain, (m1_subset_1(k3_xcmplx_0(X1,np__1),k5_numbers)|~v7_ordinal1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3229, plain, (m1_subset_1(k1_card_1(X1),k5_numbers)|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_3503, negated_conjecture, (k2_xcmplx_0(np__1,k4_nat_1(np__2,np__4))=np__9)).#trainneg
cnf(i_0_3789, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3523, negated_conjecture, (esk5_2(X1,k12_glib_001(esk1_0,esk2_0))=np__4|~r1_xxreal_0(np__4,esk5_2(X1,k12_glib_001(esk1_0,esk2_0)))|~v5_chord(k12_glib_001(esk1_0,esk2_0),X1)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(esk5_2(X1,k12_glib_001(esk1_0,esk2_0))))).#trainneg
cnf(i_0_3524, negated_conjecture, (esk5_2(X1,k12_glib_001(esk1_0,esk2_0))=np__4|~r1_xxreal_0(np__4,esk5_2(X1,k12_glib_001(esk1_0,esk2_0)))|~v5_chord(k12_glib_001(esk1_0,esk2_0),X1)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,k12_glib_001(esk1_0,esk2_0))))).#trainneg
cnf(i_0_3501, plain, (v1_int_1(k4_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3990, plain, (m1_subset_1(k3_xcmplx_0(X1,np__1),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3424, plain, (r1_xxreal_0(k2_xcmplx_0(np__2,X1),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X1))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4031, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4839, plain, (k2_xcmplx_0(np__2,X1)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__2,X1))|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4025, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(k15_glib_001(X1,X2),np__2))=k3_finseq_1(X2)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3697, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4601, plain, (X1=k2_xcmplx_0(X2,np__2)|r1_xxreal_0(X1,X2)|v1_abian(X2)|v1_abian(X1)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k2_xcmplx_0(X2,np__2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3973, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4780, plain, (k2_xcmplx_0(np__2,X1)=np__9|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v7_ordinal1(k2_xcmplx_0(np__2,X1))|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3496, plain, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_4180, plain, (k2_xcmplx_0(np__2,X1)=np__1|v1_abian(k2_xcmplx_0(np__2,X1))|r1_xxreal_0(np__1,X1)|v1_abian(X1)|~v7_ordinal1(k2_xcmplx_0(np__2,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3481, plain, (v1_abian(k3_xcmplx_0(np__2,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3250, plain, (r1_xxreal_0(np__3,np__1)|r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_3970, plain, (v1_xxreal_0(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3245, plain, (r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|v1_abian(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3248, plain, (r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_3836, plain, (k2_xcmplx_0(np__2,X1)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X1))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4731, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__4)|r1_xxreal_0(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)),X1)|v1_abian(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|v1_abian(X1)|~v1_xxreal_0(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v5_chord(k12_glib_001(esk1_0,esk2_0),X2)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v1_int_1(X1)|~v1_int_1(esk5_2(X2,k12_glib_001(esk1_0,esk2_0))))).#trainneg
cnf(i_0_3699, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~v7_ordinal1(k15_glib_001(X1,X2))|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3504, plain, (k2_xcmplx_0(np__1,k4_nat_1(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k15_glib_001(X1,X2)))).#trainneg
cnf(i_0_3485, plain, (r1_xxreal_0(np__5,np__9)|r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_4276, plain, (v1_int_1(k4_nat_1(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3760, plain, (r1_xxreal_0(np__9,np__3)|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_3826, plain, (k2_xcmplx_0(np__2,X1)=np__9|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X1))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__7,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3833, plain, (k2_xcmplx_0(np__2,X1)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X1))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_804, negated_conjecture, (k2_xcmplx_0(k4_nat_1(np__2,np__4),np__1)=np__9|~m1_subset_1(k4_nat_1(np__2,np__4),k5_numbers))).#trainneg
cnf(i_0_4279, plain, (k2_xcmplx_0(np__1,k4_nat_1(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers))).#trainneg
cnf(i_0_5197, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5042, negated_conjecture, (v1_finset_1(np__9))).#trainneg
cnf(i_0_5520, plain, (m1_subset_1(k1_card_1(X1),k5_numbers)|~v1_finset_1(X1))).#trainneg
cnf(i_0_5070, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5315, plain, (r1_xxreal_0(np__7,k3_finseq_1(X1))|esk5_2(X2,X1)=np__5|v1_abian(esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__5,esk5_2(X2,X1))|~v7_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_5074, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5498, plain, (esk5_2(X1,X2)=np__1|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__3,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5069, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5333, plain, (r1_xxreal_0(np__9,k3_finseq_1(X1))|esk5_2(X2,X1)=np__7|v1_abian(esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__7,esk5_2(X2,X1))|~v7_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_5073, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5078, plain, (r1_xxreal_0(esk5_2(X1,X2),np__1)|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5450, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5502, plain, (v7_ordinal1(k3_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5195, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(k2_xcmplx_0(X1,np__2),k5_numbers))).#trainneg
cnf(i_0_5449, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5082, plain, (r1_xxreal_0(esk5_2(X1,X2),np__3)|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5077, plain, (r1_xxreal_0(esk5_2(X1,X2),np__1)|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5499, plain, (v1_xxreal_0(k3_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5216, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(k2_xcmplx_0(X1,np__2),k1_numbers))).#trainneg
cnf(i_0_5112, plain, (r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5665, plain, (v7_ordinal1(k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4953, plain, (v3_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5866, plain, (X1=k2_xcmplx_0(X2,np__2)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5199, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|esk5_2(X2,X1)=np__3|v1_abian(esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__3,esk5_2(X2,X1))|~v7_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_4961, plain, (v3_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5342, plain, (X1=np__9|np__7=X1|v1_abian(X1)|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5903, plain, (k2_xcmplx_0(X1,np__2)=np__1|r1_xxreal_0(np__1,X1)|v1_abian(X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5966, plain, (k2_xcmplx_0(X1,X2)=np__1|v1_abian(k2_xcmplx_0(X1,X2))|r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_5006, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|v1_abian(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_5829, plain, (~v1_abian(k3_finseq_1(X2))|~v1_abian(k3_xcmplx_0(k15_glib_001(X1,X2),np__2))|~m1_subset_1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2),k5_numbers)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5662, plain, (v1_xxreal_0(k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6360, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_int_1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4957, plain, (r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__7,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5146, negated_conjecture, (k2_nat_1(np__1,k4_nat_1(np__2,np__4))=np__9|~v7_ordinal1(k4_nat_1(np__2,np__4)))).#trainneg
cnf(i_0_5114, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5867, plain, (X1=k2_xcmplx_0(X2,np__2)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X2)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5891, plain, (k2_xcmplx_0(X1,np__2)=np__9|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5690, plain, (r1_xxreal_0(k2_xcmplx_0(np__2,X1),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v1_xxreal_0(k3_finseq_1(X2))|~m1_subset_1(X1,k5_numbers)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4966, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__7)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5612, negated_conjecture, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__4,np__2))=np__9)).#trainneg
cnf(i_0_5343, plain, (r1_xxreal_0(X1,X2)|np__7=X2|v1_abian(X2)|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5724, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5043, negated_conjecture, (v1_card_1(np__9))).#trainneg
cnf(i_0_5500, plain, (v1_xreal_0(k3_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4930, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__4)|r1_xxreal_0(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)),X1)|v1_abian(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|v1_abian(X1)|~v1_xxreal_0(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v5_chord(k12_glib_001(esk1_0,esk2_0),X2)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v1_int_1(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5044, negated_conjecture, (~v1_xboole_0(np__9)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4975, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|r1_xxreal_0(np__5,k2_xcmplx_0(X1,X2))|v1_abian(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_5868, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_6006, negated_conjecture, (r1_xxreal_0(esk5_2(X1,k12_glib_001(esk1_0,esk2_0)),X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),np__4)|v1_abian(X2)|~v1_int_1(X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v5_chord(k12_glib_001(esk1_0,esk2_0),X1)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,k12_glib_001(esk1_0,esk2_0))))).#trainneg
cnf(i_0_5996, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5523, negated_conjecture, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(k4_nat_1(np__2,np__4),X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(k4_nat_1(np__2,np__4))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5165, plain, (k2_xcmplx_0(X1,np__2)=np__9|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__7,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5831, plain, (k2_xcmplx_0(np__1,k3_xcmplx_0(np__2,k15_glib_001(X1,X2)))=k3_finseq_1(X2)|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5663, plain, (v1_xreal_0(k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5964, plain, (k2_xcmplx_0(X1,X2)=np__1|r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|v1_abian(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_3733, plain, (r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_5394, plain, (np__5=X1|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5075, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__1)|r1_xxreal_0(np__3,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5257, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X3,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_6240, plain, (k2_xcmplx_0(np__2,X1)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(np__2,X1)))).#trainneg
cnf(i_0_5058, plain, (k4_nat_1(X1,X2)=np__9|r1_xxreal_0(k4_nat_1(X1,X2),np__7)|v1_abian(k4_nat_1(X1,X2))|~r1_xxreal_0(k4_nat_1(X1,X2),np__9)|~v1_int_1(k4_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5910, plain, (k2_xcmplx_0(np__2,X1)=np__1|r1_xxreal_0(np__1,X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5302, plain, (np__3=X1|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5337, plain, (k2_xcmplx_0(X1,np__2)=np__9|r1_xxreal_0(np__9,X1)|v1_abian(X1)|k2_xcmplx_0(X1,np__2)=np__7|v1_abian(k2_xcmplx_0(X1,np__2))|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__7,k2_xcmplx_0(X1,np__2))|~v7_ordinal1(k2_xcmplx_0(X1,np__2)))).#trainneg
cnf(i_0_4866, plain, (k2_xcmplx_0(X1,X2)=np__9|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(np__8,X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_3838, plain, (k2_xcmplx_0(np__2,X1)=np__5|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(np__2,X1))|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__3,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4684, plain, (v1_int_1(k1_nat_1(X1,X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6239, plain, (k2_xcmplx_0(np__2,X1)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(k2_xcmplx_0(np__2,X1)))).#trainneg
cnf(i_0_5497, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__3,esk5_2(X2,X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4730, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(np__2,X1),np__4)|r1_xxreal_0(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)),X1)|v1_abian(X1)|v1_abian(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~v1_xxreal_0(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~v1_xxreal_0(k2_xcmplx_0(np__2,X1))|~v5_chord(k12_glib_001(esk1_0,esk2_0),X2)|~m3_glib_001(k12_glib_001(esk1_0,esk2_0),X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v1_int_1(esk5_2(X2,k12_glib_001(esk1_0,esk2_0)))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5897, plain, (k2_xcmplx_0(np__2,X1)=np__9|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5730, plain, (k2_xcmplx_0(np__2,X1)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4689, plain, (v7_ordinal1(k1_nat_1(X2,X1))|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5965, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__3,k2_xcmplx_0(X2,X3))|v1_abian(k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_6523, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|esk5_2(X2,X1)=np__1|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_7556, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6765, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__9|v1_abian(esk5_2(X1,esk2_0))|~r1_xxreal_0(np__7,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6529, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|esk5_2(X1,X2)=np__5|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_6762, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|esk5_2(X1,X2)=np__1|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6541, plain, (esk5_2(X1,X2)=np__1|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v7_ordinal1(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6506, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|esk5_2(X1,X2)=np__3|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_6521, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6507, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6563, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6569, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6574, plain, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),np__7)|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__9)|~v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6672, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6657, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~r1_xxreal_0(np__3,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6562, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7137, plain, (np__7=X1|r1_xxreal_0(X2,X1)|v1_abian(X1)|np__5=X1|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(np__5,X1))).#trainneg
cnf(i_0_6540, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__3,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(k3_finseq_1(X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6508, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(k3_finseq_1(X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6897, plain, (~v1_abian(k3_finseq_1(X1))|~m1_subset_1(k15_glib_001(X2,X1),k5_numbers)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6513, plain, (r1_xxreal_0(esk5_2(X1,X2),X3)|v1_abian(X3)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6945, plain, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6566, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6518, plain, (esk5_2(X1,X2)=np__9|r1_xxreal_0(esk5_2(X1,X2),np__7)|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__9)|~v1_int_1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6819, negated_conjecture, (r1_xxreal_0(np__1,np__9)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_7023, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7559, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6658, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(k3_finseq_1(X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6596, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_6673, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(k3_finseq_1(X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6636, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7024, plain, (esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6494, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_6573, plain, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7022, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6503, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_7026, plain, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7020, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6808, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk5_2(X1,X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6739, plain, (esk4_2(X1,X2)=k2_xcmplx_0(X3,np__2)|r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v7_ordinal1(X3)|~v1_int_1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6564, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__3)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_7524, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7522, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v1_int_1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6576, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__9)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6807, plain, (k2_xcmplx_0(X1,np__2)=np__1|r1_xxreal_0(np__1,X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6797, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6470, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_6526, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_6537, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_6663, plain, (r1_xxreal_0(esk5_2(X1,X2),np__1)|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_6809, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,esk5_2(X2,X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_7025, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6814, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__9)|r1_xxreal_0(esk5_2(X2,esk2_0),X1)|v1_abian(esk5_2(X2,esk2_0))|v1_abian(X1)|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v1_int_1(X1)|~v1_int_1(esk5_2(X2,esk2_0)))).#trainneg
cnf(i_0_6924, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk5_2(X1,X2))|~r1_xxreal_0(np__3,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7525, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6869, plain, (r1_xxreal_0(np__9,np__5)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_6821, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|k2_xcmplx_0(X2,X3)=np__1|v1_abian(k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v1_int_1(X3)|~v1_int_1(X2))).#trainneg
cnf(i_0_7021, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6709, plain, (r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_6517, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(esk5_2(X1,X2),np__5)|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v1_int_1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6948, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,esk5_2(X2,X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6467, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__9)|r1_xxreal_0(esk5_2(X2,esk2_0),X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6684, plain, (r1_xxreal_0(k3_xcmplx_0(X1,np__1),X2)|v1_abian(X2)|~r1_xxreal_0(k3_xcmplx_0(X1,np__1),np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6472, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),k3_finseq_1(X2))|r1_xxreal_0(k2_xcmplx_0(esk5_2(X3,X2),np__2),np__9)|v1_abian(esk5_2(X3,X2))|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,X2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6469, plain, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),k3_finseq_1(X2))|r1_xxreal_0(esk5_2(X3,X2),X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_7606, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X3,X4)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6925, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__5,esk5_2(X2,X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6927, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__5,esk5_2(X1,X2))|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6598, plain, (v1_int_1(k3_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7546, plain, (X1=np__7|np__5=X1|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7496, plain, (k2_xcmplx_0(X1,X2)=np__1|r1_xxreal_0(np__3,k2_xcmplx_0(X1,X2))|v1_abian(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_6926, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6816, negated_conjecture, (r1_xxreal_0(np__7,np__9)|esk5_2(X1,esk2_0)=np__5|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_6825, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|k2_xcmplx_0(X1,np__2)=np__1|v1_abian(k2_xcmplx_0(X1,np__2))|~v1_xxreal_0(k2_xcmplx_0(X1,np__2))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X1,np__2))|~v1_int_1(X1))).#trainneg
cnf(i_0_6810, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(esk5_2(X1,X2),np__5)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__9,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7144, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6512, plain, (r1_xxreal_0(esk5_2(X1,X2),np__1)|r1_xxreal_0(X3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(X3,np__3)|~v1_int_1(esk5_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7235, plain, (k2_xcmplx_0(X1,np__2)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7145, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7296, negated_conjecture, (r1_xxreal_0(np__9,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(k4_nat_1(np__2,np__4),X2)|~r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7636, plain, (X1=np__1|k2_xcmplx_0(np__2,X1)=np__1|v1_abian(X1)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6491, plain, (v1_xxreal_0(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_6496, plain, (r1_xxreal_0(np__5,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),np__3)|v1_abian(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_6485, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__9)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8139, negated_conjecture, (r1_xxreal_0(np__1,np__9))).#trainneg
cnf(i_0_8103, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|r1_xxreal_0(np__9,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8132, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__7,np__7)|esk5_2(X1,esk2_0)=np__5|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_9177, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__9,np__9)|esk5_2(X1,esk2_0)=np__5|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_8131, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__7,np__7)|esk5_2(X1,esk2_0)=np__5|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_9176, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__9,np__9)|esk5_2(X1,esk2_0)=np__5|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_8494, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__9,np__9)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7921, plain, (esk5_2(X1,X2)=np__3|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|esk5_2(X1,X2)=np__1|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9002, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__7|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_8135, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__5,np__7)|esk5_2(X1,esk2_0)=np__3|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk5_2(X1,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_7980, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|esk5_2(X1,X2)=np__1|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8238, plain, (v1_int_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_8966, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__7,np__9)|esk5_2(X1,esk2_0)=np__3|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk5_2(X1,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_8014, plain, (np__5=X1|np__7=X1|r1_xxreal_0(X2,X1)|v1_abian(X1)|np__3=X1|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(np__3,X1))).#trainneg
cnf(i_0_7977, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|esk5_2(X1,X2)=np__3|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_7967, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7912, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|esk5_2(X1,X2)=np__3|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_7920, plain, (esk5_2(X1,X2)=np__1|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_9022, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|k2_xcmplx_0(esk5_2(X2,X1),np__2)=np__3|v1_abian(esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_8134, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__5,np__7)|esk5_2(X1,esk2_0)=np__3|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_8965, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__7,np__9)|esk5_2(X1,esk2_0)=np__3|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_9055, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_7906, plain, (esk5_2(X1,X2)=np__1|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_8310, negated_conjecture, (r1_xxreal_0(np__9,np__9)|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_9057, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_9056, plain, (esk5_2(X1,X2)=np__3|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_9004, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_9003, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_7969, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|r1_xxreal_0(np__3,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8236, plain, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__7|~r1_xxreal_0(esk4_2(X1,X2),np__9)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7948, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|esk5_2(X1,X2)=np__5|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_8630, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8971, plain, (r1_xxreal_0(np__7,k3_finseq_1(X1))|k2_xcmplx_0(esk5_2(X2,X1),np__2)=np__7|v1_abian(esk5_2(X2,X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__5,esk5_2(X2,X1))|~v7_ordinal1(esk5_2(X2,X1)))).#trainneg
cnf(i_0_8426, plain, (esk5_2(X1,X2)=np__1|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__1,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8716, plain, (r1_xxreal_0(X1,X2)|v1_abian(X2)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8308, plain, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|~r1_xxreal_0(esk5_2(X1,X2),np__9)|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7947, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|esk5_2(X1,X2)=np__5|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_9054, plain, (r1_xxreal_0(X1,X2)|k2_xcmplx_0(X2,np__2)=np__3|v1_abian(X2)|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8084, plain, (np__9=np__7|r1_xxreal_0(np__9,np__5)|r1_xxreal_0(np__7,np__7))).#trainneg
cnf(i_0_8415, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8144, plain, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|~r1_xxreal_0(esk4_2(X1,X2),np__9)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8123, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X2,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),np__9)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8428, plain, (esk5_2(X1,X2)=np__1|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v7_ordinal1(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8637, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_8459, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__7,np__7)|esk4_2(X1,X2)=np__5|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_8376, plain, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|~r1_xxreal_0(esk5_2(X1,X2),np__9)|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8626, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9059, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,X1),np__1)|r1_xxreal_0(X2,esk5_2(esk1_0,X1))|~r1_xxreal_0(X2,np__3)|~v1_xxreal_0(X2)|~v5_chord(X1,esk1_0)|~m3_glib_001(X1,esk1_0))).#trainneg
cnf(i_0_8136, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__1,np__7)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_8322, plain, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9052, plain, (X1=np__3|k2_xcmplx_0(X1,np__2)=np__3|v1_abian(X1)|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8344, plain, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8237, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__7|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,esk4_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8736, plain, (r1_xxreal_0(np__9,k3_finseq_1(X1))|r1_xxreal_0(np__7,np__7)|esk5_2(X2,X1)=np__5|~v7_ordinal1(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__5,esk5_2(X2,X1)))).#trainneg
cnf(i_0_8301, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7944, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__1|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_8129, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(X2,np__2),np__7)|r1_xxreal_0(esk5_2(X1,esk2_0),X2)|v1_abian(esk5_2(X1,esk2_0))|v1_abian(X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X2)|~v1_int_1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_8309, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__7|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,esk5_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8245, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8133, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__3,np__7)|esk5_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7958, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|r1_xxreal_0(np__3,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8417, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8145, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__7|esk4_2(X2,X3)=np__5|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,esk4_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8732, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(np__7,np__7)|esk5_2(X1,X2)=np__5|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2)))).#trainneg
cnf(i_0_8177, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(np__7,np__9)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9082, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|r1_xxreal_0(np__3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(np__1,esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9083, plain, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|r1_xxreal_0(np__3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8178, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~r1_xxreal_0(np__3,np__9)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9020, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__3,np__7)|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__3|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_8462, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__3,np__7)|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8224, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8269, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9014, plain, (k2_xcmplx_0(X1,np__2)=np__1|k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_9200, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9806, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10539, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(X3,esk4_2(X1,X2))|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~r1_xxreal_0(X3,np__9)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10533, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9702, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(X1,esk2_0))|r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9227, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9666, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(esk4_2(X1,X2),np__1)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10393, plain, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),np__7)|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9808, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(esk4_2(X1,X2),np__1)|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9244, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9622, plain, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_10397, plain, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9703, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10367, plain, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9393, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__9,np__9)|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9657, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10811, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9212, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_9799, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10493, plain, (r1_xxreal_0(np__3,np__7)|r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_9941, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10276, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__1,X1)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10759, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|esk5_2(X1,X2)=np__7|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9225, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9224, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10769, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__5|esk5_2(X3,X2)=np__7|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__9))).#trainneg
cnf(i_0_10389, plain, (X1=esk4_2(X2,X3)|esk4_2(X2,X3)=np__3|r1_xxreal_0(np__5,esk4_2(X2,X3))|v1_abian(X1)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10364, plain, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_10363, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__7)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_9205, plain, (esk5_2(X1,X2)=np__5|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10275, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9639, plain, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__7)|r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_10512, plain, (r1_xxreal_0(np__3,np__7)|r1_xxreal_0(np__9,np__5)|r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_9619, plain, (esk5_2(X1,X2)=np__3|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|esk5_2(X1,X2)=np__1|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9226, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10189, plain, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9665, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__1)|v1_abian(X3)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9215, plain, (r1_xxreal_0(np__7,esk5_2(X1,X2))|r1_xxreal_0(np__5,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9807, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__1)|v1_abian(X3)|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10391, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_10285, plain, (esk5_2(X1,X2)=np__1|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_10191, plain, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10510, plain, (r1_xxreal_0(np__3,np__7)|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_9589, plain, (esk5_2(X1,X2)=np__1|esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_10116, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|r1_xxreal_0(np__1,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10219, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(esk5_2(X1,X2),np__3)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__7,esk5_2(X1,X2))|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10356, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__7)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_9232, plain, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(np__5,np__5)|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9220, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(X3,np__7)|~v7_ordinal1(esk5_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9940, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9214, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(np__5,np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9231, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__5,np__7)|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10190, plain, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10875, plain, (r1_xxreal_0(X1,X2)|k2_xcmplx_0(X2,np__2)=np__3|k2_xcmplx_0(X2,np__2)=np__1|v1_abian(X2)|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_10540, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(X3,esk4_2(X1,X2))|esk4_2(X1,X2)=np__3|~r1_xxreal_0(X3,np__9)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,X2)))).#trainneg
cnf(i_0_9975, negated_conjecture, (r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__7,esk5_2(X1,esk2_0))|~v1_xreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9915, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_9888, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__7,np__7)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10284, plain, (esk5_2(X1,X2)=np__1|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_9242, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__9,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10357, plain, (r1_xxreal_0(np__7,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_9834, plain, (r1_xxreal_0(np__9,X1)|v1_abian(X1)|r1_xxreal_0(np__1,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9812, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_10117, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|r1_xxreal_0(np__1,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10790, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__3|r1_xxreal_0(np__5,esk4_2(X2,X3))|~v7_ordinal1(esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10093, plain, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,np__7)|~r1_xxreal_0(esk4_2(X1,X2),np__9)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9843, plain, (r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__1,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9656, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_9898, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(X3,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_9213, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__5,np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9805, plain, (esk4_2(X1,X2)=k2_xcmplx_0(X3,np__2)|r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9701, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__5,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10793, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~v7_ordinal1(esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10944, negated_conjecture, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__7))).#trainneg
cnf(i_0_11076, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(np__9,np__9)|esk4_2(esk1_0,esk2_0)!=np__1)).#trainneg
cnf(i_0_11094, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_11010, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_10886, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10994, negated_conjecture, (r1_xxreal_0(np__9,np__9)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_11149, plain, (r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__7,np__7))).#trainneg
cnf(i_0_11057, plain, (r1_xxreal_0(np__9,X1)|v1_abian(X1)|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__5,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11151, plain, (r1_xxreal_0(np__1,np__5)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_10910, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10911, plain, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__9)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10917, plain, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11381, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__5,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11268, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11093, plain, (r1_xxreal_0(np__1,np__5)|r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_11091, plain, (r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_10885, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__9)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10992, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10947, negated_conjecture, (r1_xxreal_0(np__5,np__1)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_11022, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11666, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11757, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__7,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11370, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__7,np__7)|r1_xxreal_0(np__5,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11585, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__7,np__7)|r1_xxreal_0(np__5,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11677, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11036, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11262, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10909, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11374, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(np__7,np__7)|r1_xxreal_0(np__5,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11822, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_11946, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_11369, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__5,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11295, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_11567, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(X2,esk2_0)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(esk5_2(X2,esk2_0),np__7)|v1_abian(esk5_2(X2,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10899, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__9)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11135, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11600, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_11392, negated_conjecture, (r1_xxreal_0(np__7,np__7)|r1_xxreal_0(np__5,np__5)|esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__9,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11016, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11136, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__3,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11216, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_12291, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__9|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_12279, negated_conjecture, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_12592, plain, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),np__7)|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13587, negated_conjecture, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__5))).#trainneg
cnf(i_0_13158, negated_conjecture, (r1_xxreal_0(np__9,np__9)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_13169, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_12242, plain, (np__9=np__1|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_12848, negated_conjecture, (r1_xxreal_0(np__5,np__5)|r1_xxreal_0(np__7,np__7)|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12192, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_12168, negated_conjecture, (np__7=np__5|r1_xxreal_0(np__5,np__5))).#trainneg
cnf(i_0_12513, negated_conjecture, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1))).#trainneg
cnf(i_0_12424, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_12531, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5))).#trainneg
cnf(i_0_13215, negated_conjecture, (r1_xxreal_0(np__5,np__5)|r1_xxreal_0(np__9,np__9)|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12157, negated_conjecture, (np__9=np__7|r1_xxreal_0(np__7,np__7))).#trainneg
cnf(i_0_12560, plain, (esk4_2(X1,X2)=np__3|esk4_2(X3,X4)=esk4_2(X1,X2)|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(esk4_2(X3,X4))|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_12558, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__7)|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_13019, plain, (esk4_2(X1,X2)=np__3|esk4_2(X3,X4)=esk4_2(X1,X2)|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(esk4_2(X3,X4))|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__9,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_13168, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__9)|~r1_xxreal_0(k4_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_12692, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__7,np__7)|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12160, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__7,np__7)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12927, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13073, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(X1,esk2_0))|r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12500, plain, (r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__1,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12540, negated_conjecture, (r1_xxreal_0(np__5,X1)|v1_abian(X1)|r1_xxreal_0(np__1,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13165, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_12415, plain, (r1_xxreal_0(np__5,X1)|v1_abian(X1)|r1_xxreal_0(np__1,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13291, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13074, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12928, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11983, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__1|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_13655, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_11748, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__5,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11546, plain, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,np__5)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10945, negated_conjecture, (r1_xxreal_0(np__5,np__3)|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_13586, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13638, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__3|r1_xxreal_0(np__3,np__5)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_11252, negated_conjecture, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(np__9,esk4_2(esk1_0,X3))|~v5_chord(X3,esk1_0)|~v5_chord(X2,X1)|~m3_glib_001(X3,esk1_0)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13592, negated_conjecture, (r1_xxreal_0(np__7,np__9)|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__3|r1_xxreal_0(np__3,np__5)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_12512, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1))).#trainneg
cnf(i_0_11368, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__5,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11550, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__7,np__7)|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12180, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_12551, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__7)|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_13350, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__3|r1_xxreal_0(np__5,np__7)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_12422, negated_conjecture, (np__9!=np__7|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_13455, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(esk5_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__5))).#trainneg
cnf(i_0_11970, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__7|esk5_2(X2,esk2_0)=np__5|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_12588, plain, (X1=esk4_2(X2,X3)|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|v1_abian(X1)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11383, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(np__5,np__5)|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13304, negated_conjecture, (r1_xxreal_0(np__7,np__9)|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__3|r1_xxreal_0(np__5,np__7)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_11213, plain, (esk5_2(X1,X2)=k3_finseq_1(X2)|~v1_xxreal_0(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__5))).#trainneg
cnf(i_0_10921, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__7,np__7)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9663, plain, (esk4_2(X1,X2)=k2_xcmplx_0(X3,np__2)|r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9394, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__9,np__9)|~r1_xxreal_0(np__7,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12202, negated_conjecture, (r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_10594, plain, (esk4_2(X1,X2)=np__3|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13591, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__5))).#trainneg
cnf(i_0_13636, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__3|r1_xxreal_0(np__3,np__5)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_10900, negated_conjecture, (r1_xxreal_0(k1_card_1(X1),np__9)|~r1_xxreal_0(k1_card_1(X1),np__1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_11497, plain, (esk5_2(X1,X2)=np__1|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_9904, negated_conjecture, (r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__7)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2)))).#trainneg
cnf(i_0_10369, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__3|r1_xxreal_0(np__5,esk4_2(X2,X3))|~v7_ordinal1(esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_13164, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_10950, negated_conjecture, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_13047, plain, (X1=esk4_2(X2,X3)|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__7|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__9,esk4_2(X2,X3))|v1_abian(X1)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11240, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__5))).#trainneg
cnf(i_0_13186, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__9,np__1)|r1_xxreal_0(np__5,np__5)|r1_xxreal_0(np__7,np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11005, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__9)|~r1_xxreal_0(k1_nat_1(X1,X2),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_12174, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_9831, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11239, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(X3,np__1)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__5))).#trainneg
cnf(i_0_10898, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__9)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_14495, negated_conjecture, (r1_xxreal_0(np__5,np__5)|r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_13760, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|X1=np__9|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_13842, negated_conjecture, (r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__5,np__5)|esk4_2(esk1_0,esk2_0)!=np__1)).#trainneg
cnf(i_0_13830, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_14655, negated_conjecture, (esk4_2(esk1_0,esk2_0)!=np__1|np__9!=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_13816, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|esk4_2(X1,X2)=np__9|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_15192, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__5)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13780, negated_conjecture, (np__7=X1|X1=np__9|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13756, negated_conjecture, (np__7=X1|X1=np__9|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15191, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__5)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13829, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13831, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_15194, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),np__7)|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_10415, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|esk4_2(X2,X3)=np__5|esk4_2(X2,X3)=np__3|~r1_xxreal_0(X1,X4)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X4,np__7)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_13656, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_9976, negated_conjecture, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__3,esk5_2(X1,esk2_0))|~v1_xreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9659, plain, (r1_xxreal_0(X1,np__3)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_15386, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__9,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9593, plain, (r1_xxreal_0(k1_card_1(X1),np__5)|r1_xxreal_0(np__7,k1_card_1(X1))|v1_abian(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_11602, plain, (esk5_2(X1,X2)=k3_finseq_1(X2)|esk5_2(X1,X2)=np__5|~v1_xxreal_0(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__7))).#trainneg
cnf(i_0_12278, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10770, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(X3,np__1)|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__9))).#trainneg
cnf(i_0_11008, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__9)|~r1_xxreal_0(k4_nat_1(X1,X2),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11599, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_13072, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__5,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11681, negated_conjecture, (r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__5,np__7)|esk5_2(X1,esk2_0)=np__1|r1_xxreal_0(np__3,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11670, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__9,np__1)|r1_xxreal_0(np__5,np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13174, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9801, plain, (r1_xxreal_0(X1,np__3)|esk4_2(X2,X3)=np__7|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__9,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10374, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(esk4_2(X3,X4),X1)|v1_abian(X2)|v1_abian(X1)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(k2_xcmplx_0(X2,np__2))|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v1_xxreal_0(X1)|~v1_xxreal_0(esk4_2(X3,X4))|~v7_ordinal1(k2_xcmplx_0(X2,np__2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_12504, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_13447, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_2(X2,X3),np__5)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_7970, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|r1_xxreal_0(np__3,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11006, negated_conjecture, (r1_xxreal_0(k1_card_1(X1),np__9)|~r1_xxreal_0(k1_card_1(X1),np__5)|~v1_finset_1(X1))).#trainneg
cnf(i_0_14884, plain, (esk5_2(X1,X2)=k3_finseq_1(X2)|~r1_xxreal_0(k3_finseq_1(X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_11384, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(np__5,np__5)|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12926, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__9,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11000, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11934, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__1,np__5)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_12549, plain, (r1_xxreal_0(esk4_2(X1,X2),k3_finseq_1(X3))|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|v1_abian(esk5_2(X4,X3))|~v1_xxreal_0(esk5_2(X4,X3))|~v1_xxreal_0(k3_finseq_1(X3))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X3,X4)|~m3_glib_001(X3,X4)|~v1_glib_000(X4)|~v1_finset_1(X4)|~v1_funct_1(X4)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X4)|~v7_ordinal1(esk5_2(X4,X3))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11372, plain, (r1_xxreal_0(np__7,k3_finseq_1(X1))|r1_xxreal_0(np__5,np__5)|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11160, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_13833, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12555, plain, (r1_xxreal_0(np__9,k3_finseq_1(X1))|r1_xxreal_0(esk4_2(X2,X3),np__7)|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|v1_abian(esk5_2(X4,X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_xxreal_0(esk4_2(X2,X3))|~v5_chord(X1,X4)|~m3_glib_001(X1,X4)|~v1_glib_000(X4)|~v1_finset_1(X4)|~v1_funct_1(X4)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X4)|~v7_ordinal1(esk5_2(X4,X1))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11238, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__5))).#trainneg
cnf(i_0_13977, negated_conjecture, (np__9=np__5|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_11513, plain, (np__9=np__7|r1_xxreal_0(np__9,np__5)|r1_xxreal_0(np__7,np__1)|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_10903, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__9)|~r1_xxreal_0(k4_nat_1(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12394, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|r1_xxreal_0(np__7,np__3)|esk5_2(X2,X1)=np__5|~v7_ordinal1(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11774, negated_conjecture, (r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__1,np__7)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_11373, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__7,np__1)|r1_xxreal_0(np__5,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9061, plain, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__7|r1_xxreal_0(np__7,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9811, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_10081, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|v1_abian(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_11150, plain, (r1_xxreal_0(np__1,np__3)|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_8442, plain, (r1_xxreal_0(np__9,esk5_2(X1,X2))|r1_xxreal_0(np__3,np__7)|esk5_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11824, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__1)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(esk5_2(X3,esk2_0))|~m1_subset_1(esk5_2(X3,esk2_0),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8707, plain, (r1_xxreal_0(np__2,X1)|v1_abian(X1)|~r1_xxreal_0(np__2,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11889, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__1,np__5)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_11092, plain, (r1_xxreal_0(np__1,np__3)|r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_12194, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_10764, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__5|esk5_2(X2,X3)=np__7|~r1_xxreal_0(X1,X4)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X4,np__9)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_16736, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15593, negated_conjecture, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|X1=np__9|X1=np__7|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_16588, plain, (r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_16174, negated_conjecture, (X1=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16890, negated_conjecture, (r1_xxreal_0(np__1,np__5)|r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_16750, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__9|v1_abian(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X2,np__5))).#trainneg
cnf(i_0_16777, negated_conjecture, (r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__9,np__9))).#trainneg
cnf(i_0_17123, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__1,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16429, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,X4)|v1_abian(X3)|v1_abian(X4)|~v1_int_1(X4)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_15577, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16428, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,X4)|v1_abian(X3)|v1_abian(X4)|~v1_int_1(X4)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_15720, negated_conjecture, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_15710, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|esk5_2(X2,X1)=np__5|esk5_2(X2,X1)=np__7|esk5_2(X2,X1)=np__9|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_15598, negated_conjecture, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_16648, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_15591, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__9|X2=np__7|v1_abian(X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_16349, negated_conjecture, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_15695, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__5|r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__3)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_15578, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__5)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15696, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__5|r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__3)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16833, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_15942, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|X3=np__9|v1_abian(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_17144, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,esk5_2(X1,X2))|X3=np__9|X3=np__7|v1_abian(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_15703, negated_conjecture, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__5|esk5_2(X2,X3)=np__7|esk5_2(X2,X3)=np__9|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_15606, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_16421, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16424, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__1)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16388, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__7,np__3)|r1_xxreal_0(np__1,np__5)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_16170, negated_conjecture, (r1_xxreal_0(k4_nat_1(X2,X1),np__9)|~r1_xxreal_0(k4_nat_1(X2,X1),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_16610, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(np__5,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_16431, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(X2,esk2_0)|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),np__9)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(esk5_2(X2,esk2_0),np__7)|v1_abian(esk5_2(X2,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v1_int_1(esk5_2(X2,esk2_0))|~m1_subset_1(esk5_2(X2,esk2_0),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16422, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__3)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16423, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__5)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16694, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk5_2(X2,esk2_0),X1)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__9|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_15964, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__5,np__7)|esk5_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_15938, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|r1_xxreal_0(X4,X3)|v1_abian(X3)|v1_abian(X4)|~r1_xxreal_0(k6_xcmplx_0(X4,np__2),np__7)|~v1_xxreal_0(k6_xcmplx_0(X4,np__2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_16352, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_16480, negated_conjecture, (esk5_2(X1,esk2_0)=np__1|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__1,np__7)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16485, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(X3,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_17138, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,esk5_2(X1,X2))|r1_xxreal_0(X4,X3)|v1_abian(X3)|v1_abian(X4)|~r1_xxreal_0(k6_xcmplx_0(X4,np__2),np__9)|~v1_xxreal_0(k6_xcmplx_0(X4,np__2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_16897, plain, (r1_xxreal_0(np__3,X1)|v1_abian(X1)|r1_xxreal_0(np__1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16565, negated_conjecture, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__7,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X3)))).#trainneg
cnf(i_0_17237, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_17261, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_17234, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_17236, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_17264, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_17537, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|X1=np__9|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_17202, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__7,np__9)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17344, plain, (esk5_2(X1,X2)=np__1|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17238, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__5,np__1)|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_17212, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__9,np__9)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17450, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_17186, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__7|X2=np__9|X2=np__5|v1_abian(X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_17521, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_17421, negated_conjecture, (esk4_2(esk1_0,X1)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,X1))|r1_xxreal_0(esk4_2(esk1_0,X1),X2)|r1_xxreal_0(X2,X3)|v1_abian(X3)|v1_abian(X2)|~v7_ordinal1(X3)|~v1_int_1(X2)|~v1_xxreal_0(X2)|~v5_chord(X1,esk1_0)|~m3_glib_001(X1,esk1_0))).#trainneg
cnf(i_0_17216, negated_conjecture, (esk5_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk5_2(X1,X2),np__7)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16447, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__3)|r1_xxreal_0(esk4_2(X1,X2),np__7)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_16776, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__5,k3_finseq_1(X2))|r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_16663, negated_conjecture, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X1,X2))|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17233, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__9,np__7)|np__7=np__5|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_17348, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__5|esk5_2(X2,X3)=np__7|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_16403, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__5,np__3)|r1_xxreal_0(np__1,np__5)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_16301, plain, (r1_xxreal_0(np__7,esk5_2(X1,X2))|r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16399, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__3,np__3)|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17690, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__9|esk5_2(X3,X2)=np__7|esk5_2(X3,X2)=np__5|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_16482, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_16537, negated_conjecture, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__3,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X3)))).#trainneg
cnf(i_0_16446, negated_conjecture, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__5)|r1_xxreal_0(esk4_2(X1,X2),np__7)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_15698, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|r1_xxreal_0(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),np__5)|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_16612, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_16843, plain, (np__7=np__3|r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_16014, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(np__3,np__3)|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_15989, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(np__9,np__3)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_16294, negated_conjecture, (esk5_2(X1,esk2_0)=np__3|r1_xxreal_0(np__3,np__7)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_18720, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_19529, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(np__7,esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19027, negated_conjecture, (r1_xxreal_0(np__9,np__1)|r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_19246, negated_conjecture, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|X1=np__9|X1=np__7|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_19499, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19052, negated_conjecture, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_19683, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__7|esk5_2(X3,X2)=np__5|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_20047, negated_conjecture, (np__9=np__3|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__9,np__1)|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_19545, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk4_2(X1,X2))|esk4_2(X1,X2)=np__5|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,X2)))).#trainneg
cnf(i_0_18776, negated_conjecture, (np__9=np__5|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_19096, negated_conjecture, (np__9=np__5|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_18770, negated_conjecture, (r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__5,np__1)|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_19065, negated_conjecture, (r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__5,np__1)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_19790, negated_conjecture, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_16309, plain, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__1)|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17197, plain, (np__7=np__1|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_19258, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,np__9)|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19548, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk4_2(X1,X2))|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16182, negated_conjecture, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__1,np__7)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2)))).#trainneg
cnf(i_0_16594, plain, (r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_18286, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(X3,esk4_2(X1,X2))|r1_xxreal_0(np__9,X3)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17184, negated_conjecture, (r1_xxreal_0(X1,np__3)|X2=np__7|X2=np__9|X2=np__5|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_17530, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|X3=np__9|v1_abian(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_16751, plain, (np__9=np__5|np__9=np__7|r1_xxreal_0(np__3,np__7)|r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_17268, negated_conjecture, (np__9=np__7|r1_xxreal_0(np__1,np__5)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_16617, negated_conjecture, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2)))).#trainneg
cnf(i_0_16473, negated_conjecture, (esk5_2(X1,esk2_0)=np__3|r1_xxreal_0(esk5_2(X1,esk2_0),np__1)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__3,np__7)|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_18132, negated_conjecture, (esk5_2(X1,esk2_0)=np__3|r1_xxreal_0(esk5_2(X1,esk2_0),np__1)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__3,np__5)|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17655, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),np__7)|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_19035, negated_conjecture, (np__9=np__1|r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_16176, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_19250, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__7,esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19351, negated_conjecture, (r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__5,np__1)|r1_xxreal_0(np__3,np__7)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_17653, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__5)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19992, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_19257, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,np__9)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17457, negated_conjecture, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16312, negated_conjecture, (r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__1)|esk5_2(X1,esk2_0)=np__1|r1_xxreal_0(np__3,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16850, negated_conjecture, (r1_xxreal_0(np__9,np__9)|esk5_2(X1,esk2_0)=np__1|r1_xxreal_0(np__3,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17593, negated_conjecture, (esk4_2(X1,X2)=X3|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|X3=np__9|v1_abian(X3)|~r1_xxreal_0(esk4_2(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_15981, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__3)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_18956, negated_conjecture, (r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_18159, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__3,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17699, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__5|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_18559, plain, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(X3,esk5_2(X1,X2))|r1_xxreal_0(np__7,X3)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_18144, negated_conjecture, (X1=esk5_2(X2,esk2_0)|esk5_2(X2,esk2_0)=np__9|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_15681, negated_conjecture, (np__9=np__1|np__7=np__1|r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_15715, negated_conjecture, (r1_xxreal_0(np__5,k3_finseq_1(X1))|r1_xxreal_0(np__3,np__3)|esk5_2(X2,X1)=np__5|esk5_2(X2,X1)=np__7|esk5_2(X2,X1)=np__9|~v7_ordinal1(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_16472, negated_conjecture, (esk5_2(X1,esk2_0)=np__3|r1_xxreal_0(esk5_2(X1,esk2_0),np__1)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__3,np__7)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16366, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__3)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_18141, negated_conjecture, (esk5_2(X1,esk2_0)=np__1|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(np__1,np__5)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_20232, negated_conjecture, (r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_20196, negated_conjecture, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__9))).#trainneg
cnf(i_0_20272, negated_conjecture, (np__7=np__3|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_20217, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__5,np__3))).#trainneg
cnf(i_0_20660, negated_conjecture, (np__9=np__7|r1_xxreal_0(np__1,np__5)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_20274, negated_conjecture, (np__5=np__1|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_20301, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(esk4_2(X2,X3))|~m1_subset_1(esk4_2(X2,X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_20297, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_21570, negated_conjecture, (esk4_2(esk1_0,X1)=np__7|esk4_2(esk1_0,X1)=np__5|r1_xxreal_0(np__3,esk4_2(X2,X3))|r1_xxreal_0(esk4_2(X2,X3),np__3)|r1_xxreal_0(np__9,esk4_2(esk1_0,X1))|~v5_chord(X1,esk1_0)|~v5_chord(X3,X2)|~m3_glib_001(X1,esk1_0)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_21241, negated_conjecture, (np__9=np__1|r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_20451, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_20275, negated_conjecture, (np__5=np__1|r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_20199, negated_conjecture, (np__5=np__1|r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_20497, negated_conjecture, (esk4_2(X1,X2)=np__3|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__7|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_20215, negated_conjecture, (np__9=np__5|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_20559, plain, (esk4_2(X1,X2)=X3|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__9,X3)|v1_abian(X3)|~r1_xxreal_0(esk4_2(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk4_2(X1,X2))|~v1_int_1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_21454, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X2,esk2_0)=esk5_2(X1,esk2_0)|r1_xxreal_0(esk5_2(X2,esk2_0),np__5)|v1_abian(esk5_2(X2,esk2_0))|esk5_2(X2,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_19498, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,X2)|v1_abian(X1)|v1_abian(X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_20563, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__5|esk4_2(X2,X3)=np__1|r1_xxreal_0(np__9,X4)|v1_abian(X4)|~r1_xxreal_0(X1,X4)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~v1_int_1(X4)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_16673, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(np__9,np__9)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16568, negated_conjecture, (r1_xxreal_0(X1,np__9)|esk5_2(X2,X3)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X3)))).#trainneg
cnf(i_0_16323, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~m1_subset_1(esk5_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_17206, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(np__5,np__7)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16498, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|v1_abian(X2)|r1_xxreal_0(np__1,np__7)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16714, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__9|~v7_ordinal1(X1)|~v7_ordinal1(esk4_2(X2,X3))|~r1_xxreal_0(esk4_2(X2,X3),np__5)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_20205, negated_conjecture, (r1_xxreal_0(np__5,np__1)|r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_16023, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(np__1,np__3)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_17533, negated_conjecture, (r1_xxreal_0(X1,np__7)|X2=np__9|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_17087, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk5_2(X1,X2))|esk4_2(X3,X4)=np__1|~r1_xxreal_0(esk4_2(X3,X4),np__9)|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_16298, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__1)|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11003, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8166, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~r1_xxreal_0(np__9,np__9)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15589, negated_conjecture, (r1_xxreal_0(X1,np__5)|X2=np__9|X2=np__7|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_16677, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_12510, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1))).#trainneg
cnf(i_0_16408, plain, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__1,np__3)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_20190, negated_conjecture, (r1_xxreal_0(np__5,np__1)|r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_15457, negated_conjecture, (r1_xxreal_0(k1_nat_1(X2,X1),np__9)|~r1_xxreal_0(k1_nat_1(X2,X1),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_10806, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__7)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16629, negated_conjecture, (r1_xxreal_0(np__4,np__9)|~r1_xxreal_0(np__4,np__5))).#trainneg
cnf(i_0_11350, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(X2,np__1)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_16427, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,X4)|v1_abian(X3)|v1_abian(X4)|~v1_int_1(X4)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k2_xcmplx_0(X4,np__2),k5_numbers))).#trainneg
cnf(i_0_18305, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__3,k3_finseq_1(X3))|~v7_ordinal1(esk5_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk5_2(X2,X3),k1_numbers)|~m1_subset_1(k3_finseq_1(X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_9060, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|r1_xxreal_0(np__7,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17529, negated_conjecture, (np__7=X1|X1=np__9|v1_abian(X1)|~r1_xxreal_0(np__7,X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12091, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__1)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(esk5_2(X3,esk2_0))|~m1_subset_1(esk5_2(X3,esk2_0),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9655, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_13453, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__5))).#trainneg
cnf(i_0_21977, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(X1,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_21808, negated_conjecture, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_23440, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_23410, negated_conjecture, (np__9=np__1|np__7=np__1|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_22850, negated_conjecture, (np__9=np__1|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_21923, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=esk4_2(X3,X4)|r1_xxreal_0(np__9,esk4_2(X3,X4))|v1_abian(esk4_2(X3,X4))|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__5|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__9,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v1_int_1(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_23469, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_21685, negated_conjecture, (np__9=np__1|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_22260, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__1,esk4_2(X1,X2))|r1_xxreal_0(np__9,X3)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_22838, negated_conjecture, (np__9=np__1|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_22003, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,X2)|v1_abian(X2)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_23471, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_21698, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|r1_xxreal_0(np__5,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22591, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__5)|v1_abian(X3)|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_22340, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_21832, negated_conjecture, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__5|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__7,esk4_2(X3,X4))|~r1_xxreal_0(esk4_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_23049, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_22281, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(X3,esk4_2(X1,X2))|r1_xxreal_0(np__9,k6_xcmplx_0(X4,np__2))|v1_abian(k6_xcmplx_0(X4,np__2))|r1_xxreal_0(X4,X3)|v1_abian(X3)|v1_abian(X4)|~v1_int_1(k6_xcmplx_0(X4,np__2))|~v1_xxreal_0(k6_xcmplx_0(X4,np__2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_22592, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__5)|v1_abian(X3)|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_21880, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_23050, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_22899, negated_conjecture, (esk4_2(X1,X2)=np__1|esk5_2(X3,X4)=np__5|esk5_2(X3,X4)=np__7|r1_xxreal_0(np__3,esk5_2(X3,X4))|~v5_chord(X4,X3)|~v5_chord(X2,X1)|~m3_glib_001(X4,X3)|~m3_glib_001(X2,X1)|~v1_glib_000(X3)|~v1_glib_000(X1)|~v1_finset_1(X3)|~v1_finset_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_23048, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk5_2(X1,X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_22964, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|r1_xxreal_0(np__9,np__1)|r1_xxreal_0(np__3,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23327, plain, (esk5_2(X1,X2)=np__7|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_22537, negated_conjecture, (r1_xxreal_0(np__1,np__7)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2)))).#trainneg
cnf(i_0_22934, negated_conjecture, (r1_xxreal_0(np__3,np__9)|~r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_21727, negated_conjecture, (np__7=np__1|r1_xxreal_0(np__1,np__5)|np__9=np__7)).#trainneg
cnf(i_0_22900, negated_conjecture, (esk4_2(X1,X2)=np__1|esk5_2(X3,X4)=np__5|esk5_2(X3,X4)=np__7|r1_xxreal_0(np__3,esk5_2(X3,X4))|~v5_chord(X4,X3)|~v5_chord(X2,X1)|~m3_glib_001(X4,X3)|~m3_glib_001(X2,X1)|~v1_glib_000(X3)|~v1_glib_000(X1)|~v1_finset_1(X3)|~v1_finset_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5))).#trainneg
cnf(i_0_21836, negated_conjecture, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X4,X5)=np__5|esk4_2(X2,X3)=np__1|esk4_2(X2,X3)=np__5|esk4_2(X2,X3)=np__3|r1_xxreal_0(np__7,esk4_2(X4,X5))|~r1_xxreal_0(X1,esk4_2(X4,X5))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X4,X5))|~v1_xxreal_0(X1)|~m1_subset_1(esk4_2(X4,X5),k1_numbers)|~v5_chord(X3,X2)|~v5_chord(X5,X4)|~m3_glib_001(X3,X2)|~m3_glib_001(X5,X4)|~v1_glib_000(X2)|~v1_glib_000(X4)|~v1_finset_1(X2)|~v1_finset_1(X4)|~v1_funct_1(X2)|~v1_funct_1(X4)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X4))).#trainneg
cnf(i_0_23119, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(X3,np__1)|r1_xxreal_0(esk5_2(X4,esk2_0),X3)|v1_abian(X3)|r1_xxreal_0(esk4_2(X1,X2),np__7)|v1_abian(esk4_2(X1,X2))|~v7_ordinal1(X3)|~v1_xxreal_0(esk5_2(X4,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(esk2_0,X4)|~m3_glib_001(esk2_0,X4)|~v1_glib_000(X4)|~v1_finset_1(X4)|~v1_funct_1(X4)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X4))).#trainneg
cnf(i_0_23333, plain, (r1_xxreal_0(X1,esk5_2(X2,X3))|k2_xcmplx_0(esk5_2(X2,X3),np__2)=np__7|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_17453, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_11626, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__5|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__7))).#trainneg
cnf(i_0_21518, negated_conjecture, (np__5=np__1|np__7=np__1|np__9=np__1|~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_11589, negated_conjecture, (esk5_2(X1,esk2_0)=np__3|r1_xxreal_0(np__3,np__7)|esk5_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_21410, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(X3,np__1)|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,X3)|v1_abian(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X3))).#trainneg
cnf(i_0_14771, negated_conjecture, (r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__1,esk5_2(X1,esk2_0))|~v1_xreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11803, plain, (np__9=np__7|r1_xxreal_0(np__9,np__5)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_20655, plain, (np__7=np__1|np__9=np__7|np__9=np__5|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_15924, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__1,esk4_2(X1,X2))|~r1_xxreal_0(np__1,np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13758, negated_conjecture, (r1_xxreal_0(X1,np__7)|X2=np__9|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21879, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_8419, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(esk5_2(X1,X2),np__5)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17575, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__5|r1_xxreal_0(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),np__3)|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_11740, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|esk5_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_8609, plain, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(np__9,np__3)|esk4_2(X1,X2)=np__7|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_7968, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__3,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24780, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_24190, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__7,esk4_2(X3,X4))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_24998, plain, (r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_24598, negated_conjecture, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_24327, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X3,X2),np__2)=np__7|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__7))).#trainneg
cnf(i_0_24191, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__3,esk5_2(X1,X2))|esk4_2(X3,X4)=np__7|r1_xxreal_0(np__9,esk4_2(X3,X4))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_24381, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_24401, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_24013, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24400, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_24413, negated_conjecture, (np__9=np__3|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|esk5_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_24369, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_24640, plain, (np__9=np__5|r1_xxreal_0(np__9,np__3)|np__9=np__7|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_24652, plain, (np__7=np__3|r1_xxreal_0(np__3,np__5)|np__9=np__5|np__9=np__7|np__7=np__1)).#trainneg
cnf(i_0_23778, negated_conjecture, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__9|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_24412, negated_conjecture, (np__9=np__3|r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__9,np__7)|esk5_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_23852, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_24631, plain, (np__9=np__5|np__9=np__7|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_24060, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk5_2(X1,X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24397, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,esk5_2(X1,X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24129, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__7|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__9,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24782, negated_conjecture, (r1_xxreal_0(X1,np__9)|esk5_2(X2,X3)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~m1_subset_1(esk5_2(X2,X3),k1_numbers)|~m1_subset_1(k3_finseq_1(X3),k1_numbers))).#trainneg
cnf(i_0_24368, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_23939, negated_conjecture, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__9|r1_xxreal_0(np__9,esk4_2(X3,X4))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_24131, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_23618, negated_conjecture, (esk4_2(esk1_0,X1)=np__3|esk4_2(esk1_0,X1)=np__5|esk4_2(esk1_0,X1)=np__1|r1_xxreal_0(np__1,esk4_2(esk1_0,X1))|r1_xxreal_0(np__9,X2)|v1_abian(X2)|~v7_ordinal1(X2)|~v5_chord(X1,esk1_0)|~m3_glib_001(X1,esk1_0))).#trainneg
cnf(i_0_12283, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_23278, negated_conjecture, (X1=esk5_2(X2,X3)|r1_xxreal_0(np__3,k3_finseq_1(X3))|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v1_int_1(X1)|~m1_subset_1(esk5_2(X2,X3),k1_numbers)|~m1_subset_1(k3_finseq_1(X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_22365, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(k2_xcmplx_0(X3,np__2),np__9)|r1_xxreal_0(esk5_2(X1,X2),X3)|v1_abian(esk5_2(X1,X2))|v1_abian(X3)|~v7_ordinal1(k3_finseq_1(X2))|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X3)|~v1_int_1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_21740, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__5)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,esk2_0))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12022, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=k2_xcmplx_0(esk5_2(X3,esk2_0),np__2)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12185, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10807, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|r1_xxreal_0(np__3,np__5)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7929, plain, (esk5_2(X1,X2)=k3_finseq_1(X2)|esk5_2(X1,X2)=np__1|~v1_xxreal_0(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__3))).#trainneg
cnf(i_0_8130, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(X2,np__2),np__7)|r1_xxreal_0(esk5_2(X1,esk2_0),X2)|v1_abian(esk5_2(X1,esk2_0))|v1_abian(X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_13836, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8316, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|r1_xxreal_0(np__9,np__3)|esk5_2(X2,X1)=np__7|~v7_ordinal1(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__7,esk5_2(X2,X1)))).#trainneg
cnf(i_0_22952, negated_conjecture, (r1_xxreal_0(np__2,np__9)|~r1_xxreal_0(np__2,np__5))).#trainneg
cnf(i_0_15972, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__3)|r1_xxreal_0(esk4_2(X1,X2),np__7)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_11684, plain, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(esk4_2(X1,X2),np__9)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_18524, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X1)|v1_abian(X1)|r1_xxreal_0(esk5_2(X3,esk2_0),X2)|v1_abian(esk5_2(X3,esk2_0))|v1_abian(X2)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~m1_subset_1(X1,k1_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(X2)|~v1_int_1(esk5_2(X3,esk2_0)))).#trainneg
cnf(i_0_12530, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5))).#trainneg
cnf(i_0_25348, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_25647, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|k2_xcmplx_0(esk5_2(X2,X1),np__2)=np__7|esk5_2(X2,X1)=np__9|esk5_2(X2,X1)=np__7|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_25969, plain, (esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_26803, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X2)|v1_abian(X3)|~v1_int_1(X3)|~v1_int_1(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25619, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_25133, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_26465, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__1|~r1_xxreal_0(np__9,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_26804, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X2)|v1_abian(X3)|~v1_int_1(X3)|~v1_int_1(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_26764, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_26797, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|r1_xxreal_0(np__7,X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~v1_int_1(X2)|~v1_xxreal_0(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_25929, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(k2_xcmplx_0(X3,np__2),np__9)|r1_xxreal_0(esk5_2(X1,X2),X3)|v1_abian(esk5_2(X1,X2))|v1_abian(X3)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X3)|~v1_int_1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_26798, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__1)|r1_xxreal_0(np__3,X2)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~v1_int_1(X2)|~v1_xxreal_0(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_25132, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_25089, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|esk5_2(X1,X2)=np__1|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_26468, plain, (esk5_2(X1,X2)=k3_finseq_1(X2)|esk5_2(X1,X2)=np__1|~r1_xxreal_0(k3_finseq_1(X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_25152, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_25131, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk5_2(X1,X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_25145, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__9,k3_finseq_1(X2))|esk5_2(X1,X2)=np__1|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_25696, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_26298, negated_conjecture, (X1=esk5_2(X2,X3)|r1_xxreal_0(np__3,k3_finseq_1(X3))|r1_xxreal_0(X1,np__7)|v1_abian(X1)|esk5_2(X2,X3)=np__7|esk5_2(X2,X3)=np__5|r1_xxreal_0(np__7,X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~m1_subset_1(esk5_2(X2,X3),k1_numbers)|~m1_subset_1(k3_finseq_1(X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_23352, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|r1_xxreal_0(np__7,np__9)|k2_xcmplx_0(esk5_2(X2,X1),np__2)=np__7|~v7_ordinal1(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_12502, negated_conjecture, (esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16439, negated_conjecture, (r1_xxreal_0(np__7,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__5)|r1_xxreal_0(esk5_2(X1,X2),np__7)|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_13166, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_16628, negated_conjecture, (r1_xxreal_0(np__0,np__9)|~r1_xxreal_0(np__0,np__5))).#trainneg
cnf(i_0_25175, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|esk5_2(X1,X2)=np__1|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_10354, plain, (r1_xxreal_0(esk4_2(X1,X2),k3_finseq_1(X3))|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(esk5_2(X4,X3))|~v1_xxreal_0(esk5_2(X4,X3))|~v1_xxreal_0(k3_finseq_1(X3))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X3,X4)|~m3_glib_001(X3,X4)|~v1_glib_000(X4)|~v1_finset_1(X4)|~v1_funct_1(X4)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X4)|~v7_ordinal1(esk5_2(X4,X3))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_22254, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__1,esk4_2(X1,X2))|r1_xxreal_0(np__9,np__9)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19195, negated_conjecture, (np__7=np__3|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_9694, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__3|esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_17857, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(X2,np__1)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_25182, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_8306, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_10372, plain, (esk4_2(X1,X2)=k2_xcmplx_0(X3,np__2)|r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(k2_xcmplx_0(X3,np__2))|~v1_int_1(X3)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(esk4_2(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X3,np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16437, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__1)|r1_xxreal_0(esk5_2(X1,X2),np__7)|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_12590, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_13940, plain, (np__7=np__5|r1_xxreal_0(np__7,np__3)|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_11890, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__1,np__5)|v1_abian(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_8318, plain, (r1_xxreal_0(np__7,k3_finseq_1(X1))|r1_xxreal_0(np__9,np__5)|esk5_2(X2,X1)=np__7|~v7_ordinal1(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__7,esk5_2(X2,X1)))).#trainneg
cnf(i_0_9591, plain, (r1_xxreal_0(k1_card_1(X1),np__7)|r1_xxreal_0(np__9,k1_card_1(X1))|v1_abian(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_22432, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__5,np__7)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_18616, plain, (np__3=np__1|r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_8639, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_25017, plain, (r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__7,np__5)|r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_13782, negated_conjecture, (r1_xxreal_0(X1,np__7)|X2=np__9|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_19584, negated_conjecture, (np__7=np__5|r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_8552, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_11437, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__5|esk4_2(X3,X4)=np__7|r1_xxreal_0(np__9,esk4_2(X3,X4))|r1_xxreal_0(X2,np__1)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_21924, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=esk4_2(X3,X4)|r1_xxreal_0(np__9,esk4_2(X3,X4))|v1_abian(esk4_2(X3,X4))|esk4_2(X3,X4)=np__1|esk4_2(X3,X4)=np__5|esk4_2(X3,X4)=np__3|esk4_2(X1,X2)=np__9|v1_abian(esk4_2(X1,X2))|~v1_int_1(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_21403, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__7|esk5_2(X2,esk2_0)=np__5|r1_xxreal_0(np__7,X1)|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9590, plain, (r1_xxreal_0(np__3,k1_card_1(X1))|r1_xxreal_0(k1_card_1(X1),np__1)|v1_abian(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_19789, negated_conjecture, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_26119, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11827, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk4_2(X2,X3),np__7)|esk4_2(X2,X3)=np__3|r1_xxreal_0(np__5,esk4_2(X2,X3))|r1_xxreal_0(esk5_2(X1,esk2_0),np__1)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk4_2(X2,X3))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_17213, negated_conjecture, (esk5_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk5_2(X1,X2),np__7)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__7,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27544, negated_conjecture, (esk4_2(X1,X2)=np__3|r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28544, negated_conjecture, (r1_xxreal_0(np__9,np__1)|r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_27346, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|esk5_2(X1,X2)=np__9|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_27444, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_27326, negated_conjecture, (r1_xxreal_0(np__7,k3_finseq_1(X1))|esk5_2(X2,X1)=np__9|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_27700, plain, (r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__3,np__5)|~r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_27443, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk5_2(X1,X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27353, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_28570, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_27350, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|esk5_2(X1,X2)=np__9|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_28547, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_28551, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,np__9))).#trainneg
cnf(i_0_28231, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_27445, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_28568, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_27354, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_27352, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_28528, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|X1=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28569, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_28571, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_26833, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_28395, plain, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__3,np__5)|~r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_27981, negated_conjecture, (r1_xxreal_0(np__1,np__5)|r1_xxreal_0(np__9,np__9)|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27025, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_28574, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_27310, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(np__9,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28233, negated_conjecture, (r1_xxreal_0(X1,np__9)|esk5_2(X2,X3)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~m1_subset_1(k3_finseq_1(X3),k1_numbers))).#trainneg
cnf(i_0_28596, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk5_2(X2,X3)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,esk5_2(X2,X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_27340, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v7_ordinal1(esk5_2(X1,X2))|~r1_xxreal_0(np__9,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27505, negated_conjecture, (r1_xxreal_0(np__3,np__9)|~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_27030, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(esk4_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_27272, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(np__7,np__9)|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27349, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(X3,esk5_2(X1,X2))|esk5_2(X1,X2)=np__9|~r1_xxreal_0(X3,np__9)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_28538, plain, (k1_card_1(X1)=np__1|v1_abian(k1_card_1(X1))|r1_xxreal_0(np__3,k1_card_1(X1))|~v7_ordinal1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_27197, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__1)|r1_xxreal_0(esk4_2(X2,X3),np__7)|r1_xxreal_0(np__3,esk4_2(X2,X3))|v1_abian(esk4_2(X2,X3))|~v1_int_1(esk4_2(X2,X3))|~m1_subset_1(esk4_2(X2,X3),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_23287, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(np__3,k3_finseq_1(X2))|r1_xxreal_0(np__9,np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_11766, plain, (esk5_2(X1,X2)=np__1|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_26802, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X2)|v1_abian(X3)|~v1_int_1(X3)|~v1_int_1(X2)|~v1_xxreal_0(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k2_xcmplx_0(X2,np__2)))).#trainneg
cnf(i_0_16496, negated_conjecture, (r1_xxreal_0(np__3,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_25076, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_27064, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|r1_xxreal_0(X2,np__7)|r1_xxreal_0(np__7,X2)|v1_abian(X2)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_24643, plain, (np__7=np__3|np__9=np__5|np__9=np__7|np__7=np__1|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_27373, negated_conjecture, (r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,np__3)|~r1_xxreal_0(np__1,esk4_2(X1,esk2_0))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_xreal_0(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_22540, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(np__1,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17180, negated_conjecture, (np__3=X1|X1=np__7|X1=np__9|X1=np__5|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12193, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_10926, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__7)|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_25689, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,esk5_2(X1,X2))|esk5_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_9278, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__9,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28642, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__5|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~v1_int_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_29760, negated_conjecture, (r1_xxreal_0(np__3,np__7))).#trainneg
cnf(i_0_28641, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__5|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_28696, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28687, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28700, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,esk2_0),np__2)=np__7|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28643, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,esk2_0))|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_28653, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__3|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_29793, negated_conjecture, (np__9=np__7|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|esk5_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_29065, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|X1=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_29791, negated_conjecture, (np__9=np__7|r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__9,np__7)|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_30060, negated_conjecture, (np__9=np__7|r1_xxreal_0(np__7,np__9)|r1_xxreal_0(np__9,np__7)|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_29995, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29985, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29129, negated_conjecture, (r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__1,np__5)|esk4_2(esk1_0,esk2_0)!=np__1)).#trainneg
cnf(i_0_28655, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__3|esk4_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__3|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29991, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__5,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_29999, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28772, negated_conjecture, (r1_xxreal_0(np__7,np__9)|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__3|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_28656, negated_conjecture, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29763, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__7)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_29041, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(np__7,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28636, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_31841, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__3,esk5_2(X1,X2))|esk4_2(X3,X4)=np__7|r1_xxreal_0(np__5,esk4_2(X3,X4))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_30339, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_30340, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_31742, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__7)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_30982, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=np__3|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v1_int_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_31821, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X3,X4),esk5_2(X1,X2))|esk4_2(X3,X4)=np__7|r1_xxreal_0(np__5,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_29581, negated_conjecture, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_28816, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_29216, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk5_2(X1,X2))|~r1_xxreal_0(np__9,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29795, negated_conjecture, (np__7=np__1|r1_xxreal_0(np__1,np__7)|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_28755, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_31753, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__7)|esk4_2(X3,X4)=np__9|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_31840, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__5)|v1_abian(X3)|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29130, negated_conjecture, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__5|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__7,esk4_2(X3,X4))|~r1_xxreal_0(esk4_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_30070, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_30342, negated_conjecture, (np__7=np__3|r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_31837, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29776, plain, (np__7=np__1|np__9=np__7|np__9=np__5|np__7=np__3|np__7=np__5|~r1_xxreal_0(np__5,np__3))).#trainneg
cnf(i_0_28644, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__1,np__5)|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_28677, negated_conjecture, (np__9=np__1|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_29273, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_28676, negated_conjecture, (np__9=np__1|r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_28942, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_29757, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,np__2),X2)|r1_xxreal_0(esk5_2(X3,esk2_0),X1)|r1_xxreal_0(X2,np__7)|v1_abian(X2)|v1_abian(X1)|~v1_int_1(X2)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(k2_xcmplx_0(X1,np__2),k5_numbers))).#trainneg
cnf(i_0_30288, negated_conjecture, (esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|k1_funct_1(esk2_0,esk5_2(X1,esk2_0))!=k1_funct_1(esk2_0,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29620, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(np__7,np__9)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_29382, negated_conjecture, (esk5_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__7)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_30981, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=np__3|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_32024, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)!=np__9|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_29448, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|r1_xxreal_0(np__7,np__9)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_28686, negated_conjecture, (r1_xxreal_0(np__9,X1)|v1_abian(X1)|r1_xxreal_0(np__3,np__9)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_29293, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(k2_xcmplx_0(X3,np__2),np__9)|r1_xxreal_0(esk5_2(X1,X2),X3)|v1_abian(esk5_2(X1,X2))|v1_abian(X3)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X3)|~v1_int_1(esk5_2(X1,X2)))).#trainneg
cnf(i_0_29066, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_11669, plain, (r1_xxreal_0(np__9,k3_finseq_1(X1))|r1_xxreal_0(np__5,np__7)|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_14370, negated_conjecture, (r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__7,np__9)|~r1_xxreal_0(np__1,esk5_2(X1,esk2_0))|~v1_xreal_0(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_22542, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__1,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_32215, negated_conjecture, (r1_xxreal_0(np__3,np__3)|r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_32292, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(np__9,np__9)|r1_xxreal_0(np__7,np__9))).#trainneg
cnf(i_0_32197, negated_conjecture, (r1_xxreal_0(np__3,np__9)|r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_33530, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__3|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__3|v1_abian(esk4_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_32291, negated_conjecture, (r1_xxreal_0(np__7,np__9)|esk5_2(esk1_0,esk2_0)!=np__7)).#trainneg
cnf(i_0_33624, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__9,np__7)|r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_33533, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__3|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_33370, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__3|esk4_2(X1,esk2_0)=np__1|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=k1_funct_1(esk2_0,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_32170, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__7)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_32217, negated_conjecture, (r1_xxreal_0(np__3,np__1)|r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_32172, negated_conjecture, (r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_33623, negated_conjecture, (np__9=np__5|r1_xxreal_0(np__9,np__3)|r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_32236, negated_conjecture, (esk5_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__7|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__5,esk4_2(X3,X4))|~r1_xxreal_0(esk5_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_33156, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|r1_xxreal_0(np__1,np__5)|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_33536, negated_conjecture, (np__9=np__7|np__9=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_32998, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)!=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_32168, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__1|esk5_2(esk1_0,X3)=np__5|esk5_2(esk1_0,X3)=np__7|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(np__3,esk5_2(esk1_0,X3))|~v5_chord(X3,esk1_0)|~v5_chord(X2,X1)|~m3_glib_001(X3,esk1_0)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_32416, negated_conjecture, (r1_xxreal_0(np__7,np__9)|k1_funct_1(esk2_0,esk5_2(X1,esk2_0))!=k1_funct_1(esk2_0,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_33856, negated_conjecture, (np__5=np__1|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_33611, negated_conjecture, (np__9=np__5|np__9=np__7|r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_33314, negated_conjecture, (np__9=np__3|np__9=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_32194, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_32191, negated_conjecture, (r1_xxreal_0(np__1,np__7)|~r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_33540, negated_conjecture, (np__9=np__7|np__9=np__5|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_33383, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|r1_xxreal_0(np__7,np__9)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_33008, negated_conjecture, (np__5=np__3|r1_xxreal_0(np__3,np__3)|np__7=np__3)).#trainneg
cnf(i_0_33357, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|k1_funct_1(esk2_0,esk5_2(X1,esk2_0))!=k1_funct_1(esk2_0,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_33853, negated_conjecture, (np__9=np__5|r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_33141, plain, (np__7=np__5|np__7=np__3|np__9=np__5|np__9=np__7|np__7=np__1|r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_28988, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__5|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2)))).#trainneg
cnf(i_0_15277, negated_conjecture, (r1_xxreal_0(np__0,np__9)|~r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_22259, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__1,esk4_2(X1,X2))|r1_xxreal_0(np__9,np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12249, plain, (np__7=np__5|r1_xxreal_0(np__1,np__5)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_28661, negated_conjecture, (r1_xxreal_0(X1,np__9)|esk4_2(X2,X3)=np__3|r1_xxreal_0(np__5,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_12387, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__5|~v7_ordinal1(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_16676, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_20010, plain, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(np__7,np__3)|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__3,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_26136, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__3,k3_finseq_1(X3))|~v7_ordinal1(esk5_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k3_finseq_1(X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8408, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16642, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__5)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11973, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X1)|v1_abian(X1)|r1_xxreal_0(esk5_2(X3,esk2_0),X2)|v1_abian(esk5_2(X3,esk2_0))|v1_abian(X2)|~v7_ordinal1(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(esk5_2(X3,esk2_0))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_17775, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_34755, negated_conjecture, (esk4_2(X1,X2)=np__7|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_35043, negated_conjecture, (np__7=np__1|np__5=np__1|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_34763, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_34765, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_35098, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_34769, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3))).#trainneg
cnf(i_0_34761, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__3)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_35052, negated_conjecture, (np__7=np__1|r1_xxreal_0(np__1,np__5)|np__5=np__1)).#trainneg
cnf(i_0_34035, negated_conjecture, (np__5=np__3|r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_35864, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_35354, plain, (np__7=np__5|r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_35920, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_34803, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__7,np__5)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_34591, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(X1,esk2_0))=np__7|esk4_2(X1,esk2_0)=np__1|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=np__3|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_35205, negated_conjecture, (np__9=np__7|np__9=np__5|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_35100, negated_conjecture, (np__7=np__1|~r1_xxreal_0(np__7,np__1)|~r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_35921, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_34867, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(np__1,np__5)|esk4_2(esk1_0,esk2_0)!=np__1)).#trainneg
cnf(i_0_35051, negated_conjecture, (np__5=np__1|np__7=np__1|~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_35896, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_35054, negated_conjecture, (np__9=np__3|np__9=np__7|np__9=np__5|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_34726, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|~r1_xxreal_0(np__9,np__7)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_35893, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|np__9=np__5)).#trainneg
cnf(i_0_35899, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_35069, negated_conjecture, (np__9=np__1|np__9=np__3|np__7=np__3|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_35867, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_35309, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_35062, negated_conjecture, (np__7=np__3|np__9=np__3|np__9=np__7|np__9=np__5|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_35095, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_35149, negated_conjecture, (np__7=np__5|np__9=np__5|np__9=np__7|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_35865, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_35142, negated_conjecture, (np__9=np__5|np__9=np__7|~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_35387, negated_conjecture, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__9,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_35860, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_35691, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__5)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_35393, negated_conjecture, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__9|r1_xxreal_0(np__9,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_35110, negated_conjecture, (np__7=np__1|r1_xxreal_0(np__1,np__5)|~r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_34804, negated_conjecture, (np__9=np__7|r1_xxreal_0(np__1,np__5)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_35396, negated_conjecture, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__9|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_34263, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(X1,esk2_0))=np__7|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__3|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=np__3|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27063, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|r1_xxreal_0(X2,np__7)|r1_xxreal_0(np__7,X2)|v1_abian(X2)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_21998, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_22984, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(X3,np__5)|v1_abian(esk5_2(X1,X2))|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(X3,esk5_2(X1,X2))|~v7_ordinal1(esk5_2(X1,X2))|~v1_xxreal_0(X3)|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29266, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_12183, negated_conjecture, (esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_23704, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__9|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__5)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_24143, negated_conjecture, (np__9=np__3|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_18527, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|esk5_2(X2,esk2_0)=np__5|v1_abian(esk5_2(X2,esk2_0))|~v1_int_1(X1)|~v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(np__5,esk5_2(X2,esk2_0))|~v1_int_1(esk5_2(X2,esk2_0))|~m1_subset_1(esk5_2(X2,esk2_0),k1_numbers))).#trainneg
cnf(i_0_36070, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_36034, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_35930, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5))).#trainneg
cnf(i_0_36023, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=k1_funct_1(esk2_0,np__3))).#trainneg
cnf(i_0_36557, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_36378, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_37029, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_36545, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_36366, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_35935, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__3),esk6_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_36093, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_36037, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__3),esk6_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_36036, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__3),k1_funct_1(esk2_0,np__1),esk6_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_35933, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),k1_nat_1(np__3,np__2)))).#trainneg
cnf(i_0_36107, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X1,X2))|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_36556, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_36377, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_35929, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|np__3!=np__1)).#trainneg
cnf(i_0_37607, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(np__3,esk4_2(X1,X2))|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_35934, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__3),k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_36092, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_36146, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_35937, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),k1_nat_1(np__2,np__3)))).#trainneg
cnf(i_0_36094, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36943, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__5|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v1_xreal_0(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_37163, negated_conjecture, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__5)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_36552, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_36219, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,esk5_2(X1,X2))|X3=np__7|v1_abian(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_36200, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_36551, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_36373, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38068, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9)).#trainneg
cnf(i_0_38150, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38104, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38434, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38435, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38395, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38394, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38087, negated_conjecture, (np__9=np__5|np__9=np__7|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_38474, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__9)|~r1_xxreal_0(k1_nat_1(X1,X2),np__7)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38273, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38248, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38212, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__7)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38090, negated_conjecture, (np__9=np__5|np__9=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_38096, negated_conjecture, (np__9=np__5|np__9=np__7|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_38211, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__7)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_38473, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__9)|~r1_xxreal_0(k1_nat_1(X1,X2),np__7)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_38461, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_38744, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38477, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__9)|~r1_xxreal_0(k4_nat_1(X1,X2),np__7)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38215, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__7)|~r1_xxreal_0(k4_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38223, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_38472, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38485, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__7)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_38089, negated_conjecture, (np__9=np__5|np__9=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_38199, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__7)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_38220, negated_conjecture, (r1_xxreal_0(np__1,np__7))).#trainneg
cnf(i_0_38210, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38469, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__7)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38480, negated_conjecture, (r1_xxreal_0(np__3,np__9))).#trainneg
cnf(i_0_38198, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__7)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_38484, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__7)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38207, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_38222, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38317, negated_conjecture, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_40283, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_40629, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_15924, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_40735, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_40854, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_40758, negated_conjecture, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1))).#trainneg
cnf(i_0_40614, negated_conjecture, (esk4_2(X1,X2)=np__7|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_40746, negated_conjecture, (esk5_2(X1,X2)=np__7|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_40801, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__7)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_40739, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_40265, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_40624, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_40255, negated_conjecture, (r1_xxreal_0(k4_nat_1(X2,X1),np__7)|~r1_xxreal_0(k4_nat_1(X2,X1),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_28395, plain, (r1_xxreal_0(np__3,np__5)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_40628, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_40395, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_40276, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8178, plain, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_40239, negated_conjecture, (r1_xxreal_0(k4_nat_1(X2,X1),np__9)|~r1_xxreal_0(k4_nat_1(X2,X1),np__7)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_40298, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__7))).#trainneg
cnf(i_0_39887, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X2,X1),np__7)|~r1_xxreal_0(k2_xcmplx_0(X2,X1),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_40134, negated_conjecture, (esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_38313, negated_conjecture, (r1_xxreal_0(np__7,np__5)|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_40154, negated_conjecture, (esk4_2(X1,X2)=np__1|esk5_2(X3,X4)=np__5|esk5_2(X3,X4)=np__7|r1_xxreal_0(np__3,esk5_2(X3,X4))|~v5_chord(X4,X3)|~v5_chord(X2,X1)|~m3_glib_001(X4,X3)|~m3_glib_001(X2,X1)|~v1_glib_000(X3)|~v1_glib_000(X1)|~v1_finset_1(X3)|~v1_finset_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7))).#trainneg
cnf(i_0_39057, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__5|esk5_2(X3,X2)=np__7|X1=np__7|v1_abian(X1)|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_38743, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38463, negated_conjecture, (r1_xxreal_0(np__9,np__9)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_36084, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_38746, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_36222, negated_conjecture, (r1_xxreal_0(X1,np__5)|X2=np__7|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_40634, negated_conjecture, (np__7=np__1|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_38471, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_39889, negated_conjecture, (np__9=np__7|np__9=np__5|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_38980, negated_conjecture, (r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,np__5)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_38852, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~m1_subset_1(esk5_2(X1,X2),k5_numbers))).#trainneg
cnf(i_0_38849, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_36372, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_40280, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_40757, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1))).#trainneg
cnf(i_0_17081, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk5_2(X1,X2))|esk5_2(X3,X4)=np__1|~r1_xxreal_0(esk5_2(X3,X4),np__9)|~v1_xxreal_0(esk5_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_12138, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(X2,np__1)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_32183, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__7)|~r1_xxreal_0(k1_nat_1(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_35712, negated_conjecture, (r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X2)|v1_abian(X3)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8223, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_27341, negated_conjecture, (esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v7_ordinal1(esk5_2(X1,X2))|~r1_xxreal_0(np__9,np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_40643, negated_conjecture, (np__7=np__1|r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_15998, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X3,np__2),np__3)|r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(esk4_2(X1,X2))|v1_abian(X3)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X3)|~v1_int_1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_41037, negated_conjecture, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|X1=np__7|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_8166, plain, (r1_xxreal_0(np__9,X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_41145, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_41095, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41067, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41023, negated_conjecture, (np__9=np__7|np__9=np__5|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_41108, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__9)|~r1_xxreal_0(k1_nat_1(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41073, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__5|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_41119, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_41851, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_41800, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__9,esk4_2(X3,X4))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_41111, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__9)|~r1_xxreal_0(k4_nat_1(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41801, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X3,X4))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_41173, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41400, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__7|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41535, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41850, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_42396, negated_conjecture, (r1_xxreal_0(k1_nat_1(X2,X1),np__7)|~r1_xxreal_0(k1_nat_1(X2,X1),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41235, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk5_2(X1,X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41972, negated_conjecture, (np__7=np__1|np__7=np__3|~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_41118, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_41107, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__9)|~r1_xxreal_0(k1_nat_1(X1,X2),np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_41106, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41407, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|esk4_2(X1,X2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_41746, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_40969, negated_conjecture, (np__9=np__7|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_41725, negated_conjecture, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41613, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__3|esk4_2(X1,esk2_0)=np__3|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41658, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_41237, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_41175, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_41103, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41174, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_41609, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__3|v1_abian(esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v1_int_1(esk4_2(X1,esk2_0))|~m1_subset_1(esk4_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_41656, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__9,k3_finseq_1(X2))|~r1_xxreal_0(np__9,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_41281, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_41122, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~m1_subset_1(esk5_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_41690, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_41846, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_42395, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X2,X1),np__7)|~r1_xxreal_0(k2_xcmplx_0(X2,X1),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41236, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_41511, plain, (r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(np__3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41285, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_41033, negated_conjecture, (r1_xxreal_0(X1,np__5)|X2=np__7|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_41112, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__9)|~r1_xxreal_0(k4_nat_1(X1,X2),np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_41274, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_41650, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_42033, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__3,np__5)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_42440, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_42730, negated_conjecture, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|X1=np__7|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_42662, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__7|X2=np__5|v1_abian(X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_43264, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_42907, negated_conjecture, (np__9=np__5|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_44043, plain, (esk4_2(X1,X2)=np__7|r1_xxreal_0(esk4_2(X1,X2),np__5)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_43108, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|esk5_2(X2,X1)=np__5|esk5_2(X2,X1)=np__7|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_43686, negated_conjecture, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_43921, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_42735, negated_conjecture, (esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_42904, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_43430, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__5,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_43098, negated_conjecture, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__5|esk5_2(X2,X3)=np__7|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_43337, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__7),k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_43341, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_43126, negated_conjecture, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_43190, negated_conjecture, (np__9=np__7|np__7=np__3|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_44186, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_43993, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X2,X1),np__7)|~r1_xxreal_0(k2_xcmplx_0(X2,X1),np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_42802, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X2,X1),np__9)|~r1_xxreal_0(k2_xcmplx_0(X2,X1),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_44204, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_43910, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_43897, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_43917, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_43896, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_42660, negated_conjecture, (r1_xxreal_0(X1,np__3)|X2=np__7|X2=np__5|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_42740, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__3)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_43834, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,np__5)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_43251, negated_conjecture, (r1_xxreal_0(np__0,np__9)|~r1_xxreal_0(np__0,np__3))).#trainneg
cnf(i_0_42905, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_42896, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_43069, negated_conjecture, (r1_xxreal_0(np__1,np__5)|r1_xxreal_0(np__3,np__5)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_43898, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_41652, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_41897, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_41649, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_41651, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_41510, plain, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(np__3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_43916, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_41681, negated_conjecture, (r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_41747, negated_conjecture, (esk4_2(X1,X2)=np__9|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5))).#trainneg
cnf(i_0_41109, negated_conjecture, (r1_xxreal_0(k1_card_1(X1),np__9)|~r1_xxreal_0(k1_card_1(X1),np__3)|~v1_finset_1(X1))).#trainneg
cnf(i_0_42484, negated_conjecture, (np__5=np__1|r1_xxreal_0(np__1,np__3)|np__7=np__1)).#trainneg
cnf(i_0_41243, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_41280, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_43812, negated_conjecture, (X1=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_43818, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_42111, negated_conjecture, (v1_abian(esk4_2(X1,esk2_0))|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__5|~v1_int_1(esk4_2(X1,esk2_0))|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_41380, plain, (np__7=np__5|r1_xxreal_0(np__3,np__5)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_44015, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_44324, negated_conjecture, (np__5=np__3|~r1_xxreal_0(np__5,np__3))).#trainneg
cnf(i_0_44841, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_45387, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_45433, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_45699, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(np__9,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_44366, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__7|X2=np__5|v1_abian(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,np__1)|~v1_int_1(X2))).#trainneg
cnf(i_0_44894, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_45446, negated_conjecture, (np__9=np__5|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_46030, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__3|v1_abian(esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_44462, negated_conjecture, (X1=np__9|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_45592, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_44327, negated_conjecture, (r1_xxreal_0(np__1,np__5)|~r1_xxreal_0(np__5,np__3))).#trainneg
cnf(i_0_46033, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__1|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_45245, negated_conjecture, (r1_xxreal_0(np__3,np__5)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_45567, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_44489, plain, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_44694, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_45378, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_44339, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_44867, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__1,np__5)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_45067, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|k2_xcmplx_0(esk5_2(X2,X1),np__2)=np__7|esk5_2(X2,X1)=np__7|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~m1_subset_1(esk5_2(X2,X1),k1_numbers))).#trainneg
cnf(i_0_45671, negated_conjecture, (np__9=np__5|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~r1_xxreal_0(np__9,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_44836, negated_conjecture, (r1_xxreal_0(np__2,np__9)|~r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_44636, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__7|esk5_2(X3,X2)=np__5|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_45583, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_45611, plain, (esk5_2(X1,X2)=np__7|r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__3,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_45656, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_45655, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v1_xreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_45558, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_45401, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_44471, negated_conjecture, (np__9=np__5|np__7=np__5|~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_45657, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_44996, negated_conjecture, (X1=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_45719, negated_conjecture, (np__3=np__1|np__7=np__1|np__5=np__1|~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_45390, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_44782, negated_conjecture, (r1_xxreal_0(np__2,np__7)|~r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_44698, negated_conjecture, (np__7=np__3|np__9=np__7|r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_41104, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_41608, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__3|v1_abian(esk4_2(X1,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,esk2_0))|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_44035, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__3,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_43726, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_43225, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk5_2(X1,X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_46044, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_44476, plain, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X3,X4),esk5_2(X1,X2))|esk4_2(X3,X4)=np__7|r1_xxreal_0(np__3,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_43347, negated_conjecture, (r1_xxreal_0(X1,np__3)|esk4_2(X2,X3)=np__7|esk4_2(X2,X3)=np__5|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_41397, negated_conjecture, (r1_xxreal_0(X1,np__5)|esk4_2(X2,X3)=np__7|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_45794, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_46321, negated_conjecture, (~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_46817, negated_conjecture, (r1_xxreal_0(np__1,np__5))).#trainneg
cnf(i_0_46832, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_46756, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,np__5)|v1_abian(X1)|X1=np__9|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_46834, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48139, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|k2_xcmplx_0(esk5_2(X2,X1),np__2)=np__7|esk5_2(X2,X1)=np__7|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_47734, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_24143, negated_conjecture, (np__9=np__3|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_46303, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_46557, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_46753, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__9|v1_abian(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_46865, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_46295, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__7),esk6_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_48392, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,np__7)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_46328, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__7,np__5)|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_46629, negated_conjecture, (r1_xxreal_0(np__9,np__7)|np__9=np__5|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_46741, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|esk4_2(X1,X2)=np__9|v1_abian(esk4_2(X1,X2))|~v7_ordinal1(X3)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_int_1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_46245, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_44701, negated_conjecture, (np__7=np__3|np__9=np__7|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_45002, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_45793, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v1_xreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_41272, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__7)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_46618, negated_conjecture, (np__7=np__5|np__9=np__5|~r1_xxreal_0(np__7,np__5)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_43367, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk4_2(X2,X3)=np__7|esk4_2(X2,X3)=np__5|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~r1_xxreal_0(esk4_2(X2,X3),np__1))).#trainneg
cnf(i_0_44909, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_43396, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(esk1_0,esk2_0))|np__9=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_43635, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_43354, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|~v7_ordinal1(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_44288, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__5,esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk4_2(X2,X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_45767, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(X3)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_44272, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__3,esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk4_2(X2,X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_45795, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_46585, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__7),esk6_2(esk1_0,esk2_0))|np__9=np__5|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_46246, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_44725, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_42421, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__3)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_45176, negated_conjecture, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__3,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_47672, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__7)|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_46886, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__3,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_45192, negated_conjecture, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__5,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_44513, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_48598, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2)=np__7)).#trainneg
cnf(i_0_48462, negated_conjecture, (r1_xxreal_0(np__7,np__5)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_48422, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_48444, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__5)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_48584, negated_conjecture, (np__9=np__3|np__7=np__3|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_48593, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_48436, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_48468, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_48601, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_48414, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_49259, negated_conjecture, (np__9=np__5|np__7=np__5|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_48457, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__5)|~r1_xxreal_0(k1_nat_1(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_48449, negated_conjecture, (r1_xxreal_0(np__0,np__5)|~r1_xxreal_0(np__0,np__3))).#trainneg
cnf(i_0_48456, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__5)|~r1_xxreal_0(k1_nat_1(X1,X2),np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_48459, negated_conjecture, (r1_xxreal_0(np__2,np__5)|~r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_48987, negated_conjecture, (np__7=np__5|np__9=np__7|np__7=np__3|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_49919, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_48495, negated_conjecture, (r1_xxreal_0(X1,np__9)|X2=np__9|X2=np__7|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_48494, negated_conjecture, (r1_xxreal_0(X1,np__5)|X2=np__9|X2=np__7|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__3)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_48911, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__7,np__5)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_48587, negated_conjecture, (r1_xxreal_0(np__5,np__7)|np__7=np__3|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_49157, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_49160, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_50228, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_49310, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_49156, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_48446, negated_conjecture, (r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_49378, plain, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_48467, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_48425, negated_conjecture, (np__5=np__1|r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_49311, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_48434, plain, (k2_xcmplx_0(X1,np__3)=k2_xcmplx_0(X2,np__3)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50616, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_50617, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_52071, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50593, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__5)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_50871, negated_conjecture, (~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_52093, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5))).#trainneg
cnf(i_0_50606, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__5)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50605, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__5)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_52024, negated_conjecture, (np__9=np__5|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_50604, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50618, negated_conjecture, (esk4_2(X1,X2)=np__5|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_52062, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50609, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__5)|~r1_xxreal_0(k4_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50603, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_50589, negated_conjecture, (np__9=np__7|np__9=np__5|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_50601, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_52061, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_51085, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X2,X1),np__5)|~r1_xxreal_0(k2_xcmplx_0(X2,X1),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50592, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__5)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_52065, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__5)|~r1_xxreal_0(k4_nat_1(X1,X2),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50546, negated_conjecture, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|v1_abian(X1)|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1))).#trainneg
cnf(i_0_50552, negated_conjecture, (r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|v1_abian(X1)|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~v1_int_1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_50611, negated_conjecture, (r1_xxreal_0(np__7,np__5)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_52073, negated_conjecture, (np__3=np__1|np__5=np__1|~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_50595, negated_conjecture, (r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_50610, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__5)|~r1_xxreal_0(k4_nat_1(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_50853, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_52098, plain, (k2_xcmplx_0(esk5_2(X2,X1),np__3)=k2_xcmplx_0(k3_finseq_1(X1),np__3)|~r1_xxreal_0(k3_finseq_1(X1),esk5_2(X2,X1))|~v1_xreal_0(k3_finseq_1(X1))|~v1_xreal_0(esk5_2(X2,X1))|~m1_subset_1(esk5_2(X2,X1),k5_numbers)|~m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_50602, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__5)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_50607, negated_conjecture, (r1_xxreal_0(k1_card_1(X1),np__5)|~r1_xxreal_0(k1_card_1(X1),np__1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_50637, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3))).#trainneg
cnf(i_0_50624, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__3)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_50549, negated_conjecture, (r1_xxreal_0(np__7,k3_finseq_1(X1))|r1_xxreal_0(esk5_2(X2,X1),esk5_2(esk1_0,esk2_0))|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_50427, negated_conjecture, (k2_xcmplx_0(np__2,esk5_2(esk1_0,esk2_0))=np__7|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_51632, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X2)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_49159, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_48689, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_48688, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__9,k1_nat_1(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_50559, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|v1_abian(esk5_2(esk1_0,esk2_0))|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,np__7)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_49158, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_50887, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_49842, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_49364, plain, (r1_xxreal_0(np__7,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_43239, negated_conjecture, (np__7=np__3|np__7=np__1|np__3=np__1|~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_44717, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_52530, negated_conjecture, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_52481, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_52585, negated_conjecture, (esk4_2(X1,X2)=np__5|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_53074, negated_conjecture, (esk5_2(X1,X2)=np__5|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__1))).#trainneg
cnf(i_0_52596, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|v1_abian(esk4_2(X1,X2))|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_52603, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_53317, negated_conjecture, (np__9=np__5|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_53118, negated_conjecture, (r1_xxreal_0(k4_nat_1(X2,X1),np__5)|~r1_xxreal_0(k4_nat_1(X2,X1),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_52768, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk4_2(X1,X2))|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_52617, negated_conjecture, (np__7=np__3|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_52526, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_52492, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_52593, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_54222, negated_conjecture, (X1=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_53533, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__3,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_52999, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X2,X1),np__5)|~r1_xxreal_0(k3_xcmplx_0(X2,X1),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_52468, negated_conjecture, (r1_xxreal_0(np__0,np__5)|~r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_53552, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_54128, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_53066, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_52548, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__5)|v1_abian(X3)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_52510, negated_conjecture, (esk5_2(X1,X2)=k3_finseq_1(X2)|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__1))).#trainneg
cnf(i_0_52998, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X2,X1),np__5)|~r1_xxreal_0(k3_xcmplx_0(X2,X1),np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_43263, negated_conjecture, (r1_xxreal_0(np__4,np__9)|~r1_xxreal_0(np__4,np__3))).#trainneg
cnf(i_0_44289, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__7|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__9,esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk4_2(X2,X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_51531, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_46296, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__7),k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_48416, negated_conjecture, (np__5=np__1|~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_29496, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__5)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_21980, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__9)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9645, plain, (r1_xxreal_0(np__5,np__7)|r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_23091, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12423, negated_conjecture, (~r1_xxreal_0(np__9,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_38470, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__7)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_53429, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__3)=k2_xcmplx_0(k3_finseq_1(X2),np__3)|~v1_xreal_0(k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__5)|~v1_xxreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_29992, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__5,esk4_2(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(esk4_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_38379, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X1,X2))|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_35665, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__3,X1)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10791, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__5,esk4_2(X2,X3))|~v7_ordinal1(esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_17984, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(X2,np__3)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_54640, negated_conjecture, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_54562, negated_conjecture, (X1=np__5|r1_xxreal_0(X1,np__3)|v1_abian(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X1,np__1))).#trainneg
cnf(i_0_54642, negated_conjecture, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_54641, negated_conjecture, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_54507, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__5|v1_abian(X2)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_55361, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__3,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X3)))).#trainneg
cnf(i_0_55360, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__3,esk5_2(X2,X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_55242, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_56206, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_54505, negated_conjecture, (r1_xxreal_0(X1,np__3)|X2=np__5|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_56208, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_56195, plain, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_54466, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_56181, plain, (r1_xxreal_0(np__7,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55071, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55235, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55986, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_54761, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(np__3,esk4_2(X1,X2))|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_55080, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__5)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55079, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__3)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55236, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_56274, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(X2,np__1)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55226, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_54497, negated_conjecture, (np__3=X1|X1=np__5|v1_abian(X1)|~r1_xxreal_0(np__3,X1)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_54702, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__5|~v7_ordinal1(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_56054, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55921, negated_conjecture, (r1_xxreal_0(np__2,np__9)|~r1_xxreal_0(np__2,np__7))).#trainneg
cnf(i_0_55939, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__3)=k2_xcmplx_0(k3_finseq_1(X2),np__3)|~r1_xxreal_0(k3_finseq_1(X2),np__5)|~m1_subset_1(k3_finseq_1(X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_55066, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_55648, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__9)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_55221, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_56169, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__1)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(esk5_2(X3,esk2_0))|~m1_subset_1(esk5_2(X3,esk2_0),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_55940, plain, (k2_xcmplx_0(esk5_2(X1,X2),np__3)=k2_xcmplx_0(k3_finseq_1(X2),np__3)|~r1_xxreal_0(k3_finseq_1(X2),np__5)|~m1_subset_1(k3_finseq_1(X2),k5_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_18089, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(X2,esk2_0)|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),np__9)|v1_abian(esk5_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X2,esk2_0),np__5)|v1_abian(esk5_2(X2,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_31826, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__7|r1_xxreal_0(np__5,esk4_2(X2,X3))|~v7_ordinal1(esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_23708, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk4_2(X2,X3),np__7)|esk4_2(X2,X3)=np__9|r1_xxreal_0(np__7,esk4_2(X2,X3))|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk4_2(X2,X3))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_49840, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_38208, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_22345, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_41239, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(X1,X2),np__7)|~m1_subset_1(esk5_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_13173, negated_conjecture, (r1_xxreal_0(np__4,np__9)|~r1_xxreal_0(np__4,np__1))).#trainneg
cnf(i_0_51972, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X2)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_16643, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X2,X1),np__9)|~r1_xxreal_0(k3_xcmplx_0(X2,X1),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_51871, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_45380, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_11393, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__1)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_40620, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_55918, negated_conjecture, (r1_xxreal_0(np__0,np__9)|~r1_xxreal_0(np__0,np__7))).#trainneg
cnf(i_0_8175, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__2,X1)|v1_abian(X1)|~r1_xxreal_0(np__2,np__9)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35464, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_11178, plain, (esk4_2(X1,X2)=np__3|esk4_2(X3,X4)=esk4_2(X1,X2)|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(esk4_2(X3,X4))|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_9917, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__3)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_16493, negated_conjecture, (r1_xxreal_0(k1_nat_1(np__2,esk4_2(X1,esk2_0)),np__7)|v1_abian(k1_nat_1(np__2,esk4_2(X1,esk2_0)))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(k1_nat_1(np__2,esk4_2(X1,esk2_0)))|~m1_subset_1(k1_nat_1(np__2,esk4_2(X1,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_21997, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_18066, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_43242, negated_conjecture, (np__7=np__3|np__7=np__1|r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_56456, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|esk5_2(X2,X1)=np__5|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_56554, negated_conjecture, (esk5_2(X1,X2)=np__5|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_56542, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|X2=np__5|v1_abian(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,np__1)|~v1_int_1(X2))).#trainneg
cnf(i_0_56474, negated_conjecture, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|esk5_2(X1,X2)=np__5|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_56438, negated_conjecture, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__5|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_57052, negated_conjecture, (esk4_2(X1,X2)=np__3|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__5|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_57761, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_57766, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_57760, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_57475, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),np__9)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_57768, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v7_ordinal1(X3)|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_56822, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_56537, negated_conjecture, (r1_xxreal_0(X1,np__3)|X2=np__5|v1_abian(X2)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_56528, negated_conjecture, (np__9=np__5|np__9=np__7|np__7=np__5|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_56830, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|esk4_2(X3,X4)=np__3|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk5_2(X1,X2),np__3)|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_57759, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_57763, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_9458, negated_conjecture, (k2_xcmplx_0(np__2,esk5_2(X1,esk2_0))=np__7|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__9|~m1_subset_1(esk5_2(X1,esk2_0),k5_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_36758, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=np__1|~r1_xxreal_0(np__9,np__5)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_55640, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__7),esk6_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_9001, plain, (r1_xxreal_0(X1,X2)|k2_xcmplx_0(X2,np__2)=np__7|v1_abian(X2)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,X2)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_59673, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_59684, negated_conjecture, (np__7=np__5|np__9=np__5|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_59688, negated_conjecture, (np__7=np__5|np__7=np__3|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_58554, negated_conjecture, (esk5_2(X1,X2)=k3_finseq_1(X2)|esk5_2(X1,X2)=np__5|~v1_xxreal_0(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__3))).#trainneg
cnf(i_0_58537, negated_conjecture, (esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_58611, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__5|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_60414, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|k1_funct_1(esk2_0,esk5_2(X1,esk2_0))!=k1_funct_1(esk2_0,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_59640, negated_conjecture, (np__7=np__5|r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_59140, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_58788, negated_conjecture, (X1=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_8463, plain, (r1_xxreal_0(np__9,esk4_2(X1,X2))|r1_xxreal_0(np__5,np__7)|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_49841, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13682, negated_conjecture, (r1_xxreal_0(np__9,X2)|v1_abian(X2)|esk5_2(X1,esk2_0)=np__7|~r1_xxreal_0(np__9,np__1)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_23581, negated_conjecture, (esk4_2(esk1_0,X1)=esk4_2(X2,X3)|esk4_2(X2,X3)=np__1|esk4_2(X2,X3)=np__5|esk4_2(X2,X3)=np__3|esk4_2(esk1_0,X1)=np__3|esk4_2(esk1_0,X1)=np__5|esk4_2(esk1_0,X1)=np__1|r1_xxreal_0(np__9,esk4_2(esk1_0,X1))|r1_xxreal_0(np__9,esk4_2(X2,X3))|~v5_chord(X1,esk1_0)|~v5_chord(X3,X2)|~m3_glib_001(X1,esk1_0)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_48514, negated_conjecture, (k2_xcmplx_0(esk5_2(X1,X2),np__2)=np__7|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(k3_finseq_1(X2)))).#trainneg
cnf(i_0_8429, plain, (r1_xxreal_0(esk5_2(X1,X2),X3)|v1_abian(X3)|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_59030, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_22342, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(esk5_2(X1,X2),np__1)|v1_abian(esk5_2(X1,X2))|~v7_ordinal1(k3_finseq_1(X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_49692, plain, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__5)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_16041, plain, (r1_xxreal_0(np__3,esk5_2(X1,X2))|r1_xxreal_0(np__9,np__1)|r1_xxreal_0(np__9,np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_38483, negated_conjecture, (r1_xxreal_0(np__4,np__9)|~r1_xxreal_0(np__4,np__7))).#trainneg
cnf(i_0_30216, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk5_2(X1,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(X1,esk2_0))|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__3|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_48684, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_54230, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X2)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_9848, plain, (r1_xxreal_0(np__5,X1)|v1_abian(X1)|r1_xxreal_0(np__1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11306, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__1)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_50215, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_7926, plain, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,X4)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X4,np__3)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_58544, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__5|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_46053, negated_conjecture, (np__9=np__5|~r1_xxreal_0(np__9,np__5)|np__7!=np__5)).#trainneg
cnf(i_0_44511, negated_conjecture, (r1_xxreal_0(X1,np__5)|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_53251, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(X1,np__7)|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v1_int_1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_44195, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_37189, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(X1,esk2_0))=np__3|k2_xcmplx_0(np__2,esk4_2(X1,esk2_0))=np__7|esk4_2(X1,esk2_0)=np__3|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29080, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_27601, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(X3,np__1)|r1_xxreal_0(esk4_2(X4,X5),X3)|v1_abian(X3)|r1_xxreal_0(np__3,esk4_2(X4,X5))|v1_abian(esk4_2(X1,X2))|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk4_2(X4,X5))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X5,X4)|~m3_glib_001(X5,X4)|~v1_glib_000(X4)|~v1_finset_1(X4)|~v1_funct_1(X4)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X4))).#trainneg
cnf(i_0_28989, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__5|~v1_int_1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,k1_nat_1(esk4_2(X1,esk2_0),np__2)))).#trainneg
cnf(i_0_18148, negated_conjecture, (r1_xxreal_0(X1,X2)|esk5_2(X3,esk2_0)=np__9|r1_xxreal_0(X2,np__5)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(X3,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_32220, negated_conjecture, (r1_xxreal_0(np__3,X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_17617, negated_conjecture, (esk5_2(X1,X2)=X3|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|X3=np__7|X3=np__9|v1_abian(X3)|~r1_xxreal_0(esk5_2(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_50615, negated_conjecture, (r1_xxreal_0(np__4,np__5)|~r1_xxreal_0(np__4,np__1))).#trainneg
cnf(i_0_26826, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11692, plain, (esk4_2(X1,X2)=np__9|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__3|~r1_xxreal_0(esk4_2(X1,X2),np__9)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_60547, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|v1_abian(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__9,np__5)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_61756, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|X3=np__5|v1_abian(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_61046, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_61323, negated_conjecture, (r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_61757, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|X3=np__5|v1_abian(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_61634, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_62001, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_61704, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__1,esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_61017, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,esk5_2(X1,esk2_0)),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__3|esk4_2(X1,esk2_0)=np__1|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1))).#trainneg
cnf(i_0_61430, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_62386, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__5)|r1_xxreal_0(np__3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_62387, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__5)|r1_xxreal_0(np__3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_61864, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__9,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_61431, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__3,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_61677, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__5,esk4_2(X1,X2))|esk4_2(X3,X4)=np__3|~r1_xxreal_0(esk4_2(X3,X4),np__3)|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__3,esk4_2(X3,X4))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_60554, negated_conjecture, (np__7=np__3|np__7=np__5|~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_61629, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(np__9,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_60565, negated_conjecture, (esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_61753, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_61488, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(np__9,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_61432, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|esk4_2(X3,X4)=np__5|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X3,X4),np__1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_62331, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(X2,np__5)|r1_xxreal_0(k2_xcmplx_0(X3,np__2),X2)|v1_abian(X2)|r1_xxreal_0(esk5_2(X1,esk2_0),X3)|v1_abian(esk5_2(X1,esk2_0))|v1_abian(X3)|~v7_ordinal1(X2)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X3)|~v1_int_1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_60921, negated_conjecture, (r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_xreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_61679, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(np__9,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_61701, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__1,esk4_2(X1,X2))|~r1_xxreal_0(np__9,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_62000, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__7)|v1_abian(esk5_2(esk1_0,esk2_0))|v1_abian(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_61754, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(X3,np__1)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_60852, plain, (r1_xxreal_0(np__3,k3_finseq_1(X2))|v1_abian(k2_xcmplx_0(esk4_2(X1,X2),np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k2_xcmplx_0(esk4_2(X1,X2),np__2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_62365, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(X2,np__5)|r1_xxreal_0(np__5,X2)|v1_abian(X2)|~v7_ordinal1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_61746, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|r1_xxreal_0(X4,X3)|v1_abian(X3)|v1_abian(X4)|~r1_xxreal_0(k6_xcmplx_0(X4,np__2),np__3)|~v1_xxreal_0(k6_xcmplx_0(X4,np__2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X4)|~v1_int_1(X3))).#trainneg
cnf(i_0_60627, negated_conjecture, (r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_60846, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__3,k3_finseq_1(X4))|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_finseq_1(X4),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_10815, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__9)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11348, plain, (X1=esk4_2(X2,X3)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_48608, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_29269, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk5_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_59407, negated_conjecture, (np__7=np__5|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_60455, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__7,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_60854, plain, (r1_xxreal_0(np__3,k3_finseq_1(X2))|v1_abian(k1_nat_1(esk4_2(X1,X2),np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k1_nat_1(esk4_2(X1,X2),np__2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_53490, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_62874, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_62869, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_63001, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_63730, negated_conjecture, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__5|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X3,X4),np__1)|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_62861, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_63195, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X1,X2),np__3)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_62862, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_62863, negated_conjecture, (r1_xxreal_0(np__5,np__7)|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_64559, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_63255, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_64962, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_abian(k1_nat_1(esk4_2(X2,X1),np__2))|~m1_subset_1(k3_finseq_1(X1),k1_numbers)|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(esk4_2(X2,X1)))).#trainneg
cnf(i_0_64041, plain, (r1_xxreal_0(np__3,k3_finseq_1(X1))|v1_abian(k2_xcmplx_0(esk4_2(X2,X1),np__2))|~m1_subset_1(k3_finseq_1(X1),k1_numbers)|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(esk4_2(X2,X1)))).#trainneg
cnf(i_0_63020, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__9,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_64046, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(X3,esk4_2(X1,X2))|r1_xxreal_0(np__5,X3)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_63021, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_63299, negated_conjecture, (np__7!=np__5|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_63079, plain, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__7,esk4_2(X3,X4))|~r1_xxreal_0(esk4_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_63153, negated_conjecture, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__3,esk5_2(X2,X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_63458, plain, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__5|esk4_2(X3,X4)=np__7|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__3,esk4_2(X3,X4))|~r1_xxreal_0(esk4_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_64824, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)!=np__1|~v5_chord(esk2_0,X1)|~v8_glib_001(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__1)|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_64358, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_64043, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64353, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_64368, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_63604, negated_conjecture, (esk5_2(X1,X2)=np__5|r1_xxreal_0(X3,k3_finseq_1(X2))|~r1_xxreal_0(X3,np__3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_64354, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~m1_subset_1(esk4_2(X1,X2),k5_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_64394, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3))).#trainneg
cnf(i_0_62812, negated_conjecture, (esk4_2(X1,X2)=X3|esk4_2(X1,X2)=np__1|X3=np__5|v1_abian(X3)|~r1_xxreal_0(esk4_2(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(X3,np__1)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_64389, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_63003, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(np__5,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_63302, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_64356, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_65170, negated_conjecture, (~r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_66052, negated_conjecture, (esk4_2(esk1_0,esk2_0)!=np__1|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_65175, negated_conjecture, (esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_65150, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__3)|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_65185, negated_conjecture, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_65174, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65149, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__3)|~r1_xxreal_0(esk5_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_65409, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|k1_funct_1(esk2_0,esk5_2(X1,esk2_0))!=k1_funct_1(esk2_0,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_65649, negated_conjecture, (r1_xxreal_0(np__5,np__7)|esk4_2(esk1_0,esk2_0)=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_65152, negated_conjecture, (r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_65648, negated_conjecture, (r1_xxreal_0(np__5,np__7)|esk4_2(esk1_0,esk2_0)=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_65653, negated_conjecture, (r1_xxreal_0(np__5,np__7)|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_65651, negated_conjecture, (r1_xxreal_0(np__5,np__7)|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_65942, negated_conjecture, (r1_xxreal_0(np__5,np__7)|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_65559, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__3,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_65173, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65870, plain, (v1_abian(esk4_2(X1,X2))|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_65650, negated_conjecture, (r1_xxreal_0(np__5,np__7)|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_65168, negated_conjecture, (r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_65561, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X3,X4),esk5_2(X1,X2))|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__3,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_65163, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__3)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_66881, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_65162, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__3)|~r1_xxreal_0(k1_nat_1(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_66876, negated_conjecture, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|esk4_2(X2,X3)=np__7|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,np__3)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_63795, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X3)|r1_xxreal_0(X3,np__5)|v1_abian(X2)|v1_abian(X3)|~v7_ordinal1(X3)|~v1_int_1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64355, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_65652, negated_conjecture, (r1_xxreal_0(np__5,np__7)|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_64145, negated_conjecture, (r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X1)|v1_abian(X1)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X2)|v1_abian(esk5_2(esk1_0,esk2_0))|v1_abian(X2)|~v7_ordinal1(X1)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v1_int_1(X2)|~v1_int_1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_66912, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__1|esk4_2(X1,esk2_0)=np__7|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_66919, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__5),esk6_2(X1,esk2_0))|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1))).#trainneg
cnf(i_0_62957, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_64752, negated_conjecture, (np__7=np__5|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_65247, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk5_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_59406, negated_conjecture, (np__7=np__5|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_44058, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38373, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X1,X2))|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__7|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__5,esk4_2(X1,X2)))).#trainneg
cnf(i_0_57671, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2),np__9)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8182, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__9)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_28848, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_11471, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X3,np__2),X4)|v1_abian(X4)|r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(esk4_2(X1,X2))|v1_abian(X3)|~v7_ordinal1(X4)|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X3)|~v1_int_1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_50873, negated_conjecture, (esk5_2(X1,X2)=k3_finseq_1(X2)|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__3))).#trainneg
cnf(i_0_65246, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk5_2(X1,esk2_0)),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_65090, negated_conjecture, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_67285, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_67239, negated_conjecture, (~r1_xxreal_0(np__7,np__1)|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_67190, negated_conjecture, (~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_67200, negated_conjecture, (np__7=np__3|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_67201, negated_conjecture, (np__9=np__5|np__9=np__7|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_67240, negated_conjecture, (np__9=np__5|~r1_xxreal_0(np__7,np__1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_67187, negated_conjecture, (np__9=np__3|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_67234, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_67094, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_67238, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_67025, negated_conjecture, (esk4_2(X1,X2)=np__3|~r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_67738, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_67596, negated_conjecture, (r1_xxreal_0(np__5,np__7))).#trainneg
cnf(i_0_67245, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__9,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_67191, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_68235, negated_conjecture, (np__7=np__5|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_68707, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|~r1_xxreal_0(np__9,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_67035, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_67086, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(np__7,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_67085, negated_conjecture, (esk4_2(X1,X2)=np__9|r1_xxreal_0(np__5,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_67661, negated_conjecture, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__3,esk4_2(X3,X4))|~r1_xxreal_0(esk4_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_67055, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__1))).#trainneg
cnf(i_0_67096, negated_conjecture, (esk4_2(X1,X2)=np__7|r1_xxreal_0(np__5,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_68296, negated_conjecture, (np__7=np__5|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_68204, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_67218, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_67246, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__9,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_67252, negated_conjecture, (esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__7,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_67247, negated_conjecture, (~r1_xxreal_0(np__9,np__1)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_67045, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|v1_abian(esk4_2(X1,X2))|~v7_ordinal1(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_68486, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__7|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X1,X2))|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_67774, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_67183, negated_conjecture, (esk5_2(X1,X2)=np__1|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_67802, negated_conjecture, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk4_2(X4,X5)=np__5|esk5_2(X2,X3)=np__5|esk5_2(X2,X3)=np__7|r1_xxreal_0(np__3,esk4_2(X4,X5))|~r1_xxreal_0(X1,esk4_2(X4,X5))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(esk4_2(X4,X5))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~v5_chord(X5,X4)|~m3_glib_001(X3,X2)|~m3_glib_001(X5,X4)|~v1_glib_000(X2)|~v1_glib_000(X4)|~v1_finset_1(X2)|~v1_finset_1(X4)|~v1_funct_1(X2)|~v1_funct_1(X4)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X4))).#trainneg
cnf(i_0_67184, negated_conjecture, (esk5_2(X1,X2)=np__1|esk5_2(X1,X2)=np__9|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_67781, negated_conjecture, (r1_xxreal_0(X1,np__5)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_67729, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_67732, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__3)|~r1_xxreal_0(k4_nat_1(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_67330, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_71187, negated_conjecture, (~r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_71071, negated_conjecture, (~r1_xxreal_0(np__9,np__1))).#trainneg
cnf(i_0_69722, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_68861, negated_conjecture, (r1_xxreal_0(X1,esk5_2(X2,X3))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_68883, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|~v1_xxreal_0(esk5_2(X2,X1))|~v1_xxreal_0(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_69723, negated_conjecture, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,X1)|v1_abian(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_68899, negated_conjecture, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_69713, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_71085, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_69717, negated_conjecture, (esk4_2(X1,X2)=np__1|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_71440, negated_conjecture, (r1_xxreal_0(k4_nat_1(X2,X1),np__3)|~r1_xxreal_0(k4_nat_1(X2,X1),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_67792, negated_conjecture, (esk5_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__5|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X3,X4))|~r1_xxreal_0(esk5_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_70174, negated_conjecture, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_70047, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_68171, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__9|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_67114, negated_conjecture, (esk4_2(X1,X2)=X3|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(X3,np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_67220, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X2)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_67730, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_62708, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__9|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~m1_subset_1(esk4_2(X1,X2),k5_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_48176, plain, (esk5_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X3,X4)=np__7|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X3,X4))|~r1_xxreal_0(esk5_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_8696, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),np__9)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(X3,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_21396, plain, (esk5_2(X1,X2)=X3|esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(np__7,X3)|v1_abian(X3)|~r1_xxreal_0(esk5_2(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk5_2(X1,X2))|~v1_int_1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_50227, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_64373, plain, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_63182, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X3,X4),esk5_2(X1,X2))|r1_xxreal_0(np__3,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_64194, negated_conjecture, (r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(X2,np__5))).#trainneg
cnf(i_0_27316, negated_conjecture, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__9|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__9,np__5)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_21999, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_11435, plain, (X1=esk4_2(X2,X3)|esk4_2(X2,X3)=np__5|esk4_2(X2,X3)=np__7|r1_xxreal_0(np__9,esk4_2(X2,X3))|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_71525, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_71505, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_71477, negated_conjecture, (r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_71506, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_71526, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_71504, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_71475, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__7)|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_71522, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_71507, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_71973, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1))).#trainneg
cnf(i_0_71644, negated_conjecture, (np__7=np__5|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_71891, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__3)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_71651, negated_conjecture, (esk4_2(X1,X2)=np__3|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_71499, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__5)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_71523, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_71577, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_71643, negated_conjecture, (esk5_2(X1,X2)=np__3|r1_xxreal_0(np__5,k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k5_numbers))).#trainneg
cnf(i_0_71474, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__7)|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_71524, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_71509, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_71716, negated_conjecture, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,np__3)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_71640, negated_conjecture, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,np__5)|v1_abian(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_71513, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__9,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_71516, negated_conjecture, (esk5_2(X1,X2)=k3_finseq_1(X2)|~v1_xxreal_0(esk5_2(X1,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(k3_finseq_1(X2),np__3))).#trainneg
cnf(i_0_71491, negated_conjecture, (r1_xxreal_0(k4_nat_1(X1,X2),np__7)|~r1_xxreal_0(k4_nat_1(X1,X2),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_71498, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__5)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71487, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__7)|~r1_xxreal_0(k1_nat_1(X1,X2),np__5)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_71512, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__3,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_72950, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_71488, negated_conjecture, (r1_xxreal_0(k1_nat_1(X1,X2),np__7)|~r1_xxreal_0(k1_nat_1(X1,X2),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_72922, negated_conjecture, (esk5_2(esk1_0,esk2_0)=esk5_2(X1,esk2_0)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_71486, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_71706, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_72030, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(X1,X2)=np__7|v1_abian(esk5_2(esk1_0,esk2_0))|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2)))).#trainneg
cnf(i_0_71510, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_72946, negated_conjecture, (esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),np__9)|v1_abian(esk5_2(X1,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|~v7_ordinal1(esk5_2(X1,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_72867, negated_conjecture, (r1_xxreal_0(X1,esk5_2(X2,X3))|esk5_2(X2,X3)=np__5|esk5_2(X2,X3)=np__7|r1_xxreal_0(np__3,esk4_2(X4,X5))|~r1_xxreal_0(X1,esk4_2(X4,X5))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(esk4_2(X4,X5))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~v5_chord(X5,X4)|~m3_glib_001(X3,X2)|~m3_glib_001(X5,X4)|~v1_glib_000(X2)|~v1_glib_000(X4)|~v1_finset_1(X2)|~v1_finset_1(X4)|~v1_funct_1(X2)|~v1_funct_1(X4)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X4))).#trainneg
cnf(i_0_71514, negated_conjecture, (esk5_2(X1,X2)=np__9|r1_xxreal_0(np__9,k3_finseq_1(X2))|~m1_subset_1(k3_finseq_1(X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_49690, plain, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__3)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X3,X4),k1_numbers)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_66069, plain, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X1,X2)=np__1|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2)))).#trainneg
cnf(i_0_27029, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_63796, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),X3)|r1_xxreal_0(X3,np__5)|v1_abian(X2)|v1_abian(X3)|~v7_ordinal1(X3)|~v1_int_1(X2)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_11015, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__5)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7985, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__9,X1)|v1_abian(X1)|~r1_xxreal_0(np__9,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_53166, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(X1,np__7)|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__7,X1)|v1_abian(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_22946, negated_conjecture, (r1_xxreal_0(k3_finseq_1(X2),np__9)|~r1_xxreal_0(k3_finseq_1(X2),np__5)|~v7_ordinal1(k3_xcmplx_0(k15_glib_001(X1,X2),np__2))|~m1_subset_1(k15_glib_001(X1,X2),k5_numbers)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_64390, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|~v1_int_1(X3)|~m1_subset_1(X3,k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5))).#trainneg
cnf(i_0_68167, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__7|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_59765, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_10373, plain, (esk4_2(X1,X2)=k2_xcmplx_0(X3,np__2)|r1_xxreal_0(esk4_2(X1,X2),X3)|v1_abian(X3)|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(k2_xcmplx_0(X3,np__2))|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v1_xxreal_0(esk4_2(X1,X2))|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X3,np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10895, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__9)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_45402, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_13049, plain, (r1_xxreal_0(X1,X2)|esk4_2(X3,X4)=np__3|esk4_2(X3,X4)=np__7|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__9,esk4_2(X3,X4))|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_73775, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_73603, negated_conjecture, (np__9=np__5|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_73703, negated_conjecture, (np__9=np__7|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_73976, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_73980, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_74035, negated_conjecture, (esk4_2(X1,X2)=np__7|~r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_74069, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_74259, negated_conjecture, (esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|~r1_xxreal_0(np__5,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_74700, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,np__5)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_74357, negated_conjecture, (esk5_2(X1,X2)=np__7|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(esk5_2(X1,X2),np__5)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_75442, negated_conjecture, (r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_74401, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_74128, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_75300, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,np__5)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_74283, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_75314, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__7,X1)|r1_xxreal_0(X1,np__7)|v1_abian(esk5_2(esk1_0,esk2_0))|v1_abian(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~v1_xxreal_0(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_74630, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__5)|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_74044, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__5)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_74174, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__3|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_74891, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_74699, negated_conjecture, (r1_xxreal_0(np__4,np__7)|~r1_xxreal_0(np__4,np__5))).#trainneg
cnf(i_0_74692, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),np__7)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__5)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_74786, negated_conjecture, (esk5_2(esk1_0,esk2_0)=esk5_2(X1,esk2_0)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9))).#trainneg
cnf(i_0_74955, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),esk5_2(esk1_0,esk2_0))|esk5_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|~v1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_73469, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_74079, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_75336, negated_conjecture, (r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_74366, negated_conjecture, (r1_xxreal_0(X1,np__7)|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk5_2(X2,X3),np__5)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_74022, negated_conjecture, (k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2)=np__7|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(X1,np__7)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_74335, negated_conjecture, (k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2)=np__7|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,np__7)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_74299, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_73470, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_73657, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_73864, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_74687, negated_conjecture, (r1_xxreal_0(np__0,np__7)|~r1_xxreal_0(np__0,np__5))).#trainneg
cnf(i_0_73729, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_73463, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|np__9!=np__7|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_23597, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__9|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_23860, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__9,np__5)|r1_xxreal_0(X3,esk4_2(X1,X2))|r1_xxreal_0(np__7,X3)|v1_abian(X3)|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28571, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_8695, negated_conjecture, (X1=esk5_2(X2,esk2_0)|r1_xxreal_0(k2_xcmplx_0(X1,np__2),np__9)|v1_abian(X1)|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_77940, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_75826, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|v1_abian(esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_75828, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|esk5_2(esk1_0,esk2_0)=np__9)).#trainneg
cnf(i_0_76849, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__5)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_76062, negated_conjecture, (np__9=np__5|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_77956, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_77960, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_75879, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_77405, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2)=np__7|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__3)|~r1_xxreal_0(np__9,np__7)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_76230, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|v1_abian(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__9,np__3)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_76300, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__7,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_76231, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|v1_abian(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__9,np__3)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_77950, negated_conjecture, (r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_76301, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_75821, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_75817, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|esk5_2(esk1_0,esk2_0)=np__9|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_75834, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__7)|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__9|~v7_ordinal1(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_77962, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__9,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_76220, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|~r1_xxreal_0(np__9,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_76451, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_76221, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_76730, negated_conjecture, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_77382, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v7_ordinal1(k3_finseq_1(X3)))).#trainneg
cnf(i_0_75818, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk5_2(X1,X2))|esk5_2(esk1_0,esk2_0)=np__9|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_75823, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__9|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_77694, negated_conjecture, (~r1_xxreal_0(np__7,np__3)|np__9!=np__7)).#trainneg
cnf(i_0_76453, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~r1_xxreal_0(np__7,esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k3_finseq_1(X2),k1_numbers))).#trainneg
cnf(i_0_73986, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_61428, plain, (esk4_2(X1,X2)=np__1|r1_xxreal_0(esk4_2(X3,X4),esk4_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_54238, negated_conjecture, (r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(k1_nat_1(esk4_2(esk1_0,esk2_0),np__2),np__7)|v1_abian(k1_nat_1(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k1_nat_1(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(k1_nat_1(esk4_2(esk1_0,esk2_0),np__2),k1_numbers))).#trainneg
cnf(i_0_74676, negated_conjecture, (r1_xxreal_0(k4_nat_1(X2,X1),np__7)|~r1_xxreal_0(k4_nat_1(X2,X1),np__5)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_76731, negated_conjecture, (esk5_2(X1,X2)=np__7|esk5_2(X1,X2)=np__5|r1_xxreal_0(esk4_2(X3,X4),esk5_2(X1,X2))|r1_xxreal_0(np__5,esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_18155, negated_conjecture, (esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__5)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2)))).#trainneg
cnf(i_0_15690, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__5|k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__5,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12136, plain, (X1=esk4_2(X2,X3)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_78242, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__5)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_78297, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers))).#trainneg
cnf(i_0_78356, negated_conjecture, (~r1_xxreal_0(np__7,k1_nat_1(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_78139, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_78298, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__7)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_78949, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__9)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_78352, negated_conjecture, (~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_78230, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_78347, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__5)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_78222, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_78907, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__7)|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_79730, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__7)|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers))).#trainneg
cnf(i_0_78225, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_79916, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__7)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_79275, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(k1_nat_1(esk4_2(esk1_0,esk2_0),np__2),np__5)|v1_abian(k1_nat_1(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(k1_nat_1(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_78708, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|~r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_78779, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_79525, negated_conjecture, (esk5_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|~r1_xxreal_0(esk5_2(X1,X2),esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_79828, plain, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__5,esk4_2(X3,X4))|~r1_xxreal_0(esk4_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_79544, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_79145, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__5)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_78273, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_78649, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1)|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_78651, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_79075, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__7,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_78291, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(X1,np__7)|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_78180, negated_conjecture, (r1_xxreal_0(X1,np__9)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_78179, negated_conjecture, (r1_xxreal_0(X1,np__5)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_79939, negated_conjecture, (esk5_2(X1,X2)=esk4_2(X3,X4)|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__5,esk4_2(X3,X4))|~r1_xxreal_0(esk5_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_78193, negated_conjecture, (r1_xxreal_0(X1,np__9)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_78182, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_81432, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_80501, negated_conjecture, (~r1_xxreal_0(np__9,np__3))).#trainneg
cnf(i_0_80383, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__5))).#trainneg
cnf(i_0_80447, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__9))).#trainneg
cnf(i_0_80444, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__7))).#trainneg
cnf(i_0_80224, negated_conjecture, (np__9=np__5|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_80189, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_80188, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,esk4_2(esk1_0,esk2_0)),k1_funct_1(esk2_0,np__1),esk6_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_80187, negated_conjecture, (~r1_xxreal_0(np__9,k1_nat_1(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_81480, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(X1,np__7)|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_81362, plain, (esk4_2(X1,X2)=esk4_2(X3,X4)|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__5,esk4_2(X3,X4))|esk4_2(X3,X4)=np__1|r1_xxreal_0(np__5,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2)))).#trainneg
cnf(i_0_80943, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_80971, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_80955, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_81442, negated_conjecture, (r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_81918, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__9)|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82044, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__7)|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_81920, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__9)|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_80745, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_80985, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__5)|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_78194, negated_conjecture, (r1_xxreal_0(X1,np__7)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_78252, negated_conjecture, (r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(X1,np__5)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_39041, negated_conjecture, (esk5_2(X1,X2)=X3|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|X3=np__7|v1_abian(X3)|~r1_xxreal_0(esk5_2(X1,X2),X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk5_2(X1,X2))|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_19920, negated_conjecture, (r1_xxreal_0(np__3,k3_finseq_1(X1))|r1_xxreal_0(esk4_2(X2,X3),np__9)|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__7|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__9,esk4_2(X2,X3))|v1_abian(esk5_2(X4,X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_xxreal_0(esk4_2(X2,X3))|~v5_chord(X1,X4)|~m3_glib_001(X1,X4)|~v1_glib_000(X4)|~v1_finset_1(X4)|~v1_funct_1(X4)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X4)|~v7_ordinal1(esk5_2(X4,X1))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10109, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_59876, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(X2,np__7)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1)|~v7_ordinal1(X2)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_77631, negated_conjecture, (r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,esk4_2(X1,esk2_0)),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__7|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1))).#trainneg
cnf(i_0_62999, negated_conjecture, (esk4_2(X1,X2)=np__9|esk4_2(X1,X2)=np__7|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk4_2(X1,X2),np__5)|~v7_ordinal1(esk4_2(X1,X2))|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_56266, plain, (X1=esk4_2(X2,X3)|r1_xxreal_0(np__5,esk4_2(X2,X3))|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_78308, negated_conjecture, (r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_57649, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(X1,X2)=np__9|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__7)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12524, negated_conjecture, (r1_xxreal_0(X1,np__9)|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~r1_xxreal_0(esk4_2(X2,X3),np__5)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_29710, plain, (r1_xxreal_0(X1,k1_card_1(X2))|k1_card_1(X2)=np__1|v1_abian(k1_card_1(X2))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(k1_card_1(X2))|~v1_xxreal_0(X1)|~v1_finset_1(X2))).#trainneg
cnf(i_0_8701, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_nat_1(esk4_2(X1,esk2_0),np__2),np__2),np__9)|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(k1_nat_1(esk4_2(X1,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_7974, negated_conjecture, (r1_xxreal_0(X1,np__5)|esk5_2(X2,esk2_0)=np__9|esk5_2(X2,esk2_0)=np__7|~r1_xxreal_0(X1,esk5_2(X2,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v1_xxreal_0(X1)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_55949, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__5)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v7_ordinal1(esk5_2(X3,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_72183, negated_conjecture, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__3,esk4_2(X2,X3))|~v7_ordinal1(esk4_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_82138, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__3)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_82129, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__3)|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82130, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82918, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82131, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_83500, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_83478, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_83460, negated_conjecture, (r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__3)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_82917, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82444, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82476, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82542, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82497, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_82459, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_82198, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_82458, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_83225, negated_conjecture, (r1_xxreal_0(X1,np__9)|r1_xxreal_0(np__7,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_83418, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__9)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_82473, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__3,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_82544, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_82456, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v1_xreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_83556, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(X2,np__1)|v1_abian(X2)|~r1_xxreal_0(X1,esk4_2(X3,X4))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~m1_subset_1(X2,k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_82557, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__3)|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_82438, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__7,np__3)|~v1_int_1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82441, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__2)=np__7|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_82436, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82199, negated_conjecture, (~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__1,np__2)|~r1_xxreal_0(np__8,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82555, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),X1)|v1_abian(X1)|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_82889, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_82864, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_83378, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(X2,np__2),np__5)|r1_xxreal_0(esk5_2(X1,esk2_0),X2)|v1_abian(esk5_2(X1,esk2_0))|v1_abian(X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(X2)|~v1_int_1(esk5_2(X1,esk2_0)))).#trainneg
cnf(i_0_82593, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_83315, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X2,esk2_0),np__5)|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),np__9)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk5_2(X1,esk2_0),k5_numbers)|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_82539, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__7)|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_82952, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|r1_glib_000(X1,k1_funct_1(esk2_0,np__1),k1_funct_1(esk2_0,np__1),esk6_2(X1,esk2_0))|esk5_2(X1,esk2_0)=np__5|~r1_xxreal_0(np__7,np__3)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_83318, negated_conjecture, (esk5_2(X1,esk2_0)=np__7|esk5_2(X1,esk2_0)=np__9|r1_xxreal_0(esk5_2(X2,esk2_0),np__5)|esk5_2(X2,esk2_0)=np__9|r1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(X2,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,esk2_0))|~v1_xxreal_0(k2_xcmplx_0(esk5_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X2)|~m3_glib_001(esk2_0,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_82556, negated_conjecture, (r1_xxreal_0(X1,np__7)|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(X1,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_83432, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk5_2(X3,esk2_0),np__1)|v1_abian(esk5_2(X3,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~v1_int_1(esk5_2(X3,esk2_0))|~m1_subset_1(esk5_2(X3,esk2_0),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_82594, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_82200, negated_conjecture, (~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,np__2)|~r1_xxreal_0(np__2,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_82595, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|r1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2),np__9)|v1_abian(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(X1,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(k2_xcmplx_0(esk5_2(esk1_0,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_52630, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_83939, negated_conjecture, (r1_xxreal_0(np__3,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_83757, negated_conjecture, (r1_xxreal_0(np__1,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_83977, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_84002, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_83989, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_84003, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_83990, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_84025, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_84274, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v1_int_1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_83975, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_83988, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_84001, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_84023, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_84272, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_84312, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_84146, negated_conjecture, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__5,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_84166, negated_conjecture, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__3,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_85393, negated_conjecture, (r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__7)|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers))).#trainneg
cnf(i_0_84142, negated_conjecture, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|r1_xxreal_0(np__3,esk4_2(X3,X4))|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_84517, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_84069, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__9,np__5)|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_83835, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_84292, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_84293, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_84453, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_xxreal_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_84308, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_85257, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(X1,esk2_0),k1_numbers))).#trainneg
cnf(i_0_84147, negated_conjecture, (r1_xxreal_0(np__7,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X3,X4),np__9)|esk4_2(X3,X4)=np__5|r1_xxreal_0(np__7,esk4_2(X3,X4))|r1_xxreal_0(esk4_2(X1,X2),np__1)|v1_abian(esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk4_2(X1,X2))|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_83848, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_84275, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__2)=np__3|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_85198, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__9)|esk5_2(esk1_0,esk2_0)=np__5|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_84514, negated_conjecture, (k2_xcmplx_0(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__3,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_11659, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_41818, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(esk1_0,esk2_0),X1)|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__7|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_38213, negated_conjecture, (r1_xxreal_0(k1_card_1(X1),np__7)|~r1_xxreal_0(k1_card_1(X1),np__1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_28223, plain, (r1_xxreal_0(X1,np__5)|esk5_2(X2,X3)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X3))|~r1_xxreal_0(X1,esk5_2(X2,X3))|~v1_xxreal_0(esk5_2(X2,X3))|~v1_xxreal_0(X1)|~m1_subset_1(k3_finseq_1(X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_71441, negated_conjecture, (r1_xxreal_0(k3_xcmplx_0(X1,X2),np__3)|~r1_xxreal_0(k3_xcmplx_0(X1,X2),np__1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_62710, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__9|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~m1_subset_1(esk4_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_8181, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__4,X1)|v1_abian(X1)|~r1_xxreal_0(np__4,np__9)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_58686, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|v1_abian(esk5_2(X1,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(X1,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_71749, negated_conjecture, (esk4_2(X1,X2)=np__5|r1_xxreal_0(X3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~r1_xxreal_0(X3,np__7)|~v7_ordinal1(esk4_2(X1,X2))|~v1_xxreal_0(X3)|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_50560, negated_conjecture, (X1=np__7|r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|v1_abian(esk5_2(esk1_0,esk2_0))|v1_abian(X1)|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(X1,np__7)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_21919, plain, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=esk4_2(X3,X4)|r1_xxreal_0(np__9,esk4_2(X3,X4))|v1_abian(esk4_2(X3,X4))|esk4_2(X3,X4)=np__1|esk4_2(X3,X4)=np__5|esk4_2(X3,X4)=np__3|~v1_int_1(esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__7)|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_43611, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(X1,esk2_0))=np__7|esk4_2(X1,esk2_0)=np__3|esk4_2(X1,esk2_0)=np__1|~m1_subset_1(k2_xcmplx_0(np__2,esk4_2(X1,esk2_0)),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_15187, negated_conjecture, (k1_nat_1(esk4_2(X1,esk2_0),np__2)=np__7|v1_abian(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~r1_xxreal_0(np__7,k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v1_xxreal_0(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v7_ordinal1(k1_nat_1(esk4_2(X1,esk2_0),np__2))|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11122, plain, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(esk4_2(X2,X3),X1)|v1_abian(X1)|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(esk4_2(X2,X3))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_72857, negated_conjecture, (esk5_2(X1,X2)=esk4_2(X3,X4)|esk5_2(X1,X2)=np__5|esk5_2(X1,X2)=np__7|r1_xxreal_0(np__3,esk4_2(X3,X4))|~r1_xxreal_0(esk5_2(X1,X2),esk4_2(X3,X4))|~v1_xxreal_0(esk4_2(X3,X4))|~v1_xxreal_0(esk5_2(X1,X2))|~v5_chord(X2,X1)|~v5_chord(X4,X3)|~m3_glib_001(X2,X1)|~m3_glib_001(X4,X3)|~v1_glib_000(X1)|~v1_glib_000(X3)|~v1_finset_1(X1)|~v1_finset_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_84012, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__7))).#trainneg
cnf(i_0_71754, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk4_2(X1,esk2_0)=np__1|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_20830, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|r1_xxreal_0(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),np__5)|v1_abian(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~v1_int_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2))|~m1_subset_1(k2_xcmplx_0(esk4_2(X1,esk2_0),np__2),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,esk2_0)))).#trainneg
cnf(i_0_82692, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|v1_abian(esk5_2(esk1_0,esk2_0))|v1_abian(X1)|~v1_xxreal_0(X1)|~v1_int_1(esk5_2(esk1_0,esk2_0))|~v1_int_1(X1))).#trainneg
cnf(i_0_15792, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__9,esk4_2(X1,X2))|esk5_2(X3,X4)=np__7|~r1_xxreal_0(esk5_2(X3,X4),np__7)|~v1_xxreal_0(esk5_2(X3,X4))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(np__7,esk5_2(X3,X4))|~v5_chord(X4,X3)|~m3_glib_001(X4,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_29678, negated_conjecture, (r1_xxreal_0(X1,k3_finseq_1(X2))|esk5_2(X3,X2)=np__9|esk5_2(X3,X2)=np__7|~v1_xxreal_0(esk5_2(X3,X2))|~v1_xxreal_0(k3_finseq_1(X2))|~v1_xxreal_0(X1)|~v5_chord(X2,X3)|~m3_glib_001(X2,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3)|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(X1,np__9))).#trainneg
cnf(i_0_10601, plain, (esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(X3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|k2_xcmplx_0(esk4_2(X1,X2),np__2)=np__3|~r1_xxreal_0(X3,np__9)|~v7_ordinal1(esk4_2(X1,X2))|~v1_xxreal_0(X3)|~r1_xxreal_0(np__1,esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_15001, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(esk5_2(X2,X3),X1)|v1_abian(X1)|r1_xxreal_0(np__7,k3_finseq_1(X3))|~v7_ordinal1(esk5_2(X2,X3))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k3_finseq_1(X3),k1_numbers)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8296, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8502, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(k2_xcmplx_0(X2,np__2),np__9)|v1_abian(X2)|~r1_xxreal_0(X1,esk5_2(X3,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk5_2(X3,esk2_0))|~v1_xxreal_0(X1)|~v1_int_1(X2)|~v1_xxreal_0(k2_xcmplx_0(X2,np__2))|~v5_chord(esk2_0,X3)|~m3_glib_001(esk2_0,X3)|~v1_glib_000(X3)|~v1_finset_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_22366, negated_conjecture, (esk5_2(X1,X2)=np__7|r1_xxreal_0(np__9,k3_finseq_1(X2))|r1_xxreal_0(k2_xcmplx_0(X3,np__2),np__9)|r1_xxreal_0(esk5_2(X1,X2),X3)|v1_abian(esk5_2(X1,X2))|v1_abian(X3)|~v7_ordinal1(k3_finseq_1(X2))|~v1_xxreal_0(k2_xcmplx_0(X3,np__2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_int_1(esk5_2(X1,X2))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_85508, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85518, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_85604, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_85619, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_85660, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_85596, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85449, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_85439, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85595, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85647, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_85673, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_85591, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__9,np__5)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85597, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85578, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,np__3))).#trainneg
cnf(i_0_85697, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85599, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_85710, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_85698, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_85738, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_86639, negated_conjecture, (r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(esk5_2(X1,esk2_0),np__7)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_85487, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__3|v1_abian(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85762, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85725, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_85708, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_85696, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_85763, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_85580, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_85709, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_85761, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_86170, negated_conjecture, (k2_xcmplx_0(X1,np__2)=np__3|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_87936, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_88033, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__3)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_87895, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_88755, negated_conjecture, (~r1_xxreal_0(np__7,np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_87886, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_88830, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_87908, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_89061, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_87943, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_88095, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_87945, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_89204, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_89176, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_87952, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_87967, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_89189, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_88111, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_89555, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_88081, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_88125, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_89060, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_88080, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__9,np__5)|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_88108, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_89569, negated_conjecture, (v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__1,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers))).#trainneg
cnf(i_0_87981, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_88056, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_88107, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_89175, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_88845, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__3|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_88996, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~r1_xxreal_0(np__9,np__5))).#trainneg
cnf(i_0_89086, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k5_numbers)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_87300, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,np__9)|~v1_xxreal_0(X1)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_89975, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_90982, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_89790, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2)))).#trainneg
cnf(i_0_89727, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_90418, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_89859, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_89753, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_89834, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_90355, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_90356, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_90409, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_90983, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_90436, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_89860, negated_conjecture, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(X2,np__3))).#trainneg
cnf(i_0_90437, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_89789, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk5_2(esk1_0,esk2_0)=np__5|v1_abian(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v1_xxreal_0(esk4_2(X1,X2)))).#trainneg
cnf(i_0_91416, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_92107, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_91727, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(X3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(X3,np__1)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(X3)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_92556, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__7)|r1_xxreal_0(X3,esk4_2(X1,X2))|v1_abian(esk4_2(X1,X2))|~r1_xxreal_0(X3,np__3)|~v1_int_1(esk4_2(X1,X2))|~v1_xxreal_0(X3)|~m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_91509, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_91510, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X2,np__1)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_91486, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_91726, negated_conjecture, (r1_xxreal_0(esk5_2(X1,X2),np__7)|r1_xxreal_0(X3,esk5_2(X1,X2))|v1_abian(esk5_2(X1,X2))|~r1_xxreal_0(X3,np__1)|~v1_int_1(esk5_2(X1,X2))|~v1_xxreal_0(X3)|~m1_subset_1(esk5_2(X1,X2),k1_numbers)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_91487, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_91773, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_93593, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_93207, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1)|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_93218, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_93342, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5))).#trainneg
cnf(i_0_93418, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_93407, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1)|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_94169, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_94170, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_94171, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_93542, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5))).#trainneg
cnf(i_0_93584, negated_conjecture, (np__7=np__5|np__5=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_93554, negated_conjecture, (np__5=np__3|np__7=np__5|esk4_2(esk1_0,esk2_0)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_94729, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_94973, negated_conjecture, (r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_94172, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_94731, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95011, negated_conjecture, (r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_93375, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_94730, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),esk5_2(esk1_0,esk2_0))|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_93373, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_93575, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_93573, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_93376, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_93576, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_94648, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_94649, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk5_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(np__7,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_93390, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_93190, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_93318, negated_conjecture, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|v1_abian(X1)|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(esk4_2(esk1_0,esk2_0))|~v1_int_1(X1))).#trainneg
cnf(i_0_95161, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7)).#trainneg
cnf(i_0_95108, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_95469, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__7)).#trainneg
cnf(i_0_95125, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95127, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95097, negated_conjecture, (np__9=np__7|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_95183, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95112, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95173, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95126, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95184, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95114, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95124, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95175, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95185, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk4_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_96442, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_96967, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2),k1_numbers)|~v1_int_1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95435, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95433, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95901, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95199, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95348, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_96052, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95890, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_97051, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9)).#trainneg
cnf(i_0_98316, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_98333, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_98332, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_98000, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|np__9=np__7)).#trainneg
cnf(i_0_98336, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_98624, negated_conjecture, (r1_xxreal_0(np__3,esk4_2(X1,X2))|r1_xxreal_0(esk4_2(X1,X2),np__9)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_97486, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_97608, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_96129, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95188, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__3,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_99224, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_98837, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_99285, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_95270, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95434, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_96285, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_98416, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_98410, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95432, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_100142, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_98334, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96842, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_98409, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_97487, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_98335, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96140, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95155, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95832, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|np__5=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_98413, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95710, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|np__5=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96385, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_98411, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96365, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_99037, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=k1_funct_1(esk2_0,np__1))).#trainneg
cnf(i_0_96613, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_96296, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X1,X2))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_95869, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|np__5=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_96387, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_97422, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_97417, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_100331, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7))).#trainneg
cnf(i_0_101872, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_100705, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_101108, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_101315, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__9,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_100706, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_100986, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_98412, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95865, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|np__5=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_96367, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95198, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_97457, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__3|esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_101272, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_98806, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__5,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_101273, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95463, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95870, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|np__5=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_101576, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_101271, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_101274, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96178, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_100987, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95900, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_101275, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_98796, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__5,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96735, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_95873, negated_conjecture, (np__5=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_101577, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_102546, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_102547, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_102651, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_102351, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=k1_funct_1(esk2_0,np__3))).#trainneg
cnf(i_0_103645, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_102590, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_102693, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_103644, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_102692, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_102167, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5))).#trainneg
cnf(i_0_102337, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__3|np__3=np__1|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_102648, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__7,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_103646, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_103059, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_102694, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_103310, negated_conjecture, (esk4_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_103902, negated_conjecture, (esk4_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_102591, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_97645, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_95872, negated_conjecture, (np__5=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_95877, negated_conjecture, (np__7=np__3|np__5=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_98476, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_96614, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_96334, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__5|esk4_2(X1,X2)=np__7|r1_xxreal_0(np__9,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96077, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),np__9)|esk4_2(X1,X2)=np__3|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|~v1_xxreal_0(esk4_2(X1,X2))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_97646, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95868, negated_conjecture, (np__5=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96768, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_95867, negated_conjecture, (np__5=np__3|np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_95280, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk4_2(X2,X3)=np__3|r1_xxreal_0(np__5,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_97655, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_97641, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_103311, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_98477, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_104417, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104609, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104611, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_104639, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|r1_xxreal_0(np__7,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104922, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_104608, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104921, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104925, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104371, negated_conjecture, (np__7=np__5|esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104717, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__5|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_104610, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_105325, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_105324, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_104363, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7))).#trainneg
cnf(i_0_102171, negated_conjecture, (r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__3),k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0)),esk6_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_105326, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_96773, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_93324, negated_conjecture, (r1_xxreal_0(np__5,X1)|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|v1_abian(X1)|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~v1_int_1(X1)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_100212, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_95358, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk4_2(X2,X3)=np__3|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_94632, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__1|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_103903, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_94600, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__7,esk4_2(X1,X2))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_105518, negated_conjecture, (esk4_2(X1,X2)=np__3|r1_xxreal_0(X3,esk5_2(esk1_0,esk2_0))|r1_xxreal_0(np__5,esk4_2(X1,X2))|~r1_xxreal_0(X3,esk4_2(X1,X2))|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_100916, negated_conjecture, (esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|esk4_2(X1,X2)=np__1|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v1_xxreal_0(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96459, negated_conjecture, (r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__9)|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_96139, negated_conjecture, (r1_xxreal_0(X1,esk5_2(esk1_0,esk2_0))|esk4_2(X2,X3)=np__5|r1_xxreal_0(np__7,esk4_2(X2,X3))|~r1_xxreal_0(X1,esk4_2(X2,X3))|~v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_105814, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_105837, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_107287, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_107154, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_105834, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_106429, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_107156, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_105833, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_105835, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_107054, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_107162, negated_conjecture, (esk4_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_106759, negated_conjecture, (np__7=np__5|np__9=np__5|np__9=np__7|k1_funct_1(esk2_0,np__5)!=k1_funct_1(esk2_0,np__1))).#trainneg
cnf(i_0_106733, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|np__9=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|np__9=np__7)).#trainneg
cnf(i_0_106307, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_106308, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_106760, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|np__9=np__5|np__9=np__7)).#trainneg
cnf(i_0_106236, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__1|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~v1_int_1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_106468, negated_conjecture, (np__5=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_106465, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v7_ordinal1(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_107163, negated_conjecture, (k2_xcmplx_0(esk4_2(X1,esk2_0),np__2)=np__7|esk5_2(X1,esk2_0)=np__5|esk5_2(X1,esk2_0)=np__7|esk4_2(X1,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,esk2_0)=np__1|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_108514, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_108533, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__5|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_108624, negated_conjecture, (np__9=np__7|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_108246, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_108532, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_107705, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7))).#trainneg
cnf(i_0_108921, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_108931, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_107970, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__7),k1_funct_1(esk2_0,np__5),esk6_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_107967, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|r1_glib_000(esk1_0,k1_funct_1(esk2_0,np__5),k1_funct_1(esk2_0,np__7),esk6_2(esk1_0,esk2_0))|esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_107875, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__5|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=k1_funct_1(esk2_0,np__5))).#trainneg
cnf(i_0_108932, negated_conjecture, (np__9=np__7|np__9=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_108538, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_109218, negated_conjecture, (np__9=np__7|np__9=np__5|esk5_2(esk1_0,esk2_0)!=np__7)).#trainneg
cnf(i_0_109217, negated_conjecture, (np__9=np__7|np__9=np__5|esk5_2(esk1_0,esk2_0)!=np__9)).#trainneg
cnf(i_0_108433, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_107703, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__5|np__5!=np__3)).#trainneg
cnf(i_0_108434, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_108539, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(esk1_0,esk2_0)=np__1|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_108634, negated_conjecture, (np__9=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_110417, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__9,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_110375, negated_conjecture, (np__7=np__5|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__5)).#trainneg
cnf(i_0_110304, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_110305, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_110455, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__7|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_110620, negated_conjecture, (np__9=np__5|np__9=np__7|~r1_xxreal_0(np__7,np__5))).#trainneg
cnf(i_0_109822, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk4_2(esk1_0,esk2_0)=np__3|esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__7))).#trainneg
cnf(i_0_110539, negated_conjecture, (np__9=np__5|np__9=np__7|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_109867, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|np__9=np__7|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_110540, negated_conjecture, (np__9=np__5|np__9=np__7|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_110153, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|np__9=np__7|esk5_2(esk1_0,esk2_0)!=np__7)).#trainneg
cnf(i_0_110074, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|np__9=np__7|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=k1_funct_1(esk2_0,np__1))).#trainneg
cnf(i_0_106464, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_106237, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_106849, negated_conjecture, (esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|v1_abian(esk4_2(esk1_0,esk2_0))|v1_abian(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_107111, negated_conjecture, (esk4_2(X1,X2)=np__7|esk4_2(X1,X2)=np__5|r1_xxreal_0(np__9,esk4_2(X1,X2))|esk4_2(X1,X2)=esk5_2(esk1_0,esk2_0)|esk4_2(X1,X2)=np__1|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1))).#trainneg
cnf(i_0_106853, negated_conjecture, (esk4_2(esk1_0,esk2_0)=esk5_2(esk1_0,esk2_0)|esk4_2(esk1_0,esk2_0)=np__5|r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__1|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0))|v1_abian(esk4_2(esk1_0,esk2_0))|v1_abian(esk5_2(esk1_0,esk2_0))|~v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v1_int_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_106460, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk4_2(esk1_0,esk2_0)=np__3|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_109139, negated_conjecture, (np__9=np__7|np__9=np__5|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=k1_funct_1(esk2_0,np__1))).#trainneg
cnf(i_0_110622, negated_conjecture, (np__9=np__5|np__9=np__7|np__7!=np__5)).#trainneg
cnf(i_0_103111, negated_conjecture, (np__7=np__3|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_98436, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|r1_xxreal_0(np__3,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_111220, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__9|np__9=np__7|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__7,esk4_2(esk1_0,esk2_0))|~v1_xreal_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_93332, negated_conjecture, (X1=np__5|r1_xxreal_0(X1,esk4_2(esk1_0,esk2_0))|v1_abian(esk4_2(esk1_0,esk2_0))|v1_abian(X1)|esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|~r1_xxreal_0(X1,np__5)|~v1_int_1(esk4_2(esk1_0,esk2_0))|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_96584, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__3|k2_xcmplx_0(np__2,esk4_2(esk1_0,esk2_0))=np__5|esk5_2(esk1_0,esk2_0)=np__9|v1_abian(esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers))).#trainneg
cnf(i_0_101527, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__1|esk5_2(esk1_0,esk2_0)=np__9|np__7=np__5|v1_abian(esk4_2(esk1_0,esk2_0))|~r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k5_numbers)|~v7_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_108643, negated_conjecture, (np__9=np__7|k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|v1_abian(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~r1_xxreal_0(np__3,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2))|~v7_ordinal1(k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_111752, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_112279, negated_conjecture, (np__9=np__7|np__9=np__5)).#trainneg
cnf(i_0_111934, negated_conjecture, (np__9=np__7|esk4_2(esk1_0,esk2_0)=np__1|esk4_2(esk1_0,esk2_0)=np__3)).#trainneg
cnf(i_0_112277, negated_conjecture, (np__9=np__7|esk5_2(esk1_0,esk2_0)=np__9|esk5_2(esk1_0,esk2_0)=np__7)).#trainneg
cnf(i_0_112094, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__5))).#trainneg
cnf(i_0_112059, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~v7_ordinal1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_113202, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__3))).#trainneg
cnf(i_0_112060, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|~r1_xxreal_0(esk5_2(esk1_0,esk2_0),np__1)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_numbers))).#trainneg
cnf(i_0_112096, negated_conjecture, (k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)=np__3|esk5_2(esk1_0,esk2_0)=np__5|esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_113488, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__5|esk5_2(esk1_0,esk2_0)!=np__7)).#trainneg
cnf(i_0_87531, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(X2,X1)|v1_abian(X1)|~r1_xxreal_0(X2,np__7)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_48466, negated_conjecture, (r1_xxreal_0(np__4,np__5)|~r1_xxreal_0(np__4,np__3))).#trainneg
cnf(i_0_67124, negated_conjecture, (r1_xxreal_0(X1,esk4_2(X2,X3))|esk4_2(X2,X3)=np__1|~r1_xxreal_0(X1,X4)|~v1_xxreal_0(esk4_2(X2,X3))|~v1_xxreal_0(X4)|~v1_xxreal_0(X1)|~r1_xxreal_0(X4,np__1)|~v5_chord(X3,X2)|~m3_glib_001(X3,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_74885, negated_conjecture, (esk4_2(X1,X2)=np__1|esk4_2(X1,X2)=np__5|~r1_xxreal_0(np__7,np__3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r1_xxreal_0(esk4_2(X1,X2),np__3)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_23349, plain, (r1_xxreal_0(np__5,k3_finseq_1(X1))|r1_xxreal_0(np__7,np__3)|k2_xcmplx_0(esk5_2(X2,X1),np__2)=np__7|~v7_ordinal1(k3_finseq_1(X1))|~v5_chord(X1,X2)|~m3_glib_001(X1,X2)|~v1_glib_000(X2)|~v1_finset_1(X2)|~v1_funct_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8999, plain, (X1=np__7|k2_xcmplx_0(X1,np__2)=np__7|v1_abian(X1)|~r1_xxreal_0(X1,np__7)|~v1_xxreal_0(X1)|~r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46822, negated_conjecture, (k2_xcmplx_0(np__2,esk4_2(X1,esk2_0))=np__7|k2_xcmplx_0(np__2,esk4_2(X1,esk2_0))=np__5|esk4_2(X1,esk2_0)=np__1|~m1_subset_1(k2_xcmplx_0(np__2,esk4_2(X1,esk2_0)),k1_numbers)|~v5_chord(esk2_0,X1)|~m3_glib_001(esk2_0,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_76747, negated_conjecture, (r1_xxreal_0(esk4_2(X1,X2),X3)|r1_xxreal_0(X3,np__7)|v1_abian(X3)|r1_xxreal_0(np__5,esk4_2(X1,X2))|~v7_ordinal1(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_27660, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(X3,np__1)|r1_xxreal_0(np__1,X3)|v1_abian(X3)|v1_abian(esk4_2(X1,X2))|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2)))).#trainneg
cnf(i_0_56362, plain, (esk4_2(X1,X2)=np__3|r1_xxreal_0(np__5,esk4_2(X1,X2))|r1_xxreal_0(X3,np__3)|r1_xxreal_0(esk4_2(X4,X5),X3)|v1_abian(X3)|r1_xxreal_0(np__3,esk4_2(X4,X5))|v1_abian(esk4_2(X1,X2))|~v1_int_1(X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk4_2(X4,X5))|~v5_chord(X2,X1)|~m3_glib_001(X2,X1)|~v1_glib_000(X1)|~v1_finset_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(esk4_2(X1,X2))|~v5_chord(X5,X4)|~m3_glib_001(X5,X4)|~v1_glib_000(X4)|~v1_finset_1(X4)|~v1_funct_1(X4)|~v4_relat_1(X4,k5_numbers)|~v1_relat_1(X4))).#trainneg
cnf(i_0_22000, negated_conjecture, (esk4_2(esk1_0,esk2_0)=np__3|r1_xxreal_0(np__5,esk4_2(esk1_0,esk2_0))|r1_xxreal_0(esk4_2(esk1_0,esk2_0),X1)|r1_xxreal_0(X1,np__1)|v1_abian(X1)|~v1_int_1(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_97051, negated_conjecture, (esk5_2(esk1_0,esk2_0)=np__7|esk4_2(esk1_0,esk2_0)=np__1)).#trainneg
cnf(i_0_3264, negated_conjecture, (k1_funct_1(esk2_0,np__7)=k1_funct_1(esk2_0,np__1))).#trainneg
cnf(i_0_120, plain, (np__8=np__6)).#trainneg
cnf(i_0_1838, negated_conjecture, (k3_finseq_1(esk2_0)=np__7)).#trainneg
cnf(i_0_114958, negated_conjecture, (~r1_xxreal_0(np__5,k2_xcmplx_0(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_462, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=np__7|esk4_2(esk1_0,esk2_0)!=np__3)).#trainneg
cnf(i_0_114962, negated_conjecture, (~r1_xxreal_0(np__5,k1_nat_1(esk4_2(esk1_0,esk2_0),np__2)))).#trainneg
cnf(i_0_3255, negated_conjecture, (k1_card_1(esk2_0)=np__7)).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 130
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 199
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 197
# Processed clauses                    : 11476
# ...of these trivial                  : 182
# ...subsumed                          : 8055
# ...remaining for further processing  : 3238
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 990
# Backward-rewritten                   : 1506
# Generated clauses                    : 115656
# ...of the previous two non-trivial   : 68535
# ...frozen by parental guidance       : 22432
# ...of these subsequently unfrozen    : 0
# Contextual simplify-reflections      : 1361
# Paramodulations                      : 115571
# Factorizations                       : 80
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 737
#    Positive orientable unit clauses  : 131
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 593
# Current number of unprocessed clauses: 51669
# ...number of literals in the above   : 575672
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5740
# Proof object given clauses           : 392
# Proof search given clauses           : 3239
# Clause-clause subsumption calls (NU) : 5919139
# Rec. Clause-clause subsumption calls : 94367
# Non-unit clause-clause subsumptions  : 9482
# Unit Clause-clause subsumption calls : 16261
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57
# BW rewrite match successes           : 57
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4346900

# -------------------------------------------------
# User time                : 65.997 s
# System time              : 1.478 s
# Total time               : 67.475 s
# Maximum resident set size: 306196 pages
65.99user 1.51system 2:36.07elapsed 43%CPU (0avgtext+0avgdata 306376maxresident)k
0inputs+1880outputs (0major+348306minor)pagefaults 0swaps
