2020-05-31 12:29:09.487557: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55
2020-05-31 12:29:09.569234: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 12:29:09.789828: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 12:29:10.658054: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 12:29:12.931104: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3443522 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.594 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', fc6_ordinal1)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:(((((~(v1_xboole_0(X1))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&m1_subset_1(X4,X1))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k3_funct_2)).
fof(t4_setlim_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,k5_numbers,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))=>(r2_hidden(X1,k2_relset_1(X2,X3))<=>?[X4]:(m1_subset_1(X4,k5_numbers)&X1=k3_funct_2(k5_numbers,X2,X3,X4))))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t4_setlim_1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k5_numbers)).
fof(t8_rinfsup1, conjecture, ![X1]:(v1_xreal_0(X1)=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_seq_2(X2)=>(X1=k2_rinfsup1(X2)<=>(![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>r1_xxreal_0(X1,k1_seq_1(X2,X3)))&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>r1_xxreal_0(k2_xcmplx_0(X1,X3),k1_seq_1(X2,X4))))))))))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t8_rinfsup1)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k1_seq_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', cc1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', cc2_relset_1)).
fof(fc63_valued_0, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>v1_xreal_0(k1_funct_1(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', fc63_valued_0)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', cc18_valued_0)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k2_relset_1)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', cc21_membered)).
fof(dt_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', dt_k2_relset_1)).
fof(redefinition_k5_seq_4, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(k1_numbers))=>k5_seq_4(X1)=k3_seq_4(X1)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k5_seq_4)).
fof(d2_rinfsup1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>k2_rinfsup1(X1)=k5_seq_4(k2_relset_1(k1_numbers,X1))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', d2_rinfsup1)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', fc3_membered)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t4_subset)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', dt_k1_relset_1)).
fof(d2_seq_4, axiom, ![X1]:(v3_membered(X1)=>(v3_xxreal_2(X1)=>(v1_xboole_0(X1)|![X2]:(v1_xreal_0(X2)=>(X2=k3_seq_4(X1)<=>(![X3]:(v1_xreal_0(X3)=>(r2_hidden(X3,X1)=>r1_xxreal_0(X2,X3)))&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(v1_xreal_0(X4)=>~((r2_hidden(X4,X1)&~(r1_xxreal_0(k9_binop_2(X2,X3),X4)))))))))))))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', d2_seq_4)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t7_boole)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', fc1_numbers)).
fof(t6_rinfsup1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_seq_2(X1)<=>v3_xxreal_2(k2_relset_1(k1_numbers,X1)))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t6_rinfsup1)).
fof(d3_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X1))&X3=k1_funct_1(X1,X4))))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', d3_funct_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k1_relset_1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_m2_subset_1)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', commutativity_k2_xcmplx_0)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', cc3_xreal_0)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', dt_k5_numbers)).
fof(redefinition_k9_binop_2, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>k9_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k9_binop_2)).
fof(commutativity_k9_binop_2, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>k9_binop_2(X1,X2)=k9_binop_2(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', commutativity_k9_binop_2)).
fof(t1_numerals, axiom, m1_subset_1(k1_xboole_0,k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t1_numerals)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', redefinition_k6_numbers)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t1_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t8_rinfsup1', t2_subset)).
fof(c_0_34, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_35, plain, ![X41, X42, X43, X44]:(v1_xboole_0(X41)|~v1_funct_1(X43)|~v1_funct_2(X43,X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))|~m1_subset_1(X44,X41)|k3_funct_2(X41,X42,X43,X44)=k1_funct_1(X43,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
fof(c_0_36, plain, ![X35, X36, X37, X39]:(((m1_subset_1(esk12_3(X35,X36,X37),k5_numbers)|~r2_hidden(X35,k2_relset_1(X36,X37))|(~v1_funct_1(X37)|~v1_funct_2(X37,k5_numbers,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X36))))|v1_xboole_0(X36))&(X35=k3_funct_2(k5_numbers,X36,X37,esk12_3(X35,X36,X37))|~r2_hidden(X35,k2_relset_1(X36,X37))|(~v1_funct_1(X37)|~v1_funct_2(X37,k5_numbers,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X36))))|v1_xboole_0(X36)))&(~m1_subset_1(X39,k5_numbers)|X35!=k3_funct_2(k5_numbers,X36,X37,X39)|r2_hidden(X35,k2_relset_1(X36,X37))|(~v1_funct_1(X37)|~v1_funct_2(X37,k5_numbers,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X36))))|v1_xboole_0(X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_setlim_1])])])])])])).
cnf(c_0_37, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_38, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_39, negated_conjecture, ~(![X1]:(v1_xreal_0(X1)=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_seq_2(X2)=>(X1=k2_rinfsup1(X2)<=>(![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>r1_xxreal_0(X1,k1_seq_1(X2,X3)))&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>r1_xxreal_0(k2_xcmplx_0(X1,X3),k1_seq_1(X2,X4)))))))))))), inference(assume_negation,[status(cth)],[t8_rinfsup1])).
cnf(c_0_40, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (X1=k3_funct_2(k5_numbers,X2,X3,esk12_3(X1,X2,X3))|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_37])).
cnf(c_0_43, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_38])).
fof(c_0_44, negated_conjecture, ![X9, X10, X11]:(v1_xreal_0(esk1_0)&(((v1_funct_1(esk2_0)&v1_funct_2(esk2_0,k5_numbers,k1_numbers))&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))&(v2_seq_2(esk2_0)&((((v1_xreal_0(esk4_0)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|esk1_0!=k2_rinfsup1(esk2_0))&((~r1_xxreal_0(esk4_0,k6_numbers)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|esk1_0!=k2_rinfsup1(esk2_0))&(~m2_subset_1(X9,k1_numbers,k5_numbers)|r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X9))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|esk1_0!=k2_rinfsup1(esk2_0))))&((v1_xreal_0(esk4_0)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|esk1_0!=k2_rinfsup1(esk2_0))&((~r1_xxreal_0(esk4_0,k6_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|esk1_0!=k2_rinfsup1(esk2_0))&(~m2_subset_1(X9,k1_numbers,k5_numbers)|r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X9))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|esk1_0!=k2_rinfsup1(esk2_0)))))&((~m2_subset_1(X10,k1_numbers,k5_numbers)|r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,X10))|esk1_0=k2_rinfsup1(esk2_0))&((m2_subset_1(esk5_1(X11),k1_numbers,k5_numbers)|r1_xxreal_0(X11,k6_numbers)|~v1_xreal_0(X11)|esk1_0=k2_rinfsup1(esk2_0))&(~r1_xxreal_0(k2_xcmplx_0(esk1_0,X11),k1_seq_1(esk2_0,esk5_1(X11)))|r1_xxreal_0(X11,k6_numbers)|~v1_xreal_0(X11)|esk1_0=k2_rinfsup1(esk2_0)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_39])])])])])])).
fof(c_0_45, plain, ![X57, X58]:(~v1_relat_1(X57)|~v1_funct_1(X57)|~v3_valued_0(X57)|k1_seq_1(X57,X58)=k1_funct_1(X57,X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
fof(c_0_46, plain, ![X52, X53, X54]:(~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X53)))|v1_relat_1(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_47, plain, ![X68, X69, X70]:((v4_relat_1(X70,X68)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))&(v5_relat_1(X70,X69)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_48, plain, ![X86, X87]:(~v1_relat_1(X86)|~v1_funct_1(X86)|~v3_valued_0(X86)|v1_xreal_0(k1_funct_1(X86,X87))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc63_valued_0])])).
cnf(c_0_49, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (k3_funct_2(k5_numbers,X2,X3,esk12_3(X1,X2,X3))=X1|v1_xboole_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, plain, (m1_subset_1(esk12_3(X1,X2,X3),k5_numbers)|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|esk1_0!=k2_rinfsup1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_58, plain, ![X49, X50, X51]:(~v3_membered(X50)|(~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X49,X50)))|v3_valued_0(X51))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
fof(c_0_59, plain, ![X55, X56]:(~v1_relat_1(X56)|~v5_relat_1(X56,X55)|k2_relset_1(X55,X56)=k10_xtuple_0(X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_60, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_61, plain, ![X117, X118]:(~v3_membered(X117)|(~m1_subset_1(X118,k1_zfmisc_1(X117))|v3_membered(X118))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
fof(c_0_62, plain, ![X80, X81]:(~v1_relat_1(X81)|~v5_relat_1(X81,X80)|m1_subset_1(k2_relset_1(X80,X81),k1_zfmisc_1(X80))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_relset_1])])).
fof(c_0_63, plain, ![X94]:(~m1_subset_1(X94,k1_zfmisc_1(k1_numbers))|k5_seq_4(X94)=k3_seq_4(X94)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_seq_4])])).
fof(c_0_64, plain, ![X79]:(~v1_funct_1(X79)|~v1_funct_2(X79,k5_numbers,k1_numbers)|~m1_subset_1(X79,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k2_rinfsup1(X79)=k5_seq_4(k2_relset_1(k1_numbers,X79))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_rinfsup1])])).
cnf(c_0_65, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_66, plain, (k1_funct_1(X2,esk12_3(X3,X1,X2))=X3|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X3,X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~r2_hidden(X3,k2_relset_1(X1,X2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])])).
cnf(c_0_67, plain, (v1_xboole_0(X2)|m1_subset_1(esk12_3(X1,X2,X3),k5_numbers)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_68, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|esk1_0!=k2_rinfsup1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_69, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_70, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_71, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_72, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_73, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_74, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_75, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_76, plain, ![X114, X115, X116]:(~r2_hidden(X114,X115)|~m1_subset_1(X115,k1_zfmisc_1(X116))|m1_subset_1(X114,X116)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_77, plain, ![X75, X76]:(~v1_relat_1(X76)|~v4_relat_1(X76,X75)|m1_subset_1(k1_relset_1(X75,X76),k1_zfmisc_1(X75))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
fof(c_0_78, plain, ![X13, X14, X15, X16, X20]:(((~v1_xreal_0(X15)|(~r2_hidden(X15,X13)|r1_xxreal_0(X14,X15))|X14!=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&((v1_xreal_0(esk6_3(X13,X14,X16))|r1_xxreal_0(X16,k6_numbers)|~v1_xreal_0(X16)|X14!=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&((r2_hidden(esk6_3(X13,X14,X16),X13)|r1_xxreal_0(X16,k6_numbers)|~v1_xreal_0(X16)|X14!=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&(~r1_xxreal_0(k9_binop_2(X14,X16),esk6_3(X13,X14,X16))|r1_xxreal_0(X16,k6_numbers)|~v1_xreal_0(X16)|X14!=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13)))))&(((v1_xreal_0(esk8_2(X13,X14))|v1_xreal_0(esk7_2(X13,X14))|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&((~r1_xxreal_0(esk8_2(X13,X14),k6_numbers)|v1_xreal_0(esk7_2(X13,X14))|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&(~v1_xreal_0(X20)|(~r2_hidden(X20,X13)|r1_xxreal_0(k9_binop_2(X14,esk8_2(X13,X14)),X20))|v1_xreal_0(esk7_2(X13,X14))|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))))&(((v1_xreal_0(esk8_2(X13,X14))|r2_hidden(esk7_2(X13,X14),X13)|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&((~r1_xxreal_0(esk8_2(X13,X14),k6_numbers)|r2_hidden(esk7_2(X13,X14),X13)|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&(~v1_xreal_0(X20)|(~r2_hidden(X20,X13)|r1_xxreal_0(k9_binop_2(X14,esk8_2(X13,X14)),X20))|r2_hidden(esk7_2(X13,X14),X13)|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))))&((v1_xreal_0(esk8_2(X13,X14))|~r1_xxreal_0(X14,esk7_2(X13,X14))|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&((~r1_xxreal_0(esk8_2(X13,X14),k6_numbers)|~r1_xxreal_0(X14,esk7_2(X13,X14))|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))&(~v1_xreal_0(X20)|(~r2_hidden(X20,X13)|r1_xxreal_0(k9_binop_2(X14,esk8_2(X13,X14)),X20))|~r1_xxreal_0(X14,esk7_2(X13,X14))|X14=k3_seq_4(X13)|~v1_xreal_0(X14)|v1_xboole_0(X13)|~v3_xxreal_2(X13)|~v3_membered(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_seq_4])])])])])])).
fof(c_0_79, plain, ![X131, X132]:(~r2_hidden(X131,X132)|~v1_xboole_0(X132)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_80, plain, (r2_hidden(X2,k2_relset_1(X3,X4))|v1_xboole_0(X3)|~m1_subset_1(X1,k5_numbers)|X2!=k3_funct_2(k5_numbers,X3,X4,X1)|~v1_funct_1(X4)|~v1_funct_2(X4,k5_numbers,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_81, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_82, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_60])).
fof(c_0_83, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
fof(c_0_84, plain, ![X31]:((~v2_seq_2(X31)|v3_xxreal_2(k2_relset_1(k1_numbers,X31))|(~v1_funct_1(X31)|~v1_funct_2(X31,k5_numbers,k1_numbers)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~v3_xxreal_2(k2_relset_1(k1_numbers,X31))|v2_seq_2(X31)|(~v1_funct_1(X31)|~v1_funct_2(X31,k5_numbers,k1_numbers)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_rinfsup1])])])).
cnf(c_0_85, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_86, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_87, plain, (k5_seq_4(X1)=k3_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_88, plain, (k2_rinfsup1(X1)=k5_seq_4(k2_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_89, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_90, plain, (k1_funct_1(X1,esk12_3(X2,X3,X1))=X2|v1_xboole_0(X3)|~r2_hidden(X2,k2_relset_1(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_91, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_92, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])])])).
cnf(c_0_93, negated_conjecture, (v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73])])).
cnf(c_0_94, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_95, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_75])).
fof(c_0_96, plain, ![X21, X22, X23, X25, X26, X27, X29]:((((r2_hidden(esk9_3(X21,X22,X23),k9_xtuple_0(X21))|~r2_hidden(X23,X22)|X22!=k10_xtuple_0(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21)))&(X23=k1_funct_1(X21,esk9_3(X21,X22,X23))|~r2_hidden(X23,X22)|X22!=k10_xtuple_0(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21))))&(~r2_hidden(X26,k9_xtuple_0(X21))|X25!=k1_funct_1(X21,X26)|r2_hidden(X25,X22)|X22!=k10_xtuple_0(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21))))&((~r2_hidden(esk10_2(X21,X27),X27)|(~r2_hidden(X29,k9_xtuple_0(X21))|esk10_2(X21,X27)!=k1_funct_1(X21,X29))|X27=k10_xtuple_0(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21)))&((r2_hidden(esk11_2(X21,X27),k9_xtuple_0(X21))|r2_hidden(esk10_2(X21,X27),X27)|X27=k10_xtuple_0(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21)))&(esk10_2(X21,X27)=k1_funct_1(X21,esk11_2(X21,X27))|r2_hidden(esk10_2(X21,X27),X27)|X27=k10_xtuple_0(X21)|(~v1_relat_1(X21)|~v1_funct_1(X21)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_1])])])])])])).
cnf(c_0_97, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_98, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_99, plain, ![X59, X60]:(~v1_relat_1(X60)|~v4_relat_1(X60,X59)|k1_relset_1(X59,X60)=k9_xtuple_0(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_100, plain, (r1_xxreal_0(X3,X1)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,X2)|X3!=k3_seq_4(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X2)|~v3_membered(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_101, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_102, plain, (v1_xboole_0(X3)|r2_hidden(X2,k2_relset_1(X3,X4))|X2!=k3_funct_2(k5_numbers,X3,X4,X1)|~v1_funct_1(X4)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X4,k5_numbers,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_103, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_104, negated_conjecture, (v5_relat_1(esk2_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_73])])).
cnf(c_0_105, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_106, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_107, plain, (v3_xxreal_2(k2_relset_1(k1_numbers,X1))|~v2_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_108, negated_conjecture, (v2_seq_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_109, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_110, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_111, plain, (k5_seq_4(X1)=k3_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_112, plain, (k5_seq_4(k2_relset_1(k1_numbers,X1))=k2_rinfsup1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_113, plain, (v1_xreal_0(X2)|v1_xboole_0(X3)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k2_relset_1(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_114, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_70]), c_0_71])])])).
cnf(c_0_115, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~v3_valued_0(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_116, negated_conjecture, (v3_valued_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_73]), c_0_95])])])).
cnf(c_0_117, plain, (X1=k1_funct_1(X2,esk9_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_96])).
fof(c_0_118, plain, ![X32, X33, X34]:((~m2_subset_1(X34,X32,X33)|m1_subset_1(X34,X33)|(v1_xboole_0(X32)|v1_xboole_0(X33)|~m1_subset_1(X33,k1_zfmisc_1(X32))))&(~m1_subset_1(X34,X33)|m2_subset_1(X34,X32,X33)|(v1_xboole_0(X32)|v1_xboole_0(X33)|~m1_subset_1(X33,k1_zfmisc_1(X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_119, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_120, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_121, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_122, plain, ![X66, X67]:(~v1_xcmplx_0(X66)|~v1_xcmplx_0(X67)|k2_xcmplx_0(X66,X67)=k2_xcmplx_0(X67,X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
fof(c_0_123, plain, ![X82]:(~v1_xreal_0(X82)|v1_xcmplx_0(X82)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_124, plain, (v1_xboole_0(X2)|r1_xxreal_0(X3,X1)|X3!=k3_seq_4(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_125, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_126, plain, (r2_hidden(k3_funct_2(k5_numbers,X1,X2,X3),k2_relset_1(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_102])])).
cnf(c_0_127, negated_conjecture, (k2_relset_1(k1_numbers,esk2_0)=k10_xtuple_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_93])])])).
cnf(c_0_128, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_129, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_130, plain, (v3_xxreal_2(k2_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~v2_seq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_131, negated_conjecture, (v2_seq_2(esk2_0)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_132, plain, (v3_membered(k2_relset_1(X1,X2))|~v3_membered(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_133, plain, (k3_seq_4(k2_relset_1(k1_numbers,X1))=k2_rinfsup1(X1)|~m1_subset_1(k2_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_134, plain, (v1_xboole_0(X1)|v1_xreal_0(X2)|~v3_valued_0(X3)|~r2_hidden(X2,k2_relset_1(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_113, c_0_72])).
cnf(c_0_135, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~v3_valued_0(esk2_0)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_93])])).
cnf(c_0_136, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_116])])])).
cnf(c_0_137, plain, (k1_funct_1(X2,esk9_3(X2,X3,X1))=X1|X3!=k10_xtuple_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_138, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_139, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_140, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_141, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_142, plain, (r2_hidden(esk9_3(X1,X2,X3),k9_xtuple_0(X1))|~r2_hidden(X3,X2)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_143, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_144, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_145, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_146, negated_conjecture, (v1_xreal_0(esk4_0)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|esk1_0!=k2_rinfsup1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_147, plain, (r1_xxreal_0(X1,X2)|X1!=k3_seq_4(X3)|~r2_hidden(X2,X3)|~v3_xxreal_2(X3)|~v3_membered(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_148, negated_conjecture, (r2_hidden(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1),k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_73]), c_0_128]), c_0_71])]), c_0_129])])).
cnf(c_0_149, negated_conjecture, (v3_xxreal_2(k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_127]), c_0_131]), c_0_73]), c_0_128]), c_0_71])])])).
cnf(c_0_150, negated_conjecture, (v3_membered(k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_127]), c_0_104]), c_0_93]), c_0_95])])])).
cnf(c_0_151, negated_conjecture, (k3_seq_4(k10_xtuple_0(esk2_0))=k2_rinfsup1(esk2_0)|~m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_127]), c_0_73]), c_0_128]), c_0_71])])])).
cnf(c_0_152, negated_conjecture, (m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_127]), c_0_104]), c_0_93])])])).
cnf(c_0_153, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_127]), c_0_116]), c_0_73]), c_0_128]), c_0_71])]), c_0_129])])).
cnf(c_0_154, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_116])])])).
cnf(c_0_155, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X2)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_93]), c_0_71])])])).
cnf(c_0_156, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_157, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_158, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141])])).
cnf(c_0_159, plain, (r2_hidden(esk9_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_160, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_143])).
fof(c_0_161, plain, ![X73, X74]:(~v1_xreal_0(X73)|~v1_xreal_0(X74)|k9_binop_2(X73,X74)=k2_xcmplx_0(X73,X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_binop_2])])).
cnf(c_0_162, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_163, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_164, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_165, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|X1!=k3_seq_4(k10_xtuple_0(esk2_0))|~v1_xreal_0(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_149]), c_0_150])])])).
cnf(c_0_166, negated_conjecture, (k3_seq_4(k10_xtuple_0(esk2_0))=k2_rinfsup1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_151, c_0_152])])])).
cnf(c_0_167, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_148])])).
cnf(c_0_168, negated_conjecture, (~r1_xxreal_0(esk4_0,k6_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|esk1_0!=k2_rinfsup1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_169, plain, (r2_hidden(X1,k2_relset_1(X2,X3))|v1_xboole_0(X2)|X1!=k1_funct_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_49]), c_0_51])])).
cnf(c_0_170, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_137]), c_0_93]), c_0_71])])])).
cnf(c_0_171, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_157])]), c_0_51]), c_0_129])])).
cnf(c_0_172, plain, (m1_subset_1(esk9_3(X1,X2,X3),X4)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159])])).
cnf(c_0_173, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_73])])).
cnf(c_0_174, plain, (k9_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_175, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
cnf(c_0_176, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_70]), c_0_71])])])).
cnf(c_0_177, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_165, c_0_166]), c_0_167])).
cnf(c_0_178, negated_conjecture, (v1_xreal_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_179, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_180, negated_conjecture, (v1_xreal_0(esk4_0)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|esk1_0!=k2_rinfsup1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_181, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[c_0_168])).
cnf(c_0_182, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relset_1(X3,X1))|v1_xboole_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_169])])).
cnf(c_0_183, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_156]), c_0_157])]), c_0_51]), c_0_129])])).
fof(c_0_184, plain, ![X89, X90]:(~v1_xreal_0(X89)|~v1_xreal_0(X90)|k9_binop_2(X89,X90)=k9_binop_2(X90,X89)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k9_binop_2])])).
cnf(c_0_185, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_173]), c_0_93]), c_0_71])])])).
cnf(c_0_186, plain, (k2_xcmplx_0(X1,X2)=k9_binop_2(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_187, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_163])])).
cnf(c_0_188, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_176, c_0_116]), c_0_93])])).
cnf(c_0_189, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_49]), c_0_73]), c_0_128]), c_0_71])]), c_0_51])])).
cnf(c_0_190, negated_conjecture, (v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_191, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_192, negated_conjecture, (v1_xreal_0(esk4_0)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_193, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))|~r1_xxreal_0(esk4_0,k6_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_70]), c_0_71])])])).
cnf(c_0_194, negated_conjecture, (r2_hidden(k1_funct_1(esk2_0,X1),k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_127]), c_0_73]), c_0_128]), c_0_71])]), c_0_129])])).
cnf(c_0_195, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_172]), c_0_173]), c_0_93]), c_0_71])])])).
cnf(c_0_196, negated_conjecture, (r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,X1))|esk1_0=k2_rinfsup1(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_197, plain, (r1_xxreal_0(X2,k6_numbers)|v1_xboole_0(X3)|~r1_xxreal_0(k9_binop_2(X1,X2),esk6_3(X3,X1,X2))|~v1_xreal_0(X2)|X1!=k3_seq_4(X3)|~v1_xreal_0(X1)|~v3_xxreal_2(X3)|~v3_membered(X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_198, plain, (k9_binop_2(X1,X2)=k9_binop_2(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_184])).
cnf(c_0_199, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_185])])).
cnf(c_0_200, plain, (k2_xcmplx_0(X2,X1)=k9_binop_2(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187])])).
cnf(c_0_201, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~m1_subset_1(esk3_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189]), c_0_190])])])).
cnf(c_0_202, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)|v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_157])]), c_0_129]), c_0_51])])).
cnf(c_0_203, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~v3_valued_0(esk2_0)|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))|~r1_xxreal_0(esk4_0,k6_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_193, c_0_93])])).
cnf(c_0_204, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|X1!=k3_seq_4(k10_xtuple_0(esk2_0))|~v1_xreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_194]), c_0_149]), c_0_150])])])).
cnf(c_0_205, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|~r1_xxreal_0(esk4_0,k6_numbers)|esk1_0!=k2_rinfsup1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_206, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_195])])).
cnf(c_0_207, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_208, plain, (v1_xboole_0(X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k3_seq_4(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_membered(X3)|~v3_xxreal_2(X3)|~r1_xxreal_0(k9_binop_2(X1,X2),esk6_3(X3,X1,X2))), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_209, plain, (k9_binop_2(X1,X2)=k9_binop_2(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_210, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~v1_xreal_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_190])])])).
cnf(c_0_211, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0), inference(csr,[status(thm)],[c_0_201, c_0_202])).
cnf(c_0_212, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))|~r1_xxreal_0(esk4_0,k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_203, c_0_116])])])).
cnf(c_0_213, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(k1_funct_1(esk2_0,X2))|~v1_xreal_0(X1)), inference(rw,[status(thm)],[c_0_204, c_0_166])).
cnf(c_0_214, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers)), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_215, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_200]), c_0_190])])])).
cnf(c_0_216, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_70]), c_0_71])])])).
cnf(c_0_217, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|esk1_0=k2_rinfsup1(esk2_0)|~r1_xxreal_0(k2_xcmplx_0(esk1_0,X1),k1_seq_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_218, plain, (r1_xxreal_0(k9_binop_2(X3,esk8_2(X2,X3)),X1)|r2_hidden(esk7_2(X2,X3),X2)|X3=k3_seq_4(X2)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X2)|~v3_membered(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_219, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,k6_numbers)|X3!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k9_binop_2(X2,X3),esk6_3(X1,X3,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_220, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))), inference(csr,[status(thm)],[c_0_210, c_0_211])).
cnf(c_0_221, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers)|~m1_subset_1(esk3_0,k5_numbers)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_190])])])).
cnf(c_0_222, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_194])])).
cnf(c_0_223, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_214]), c_0_157])]), c_0_129]), c_0_51])])).
cnf(c_0_224, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(csr,[status(thm)],[c_0_215, c_0_211])).
cnf(c_0_225, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|~v3_valued_0(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_216, c_0_93])])])).
cnf(c_0_226, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(esk1_0,X1),k1_seq_1(esk2_0,esk5_1(X1)))), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_227, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|r2_hidden(esk7_2(X2,X3),X2)|r1_xxreal_0(k9_binop_2(X3,esk8_2(X2,X3)),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_228, plain, (r1_xxreal_0(k9_binop_2(X3,esk8_2(X2,X3)),X1)|X3=k3_seq_4(X2)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~r2_hidden(X1,X2)|~r1_xxreal_0(X3,esk7_2(X2,X3))|~v1_xreal_0(X3)|~v3_xxreal_2(X2)|~v3_membered(X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_229, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(esk4_0,k6_numbers)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k3_seq_4(X1)!=esk1_0|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk6_3(X1,esk1_0,esk4_0),k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220]), c_0_190])])])).
cnf(c_0_230, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_221, c_0_222]), c_0_223])).
cnf(c_0_231, plain, (r2_hidden(esk6_3(X1,X2,X3),X1)|r1_xxreal_0(X3,k6_numbers)|v1_xboole_0(X1)|~v1_xreal_0(X3)|X2!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_232, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(esk4_0,k6_numbers)|k3_seq_4(X1)!=esk1_0|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk6_3(X1,esk1_0,esk4_0),k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_224]), c_0_190])])])).
cnf(c_0_233, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_225, c_0_116])])).
cnf(c_0_234, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k9_binop_2(esk1_0,X1),k1_seq_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_186]), c_0_190])])])).
cnf(c_0_235, plain, (X1=k3_seq_4(X2)|r2_hidden(esk7_2(X2,X1),X2)|r1_xxreal_0(k9_binop_2(X1,esk8_2(X2,X1)),X3)|~r2_hidden(X3,X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_227, c_0_125])).
cnf(c_0_236, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|r1_xxreal_0(k9_binop_2(X3,esk8_2(X2,X3)),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~r2_hidden(X1,X2)|~r1_xxreal_0(X3,esk7_2(X2,X3))), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_237, negated_conjecture, (v1_xboole_0(X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|k3_seq_4(X1)!=esk1_0|~r2_hidden(esk6_3(X1,esk1_0,esk4_0),k10_xtuple_0(esk2_0))|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_229, c_0_211]), c_0_230])).
cnf(c_0_238, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|r2_hidden(esk6_3(X1,X2,X3),X1)|X2!=k3_seq_4(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_239, negated_conjecture, (v1_xboole_0(X1)|k2_rinfsup1(esk2_0)!=esk1_0|k3_seq_4(X1)!=esk1_0|~r2_hidden(esk6_3(X1,esk1_0,esk4_0),k10_xtuple_0(esk2_0))|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_232, c_0_211]), c_0_230])).
cnf(c_0_240, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X2)|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_137]), c_0_93]), c_0_71])])])).
cnf(c_0_241, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k9_binop_2(esk1_0,X1),k1_funct_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1)|~v3_valued_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_70]), c_0_93]), c_0_71])])])).
cnf(c_0_242, negated_conjecture, (X1=k3_seq_4(k10_xtuple_0(esk2_0))|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|r1_xxreal_0(k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1)),k1_funct_1(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_194]), c_0_149]), c_0_150])])])).
cnf(c_0_243, negated_conjecture, (m2_subset_1(esk5_1(X1),k1_numbers,k5_numbers)|r1_xxreal_0(X1,k6_numbers)|esk1_0=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_244, plain, (X1=k3_seq_4(X2)|r1_xxreal_0(k9_binop_2(X1,esk8_2(X2,X1)),X3)|~r2_hidden(X3,X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(X1,esk7_2(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X3)), inference(csr,[status(thm)],[c_0_236, c_0_125])).
cnf(c_0_245, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_relset_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_246, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|r1_xxreal_0(esk4_0,k6_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~v1_xreal_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_166]), c_0_149]), c_0_150]), c_0_190])])])).
cnf(c_0_247, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(split_conjunct,[status(thm)],[t1_numerals])).
cnf(c_0_248, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_249, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk4_0,k6_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_238]), c_0_166]), c_0_149]), c_0_150]), c_0_190])])])).
cnf(c_0_250, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_156]), c_0_157])]), c_0_51]), c_0_129])])).
cnf(c_0_251, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k9_binop_2(esk1_0,X1),k1_funct_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_241, c_0_116])])).
cnf(c_0_252, negated_conjecture, (X1=k2_rinfsup1(esk2_0)|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|r1_xxreal_0(k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1)),k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_242, c_0_166]), c_0_222])).
cnf(c_0_253, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk5_1(X1),k1_numbers,k5_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_254, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(X1,esk1_0),k6_numbers)|k3_seq_4(X1)=esk1_0|~v1_xreal_0(esk8_2(X1,esk1_0))|~r2_hidden(k1_seq_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))),X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,esk7_2(X1,esk1_0))|~v1_xreal_0(k1_seq_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_244]), c_0_190])])])).
cnf(c_0_255, negated_conjecture, (~v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_127]), c_0_73]), c_0_128]), c_0_71])]), c_0_129])])).
cnf(c_0_256, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_246, c_0_211]), c_0_230])).
cnf(c_0_257, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_258, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_259, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_249, c_0_211]), c_0_230])).
fof(c_0_260, plain, ![X103, X104]:(~r2_hidden(X103,X104)|m1_subset_1(X103,X104)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_261, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_172]), c_0_173]), c_0_93]), c_0_71])])])).
cnf(c_0_262, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k10_xtuple_0(esk2_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))|~m1_subset_1(esk5_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252]), c_0_190])])])).
cnf(c_0_263, negated_conjecture, (m1_subset_1(esk5_1(X1),k5_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_253]), c_0_157])]), c_0_129]), c_0_51])])).
cnf(c_0_264, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(X1)=esk1_0|r1_xxreal_0(esk8_2(X1,esk1_0),k6_numbers)|~r2_hidden(k1_funct_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))),X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,esk7_2(X1,esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))))|~v1_xreal_0(esk8_2(X1,esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_70]), c_0_116]), c_0_93]), c_0_71])])])).
cnf(c_0_265, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|v1_xboole_0(X2)|~m2_subset_1(esk12_3(X1,X2,esk2_0),k1_numbers,k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_90]), c_0_71])])])).
cnf(c_0_266, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256])])).
cnf(c_0_267, plain, (m1_subset_1(k6_numbers,k5_numbers)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_257, c_0_258]), c_0_43])).
cnf(c_0_268, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_70]), c_0_116]), c_0_93]), c_0_71])])])).
fof(c_0_269, plain, ![X143, X144]:(~m1_subset_1(X143,X144)|(v1_xboole_0(X144)|r2_hidden(X143,X144))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_270, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_260])).
cnf(c_0_271, plain, (r2_hidden(esk7_2(X1,X2),X1)|X2=k3_seq_4(X1)|v1_xboole_0(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_272, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_261])])).
cnf(c_0_273, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k10_xtuple_0(esk2_0))|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(csr,[status(thm)],[c_0_262, c_0_263])).
cnf(c_0_274, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk5_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0))))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))|~m1_subset_1(esk5_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0)),k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_194]), c_0_166]), c_0_149]), c_0_150])])])).
cnf(c_0_275, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,X2)|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk12_3(X2,X1,esk2_0),k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_156]), c_0_157])]), c_0_51]), c_0_129])])).
cnf(c_0_276, plain, (v1_xreal_0(esk8_2(X1,X2))|r2_hidden(esk7_2(X1,X2),X1)|X2=k3_seq_4(X1)|v1_xboole_0(X1)|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_277, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267])])).
cnf(c_0_278, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|k2_rinfsup1(esk2_0)!=esk1_0|~m1_subset_1(esk3_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_189]), c_0_190])])])).
cnf(c_0_279, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_269])).
cnf(c_0_280, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_270])).
cnf(c_0_281, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|r2_hidden(esk7_2(X1,X2),X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)), inference(evalgc,[status(thm)],[c_0_271])).
cnf(c_0_282, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_273])])).
cnf(c_0_283, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_274, c_0_222]), c_0_263])).
cnf(c_0_284, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,X2)|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_67]), c_0_71])])])).
cnf(c_0_285, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,X2))|r2_hidden(esk7_2(X1,X2),X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_286, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_277]), c_0_157])]), c_0_129]), c_0_51])])).
cnf(c_0_287, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~m1_subset_1(esk3_0,k5_numbers)), inference(csr,[status(thm)],[c_0_278, c_0_255])).
cnf(c_0_288, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_289, plain, (m1_subset_1(esk7_2(X1,X2),X1)|X2=k3_seq_4(X1)|v1_xboole_0(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281])])).
cnf(c_0_290, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(csr,[status(thm)],[c_0_282, c_0_283])).
cnf(c_0_291, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)|~r1_xxreal_0(X2,esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_292, plain, (v1_xreal_0(esk8_2(X1,X2))|X2=k3_seq_4(X1)|v1_xboole_0(X1)|~r1_xxreal_0(X2,esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_293, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_285])])).
cnf(c_0_294, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0), inference(csr,[status(thm)],[c_0_286, c_0_287])).
cnf(c_0_295, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_288])])).
cnf(c_0_296, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|m1_subset_1(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k10_xtuple_0(esk2_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290]), c_0_166]), c_0_149]), c_0_150]), c_0_190])])])).
cnf(c_0_297, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)|~r1_xxreal_0(X2,esk7_2(X1,X2))|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_298, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,X2))|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)|~r1_xxreal_0(X2,esk7_2(X1,X2))), inference(evalgc,[status(thm)],[c_0_292])).
cnf(c_0_299, negated_conjecture, (X1=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X2,esk2_0),X1))|v1_xreal_0(esk8_2(k2_relset_1(X2,esk2_0),X1))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X1)), inference(sr,[status(thm)],[c_0_293, c_0_294])).
cnf(c_0_300, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296])])).
cnf(c_0_301, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_283]), c_0_166]), c_0_149]), c_0_150]), c_0_190])])])).
cnf(c_0_302, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299]), c_0_190])])])).
cnf(c_0_303, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(csr,[status(thm)],[c_0_300, c_0_301])).
cnf(c_0_304, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_127]), c_0_166]), c_0_149]), c_0_150]), c_0_73]), c_0_128])]), c_0_294]), c_0_129])])).
cnf(c_0_305, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304]), c_0_294])])).
cnf(c_0_306, negated_conjecture, (~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_255, c_0_305])])])).
cnf(c_0_307, plain, ($false), inference(sr,[status(thm)],[c_0_267, c_0_306]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 308
# Proof object clause steps            : 244
# Proof object formula steps           : 64
# Proof object conjectures             : 131
# Proof object clause conjectures      : 128
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 60
# Proof object initial formulas used   : 34
# Proof object generating inferences   : 86
# Proof object simplifying inferences  : 250
# Training examples: 152 positive, 1311 negative
# Training: Positive examples begin
cnf(c_0_43, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_258, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_190, negated_conjecture, (v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_71, negated_conjecture, (v1_funct_1(esk2_0))).# trainpos
cnf(c_0_131, negated_conjecture, (v2_seq_2(esk2_0))).# trainpos
cnf(c_0_95, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_181, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_128, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers))).# trainpos
cnf(c_0_257, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).# trainpos
cnf(c_0_129, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_42, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_163, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_157, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_73, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_214, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers))).# trainpos
cnf(c_0_192, negated_conjecture, (v1_xreal_0(esk4_0)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0)).# trainpos
cnf(c_0_125, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_164, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_288, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_280, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_111, plain, (k5_seq_4(X1)=k3_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_72, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_109, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_82, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_160, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_162, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_91, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_207, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_69, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_226, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(esk1_0,X1),k1_seq_1(esk2_0,esk5_1(X1))))).# trainpos
cnf(c_0_186, plain, (k2_xcmplx_0(X1,X2)=k9_binop_2(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_103, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_253, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|m2_subset_1(esk5_1(X1),k1_numbers,k5_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_209, plain, (k9_binop_2(X1,X2)=k9_binop_2(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_141, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_119, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_94, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_89, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_110, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).# trainpos
cnf(c_0_120, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_70, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_112, plain, (k5_seq_4(k2_relset_1(k1_numbers,X1))=k2_rinfsup1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_130, plain, (v3_xxreal_2(k2_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~v2_seq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_208, plain, (v1_xboole_0(X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k3_seq_4(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_membered(X3)|~v3_xxreal_2(X3)|~r1_xxreal_0(k9_binop_2(X1,X2),esk6_3(X3,X1,X2)))).# trainpos
cnf(c_0_124, plain, (v1_xboole_0(X2)|r1_xxreal_0(X3,X1)|X3!=k3_seq_4(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_191, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_156, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_285, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,X2))|r2_hidden(esk7_2(X1,X2),X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1))).# trainpos
cnf(c_0_137, plain, (k1_funct_1(X2,esk9_3(X2,X3,X1))=X1|X3!=k10_xtuple_0(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_159, plain, (r2_hidden(esk9_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).# trainpos
cnf(c_0_298, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,X2))|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)|~r1_xxreal_0(X2,esk7_2(X1,X2)))).# trainpos
cnf(c_0_281, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|r2_hidden(esk7_2(X1,X2),X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers))).# trainpos
cnf(c_0_297, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)|~r1_xxreal_0(X2,esk7_2(X1,X2))|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers))).# trainpos
cnf(c_0_238, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|r2_hidden(esk6_3(X1,X2,X3),X1)|X2!=k3_seq_4(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1))).# trainpos
cnf(c_0_102, plain, (v1_xboole_0(X3)|r2_hidden(X2,k2_relset_1(X3,X4))|X2!=k3_funct_2(k5_numbers,X3,X4,X1)|~v1_funct_1(X4)|~m1_subset_1(X1,k5_numbers)|~v1_funct_2(X4,k5_numbers,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3))))).# trainpos
cnf(c_0_67, plain, (v1_xboole_0(X2)|m1_subset_1(esk12_3(X1,X2,X3),k5_numbers)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).# trainpos
cnf(c_0_49, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_50, plain, (k3_funct_2(k5_numbers,X2,X3,esk12_3(X1,X2,X3))=X1|v1_xboole_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).# trainpos
cnf(c_0_236, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|r1_xxreal_0(k9_binop_2(X3,esk8_2(X2,X3)),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~r2_hidden(X1,X2)|~r1_xxreal_0(X3,esk7_2(X2,X3)))).# trainpos
cnf(c_0_227, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|r2_hidden(esk7_2(X2,X3),X2)|r1_xxreal_0(k9_binop_2(X3,esk8_2(X2,X3)),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_216, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).# trainpos
cnf(c_0_133, plain, (k3_seq_4(k2_relset_1(k1_numbers,X1))=k2_rinfsup1(X1)|~m1_subset_1(k2_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_93, negated_conjecture, (v1_relat_1(esk2_0))).# trainpos
cnf(c_0_114, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).# trainpos
cnf(c_0_92, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).# trainpos
cnf(c_0_175, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_193, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))|~r1_xxreal_0(esk4_0,k6_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).# trainpos
cnf(c_0_234, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k9_binop_2(esk1_0,X1),k1_seq_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_223, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers))).# trainpos
cnf(c_0_104, negated_conjecture, (v5_relat_1(esk2_0,k1_numbers))).# trainpos
cnf(c_0_116, negated_conjecture, (v3_valued_0(esk2_0))).# trainpos
cnf(c_0_173, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_176, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).# trainpos
cnf(c_0_66, plain, (k1_funct_1(X2,esk12_3(X3,X1,X2))=X3|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X3,X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~r2_hidden(X3,k2_relset_1(X1,X2)))).# trainpos
cnf(c_0_169, plain, (r2_hidden(X1,k2_relset_1(X2,X3))|v1_xboole_0(X2)|X1!=k1_funct_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).# trainpos
cnf(c_0_126, plain, (r2_hidden(k3_funct_2(k5_numbers,X1,X2,X3),k2_relset_1(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).# trainpos
cnf(c_0_219, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,k6_numbers)|X3!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k9_binop_2(X2,X3),esk6_3(X1,X3,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X3))).# trainpos
cnf(c_0_263, negated_conjecture, (m1_subset_1(esk5_1(X1),k5_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_140, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3))).# trainpos
cnf(c_0_202, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)|v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0)).# trainpos
cnf(c_0_132, plain, (v3_membered(k2_relset_1(X1,X2))|~v3_membered(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).# trainpos
cnf(c_0_289, plain, (m1_subset_1(esk7_2(X1,X2),X1)|X2=k3_seq_4(X1)|v1_xboole_0(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_127, negated_conjecture, (k2_relset_1(k1_numbers,esk2_0)=k10_xtuple_0(esk2_0))).# trainpos
cnf(c_0_225, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|~v3_valued_0(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_187, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_154, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_182, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relset_1(X3,X1))|v1_xboole_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1))).# trainpos
cnf(c_0_212, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0))|~r1_xxreal_0(esk4_0,k6_numbers))).# trainpos
cnf(c_0_136, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_245, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_relset_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).# trainpos
cnf(c_0_241, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k9_binop_2(esk1_0,X1),k1_funct_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1)|~v3_valued_0(esk2_0))).# trainpos
cnf(c_0_254, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(X1,esk1_0),k6_numbers)|k3_seq_4(X1)=esk1_0|~v1_xreal_0(esk8_2(X1,esk1_0))|~r2_hidden(k1_seq_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))),X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,esk7_2(X1,esk1_0))|~v1_xreal_0(k1_seq_1(esk2_0,esk5_1(esk8_2(X1,esk1_0)))))).# trainpos
cnf(c_0_158, plain, (m1_subset_1(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3)))).# trainpos
cnf(c_0_113, plain, (v1_xreal_0(X2)|v1_xboole_0(X3)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k2_relset_1(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3))).# trainpos
cnf(c_0_150, negated_conjecture, (v3_membered(k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_149, negated_conjecture, (v3_xxreal_2(k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_151, negated_conjecture, (k3_seq_4(k10_xtuple_0(esk2_0))=k2_rinfsup1(esk2_0)|~m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_200, plain, (k2_xcmplx_0(X2,X1)=k9_binop_2(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_148, negated_conjecture, (r2_hidden(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1),k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_194, negated_conjecture, (r2_hidden(k1_funct_1(esk2_0,X1),k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_172, plain, (m1_subset_1(esk9_3(X1,X2,X3),X4)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).# trainpos
cnf(c_0_152, negated_conjecture, (m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_240, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X2)|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).# trainpos
cnf(c_0_155, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X2)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).# trainpos
cnf(c_0_170, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).# trainpos
cnf(c_0_265, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|v1_xboole_0(X2)|~m2_subset_1(esk12_3(X1,X2,esk2_0),k1_numbers,k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).# trainpos
cnf(c_0_255, negated_conjecture, (~v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_153, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_166, negated_conjecture, (k3_seq_4(k10_xtuple_0(esk2_0))=k2_rinfsup1(esk2_0))).# trainpos
cnf(c_0_250, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers))).# trainpos
cnf(c_0_183, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers))).# trainpos
cnf(c_0_171, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers))).# trainpos
cnf(c_0_275, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,X2)|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk12_3(X2,X1,esk2_0),k5_numbers))).# trainpos
cnf(c_0_204, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|X1!=k3_seq_4(k10_xtuple_0(esk2_0))|~v1_xreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_167, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_222, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_242, negated_conjecture, (X1=k3_seq_4(k10_xtuple_0(esk2_0))|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|r1_xxreal_0(k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1)),k1_funct_1(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_165, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|X1!=k3_seq_4(k10_xtuple_0(esk2_0))|~v1_xreal_0(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_221, negated_conjecture, (k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers)|~m1_subset_1(esk3_0,k5_numbers)|~v1_xreal_0(k1_funct_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_261, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_189, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_284, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,X2)|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).# trainpos
cnf(c_0_262, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k10_xtuple_0(esk2_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))|~m1_subset_1(esk5_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0)),k5_numbers))).# trainpos
cnf(c_0_195, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_185, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_264, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(X1)=esk1_0|r1_xxreal_0(esk8_2(X1,esk1_0),k6_numbers)|~r2_hidden(k1_funct_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))),X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,esk7_2(X1,esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))))|~v1_xreal_0(esk8_2(X1,esk1_0)))).# trainpos
cnf(c_0_201, negated_conjecture, (v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~m1_subset_1(esk3_0,k5_numbers))).# trainpos
cnf(c_0_272, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_274, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk5_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0))))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0))|~m1_subset_1(esk5_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0)),k5_numbers))).# trainpos
cnf(c_0_206, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_295, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_301, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).# trainpos
cnf(c_0_282, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).# trainpos
cnf(c_0_199, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_296, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|m1_subset_1(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k10_xtuple_0(esk2_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).# trainpos
cnf(c_0_300, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).# trainpos
cnf(c_0_215, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).# trainpos
cnf(c_0_210, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~v1_xreal_0(esk4_0))).# trainpos
cnf(c_0_229, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(esk4_0,k6_numbers)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k3_seq_4(X1)!=esk1_0|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk6_3(X1,esk1_0,esk4_0),k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_232, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(esk4_0,k6_numbers)|k3_seq_4(X1)!=esk1_0|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk6_3(X1,esk1_0,esk4_0),k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_246, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|r1_xxreal_0(esk4_0,k6_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~v1_xreal_0(esk4_0))).# trainpos
cnf(c_0_249, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk4_0,k6_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).# trainpos
cnf(c_0_268, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk3_0)))).# trainpos
cnf(c_0_266, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0)).# trainpos
cnf(c_0_278, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|k2_rinfsup1(esk2_0)!=esk1_0|~m1_subset_1(esk3_0,k5_numbers))).# trainpos
cnf(c_0_277, negated_conjecture, (m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0)).# trainpos
cnf(c_0_286, negated_conjecture, (m1_subset_1(esk3_0,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0)).# trainpos
cnf(c_0_293, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X2))).# trainpos
cnf(c_0_302, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).# trainpos
cnf(c_0_304, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).# trainpos
cnf(c_0_305, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0)))).# trainpos
cnf(c_0_306, negated_conjecture, (~m1_subset_1(X1,k5_numbers))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_58, plain, (v1_xreal_0(esk13_0))).#trainneg
cnf(i_0_84, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_98, plain, (v1_xboole_0(o_0_0_xboole_0))).#trainneg
cnf(i_0_61, plain, (v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_96, plain, (v1_xboole_0(esk14_0))).#trainneg
cnf(i_0_48, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_59, plain, (v1_xxreal_0(esk13_0))).#trainneg
cnf(i_0_42, plain, (X1=k1_xboole_0|k10_xtuple_0(X1)!=k1_xboole_0|~v1_relat_1(X1))).#trainneg
cnf(i_0_95, plain, (v1_xxreal_0(esk14_0))).#trainneg
cnf(i_0_57, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_60, plain, (v1_xcmplx_0(esk13_0))).#trainneg
cnf(i_0_63, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_104, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_90, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_87, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_107, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_77, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_97, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_108, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_83, plain, (v1_xreal_0(k3_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_106, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_67, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_43, plain, (X1=k1_xboole_0|k9_xtuple_0(X1)!=k1_xboole_0|~v1_relat_1(X1))).#trainneg
cnf(i_0_114, plain, (v3_membered(k10_xtuple_0(X1))|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_103, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_115, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_76, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_79, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_92, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_93, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_82, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_72, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_91, plain, (m1_subset_1(k5_seq_4(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_69, plain, (m1_subset_1(k2_rinfsup1(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_105, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_62, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_109, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_112, plain, (m1_subset_1(k9_binop_2(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_102, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_113, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_117, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_86, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_34, plain, (v2_seq_2(X1)|~v1_funct_1(X1)|~v3_xxreal_2(k2_relset_1(k1_numbers,X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_119, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_31, plain, (r2_hidden(X3,X4)|X4!=k10_xtuple_0(X2)|X3!=k1_funct_1(X2,X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2)))).#trainneg
cnf(i_0_99, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_23, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk7_2(X1,X2))|v1_xreal_0(esk8_2(X1,X2))|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1))).#trainneg
cnf(i_0_111, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_29, plain, (X2=k10_xtuple_0(X1)|r2_hidden(esk10_2(X1,X2),X2)|r2_hidden(esk11_2(X1,X2),k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28, plain, (X2=k10_xtuple_0(X1)|k1_funct_1(X1,esk11_2(X1,X2))=esk10_2(X1,X2)|r2_hidden(esk10_2(X1,X2),X2)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_22, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk7_2(X1,X2))|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers))).#trainneg
cnf(i_0_30, plain, (X2=k10_xtuple_0(X1)|esk10_2(X1,X2)!=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X3,k9_xtuple_0(X1))|~r2_hidden(esk10_2(X1,X2),X2))).#trainneg
cnf(i_0_26, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|v1_xreal_0(esk6_3(X1,X2,X3))|X2!=k3_seq_4(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v3_membered(X1)|~v3_xxreal_2(X1))).#trainneg
cnf(i_0_46, plain, (X3=k1_xboole_0|r2_hidden(k1_funct_1(X1,X4),k2_relset_1(X3,X1))|~v1_funct_1(X1)|~r2_hidden(X4,X2)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_21, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk7_2(X2,X3))|r1_xxreal_0(k9_binop_2(X3,esk8_2(X2,X3)),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_88, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_141, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_121, plain, (esk14_0=k6_numbers)).#trainneg
cnf(i_0_122, plain, (esk13_0=k6_numbers)).#trainneg
cnf(i_0_158, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_148, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_157, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),k1_seq_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_202, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k2_xcmplx_0(esk1_0,X1),k1_funct_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_128, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_124, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_123, plain, (o_0_0_xboole_0=k6_numbers)).#trainneg
cnf(i_0_125, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_221, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_176, negated_conjecture, (v1_valued_0(esk2_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_127, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_129, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_168, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_164, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_133, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_201, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_160, plain, (v1_xreal_0(k2_rinfsup1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_135, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_143, plain, (v1_xreal_0(k5_seq_4(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_234, negated_conjecture, (v1_partfun1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_228, negated_conjecture, (m1_subset_1(esk3_0,k1_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk4_0,k6_numbers))).#trainneg
cnf(i_0_137, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_138, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_139, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_142, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_163, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_174, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_140, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_144, plain, (v7_ordinal1(k5_seq_4(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_226, negated_conjecture, (m1_subset_1(esk5_1(X1),k1_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_274, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_169, plain, (v1_xreal_0(k9_binop_2(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_257, plain, (v1_xreal_0(esk12_3(X1,X2,X3))|v1_xboole_0(X2)|~v3_membered(k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_224, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|X1!=k1_funct_1(X3,X4)|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X4,k9_xtuple_0(X3)))).#trainneg
cnf(i_0_240, plain, (~v1_xboole_0(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_199, plain, (v3_valued_0(k1_relset_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_208, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_190, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_145, plain, (m1_subset_1(k3_seq_4(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_227, negated_conjecture, (m1_subset_1(esk3_0,k1_numbers)|v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0)).#trainneg
cnf(i_0_147, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_173, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_253, negated_conjecture, (k6_numbers=k1_numbers|r2_hidden(k1_funct_1(esk2_0,X1),k2_relset_1(k1_numbers,esk2_0))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_161, plain, (v7_ordinal1(k2_rinfsup1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_177, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_165, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_170, plain, (v7_ordinal1(k9_binop_2(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_178, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_252, plain, (r1_xxreal_0(X1,esk6_3(X2,X3,X4))|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|X1!=k3_seq_4(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X2,X3,X4))|X3!=k3_seq_4(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_258, plain, (v7_ordinal1(esk12_3(X1,X2,X3))|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_198, plain, (v3_membered(k1_relset_1(X1,X2))|~v3_membered(X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_241, plain, (m1_subset_1(esk9_3(X1,X2,X3),k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_251, plain, (m1_subset_1(esk6_3(X1,X2,X3),X1)|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_276, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X3,X4),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_189, plain, (v3_valued_0(k2_relset_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_242, plain, (r1_xxreal_0(X1,esk9_3(X2,X3,X4))|X1!=k3_seq_4(k9_xtuple_0(X2))|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(esk9_3(X2,X3,X4))|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_217, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(X3)|X1!=k3_seq_4(X3)|~v3_xxreal_2(X3)|~v3_membered(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_278, plain, (m1_subset_1(X3,X1)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X3,X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~r2_hidden(X3,k2_relset_1(X1,X2)))).#trainneg
cnf(i_0_237, plain, (v1_xreal_0(X3)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_191, plain, (v1_xreal_0(k1_relset_1(X1,X2))|~v3_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_182, plain, (v7_ordinal1(k2_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_243, plain, (r2_hidden(X1,X2)|X1!=k1_funct_1(X3,esk9_3(X3,X4,X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|X4!=k10_xtuple_0(X3)|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_179, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)=k1_funct_1(X1,X2)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_238, plain, (v1_xxreal_0(X3)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_239, plain, (v1_xcmplx_0(X3)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_181, plain, (v1_xreal_0(k2_relset_1(X1,X2))|~v3_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_229, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_192, plain, (v7_ordinal1(k1_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_180, plain, (k2_xcmplx_0(X1,k1_funct_1(X2,X3))=k2_xcmplx_0(k1_funct_1(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_183, plain, (v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_249, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|r1_xxreal_0(X5,k6_numbers)|X1!=k1_funct_1(X3,esk6_3(k9_xtuple_0(X3),X4,X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|X4!=k3_seq_4(k9_xtuple_0(X3))|~v3_xxreal_2(k9_xtuple_0(X3))|~v3_membered(k9_xtuple_0(X3))|~v1_xreal_0(X5)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_287, plain, (X1=k3_seq_4(X2)|r2_hidden(esk7_2(X2,X1),X2)|r1_xxreal_0(k9_binop_2(X1,esk8_2(X2,X1)),X3)|v1_xboole_0(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_275, plain, (v7_ordinal1(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_232, plain, (m1_subset_1(esk7_2(X1,X2),X1)|X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,X2))|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_195, plain, (v3_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_197, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_271, plain, (v3_membered(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_272, plain, (v3_valued_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_286, plain, (X1=k3_seq_4(X2)|r2_hidden(esk7_2(X2,X1),X2)|r1_xxreal_0(k9_binop_2(X1,esk8_2(X2,X1)),esk6_3(X2,X3,X4))|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(esk6_3(X2,X3,X4))|X3!=k3_seq_4(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_185, plain, (v3_valued_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_283, plain, (X1=k3_seq_4(k9_xtuple_0(X2))|r2_hidden(esk7_2(k9_xtuple_0(X2),X1),k9_xtuple_0(X2))|r1_xxreal_0(k9_binop_2(X1,esk8_2(k9_xtuple_0(X2),X1)),esk9_3(X2,X3,X4))|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(esk9_3(X2,X3,X4))|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_193, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_273, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~r2_hidden(X1,k3_funct_2(X3,k1_zfmisc_1(X2),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(X2))))|~m1_subset_1(X5,X3)|~v1_funct_2(X4,X3,k1_zfmisc_1(X2))|~v1_funct_1(X4))).#trainneg
cnf(i_0_184, plain, (v4_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_186, plain, (v1_valued_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_244, plain, (X1=k10_xtuple_0(X2)|v1_xboole_0(X1)|esk10_2(X2,X1)!=k1_funct_1(X2,X3)|~v1_relat_1(X2)|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X2)|~m1_subset_1(esk10_2(X2,X1),X1))).#trainneg
cnf(i_0_187, plain, (v5_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_196, plain, (v1_valued_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_194, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_263, plain, (v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_267, plain, (v3_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_210, plain, (v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_211, plain, (v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_223, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_222, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_245, plain, (r2_hidden(X1,X2)|X4=k3_seq_4(k9_xtuple_0(X3))|v1_xboole_0(k9_xtuple_0(X3))|X1!=k1_funct_1(X3,esk7_2(k9_xtuple_0(X3),X4))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v3_xxreal_2(k9_xtuple_0(X3))|~v3_membered(k9_xtuple_0(X3))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X3),X4),k6_numbers)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_285, plain, (X1=k3_seq_4(X2)|r2_hidden(esk7_2(X2,X1),X2)|r1_xxreal_0(k9_binop_2(X1,esk8_2(X2,X1)),esk7_2(X2,X3))|X3=k3_seq_4(X2)|v1_xboole_0(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(X2,X3))|~r1_xxreal_0(esk8_2(X2,X3),k6_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_230, plain, (r2_hidden(X1,X2)|X4=k3_seq_4(k9_xtuple_0(X3))|v1_xboole_0(k9_xtuple_0(X3))|v1_xreal_0(esk8_2(k9_xtuple_0(X3),X4))|X1!=k1_funct_1(X3,esk7_2(k9_xtuple_0(X3),X4))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v3_xxreal_2(k9_xtuple_0(X3))|~v3_membered(k9_xtuple_0(X3))|~v1_xreal_0(X4))).#trainneg
cnf(i_0_262, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xboole_0(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X4,X5))|~v1_funct_2(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X4,X5),X1,X2)|~v1_funct_1(k3_funct_2(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))|~m1_subset_1(X5,X3)|~v1_funct_2(X4,X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_284, plain, (X1=k3_seq_4(X2)|r2_hidden(esk7_2(X2,X1),X2)|r1_xxreal_0(k9_binop_2(X1,esk8_2(X2,X1)),esk7_2(X2,X3))|X3=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk8_2(X2,X3))|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_248, plain, (r1_xxreal_0(X1,esk7_2(X2,X3))|X3=k3_seq_4(X2)|v1_xboole_0(X2)|X1!=k3_seq_4(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(X2,X3))|~r1_xxreal_0(esk8_2(X2,X3),k6_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_233, plain, (r1_xxreal_0(X1,esk7_2(X2,X3))|X3=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk8_2(X2,X3))|X1!=k3_seq_4(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_236, plain, (v1_partfun1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|v1_xboole_0(X2)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_269, plain, (v5_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X3)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_255, plain, (X1=k6_numbers|r2_hidden(k1_funct_1(k2_relset_1(k2_zfmisc_1(X2,X1),X3),X4),k2_relset_1(X1,k2_relset_1(k2_zfmisc_1(X2,X1),X3)))|~r2_hidden(X4,X2)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X2,X1),X3),X2,X1)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X2,X1),X3))|~v5_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_235, plain, (v1_partfun1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1)|v1_xboole_0(X2)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_264, plain, (v4_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_268, plain, (v1_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|v1_xboole_0(X1)|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_282, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk8_2(X2,X3),k6_numbers)|X3=k3_seq_4(X2)|X3!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(esk8_2(X2,X3))|~v1_xreal_0(X3)|~r2_hidden(esk6_3(X1,X3,esk8_2(X2,X3)),X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(X3,esk7_2(X2,X3))|~v1_xreal_0(esk6_3(X1,X3,esk8_2(X2,X3))))).#trainneg
cnf(i_0_254, plain, (X1=k6_numbers|r2_hidden(k1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X1),X3),X4),k2_relset_1(X1,k1_relset_1(k2_zfmisc_1(X2,X1),X3)))|~r2_hidden(X4,X2)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X2,X1),X3),X2,X1)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X1),X3))|~v4_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_265, plain, (v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2,X3)|v1_xboole_0(X1)|~v1_partfun1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_261, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk8_2(X2,X3),k6_numbers)|X3=k3_seq_4(X2)|v1_xreal_0(esk7_2(X2,X3))|X3!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(esk8_2(X2,X3))|~v1_xreal_0(X3)|~r2_hidden(esk6_3(X1,X3,esk8_2(X2,X3)),X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(esk6_3(X1,X3,esk8_2(X2,X3))))).#trainneg
cnf(i_0_266, plain, (v1_partfun1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2)|v1_xboole_0(X3)|v1_xboole_0(X1)|~v1_funct_2(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5),X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_270, plain, (X1=k6_numbers|r2_hidden(k1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5),X6),k2_relset_1(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5)))|v1_xboole_0(X2)|~r2_hidden(X6,X3)|~v1_funct_2(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5),X3,X1)|~v1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_260, plain, (k9_binop_2(X1,esk8_2(X2,X1))=k3_seq_4(X3)|v1_xboole_0(X3)|v1_xreal_0(esk8_2(X3,k9_binop_2(X1,esk8_2(X2,X1))))|X1=k3_seq_4(X2)|v1_xreal_0(esk7_2(X2,X1))|~v3_xxreal_2(X3)|~v3_membered(X3)|~v1_xreal_0(k9_binop_2(X1,esk8_2(X2,X1)))|~r2_hidden(esk7_2(X3,k9_binop_2(X1,esk8_2(X2,X1))),X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(X3,k9_binop_2(X1,esk8_2(X2,X1)))))).#trainneg
cnf(i_0_281, plain, (k9_binop_2(X1,esk8_2(X2,X1))=k3_seq_4(X3)|v1_xboole_0(X3)|v1_xreal_0(esk8_2(X3,k9_binop_2(X1,esk8_2(X2,X1))))|X1=k3_seq_4(X2)|~v3_xxreal_2(X3)|~v3_membered(X3)|~v1_xreal_0(k9_binop_2(X1,esk8_2(X2,X1)))|~r2_hidden(esk7_2(X3,k9_binop_2(X1,esk8_2(X2,X1))),X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(X1,esk7_2(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(X3,k9_binop_2(X1,esk8_2(X2,X1)))))).#trainneg
cnf(i_0_290, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_292, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),k1_funct_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_365, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_289, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_60, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_342, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_294, negated_conjecture, (X1=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,X1),k9_xtuple_0(esk2_0))|r2_hidden(esk10_2(esk2_0,X1),X1))).#trainneg
cnf(i_0_306, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k9_binop_2(X1,esk1_0),k1_seq_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_538, plain, (k2_xcmplx_0(X1,X3)=k2_xcmplx_0(X3,X1)|v1_xboole_0(X4)|~v1_valued_0(X2)|~v1_xcmplx_0(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X3,k2_relset_1(X4,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_331, plain, (r2_hidden(X1,k2_relset_1(X2,X3))|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(esk9_3(X3,X4,X1),k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|X4!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_328, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_386, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_202, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k2_xcmplx_0(esk1_0,X1),k1_funct_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_344, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_380, plain, (v1_xreal_0(k3_seq_4(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_522, negated_conjecture, (v1_xreal_0(k2_relset_1(k1_numbers,esk2_0))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_382, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_295, negated_conjecture, (k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1))).#trainneg
cnf(i_0_320, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_416, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_381, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_368, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_343, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_300, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),X3)=k2_xcmplx_0(X3,k1_funct_1(X1,X2))|~v1_xreal_0(X3)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_314, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_489, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(k5_numbers,X1,X2,X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_326, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_321, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_465, plain, (r1_xxreal_0(k3_seq_4(k9_xtuple_0(X1)),esk9_3(X1,X2,X3))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(esk9_3(X1,X2,X3))|~v1_xreal_0(k3_seq_4(k9_xtuple_0(X1))))).#trainneg
cnf(i_0_318, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(k6_numbers,k6_numbers)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk5_1(k6_numbers)))|~v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_319, plain, (k9_binop_2(X1,k6_numbers)=X1|~v1_xreal_0(k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_338, plain, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,X2,X3,X4))|v1_xboole_0(X2)|X1!=k3_seq_4(k2_relset_1(X2,X3))|~v3_xxreal_2(k2_relset_1(X2,X3))|~v3_membered(k2_relset_1(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_362, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,k6_numbers)|X2=k3_seq_4(X3)|esk8_2(X3,X2)!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X2)|~v1_xreal_0(esk8_2(X3,X2))|~r2_hidden(esk6_3(X1,esk8_2(X3,X2),X2),X3)|~v3_xxreal_2(X3)|~v3_membered(X3)|~r1_xxreal_0(X2,esk7_2(X3,X2))|~v1_xreal_0(esk6_3(X1,esk8_2(X3,X2),X2)))).#trainneg
cnf(i_0_417, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_418, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_492, plain, (v1_xreal_0(esk9_3(X1,X2,X3))|k9_xtuple_0(X1)!=k10_xtuple_0(X4)|~v3_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_315, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_366, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_388, plain, (v1_xxreal_0(k5_seq_4(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_327, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_325, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_565, plain, (v1_xreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_394, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_330, plain, (r2_hidden(X1,k2_relset_1(X2,X3))|v1_xboole_0(X2)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(esk12_3(X1,X4,X3),k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|~r2_hidden(X1,k2_relset_1(X4,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X3,k5_numbers,X4))).#trainneg
cnf(i_0_336, plain, (m1_subset_1(k3_funct_2(k5_numbers,X1,X2,X3),k2_relset_1(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_506, plain, (k2_xcmplx_0(X2,k6_numbers)=X2|v1_xboole_0(X3)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k2_relset_1(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3))).#trainneg
cnf(i_0_459, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_539, plain, (k2_xcmplx_0(X1,X4)=k2_xcmplx_0(X4,X1)|~v1_valued_0(X2)|~v1_xcmplx_0(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|X3!=k10_xtuple_0(X2)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_369, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_480, plain, (r1_xxreal_0(X1,k6_numbers)|X1!=k3_seq_4(k5_numbers)|~v3_xxreal_2(k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_341, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k1_funct_1(X1,X2))|~v1_xreal_0(X3)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_390, plain, (v7_ordinal1(k3_seq_4(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_373, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_374, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_345, plain, (v1_xreal_0(k9_xtuple_0(X1))|~v3_membered(k1_zfmisc_1(X2))|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_346, plain, (v7_ordinal1(k9_xtuple_0(X1))|~v6_membered(k1_zfmisc_1(X2))|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_544, plain, (r2_hidden(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)),X4)|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(k9_xtuple_0(X1))|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_356, plain, (v3_membered(k9_xtuple_0(X1))|~v3_membered(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_537, plain, (k2_xcmplx_0(k1_funct_1(X2,X3),X1)=k9_binop_2(X1,k1_funct_1(X2,X3))|~v1_xreal_0(k1_funct_1(X2,X3))|~v1_xreal_0(X1)|~v1_valued_0(X2)|~v1_xcmplx_0(X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_498, plain, (v1_xreal_0(X1)|v1_xboole_0(X2)|X2!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_632, plain, (X1=k3_seq_4(X2)|v1_xboole_0(X3)|r1_xxreal_0(esk8_2(X2,X1),k6_numbers)|v1_xboole_0(X2)|X1!=k3_seq_4(X3)|~v3_xxreal_2(X3)|~v3_xxreal_2(X2)|~v3_membered(X3)|~v3_membered(X2)|~r1_xxreal_0(X1,esk7_2(X2,X1))|~v1_xreal_0(esk8_2(X2,X1))|~v1_xreal_0(X1)|~m1_subset_1(esk6_3(X3,X1,esk8_2(X2,X1)),X2))).#trainneg
cnf(i_0_534, plain, (k2_xcmplx_0(k6_numbers,k1_funct_1(X1,X2))=k1_funct_1(X1,X2)|~v1_xreal_0(k1_funct_1(X1,X2))|~v1_valued_0(X1)|~v1_xcmplx_0(k6_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_502, plain, (r2_hidden(k1_funct_1(X1,esk9_3(X1,X2,X3)),X4)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_393, plain, (v1_xreal_0(X3)|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X3,X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~r2_hidden(X3,k2_relset_1(X1,X2)))).#trainneg
cnf(i_0_363, negated_conjecture, (v1_xreal_0(esk5_1(X1))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_427, plain, (v1_xxreal_0(k2_rinfsup1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_432, plain, (v1_xxreal_0(k9_binop_2(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_253, negated_conjecture, (k6_numbers=k1_numbers|r2_hidden(k1_funct_1(esk2_0,X1),k10_xtuple_0(esk2_0))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_347, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k9_xtuple_0(X3))|~v1_funct_2(k9_xtuple_0(X3),X1,X2)|~v1_funct_1(k9_xtuple_0(X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_716, plain, (r1_xxreal_0(k6_numbers,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_793, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_938, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_727, negated_conjecture, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_348, plain, (v1_relat_1(k9_xtuple_0(X1))|~v4_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1))).#trainneg
cnf(i_0_349, plain, (v4_relat_1(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1))).#trainneg
cnf(i_0_681, negated_conjecture, (k5_seq_4(k10_xtuple_0(esk2_0))=k2_rinfsup1(esk2_0))).#trainneg
cnf(i_0_789, plain, (m1_subset_1(esk12_3(X1,X2,X3),k1_numbers)|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_913, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_715, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_688, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,esk1_0),k1_seq_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1098, negated_conjecture, (v3_membered(k9_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_350, plain, (v1_funct_2(k9_xtuple_0(X1),X2,X3)|~v1_partfun1(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1))).#trainneg
cnf(i_0_882, plain, (r1_xxreal_0(k6_numbers,k6_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_675, negated_conjecture, (v7_ordinal1(k10_xtuple_0(esk2_0))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1062, negated_conjecture, (v1_xreal_0(k9_xtuple_0(esk2_0))|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_857, negated_conjecture, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_719, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_351, plain, (v1_partfun1(k9_xtuple_0(X1),X2)|v1_xboole_0(X3)|~v1_funct_2(k9_xtuple_0(X1),X2,X3)|~v4_relat_1(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1))).#trainneg
cnf(i_0_352, plain, (v3_valued_0(k9_xtuple_0(X1))|~v3_membered(X2)|~v4_relat_1(X1,k2_zfmisc_1(X3,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_803, negated_conjecture, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_931, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_936, plain, (k9_binop_2(k6_numbers,X1)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_987, negated_conjecture, (r2_hidden(X1,k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_717, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_796, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,esk1_0),k1_funct_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_721, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_353, plain, (v1_valued_0(k9_xtuple_0(X1))|~v1_membered(X2)|~v4_relat_1(X1,k2_zfmisc_1(X3,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_940, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_677, negated_conjecture, (v3_valued_0(k10_xtuple_0(esk2_0))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_933, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(k6_numbers,k6_numbers)|~m2_subset_1(esk5_1(k6_numbers),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_354, plain, (v5_relat_1(k9_xtuple_0(X1),X2)|~v4_relat_1(X1,k2_zfmisc_1(X3,X2))|~v1_relat_1(X1))).#trainneg
cnf(i_0_355, plain, (X1=k6_numbers|r2_hidden(k1_funct_1(k9_xtuple_0(X2),X3),k2_relset_1(X1,k9_xtuple_0(X2)))|~r2_hidden(X3,X4)|~v1_funct_2(k9_xtuple_0(X2),X4,X1)|~v1_funct_1(k9_xtuple_0(X2))|~v4_relat_1(X2,k2_zfmisc_1(X4,X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_760, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|v1_xboole_0(X2)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk12_3(X1,X2,esk2_0),k1_numbers,k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_993, plain, (r2_hidden(k3_funct_2(k5_numbers,X1,X2,X3),k2_relset_1(X4,X2))|v1_xboole_0(X1)|v1_xboole_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_766, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_valued_0(esk2_0)|~v1_xcmplx_0(X1)|~r2_hidden(X2,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_925, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_698, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X1)|v1_xboole_0(X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(esk12_3(X1,X2,esk2_0),k1_numbers,k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_679, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_763, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~r1_xxreal_0(k9_binop_2(X1,esk1_0),k1_funct_1(esk2_0,esk5_1(X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_794, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(k6_numbers,k6_numbers)|~r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,esk5_1(k6_numbers))))).#trainneg
cnf(i_0_357, plain, (v3_valued_0(k9_xtuple_0(X1))|~v3_valued_0(X2)|~v1_relat_1(X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1059, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_367, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_370, plain, (v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_partfun1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_714, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X2)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_697, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_713, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|v1_xboole_0(X2)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk12_3(X1,X2,esk2_0),k1_numbers,k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_690, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_941, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,k6_numbers)|k3_seq_4(X1)!=k6_numbers|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(X2,esk6_3(X1,k6_numbers,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_877, plain, (k2_xcmplx_0(X3,X4)=k2_xcmplx_0(X4,X3)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xreal_0(X4)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_707, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|v1_xboole_0(X4)|X1!=k3_seq_4(k2_relset_1(X4,X2))|~v3_xxreal_2(k2_relset_1(X4,X2))|~v3_membered(k2_relset_1(X4,X2))|~v1_xreal_0(X1)|~v1_xreal_0(k1_funct_1(X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_1(X2))).#trainneg
cnf(i_0_949, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_722, plain, (k2_xcmplx_0(k6_numbers,X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1127, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,X1))|k10_xtuple_0(X2)!=k1_numbers|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_371, plain, (v1_partfun1(k2_zfmisc_1(X1,X2),X1)|v1_xboole_0(X2)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2))).#trainneg
cnf(i_0_928, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1135, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|v1_xreal_0(k5_numbers)|k10_xtuple_0(X1)!=k1_zfmisc_1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_718, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_951, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_375, plain, (X1=k6_numbers|r2_hidden(k1_funct_1(k2_zfmisc_1(X2,X1),X3),k2_relset_1(X1,k2_zfmisc_1(X2,X1)))|~r2_hidden(X3,X2)|~v1_funct_2(k2_zfmisc_1(X2,X1),X2,X1)|~v1_funct_1(k2_zfmisc_1(X2,X1)))).#trainneg
cnf(i_0_1006, negated_conjecture, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_valued_0(esk2_0)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_876, plain, (k2_xcmplx_0(X2,X4)=k2_xcmplx_0(X4,X2)|v1_xboole_0(X3)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xreal_0(X4)|~r2_hidden(X2,k2_relset_1(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3))).#trainneg
cnf(i_0_696, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X1)|v1_xboole_0(X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(esk12_3(X1,X2,esk2_0),k1_numbers,k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_901, plain, (m1_subset_1(k3_funct_2(k5_numbers,k5_numbers,X1,X2),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1176, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(X2,X3))|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_978, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(X3,X4))|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_963, plain, (v1_xxreal_0(k3_seq_4(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_918, negated_conjecture, (r1_xxreal_0(X1,k10_xtuple_0(esk2_0))|r1_xxreal_0(k10_xtuple_0(esk2_0),X1)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_930, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1120, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|v1_xreal_0(esk2_0)|k10_xtuple_0(X1)!=k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_723, plain, (k2_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_397, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|v1_xboole_0(X4)|X1!=k1_funct_1(X3,k3_funct_2(X4,k9_xtuple_0(X3),X5,X6))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k9_xtuple_0(X3))))|~m1_subset_1(X6,X4)|~v1_funct_2(X5,X4,k9_xtuple_0(X3))|~v1_funct_1(X5))).#trainneg
cnf(i_0_400, plain, (v7_ordinal1(k1_funct_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1282, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),k1_seq_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1284, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),k1_funct_1(esk2_0,X1))|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1281, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),k1_seq_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1219, negated_conjecture, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_1328, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,X1))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_402, plain, (m1_subset_1(X3,k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_1244, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_1301, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_404, plain, (m1_subset_1(esk7_2(k1_relset_1(X1,X2),X3),X1)|X3=k3_seq_4(k1_relset_1(X1,X2))|v1_xboole_0(k1_relset_1(X1,X2))|v1_xreal_0(esk8_2(k1_relset_1(X1,X2),X3))|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v3_xxreal_2(k1_relset_1(X1,X2))|~v3_membered(k1_relset_1(X1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_405, plain, (m1_subset_1(esk7_2(k1_relset_1(X1,X2),X3),X1)|X3=k3_seq_4(k1_relset_1(X1,X2))|v1_xboole_0(k1_relset_1(X1,X2))|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v3_xxreal_2(k1_relset_1(X1,X2))|~v3_membered(k1_relset_1(X1,X2))|~r1_xxreal_0(esk8_2(k1_relset_1(X1,X2),X3),k6_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1283, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),k1_funct_1(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_1353, plain, (m1_subset_1(esk9_3(X1,X2,X3),k1_numbers)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1202, negated_conjecture, (v1_xreal_0(k1_funct_1(esk2_0,X1))|k6_numbers=k1_numbers|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_407, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k1_relset_1(X2,X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k1_relset_1(X2,X3)))).#trainneg
cnf(i_0_1374, negated_conjecture, (v1_xreal_0(k10_xtuple_0(esk2_0))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1375, negated_conjecture, (v1_xxreal_0(k10_xtuple_0(esk2_0))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1311, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|r1_xxreal_0(k6_numbers,esk1_0))).#trainneg
cnf(i_0_406, plain, (m1_subset_1(esk6_3(k1_relset_1(X1,X2),X3,X4),X1)|v1_xboole_0(k1_relset_1(X1,X2))|r1_xxreal_0(X4,k6_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|X3!=k3_seq_4(k1_relset_1(X1,X2))|~v3_xxreal_2(k1_relset_1(X1,X2))|~v3_membered(k1_relset_1(X1,X2))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_408, plain, (m1_subset_1(esk7_2(k2_relset_1(X1,X2),X3),X1)|X3=k3_seq_4(k2_relset_1(X1,X2))|v1_xboole_0(k2_relset_1(X1,X2))|v1_xreal_0(esk8_2(k2_relset_1(X1,X2),X3))|~v5_relat_1(X2,X1)|~v1_relat_1(X2)|~v3_xxreal_2(k2_relset_1(X1,X2))|~v3_membered(k2_relset_1(X1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1484, plain, (k9_binop_2(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_1318, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_412, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k2_relset_1(X2,X3))|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k2_relset_1(X2,X3)))).#trainneg
cnf(i_0_409, plain, (m1_subset_1(esk7_2(k2_relset_1(X1,X2),X3),X1)|X3=k3_seq_4(k2_relset_1(X1,X2))|v1_xboole_0(k2_relset_1(X1,X2))|~v5_relat_1(X2,X1)|~v1_relat_1(X2)|~v3_xxreal_2(k2_relset_1(X1,X2))|~v3_membered(k2_relset_1(X1,X2))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,X2),X3),k6_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_410, plain, (m1_subset_1(esk6_3(k2_relset_1(X1,X2),X3,X4),X1)|v1_xboole_0(k2_relset_1(X1,X2))|r1_xxreal_0(X4,k6_numbers)|~v5_relat_1(X2,X1)|~v1_relat_1(X2)|X3!=k3_seq_4(k2_relset_1(X1,X2))|~v3_xxreal_2(k2_relset_1(X1,X2))|~v3_membered(k2_relset_1(X1,X2))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1356, plain, (v1_xreal_0(esk9_3(X1,X2,X3))|~v3_membered(X4)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1549, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1),k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1327, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1317, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1),k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1334, negated_conjecture, (r2_hidden(X1,k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|~m1_subset_1(esk12_3(X1,X2,esk2_0),k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_1545, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk12_3(X2,X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_1252, plain, (m1_subset_1(k3_seq_4(X1),k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_1314, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1559, negated_conjecture, (v3_valued_0(k9_xtuple_0(esk2_0))|~v3_valued_0(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_429, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1272, negated_conjecture, (m1_subset_1(k2_rinfsup1(esk2_0),k1_numbers)|~m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1502, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X2)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk12_3(X2,X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_1386, negated_conjecture, (r1_xxreal_0(k6_numbers,esk5_1(X1))|r1_xxreal_0(esk5_1(X1),k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1564, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers))).#trainneg
cnf(i_0_1689, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1315, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1250, negated_conjecture, (m1_subset_1(k10_xtuple_0(esk2_0),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1401, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1273, negated_conjecture, (v1_xreal_0(k2_rinfsup1(esk2_0))|~m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1274, negated_conjecture, (v7_ordinal1(k2_rinfsup1(esk2_0))|~v6_membered(k1_numbers)|~m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_953, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1565, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk9_3(esk2_0,X2,X1),k5_numbers))).#trainneg
cnf(i_0_1519, negated_conjecture, (r2_hidden(k3_funct_2(k5_numbers,X1,esk2_0,X2),k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_1683, plain, (v1_xreal_0(k3_funct_2(k5_numbers,k5_numbers,X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1461, negated_conjecture, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~v5_relat_1(esk2_0,X2)|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_1604, plain, (k2_xcmplx_0(k6_numbers,X1)=k9_binop_2(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1656, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_valued_0(esk2_0)|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1381, negated_conjecture, (m1_subset_1(k9_xtuple_0(esk2_0),k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_438, plain, (v1_xxreal_0(esk12_3(X1,X2,X3))|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1275, negated_conjecture, (v1_xxreal_0(k2_rinfsup1(esk2_0))|~v6_membered(k1_numbers)|~m1_subset_1(k10_xtuple_0(esk2_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1307, negated_conjecture, (r1_xxreal_0(esk1_0,k10_xtuple_0(esk2_0))|r1_xxreal_0(k10_xtuple_0(esk2_0),esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_923, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_440, plain, (v1_xreal_0(esk9_3(X1,X2,X3))|~v3_membered(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_441, plain, (v7_ordinal1(esk9_3(X1,X2,X3))|~v6_membered(k9_xtuple_0(X1))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1300, negated_conjecture, (r1_xxreal_0(esk1_0,esk5_1(X1))|r1_xxreal_0(esk5_1(X1),esk1_0)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1414, negated_conjecture, (r1_xxreal_0(X1,esk5_1(X2))|r1_xxreal_0(esk5_1(X2),X1)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X2,k6_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1614, plain, (r1_xxreal_0(k6_numbers,esk9_3(X1,X2,X3))|r1_xxreal_0(esk9_3(X1,X2,X3),k6_numbers)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1662, plain, (k2_xcmplx_0(k3_funct_2(k5_numbers,X1,X2,X3),X4)=k2_xcmplx_0(X4,k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|~v1_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~v1_xreal_0(X4)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_1597, plain, (v1_xboole_0(X1)|r1_xxreal_0(k3_seq_4(k2_relset_1(X1,X2)),k1_funct_1(X2,X3))|~v3_xxreal_2(k2_relset_1(X1,X2))|~v3_membered(k2_relset_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~v1_xreal_0(k1_funct_1(X2,X3))|~v1_xreal_0(k3_seq_4(k2_relset_1(X1,X2))))).#trainneg
cnf(i_0_1383, plain, (r1_xxreal_0(k6_numbers,esk12_3(X1,X2,X3))|r1_xxreal_0(esk12_3(X1,X2,X3),k6_numbers)|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1451, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0),k6_numbers)|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),esk1_0))|~v1_xreal_0(k1_seq_1(esk2_0,esk5_1(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))))|~v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))|~r2_hidden(k1_seq_1(esk2_0,esk5_1(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))),k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_1380, negated_conjecture, (r1_xxreal_0(esk1_0,k9_xtuple_0(esk2_0))|r1_xxreal_0(k9_xtuple_0(esk2_0),esk1_0)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1603, plain, (k2_xcmplx_0(k6_numbers,X1)=k9_binop_2(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1393, negated_conjecture, (r1_xxreal_0(k6_numbers,k9_xtuple_0(esk2_0))|r1_xxreal_0(k9_xtuple_0(esk2_0),k6_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1394, negated_conjecture, (r1_xxreal_0(k6_numbers,k10_xtuple_0(esk2_0))|r1_xxreal_0(k10_xtuple_0(esk2_0),k6_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1312, negated_conjecture, (r1_xxreal_0(esk1_0,k1_seq_1(X1,X2))|r1_xxreal_0(k1_seq_1(X1,X2),esk1_0)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_444, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|~v1_xboole_0(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|~v1_funct_2(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|~v1_funct_1(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_445, plain, (v1_relat_1(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1550, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk9_3(X1,X2,X3)),k1_numbers)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1566, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X2)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk12_3(X2,X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_1673, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk12_3(X2,X1,esk2_0),k5_numbers))).#trainneg
cnf(i_0_1650, negated_conjecture, (k2_xcmplx_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1),k6_numbers)=k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1)|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1617, plain, (r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k5_numbers,X1,X2))|r1_xxreal_0(k3_funct_2(k5_numbers,k5_numbers,X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1474, plain, (r1_xxreal_0(k2_rinfsup1(X1),k2_rinfsup1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1579, negated_conjecture, (k2_xcmplx_0(esk10_2(esk2_0,X1),X2)=k2_xcmplx_0(X2,esk10_2(esk2_0,X1))|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|X1!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1609, plain, (k2_xcmplx_0(X1,k6_numbers)=k9_binop_2(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1200, negated_conjecture, (k6_numbers=k1_numbers|~v1_xboole_0(k10_xtuple_0(esk2_0))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1421, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|v1_xreal_0(k10_xtuple_0(esk2_0))|k10_xtuple_0(X1)!=k1_zfmisc_1(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1399, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2))|r1_xxreal_0(k1_seq_1(X1,X2),k6_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1313, negated_conjecture, (r1_xxreal_0(esk1_0,k2_rinfsup1(X1))|r1_xxreal_0(k2_rinfsup1(X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_446, plain, (v4_relat_1(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_447, plain, (v1_funct_2(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|~v1_partfun1(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1077, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(X2,X3))|~v3_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_779, negated_conjecture, (r1_xxreal_0(esk5_1(X1),X2)|r1_xxreal_0(X2,esk5_1(X1))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1161, plain, (r1_xxreal_0(X1,k1_funct_1(X2,esk9_3(X2,X3,X4)))|X1!=k3_seq_4(X5)|~v3_xxreal_2(X5)|~v3_membered(X5)|~v1_xreal_0(X1)|~v1_xreal_0(k1_funct_1(X2,esk9_3(X2,X3,X4)))|X5!=k10_xtuple_0(X2)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1895, negated_conjecture, (v1_xreal_0(esk9_3(esk2_0,X1,X2))|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_457, plain, (v7_ordinal1(esk6_3(X1,X2,X3))|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|~v6_membered(X1)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1777, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_448, plain, (v1_partfun1(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1)|v1_xboole_0(X2)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|~v1_funct_2(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X1,X2)|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_449, plain, (v3_valued_0(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|~v3_membered(X2)|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1791, plain, (v1_xreal_0(esk9_3(X1,X2,X3))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1996, negated_conjecture, (r1_xxreal_0(X1,X3)|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk9_3(esk2_0,X2,X3),k5_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_2147, plain, (v1_xboole_0(X1)|r1_xxreal_0(k3_seq_4(k2_relset_1(X1,X2)),X4)|~v3_xxreal_2(k2_relset_1(X1,X2))|~v3_membered(k2_relset_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk9_3(X2,X3,X4),k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~v1_xreal_0(X4)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3))).#trainneg
cnf(i_0_2151, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0),k6_numbers)|~r2_hidden(k1_funct_1(esk2_0,esk5_1(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0)))).#trainneg
cnf(i_0_2268, plain, (r1_xxreal_0(k3_seq_4(X1),k1_funct_1(X2,esk9_3(X2,X3,X4)))|X1!=k10_xtuple_0(X2)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2)|~v1_xreal_0(k1_funct_1(X2,esk9_3(X2,X3,X4)))|~v1_xreal_0(k3_seq_4(X1)))).#trainneg
cnf(i_0_2267, negated_conjecture, (r1_xxreal_0(X1,X2)|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk12_3(X2,k1_numbers,esk2_0),k5_numbers)|~v1_xreal_0(X1)|~r2_hidden(X2,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_463, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k9_xtuple_0(X2),X3),X4)|~m1_subset_1(X3,X1)|~v1_funct_2(k9_xtuple_0(X2),X1,X4)|~v1_funct_1(k9_xtuple_0(X2))|~v4_relat_1(X2,k2_zfmisc_1(X1,X4))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1967, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|r1_xxreal_0(X1,k2_rinfsup1(esk2_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_450, plain, (v1_valued_0(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4))|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|~v1_membered(X2)|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_451, plain, (v5_relat_1(esk6_3(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4),X2)|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r1_xxreal_0(X4,k6_numbers)|X3!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1914, negated_conjecture, (m1_subset_1(X2,k10_xtuple_0(esk2_0))|~m1_subset_1(esk9_3(esk2_0,X1,X2),k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2148, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(k1_funct_1(esk2_0,X1)))).#trainneg
cnf(i_0_1923, negated_conjecture, (r2_hidden(X1,k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_2150, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~r2_hidden(k1_seq_1(esk2_0,esk5_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0))),k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_464, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k2_zfmisc_1(X1,X2),X3),X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2097, negated_conjecture, (r1_xxreal_0(k6_numbers,k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1999, negated_conjecture, (v7_ordinal1(k10_xtuple_0(esk2_0))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2083, negated_conjecture, (v7_ordinal1(k9_xtuple_0(esk2_0))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1966, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),esk1_0)|r1_xxreal_0(esk1_0,k2_rinfsup1(esk2_0)))).#trainneg
cnf(i_0_453, plain, (v3_membered(esk6_3(k1_zfmisc_1(X1),X2,X3))|v1_xboole_0(k1_zfmisc_1(X1))|r1_xxreal_0(X3,k6_numbers)|~v3_membered(X1)|X2!=k3_seq_4(k1_zfmisc_1(X1))|~v3_xxreal_2(k1_zfmisc_1(X1))|~v3_membered(k1_zfmisc_1(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_454, plain, (v3_valued_0(esk6_3(k1_zfmisc_1(X1),X2,X3))|v1_xboole_0(k1_zfmisc_1(X1))|r1_xxreal_0(X3,k6_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|X2!=k3_seq_4(k1_zfmisc_1(X1))|~v3_xxreal_2(k1_zfmisc_1(X1))|~v3_membered(k1_zfmisc_1(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2040, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,X1,esk2_0,X2),k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_2270, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(X2,esk9_3(X2,X3,X4)))|X1!=k2_rinfsup1(esk2_0)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2)|~v1_xreal_0(k1_funct_1(X2,esk9_3(X2,X3,X4)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2046, negated_conjecture, (r1_xxreal_0(X1,k3_funct_2(k5_numbers,X2,esk2_0,X3))|v1_xboole_0(X2)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,X2,esk2_0,X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_2067, plain, (k9_binop_2(k6_numbers,X1)=k9_binop_2(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2027, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|r1_xxreal_0(k6_numbers,k2_rinfsup1(esk2_0)))).#trainneg
cnf(i_0_1610, plain, (k2_xcmplx_0(X1,k6_numbers)=k9_binop_2(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2096, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k2_rinfsup1(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1804, negated_conjecture, (r1_xxreal_0(k10_xtuple_0(esk2_0),k10_xtuple_0(esk2_0))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_455, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k1_zfmisc_1(X2))|r1_xxreal_0(X4,k6_numbers)|~r2_hidden(X1,esk6_3(k1_zfmisc_1(X2),X3,X4))|X3!=k3_seq_4(k1_zfmisc_1(X2))|~v3_xxreal_2(k1_zfmisc_1(X2))|~v3_membered(k1_zfmisc_1(X2))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_460, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4),X2)|~m1_subset_1(X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_2018, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk2_0,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2049, negated_conjecture, (v1_xreal_0(k3_funct_2(k5_numbers,X1,esk2_0,X2))|v1_xboole_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_2262, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(k2_rinfsup1(esk2_0),esk5_1(X1))|r1_xxreal_0(esk5_1(X1),k2_rinfsup1(esk2_0))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1798, negated_conjecture, (k6_numbers=k1_numbers|v1_xreal_0(k1_funct_1(esk2_0,X1))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_500, plain, (v1_xreal_0(k2_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1805, negated_conjecture, (r1_xxreal_0(k6_numbers,k10_xtuple_0(esk2_0))|r1_xxreal_0(k10_xtuple_0(esk2_0),k6_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2102, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),X2)|r1_xxreal_0(X2,k1_seq_1(esk2_0,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_461, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X6),X3)|v1_xboole_0(X2)|~m1_subset_1(X6,X1)|~v1_funct_2(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X1,X3)|~v1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_462, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4),X2)|~m1_subset_1(X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_2081, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_valued_0(esk2_0)|~v1_xreal_0(X2)|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_2006, plain, (r1_xxreal_0(esk9_3(X1,X2,X3),k6_numbers)|r1_xxreal_0(k6_numbers,esk9_3(X1,X2,X3))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k1_numbers)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2152, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0),k6_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))|~m1_subset_1(k1_seq_1(esk2_0,esk5_1(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_1913, negated_conjecture, (m1_subset_1(X1,k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|~m1_subset_1(esk12_3(X1,X2,esk2_0),k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_1970, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1926, plain, (v7_ordinal1(k3_seq_4(X1))|~v6_membered(k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_501, plain, (v1_xxreal_0(k2_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2098, negated_conjecture, (r1_xxreal_0(X1,k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2100, negated_conjecture, (r1_xxreal_0(X1,k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|~v1_xboole_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2192, negated_conjecture, (v1_xreal_0(X2)|~m1_subset_1(esk9_3(esk2_0,X1,X2),k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2123, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk5_1(X1),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1073, plain, (v1_xboole_0(X1)|v1_xreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|~v3_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2037, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2)|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_1019, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1909, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,X1))|~v6_membered(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2101, negated_conjecture, (r1_xxreal_0(X1,k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|~m1_subset_1(X1,k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_496, plain, (v1_xreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|k2_relset_1(X1,X2)!=k10_xtuple_0(X4)|~v3_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2264, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,esk5_1(X2))|r1_xxreal_0(esk5_1(X2),X1)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1515, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(k5_numbers,X2,X3,X4),X1)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2))).#trainneg
cnf(i_0_709, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X2,X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2188, negated_conjecture, (r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|r1_xxreal_0(k1_funct_1(esk2_0,X1),esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1554, negated_conjecture, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_1655, negated_conjecture, (k2_xcmplx_0(X1,k6_numbers)=X1|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_507, plain, (k2_xcmplx_0(X3,k6_numbers)=X3|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_508, plain, (v1_xxreal_0(esk9_3(X1,X2,X3))|k9_xtuple_0(X1)!=k10_xtuple_0(X4)|~v2_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_509, plain, (v1_xxreal_0(esk7_2(X1,X2))|X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,X2))|X1!=k10_xtuple_0(X3)|~v2_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2050, negated_conjecture, (m1_subset_1(k3_funct_2(k5_numbers,X1,esk2_0,X2),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_1924, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),X2)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_1002, plain, (v1_xboole_0(k2_relset_1(X1,X2))|r1_xxreal_0(X3,k3_funct_2(k5_numbers,X1,X2,X4))|v1_xboole_0(X1)|X3!=k3_seq_4(k2_relset_1(X1,X2))|~v3_xxreal_2(k2_relset_1(X1,X2))|~v3_membered(k2_relset_1(X1,X2))|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2014, plain, (r1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,X1,X2),k6_numbers)|r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1066, plain, (v1_xreal_0(k9_xtuple_0(k2_zfmisc_1(X1,X2)))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1262, plain, (v1_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_2230, negated_conjecture, (k2_rinfsup1(esk2_0)=k6_numbers|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),k6_numbers),k10_xtuple_0(esk2_0))|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),k6_numbers),k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),k6_numbers)))).#trainneg
cnf(i_0_514, plain, (v1_xxreal_0(X1)|v1_xboole_0(X2)|X2!=k10_xtuple_0(X3)|~v2_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_510, plain, (v1_xxreal_0(esk7_2(X1,X2))|X2=k3_seq_4(X1)|v1_xboole_0(X1)|X1!=k10_xtuple_0(X3)|~v2_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_511, plain, (v1_xxreal_0(esk6_3(X1,X2,X3))|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X1!=k10_xtuple_0(X4)|~v2_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2332, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k2_rinfsup1(esk2_0)))).#trainneg
cnf(i_0_2274, negated_conjecture, (m1_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2341, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),X2)|~m1_subset_1(esk9_3(esk2_0,X1,X2),k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2287, negated_conjecture, (r1_xxreal_0(X1,X2)|X1!=k2_rinfsup1(esk2_0)|X3!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2389, plain, (v1_xreal_0(k9_binop_2(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2399, plain, (v1_xreal_0(k9_binop_2(X1,k6_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2701, negated_conjecture, (k2_xcmplx_0(esk10_2(esk2_0,X1),k6_numbers)=esk10_2(esk2_0,X1)|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|X1!=k10_xtuple_0(X2)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_515, plain, (v1_xcmplx_0(esk9_3(X1,X2,X3))|k9_xtuple_0(X1)!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2372, negated_conjecture, (v1_xxreal_0(k10_xtuple_0(esk2_0))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2374, negated_conjecture, (v1_xxreal_0(k9_xtuple_0(esk2_0))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2474, negated_conjecture, (r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,X1))|r1_xxreal_0(k1_seq_1(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_2490, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,X1))|r1_xxreal_0(k1_seq_1(esk2_0,X1),k6_numbers))).#trainneg
cnf(i_0_512, plain, (v1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|k2_relset_1(X1,X2)!=k10_xtuple_0(X4)|~v2_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_516, plain, (v1_xcmplx_0(esk7_2(X1,X2))|X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,X2))|X1!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2339, negated_conjecture, (m1_subset_1(X1,k10_xtuple_0(esk2_0))|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2335, negated_conjecture, (r1_xxreal_0(X1,k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2304, negated_conjecture, (r1_xxreal_0(X1,X2)|X1!=k2_rinfsup1(esk2_0)|~r2_hidden(X2,k10_xtuple_0(esk2_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2447, negated_conjecture, (r1_xxreal_0(k6_numbers,X2)|r1_xxreal_0(X2,k6_numbers)|~m1_subset_1(esk9_3(esk2_0,X1,X2),k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2362, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,k3_funct_2(k5_numbers,X1,esk2_0,X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2302, negated_conjecture, (r1_xxreal_0(X1,X3)|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk9_3(esk2_0,X2,X3),k5_numbers)|~v1_xreal_0(X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_2360, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2685, negated_conjecture, (r1_xxreal_0(X2,esk1_0)|r1_xxreal_0(esk1_0,X2)|~m1_subset_1(esk9_3(esk2_0,X1,X2),k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2388, plain, (m1_subset_1(k9_binop_2(k6_numbers,X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2325, negated_conjecture, (r1_xxreal_0(k9_xtuple_0(esk2_0),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k9_xtuple_0(esk2_0))|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2525, negated_conjecture, (v1_xreal_0(X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_521, plain, (v1_xcmplx_0(X1)|v1_xboole_0(X2)|X2!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_2326, negated_conjecture, (r1_xxreal_0(k10_xtuple_0(esk2_0),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k10_xtuple_0(esk2_0))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2432, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_2334, negated_conjecture, (r1_xxreal_0(X1,k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2494, negated_conjecture, (r1_xxreal_0(X1,k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_517, plain, (v1_xcmplx_0(esk7_2(X1,X2))|X2=k3_seq_4(X1)|v1_xboole_0(X1)|X1!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_518, plain, (v1_xcmplx_0(esk6_3(X1,X2,X3))|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X1!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2681, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0),k6_numbers)|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk5_1(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))))|~v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk5_1(esk8_2(k2_relset_1(X1,esk2_0),esk1_0)),k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_2367, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,X2)|v1_xboole_0(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k2_relset_1(X1,esk2_0)))).#trainneg
cnf(i_0_2385, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(X2,X3)|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X3,X1,esk2_0),k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r2_hidden(X3,k2_relset_1(X1,esk2_0)))).#trainneg
cnf(i_0_2680, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(X1)=esk1_0|r1_xxreal_0(esk8_2(X1,esk1_0),k6_numbers)|v1_xboole_0(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,esk7_2(X1,esk1_0))|~v1_xreal_0(k1_funct_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))))|~v1_xreal_0(esk8_2(X1,esk1_0))|~m1_subset_1(k1_funct_1(esk2_0,esk5_1(esk8_2(X1,esk1_0))),X1))).#trainneg
cnf(i_0_2563, plain, (r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,X1,X2,X3))|r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k6_numbers)|v1_xboole_0(X1)|~v3_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2498, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0))|~v1_xreal_0(X2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2301, negated_conjecture, (r1_xxreal_0(X1,X2)|v1_xboole_0(X3)|X1!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk12_3(X2,X3,esk2_0),k5_numbers)|~v1_xreal_0(X1)|~r2_hidden(X2,k2_relset_1(X3,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X3))).#trainneg
cnf(i_0_2398, plain, (m1_subset_1(k9_binop_2(X1,k6_numbers),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2442, negated_conjecture, (v1_xreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_2443, negated_conjecture, (m1_subset_1(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k1_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_2433, negated_conjecture, (m1_subset_1(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k10_xtuple_0(esk2_0))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_2495, negated_conjecture, (r1_xxreal_0(X1,k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2492, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k1_seq_1(esk2_0,X1))|r1_xxreal_0(k1_seq_1(esk2_0,X1),k2_rinfsup1(esk2_0)))).#trainneg
cnf(i_0_519, plain, (v1_xcmplx_0(k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|k2_relset_1(X1,X2)!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_523, plain, (v1_xboole_0(esk6_3(k1_zfmisc_1(X1),X2,X3))|v1_xboole_0(X1)|m1_subset_1(X4,X1)|v1_xboole_0(k1_zfmisc_1(X1))|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(X4,esk6_3(k1_zfmisc_1(X1),X2,X3))|X2!=k3_seq_4(k1_zfmisc_1(X1))|~v3_xxreal_2(k1_zfmisc_1(X1))|~v3_membered(k1_zfmisc_1(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2366, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2293, plain, (r1_xxreal_0(k3_seq_4(X1),X4)|X1!=k10_xtuple_0(X2)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_2380, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k3_funct_2(k5_numbers,X2,esk2_0,X3))|v1_xboole_0(X2)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_2622, plain, (r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k6_numbers)|r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2687, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|m1_subset_1(esk9_3(X1,X2,X3),k1_numbers)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k10_xtuple_0(esk2_0))|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2347, negated_conjecture, (r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),X2),k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2351, negated_conjecture, (r2_hidden(k1_funct_1(X1,esk9_3(X1,X2,X3)),k10_xtuple_0(esk2_0))|v1_xboole_0(X4)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(esk2_0,k5_numbers,X4)|k2_relset_1(X4,esk2_0)!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2514, plain, (v1_xxreal_0(k3_seq_4(X1))|~v6_membered(k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_2390, plain, (v7_ordinal1(k9_binop_2(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2441, negated_conjecture, (k2_xcmplx_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)=esk7_2(k10_xtuple_0(esk2_0),esk1_0)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|~v1_valued_0(esk2_0)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_1770, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),k1_numbers)|k6_numbers=k1_numbers|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_2776, negated_conjecture, (v1_xxreal_0(esk1_0)|k10_xtuple_0(X1)!=k1_numbers|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2491, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(esk2_0,X3))|r1_xxreal_0(k1_seq_1(esk2_0,X3),k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2113, plain, (r1_xxreal_0(k1_seq_1(k5_numbers,X1),X2)|r1_xxreal_0(X2,k1_seq_1(k5_numbers,X1))|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v1_xreal_0(X2)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_526, plain, (v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|v1_xboole_0(X2)|m1_subset_1(X5,X2)|v1_xboole_0(X1)|~m1_subset_1(X5,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_528, plain, (v1_xboole_0(k9_xtuple_0(X1))|v1_xboole_0(X2)|m1_subset_1(X3,X2)|~m1_subset_1(X3,k9_xtuple_0(X1))|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2628, plain, (r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,X1,X2,X3))|r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k6_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2295, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k1_funct_1(X1,esk9_3(X1,X2,X3)))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_funct_1(X1,esk9_3(X1,X2,X3))))).#trainneg
cnf(i_0_2559, negated_conjecture, (r1_xxreal_0(esk1_0,k3_funct_2(k5_numbers,X1,X2,X3))|r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),esk1_0)|v1_xboole_0(X1)|~v3_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2340, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|v1_xboole_0(X2)|~m1_subset_1(esk12_3(X1,X2,esk2_0),k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_2392, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,k6_numbers)|k3_seq_4(X1)!=k6_numbers|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k9_binop_2(k6_numbers,X2),esk6_3(X1,k6_numbers,X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2391, plain, (v1_xxreal_0(k9_binop_2(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2485, negated_conjecture, (r1_xxreal_0(k9_xtuple_0(esk2_0),k1_seq_1(esk2_0,X1))|r1_xxreal_0(k1_seq_1(esk2_0,X1),k9_xtuple_0(esk2_0))|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_531, plain, (v1_xreal_0(k1_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2805, negated_conjecture, (v1_xxreal_0(k2_rinfsup1(esk2_0))|k10_xtuple_0(X1)!=k1_numbers|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2796, negated_conjecture, (v1_xxreal_0(k9_xtuple_0(esk2_0))|k10_xtuple_0(X1)!=k1_numbers|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1956, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_532, plain, (v1_xxreal_0(k1_relset_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2331, negated_conjecture, (r1_xxreal_0(k1_seq_1(X1,X2),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2275, negated_conjecture, (r1_xxreal_0(k6_numbers,esk9_3(esk2_0,X1,X2))|r1_xxreal_0(esk9_3(esk2_0,X1,X2),k6_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2300, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|k2_xcmplx_0(X1,esk1_0)!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk5_1(X1),k5_numbers)|~v1_xreal_0(k2_xcmplx_0(X1,esk1_0)))).#trainneg
cnf(i_0_886, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|X1!=k1_funct_1(X3,k3_funct_2(k5_numbers,k9_xtuple_0(X3),X4,X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k9_xtuple_0(X3))))|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,k9_xtuple_0(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2400, plain, (v7_ordinal1(k9_binop_2(X1,k6_numbers))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2401, plain, (v1_xxreal_0(k9_binop_2(X1,k6_numbers))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2798, negated_conjecture, (v1_xxreal_0(k10_xtuple_0(esk2_0))|k10_xtuple_0(X1)!=k1_numbers|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2807, plain, (v1_xxreal_0(k5_numbers)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(X1)!=k1_zfmisc_1(k1_numbers)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1958, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~v1_xboole_0(X1)|~v6_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1959, plain, (r1_xxreal_0(X1,k1_seq_1(X2,X3))|r1_xxreal_0(k1_seq_1(X2,X3),X1)|~m1_subset_1(X1,k5_numbers)|~v6_membered(k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_536, plain, (k2_xcmplx_0(k6_numbers,k1_funct_1(X1,X2))=k1_funct_1(X1,X2)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_542, plain, (X1=k10_xtuple_0(k2_relset_1(k2_zfmisc_1(X2,X3),X4))|r2_hidden(esk11_2(k2_relset_1(k2_zfmisc_1(X2,X3),X4),X1),k9_xtuple_0(k2_relset_1(k2_zfmisc_1(X2,X3),X4)))|r2_hidden(esk10_2(k2_relset_1(k2_zfmisc_1(X2,X3),X4),X1),X1)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X2,X3),X4))|~v5_relat_1(X4,k2_zfmisc_1(X2,X3))|~v1_relat_1(X4))).#trainneg
cnf(i_0_1321, negated_conjecture, (X1=k3_seq_4(k10_xtuple_0(esk2_0))|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|r1_xxreal_0(k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1)),k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1548, negated_conjecture, (m1_subset_1(esk6_3(k10_xtuple_0(esk2_0),X1,X2),k1_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k3_seq_4(k10_xtuple_0(esk2_0))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1505, plain, (m1_subset_1(k3_funct_2(k5_numbers,X1,X2,X3),k2_relset_1(X4,X2))|v1_xboole_0(X4)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3082, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2837, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_566, plain, (v1_xxreal_0(k3_funct_2(X1,X2,X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2878, negated_conjecture, (r1_xxreal_0(k10_xtuple_0(esk2_0),k10_xtuple_0(esk2_0))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2884, negated_conjecture, (r1_xxreal_0(k9_xtuple_0(esk2_0),k9_xtuple_0(esk2_0))|~v6_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_543, plain, (k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),esk11_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4))=esk10_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4)|X4=k10_xtuple_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|r2_hidden(esk10_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4),X4)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_568, plain, (v7_ordinal1(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3190, negated_conjecture, (k2_xcmplx_0(esk7_2(k2_relset_1(X1,esk2_0),X2),X3)=k2_xcmplx_0(X3,esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v1_valued_0(esk2_0)|~v1_xreal_0(X3)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3057, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|v1_xboole_0(k2_relset_1(X1,esk2_0))|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0)))).#trainneg
cnf(i_0_3343, plain, (k2_xcmplx_0(k6_numbers,X2)=X2|v1_xboole_0(X3)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k2_relset_1(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3))).#trainneg
cnf(i_0_3344, plain, (k2_xcmplx_0(k6_numbers,X3)=X3|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_2916, negated_conjecture, (r1_xxreal_0(X1,X2)|v1_xboole_0(k10_xtuple_0(esk2_0))|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~m1_subset_1(X2,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_2821, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_569, plain, (X1=k3_seq_4(k2_relset_1(k1_numbers,X2))|v1_xboole_0(k2_relset_1(k1_numbers,X2))|m1_subset_1(esk7_2(k2_relset_1(k1_numbers,X2),X1),k2_relset_1(k1_numbers,X2))|v1_xreal_0(esk8_2(k2_relset_1(k1_numbers,X2),X1))|~v3_membered(k2_relset_1(k1_numbers,X2))|~v1_xreal_0(X1)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_571, plain, (v3_valued_0(k1_relset_1(k2_zfmisc_1(X1,k1_numbers),X2))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,k1_numbers),X2))|~v4_relat_1(X2,k2_zfmisc_1(X1,k1_numbers))|~v1_relat_1(X2))).#trainneg
cnf(i_0_3357, negated_conjecture, (r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2),k6_numbers)|r1_xxreal_0(k6_numbers,esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3175, negated_conjecture, (X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v1_xboole_0(k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3192, negated_conjecture, (k2_xcmplx_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X3))=k2_xcmplx_0(esk7_2(k2_relset_1(X2,esk2_0),X3),X1)|X3=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|~v1_valued_0(esk2_0)|~v1_xcmplx_0(X1)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X2,esk2_0),X3),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3184, negated_conjecture, (r1_xxreal_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X3))|X3=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(k2_relset_1(X2,esk2_0),X3))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X2,esk2_0),X3),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2920, negated_conjecture, (r1_xxreal_0(X1,X2)|X1!=k2_rinfsup1(esk2_0)|X3!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2909, negated_conjecture, (r1_xxreal_0(X1,esk6_3(k10_xtuple_0(esk2_0),X2,X3))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2922, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_3355, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),esk6_3(k10_xtuple_0(esk2_0),X1,X2))|r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2),k2_rinfsup1(esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3195, negated_conjecture, (X1=k2_rinfsup1(esk2_0)|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),X1),k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_572, plain, (k2_relset_1(X1,k1_relset_1(k2_zfmisc_1(X2,X1),X3))=k10_xtuple_0(k1_relset_1(k2_zfmisc_1(X2,X1),X3))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X2,X1),X3))|~v4_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_573, plain, (v1_xreal_0(k2_relset_1(X1,k1_relset_1(k2_zfmisc_1(X2,X1),X3)))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X2,X1),X3))|~v3_membered(k1_zfmisc_1(X1))|~v4_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3260, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(esk2_0),X2)|~r2_hidden(X2,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_3224, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_3079, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(X1)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk8_2(X1,esk1_0),k6_numbers)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,esk7_2(X1,esk1_0))|~v1_xreal_0(esk8_2(X1,esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk5_1(esk8_2(X1,esk1_0)),k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_2816, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3348, negated_conjecture, (v1_xreal_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2815, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3083, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(X2,X3)|X2!=k2_rinfsup1(esk2_0)|~r2_hidden(X3,k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3187, negated_conjecture, (k2_xcmplx_0(esk7_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)=esk7_2(k2_relset_1(X1,esk2_0),X2)|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v1_valued_0(esk2_0)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_576, plain, (v1_xboole_0(X1)|v3_membered(k1_funct_1(X3,X4))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1552, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),k1_numbers)|k6_numbers=k1_numbers|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_3104, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_3235, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k1_seq_1(k5_numbers,X2))|r1_xxreal_0(k1_seq_1(k5_numbers,X2),k1_seq_1(esk2_0,X1))|~v1_relat_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1954, plain, (r1_xxreal_0(k1_seq_1(X1,X2),k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_578, plain, (v3_valued_0(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_581, plain, (X1=k10_xtuple_0(k1_relset_1(k2_zfmisc_1(X2,X3),X4))|r2_hidden(esk11_2(k1_relset_1(k2_zfmisc_1(X2,X3),X4),X1),k9_xtuple_0(k1_relset_1(k2_zfmisc_1(X2,X3),X4)))|r2_hidden(esk10_2(k1_relset_1(k2_zfmisc_1(X2,X3),X4),X1),X1)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X3),X4))|~v4_relat_1(X4,k2_zfmisc_1(X2,X3))|~v1_relat_1(X4))).#trainneg
cnf(i_0_3149, plain, (r1_xxreal_0(k3_seq_4(X1),esk6_3(X2,X3,X4))|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|X1!=k10_xtuple_0(X5)|X2!=k10_xtuple_0(X5)|~v1_relat_1(X5)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X5)|~v1_xreal_0(esk6_3(X2,X3,X4))|X3!=k3_seq_4(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3067, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,esk6_3(k2_relset_1(X1,esk2_0),X2,X3))|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|X2!=k3_seq_4(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3069, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(esk1_0,k3_funct_2(k5_numbers,k2_relset_1(X1,esk2_0),X2,X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_relset_1(X1,esk2_0))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k2_relset_1(X1,esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3176, negated_conjecture, (m1_subset_1(esk7_2(k2_relset_1(X1,esk2_0),X2),k10_xtuple_0(esk2_0))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2917, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_3356, negated_conjecture, (r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2),esk1_0)|r1_xxreal_0(esk1_0,esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2383, negated_conjecture, (r1_xxreal_0(X1,X4)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|X1!=k2_rinfsup1(esk2_0)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3077, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(k1_numbers)=esk1_0|r1_xxreal_0(esk8_2(k1_numbers,esk1_0),k6_numbers)|~v3_xxreal_2(k1_numbers)|~r1_xxreal_0(esk1_0,esk7_2(k1_numbers,esk1_0))|~v1_xreal_0(esk8_2(k1_numbers,esk1_0)))).#trainneg
cnf(i_0_1400, plain, (r1_xxreal_0(k6_numbers,k2_rinfsup1(X1))|r1_xxreal_0(k2_rinfsup1(X1),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_582, plain, (k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),esk11_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4))=esk10_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4)|X4=k10_xtuple_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|r2_hidden(esk10_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4),X4)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_585, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~r2_hidden(X2,k1_funct_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X3))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3164, plain, (r1_xxreal_0(k3_seq_4(X1),X2)|v1_xboole_0(X4)|X1!=k10_xtuple_0(X3)|X4!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X3)|~v1_xreal_0(X2)|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_3373, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X4,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X4,X1,X3),k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3)|~r2_hidden(X4,k2_relset_1(X1,X3)))).#trainneg
cnf(i_0_904, plain, (v1_xreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1385, plain, (r1_xxreal_0(k6_numbers,esk6_3(X1,X2,X3))|r1_xxreal_0(esk6_3(X1,X2,X3),k6_numbers)|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_591, plain, (v3_valued_0(k2_relset_1(k2_zfmisc_1(X1,k1_numbers),X2))|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,k1_numbers),X2))|~v5_relat_1(X2,k2_zfmisc_1(X1,k1_numbers))|~v1_relat_1(X2))).#trainneg
cnf(i_0_3074, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k3_seq_4(k2_zfmisc_1(k5_numbers,k1_numbers))=esk1_0|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|r1_xxreal_0(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0),k6_numbers)|~v3_xxreal_2(k2_zfmisc_1(k5_numbers,k1_numbers))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~r1_xxreal_0(esk1_0,esk7_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))|~v1_xreal_0(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))|~r2_hidden(k1_funct_1(esk2_0,esk5_1(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))),esk2_0))).#trainneg
cnf(i_0_2111, plain, (r1_xxreal_0(k1_seq_1(k2_zfmisc_1(X1,k1_numbers),X2),X3)|r1_xxreal_0(X3,k1_seq_1(k2_zfmisc_1(X1,k1_numbers),X2))|~v1_funct_1(k2_zfmisc_1(X1,k1_numbers))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_589, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|v1_xboole_0(k3_funct_2(X1,k1_zfmisc_1(X3),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X3))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(X3))|~v1_funct_1(X4)|~m1_subset_1(X2,k3_funct_2(X1,k1_zfmisc_1(X3),X4,X5)))).#trainneg
cnf(i_0_590, negated_conjecture, (k10_xtuple_0(X1)=k2_zfmisc_1(k5_numbers,k1_numbers)|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|esk10_2(X1,k2_zfmisc_1(k5_numbers,k1_numbers))!=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~r2_hidden(esk10_2(X1,k2_zfmisc_1(k5_numbers,k1_numbers)),esk2_0))).#trainneg
cnf(i_0_2074, negated_conjecture, (k2_xcmplx_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2),X3)=k2_xcmplx_0(X3,esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|~v1_valued_0(esk2_0)|~v1_xreal_0(X3)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3832, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(X2,k2_relset_1(X1,esk2_0))|~r2_hidden(X2,k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_3669, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_3762, negated_conjecture, (k3_seq_4(k1_numbers)=esk1_0|k2_rinfsup1(esk2_0)=esk1_0|~v3_xxreal_2(k1_numbers)|~r1_xxreal_0(esk1_0,esk7_2(k1_numbers,esk1_0))|~v1_xreal_0(esk8_2(k1_numbers,esk1_0)))).#trainneg
cnf(i_0_592, plain, (k2_relset_1(X1,k2_relset_1(k2_zfmisc_1(X2,X1),X3))=k10_xtuple_0(k2_relset_1(k2_zfmisc_1(X2,X1),X3))|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X2,X1),X3))|~v5_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3643, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X3),k1_seq_1(esk2_0,X3)))).#trainneg
cnf(i_0_593, plain, (v1_xreal_0(k2_relset_1(X1,k2_relset_1(k2_zfmisc_1(X2,X1),X3)))|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X2,X1),X3))|~v3_membered(k1_zfmisc_1(X1))|~v5_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3))).#trainneg
cnf(i_0_3681, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(X2,X3)|v1_xboole_0(k2_relset_1(X1,esk2_0))|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2)|~m1_subset_1(X3,k2_relset_1(X1,esk2_0)))).#trainneg
cnf(i_0_3497, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|v1_xboole_0(X2)|X2!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_3548, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_3421, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_3726, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_3476, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk10_2(esk2_0,X1))=esk10_2(esk2_0,X1)|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|X1!=k10_xtuple_0(X2)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3708, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,X1))|k6_numbers=k1_numbers|~v6_membered(k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_3233, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k1_seq_1(k2_zfmisc_1(X2,k1_numbers),X3))|r1_xxreal_0(k1_seq_1(k2_zfmisc_1(X2,k1_numbers),X3),k1_seq_1(esk2_0,X1))|~v1_funct_1(k2_zfmisc_1(X2,k1_numbers)))).#trainneg
cnf(i_0_599, plain, (v1_relat_1(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_601, plain, (v3_valued_0(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3641, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(esk2_0),X2)|v1_xboole_0(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k2_relset_1(X1,esk2_0)))).#trainneg
cnf(i_0_3858, plain, (m1_subset_1(k3_funct_2(k5_numbers,X1,X2,X3),k1_numbers)|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3545, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k6_numbers)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|k9_binop_2(X2,X1)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(k9_binop_2(X2,X1)))).#trainneg
cnf(i_0_3874, negated_conjecture, (k3_seq_4(k2_zfmisc_1(k5_numbers,k1_numbers))=esk1_0|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|r1_xxreal_0(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0),k6_numbers)|~v3_xxreal_2(k2_zfmisc_1(k5_numbers,k1_numbers))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~r1_xxreal_0(esk1_0,esk7_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))|~v1_xreal_0(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))|~m1_subset_1(k1_funct_1(esk2_0,esk5_1(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))),esk2_0))).#trainneg
cnf(i_0_2953, negated_conjecture, (v1_xcmplx_0(esk1_0)|k10_xtuple_0(X1)!=k1_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1187, plain, (r1_xxreal_0(X1,k2_rinfsup1(X2))|r1_xxreal_0(k2_rinfsup1(X2),X1)|~v1_xreal_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2112, plain, (r1_xxreal_0(k1_seq_1(k2_zfmisc_1(X1,X2),X3),X4)|r1_xxreal_0(X4,k1_seq_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xreal_0(X4)|~v3_membered(X2))).#trainneg
cnf(i_0_604, plain, (X1=k3_seq_4(k9_xtuple_0(X2))|r2_hidden(k1_funct_1(X2,esk7_2(k9_xtuple_0(X2),X1)),X3)|v1_xboole_0(k9_xtuple_0(X2))|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X2),X1),k6_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3546, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k6_numbers)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|k9_binop_2(X1,X2)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(k9_binop_2(X1,X2)))).#trainneg
cnf(i_0_3814, plain, (v1_xboole_0(X1)|r1_xxreal_0(k3_seq_4(X2),k1_funct_1(X3,X4))|X2!=k10_xtuple_0(X5)|X1!=k10_xtuple_0(X5)|~v1_relat_1(X5)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_funct_1(X5)|~v1_xreal_0(k1_funct_1(X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3635, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(esk2_0),esk6_3(k2_relset_1(X1,esk2_0),X2,X3))|v1_xboole_0(k2_relset_1(X1,esk2_0))|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|X2!=k3_seq_4(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3471, plain, (k2_xcmplx_0(k6_numbers,k1_funct_1(X1,X2))=k1_funct_1(X1,X2)|v1_xboole_0(X3)|~v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3457, negated_conjecture, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_valued_0(esk2_0)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_3655, negated_conjecture, (k3_seq_4(X1)=esk1_0|v1_xboole_0(X1)|k2_rinfsup1(esk2_0)=esk1_0|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk1_0,esk7_2(X1,esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(esk8_2(X1,esk1_0)))).#trainneg
cnf(i_0_3463, plain, (k2_xcmplx_0(k6_numbers,k3_funct_2(k5_numbers,X1,X2,X3))=k3_funct_2(k5_numbers,X1,X2,X3)|v1_xboole_0(X1)|~v1_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3794, negated_conjecture, (r1_xxreal_0(k3_seq_4(X1),esk1_0)|X1!=k10_xtuple_0(X2)|k10_xtuple_0(X2)!=k1_numbers|~v1_relat_1(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_615, plain, (v5_relat_1(k1_funct_1(X4,X5),X3)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_2983, negated_conjecture, (v1_xcmplx_0(k2_rinfsup1(esk2_0))|k10_xtuple_0(X1)!=k1_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2973, negated_conjecture, (v1_xcmplx_0(k9_xtuple_0(esk2_0))|k10_xtuple_0(X1)!=k1_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1186, plain, (r1_xxreal_0(X1,k2_rinfsup1(X2))|r1_xxreal_0(k2_rinfsup1(X2),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1188, plain, (r1_xxreal_0(X1,k2_rinfsup1(X2))|r1_xxreal_0(k2_rinfsup1(X2),X1)|~v1_xboole_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_606, plain, (X1=k3_seq_4(k9_xtuple_0(X2))|r2_hidden(k1_funct_1(X2,esk7_2(k9_xtuple_0(X2),X1)),X3)|v1_xboole_0(k9_xtuple_0(X2))|v1_xreal_0(esk8_2(k9_xtuple_0(X2),X1))|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_608, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xboole_0(k1_funct_1(X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))|~m1_subset_1(X5,X1)|~v1_funct_2(k1_funct_1(X4,X5),X3,X2)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(k1_funct_1(X4,X5))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3634, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(esk2_0),esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3285, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k9_xtuple_0(esk2_0))|r1_xxreal_0(k9_xtuple_0(esk2_0),k1_seq_1(esk2_0,X1))|~m1_subset_1(k1_zfmisc_1(k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3824, plain, (r1_xxreal_0(k3_seq_4(X1),k6_numbers)|X1!=k10_xtuple_0(X2)|k10_xtuple_0(X2)!=k1_numbers|~v1_relat_1(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2975, negated_conjecture, (v1_xcmplx_0(k10_xtuple_0(esk2_0))|k10_xtuple_0(X1)!=k1_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2985, plain, (v1_xcmplx_0(k5_numbers)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(X1)!=k1_zfmisc_1(k1_numbers)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1189, plain, (r1_xxreal_0(X1,k2_rinfsup1(X2))|r1_xxreal_0(k2_rinfsup1(X2),X1)|~m1_subset_1(X1,k5_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3234, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k1_seq_1(k2_zfmisc_1(X2,X3),X4))|r1_xxreal_0(k1_seq_1(k2_zfmisc_1(X2,X3),X4),k1_seq_1(esk2_0,X1))|~v1_funct_1(k2_zfmisc_1(X2,X3))|~v3_membered(X3))).#trainneg
cnf(i_0_611, plain, (v3_valued_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),X3,X4))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)),X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_612, plain, (k2_relset_1(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5))=k10_xtuple_0(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5))|v1_xboole_0(X2)|~v1_relat_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3662, negated_conjecture, (r1_xxreal_0(X1,esk6_3(k10_xtuple_0(esk2_0),X2,X3))|r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X2,X3),X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X1)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4062, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3967, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X1)))).#trainneg
cnf(i_0_4159, negated_conjecture, (v1_xcmplx_0(esk1_0)|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3914, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|~m1_subset_1(X1,k5_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_629, plain, (v4_relat_1(k1_funct_1(X4,X5),X2)|v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4355, negated_conjecture, (v1_xcmplx_0(k5_numbers)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(esk2_0)!=k1_zfmisc_1(k1_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4277, negated_conjecture, (v1_xcmplx_0(k2_rinfsup1(esk2_0))|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4283, negated_conjecture, (v1_xcmplx_0(k9_xtuple_0(esk2_0))|k10_xtuple_0(esk2_0)!=k1_numbers|~v3_membered(k1_zfmisc_1(k5_numbers))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_924, plain, (r1_xxreal_0(X1,k2_rinfsup1(X2))|r1_xxreal_0(k2_rinfsup1(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_613, plain, (v1_xreal_0(k2_relset_1(X1,k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5)))|v1_xboole_0(X2)|~v1_relat_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)),X4,X5))|~v3_membered(k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_617, plain, (m1_subset_1(k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4),k2_relset_1(X2,k2_relset_1(k2_zfmisc_1(X1,X2),X3)))|X2=k6_numbers|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_4228, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=k6_numbers|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_4063, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4191, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(X2,esk7_2(k9_xtuple_0(X2),X3)))|X3=k3_seq_4(k9_xtuple_0(X2))|v1_xboole_0(k9_xtuple_0(X2))|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X2),X3),k6_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_4255, negated_conjecture, (k2_xcmplx_0(k6_numbers,X1)=X1|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4142, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k3_funct_2(k5_numbers,X1,X2,X3))|r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k2_rinfsup1(esk2_0))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4144, plain, (r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k6_numbers)|r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4128, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(esk2_0),k3_funct_2(k5_numbers,k2_relset_1(X1,esk2_0),X2,X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_relset_1(X1,esk2_0))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k2_relset_1(X1,esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3908, negated_conjecture, (k10_xtuple_0(esk2_0)=k6_numbers|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_3913, negated_conjecture, (k6_numbers=k1_numbers|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|~r2_hidden(X1,k5_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_4081, negated_conjecture, (v1_xxreal_0(k1_funct_1(esk2_0,X1))|k6_numbers=k1_numbers|~v6_membered(k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_4349, negated_conjecture, (v1_xcmplx_0(k10_xtuple_0(esk2_0))|k10_xtuple_0(esk2_0)!=k1_numbers|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3909, negated_conjecture, (v1_xreal_0(k10_xtuple_0(esk2_0))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_2333, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k2_rinfsup1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_618, plain, (v1_xreal_0(k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|X2=k6_numbers|k2_relset_1(X2,k2_relset_1(k2_zfmisc_1(X1,X2),X3))!=k10_xtuple_0(X5)|~v3_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_619, plain, (v1_xxreal_0(k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|X2=k6_numbers|k2_relset_1(X2,k2_relset_1(k2_zfmisc_1(X1,X2),X3))!=k10_xtuple_0(X5)|~v2_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_3988, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X2,k3_funct_2(k5_numbers,k2_relset_1(X1,esk2_0),X3,X4))|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_relset_1(X1,esk2_0))))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,k2_relset_1(X1,esk2_0))|~v1_funct_1(X3))).#trainneg
cnf(i_0_3986, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X2,esk6_3(k2_relset_1(X1,esk2_0),X3,X4))|r1_xxreal_0(X4,k6_numbers)|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2)|X3!=k3_seq_4(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3679, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(X2,k1_funct_1(X3,esk6_3(k9_xtuple_0(X3),X4,X5)))|v1_xboole_0(k9_xtuple_0(X3))|r1_xxreal_0(X5,k6_numbers)|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2)|X4!=k3_seq_4(k9_xtuple_0(X3))|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v3_xxreal_2(k9_xtuple_0(X3))|~v3_membered(k9_xtuple_0(X3))|~v1_funct_1(X3)|~v1_xreal_0(X5)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_3472, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|v1_xboole_0(X2)|v1_xboole_0(k2_relset_1(X2,X3))|~v1_valued_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|~m1_subset_1(X1,k2_relset_1(X2,X3)))).#trainneg
cnf(i_0_4186, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k1_funct_1(X1,esk7_2(k9_xtuple_0(X1),X2)))|X2=k3_seq_4(k9_xtuple_0(X1))|v1_xboole_0(k9_xtuple_0(X1))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X1),X2),k6_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4131, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(esk2_0),k3_funct_2(k5_numbers,X1,esk2_0,X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4064, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k3_funct_2(k5_numbers,X1,esk2_0,X2))|v1_xboole_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_3911, negated_conjecture, (r1_xxreal_0(k6_numbers,k10_xtuple_0(esk2_0))|r1_xxreal_0(k10_xtuple_0(esk2_0),k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_3825, plain, (r1_xxreal_0(k3_seq_4(X1),k6_numbers)|X1!=k10_xtuple_0(X2)|k10_xtuple_0(X2)!=k5_numbers|~v1_relat_1(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4074, negated_conjecture, (m1_subset_1(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k1_numbers)|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4073, negated_conjecture, (r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),esk1_0)|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_2493, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),k2_rinfsup1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_620, plain, (v1_xcmplx_0(k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|X2=k6_numbers|k2_relset_1(X2,k2_relset_1(k2_zfmisc_1(X1,X2),X3))!=k10_xtuple_0(X5)|~v1_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_623, plain, (m1_subset_1(k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4),X2)|X2=k6_numbers|~v5_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_3847, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(k1_funct_1(X3,X4),k2_relset_1(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_3661, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk5_1(X1),esk5_1(esk6_3(k10_xtuple_0(esk2_0),X2,X3)))|r1_xxreal_0(esk5_1(esk6_3(k10_xtuple_0(esk2_0),X2,X3)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X2,X3),k6_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X1)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3678, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(X2,k1_funct_1(X3,esk9_3(X3,X4,X5)))|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2)|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(X3)|X4!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~r2_hidden(X5,X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3508, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k1_funct_1(X2,X3))|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k10_xtuple_0(esk2_0))))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k10_xtuple_0(esk2_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_3645, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk2_0,X1),k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),k1_funct_1(esk2_0,X1)))).#trainneg
cnf(i_0_4077, negated_conjecture, (r1_xxreal_0(k6_numbers,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4076, negated_conjecture, (r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),esk7_2(k10_xtuple_0(esk2_0),esk1_0))|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4157, negated_conjecture, (k3_seq_4(k2_zfmisc_1(k5_numbers,k1_numbers))=esk1_0|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|r1_xxreal_0(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0),k6_numbers)|~v3_xxreal_2(k2_zfmisc_1(k5_numbers,k1_numbers))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~r1_xxreal_0(esk1_0,esk7_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))|~v1_xreal_0(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,esk2_0)))|~m1_subset_1(esk5_1(esk8_2(k2_zfmisc_1(k5_numbers,k1_numbers),esk1_0)),k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,esk2_0))).#trainneg
cnf(i_0_3292, plain, (v1_xxreal_0(k9_xtuple_0(X2))|~v6_membered(k1_zfmisc_1(X1))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3526, plain, (v3_valued_0(k9_xtuple_0(X2))|~v4_relat_1(X2,k2_zfmisc_1(X1,k1_numbers))|~v1_relat_1(X2))).#trainneg
cnf(i_0_626, plain, (X1=k6_numbers|r2_hidden(X5,k2_relset_1(X1,k2_relset_1(k2_zfmisc_1(X2,X1),X3)))|~v5_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3)|~r2_hidden(esk9_3(k2_relset_1(k2_zfmisc_1(X2,X1),X3),X4,X5),X2)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X2,X1),X3),X2,X1)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X2,X1),X3))|X4!=k10_xtuple_0(k2_relset_1(k2_zfmisc_1(X2,X1),X3))|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X2,X1),X3))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_631, plain, (v1_valued_0(k1_funct_1(X4,X5))|v1_xboole_0(X1)|~v1_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_3547, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers|~v1_xreal_0(X1)|X1!=k2_rinfsup1(esk2_0))).#trainneg
cnf(i_0_3759, negated_conjecture, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4069, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4247, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk6_3(k10_xtuple_0(esk2_0),X1,X2))=esk6_3(k10_xtuple_0(esk2_0),X1,X2)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|~v1_valued_0(esk2_0)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3665, negated_conjecture, (r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2),k1_seq_1(esk2_0,X3))|r1_xxreal_0(k1_seq_1(esk2_0,X3),esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3855, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk5_1(X1),esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2675, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,k1_funct_1(X3,X4))|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2149, plain, (r1_xxreal_0(k2_rinfsup1(X1),k1_funct_1(X1,X2))|~v3_xxreal_2(k2_relset_1(k1_numbers,X1))|~v3_membered(k2_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k1_funct_1(X1,X2))|~m1_subset_1(k2_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1470, plain, (r1_xxreal_0(k5_seq_4(X1),k5_seq_4(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1102, plain, (v3_membered(k9_xtuple_0(k2_zfmisc_1(X1,X2)))|~v3_membered(X1))).#trainneg
cnf(i_0_2698, negated_conjecture, (k2_xcmplx_0(X1,k6_numbers)=X1|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3827, negated_conjecture, (r1_xxreal_0(k3_seq_4(X1),k2_rinfsup1(esk2_0))|X1!=k10_xtuple_0(X2)|k10_xtuple_0(X2)!=k1_numbers|~v1_relat_1(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3198, plain, (r1_xxreal_0(k6_numbers,k3_seq_4(X1))|r1_xxreal_0(k3_seq_4(X1),k6_numbers)|~v6_membered(k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_4008, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(X1)!=k1_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2797, negated_conjecture, (v1_xxreal_0(k10_xtuple_0(esk2_0))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(X1)!=k1_zfmisc_1(k1_numbers)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1639, plain, (m1_subset_1(k1_funct_1(k2_zfmisc_1(X1,X2),X3),X2)|X2=k6_numbers|~r2_hidden(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_986, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_635, plain, (m1_subset_1(k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4),k2_relset_1(X2,k1_relset_1(k2_zfmisc_1(X1,X2),X3)))|X2=k6_numbers|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_636, plain, (v1_xreal_0(k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|X2=k6_numbers|k2_relset_1(X2,k1_relset_1(k2_zfmisc_1(X1,X2),X3))!=k10_xtuple_0(X5)|~v3_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_2696, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k10_xtuple_0(esk2_0))))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k10_xtuple_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_1242, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_numbers)|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3668, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2),esk5_1(X3))|r1_xxreal_0(esk5_1(X3),esk6_3(k10_xtuple_0(esk2_0),X1,X2))|r1_xxreal_0(X3,k6_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X3)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1786, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk4_0,esk1_0),X2)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk9_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_1709, plain, (r1_xxreal_0(k3_seq_4(X1),k3_seq_4(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2762, negated_conjecture, (k2_rinfsup1(esk2_0)=k6_numbers|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),k6_numbers),k10_xtuple_0(esk2_0))|r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),k6_numbers),k1_funct_1(esk2_0,X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(esk8_2(k10_xtuple_0(esk2_0),k6_numbers),k1_numbers))).#trainneg
cnf(i_0_3816, negated_conjecture, (r1_xxreal_0(k3_seq_4(X1),k9_xtuple_0(esk2_0))|X1!=k10_xtuple_0(X2)|k10_xtuple_0(X2)!=k1_numbers|~v1_relat_1(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2)|~v1_xreal_0(k9_xtuple_0(esk2_0))|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_739, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,k1_relset_1(X1,X2)),X1)|k1_relset_1(X1,X2)=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,k1_relset_1(X1,X2)),k9_xtuple_0(esk2_0))|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1733, plain, (v7_ordinal1(X3)|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_2974, negated_conjecture, (v1_xcmplx_0(k10_xtuple_0(esk2_0))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(X1)!=k1_zfmisc_1(k1_numbers)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4041, negated_conjecture, (k2_xcmplx_0(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(X1)!=k1_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1708, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(k2_zfmisc_1(X1,X2),X3))|~v6_membered(X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3429, plain, (v1_xxreal_0(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_637, plain, (v1_xxreal_0(k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|X2=k6_numbers|k2_relset_1(X2,k1_relset_1(k2_zfmisc_1(X1,X2),X3))!=k10_xtuple_0(X5)|~v2_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_638, plain, (v1_xcmplx_0(k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|X2=k6_numbers|k2_relset_1(X2,k1_relset_1(k2_zfmisc_1(X1,X2),X3))!=k10_xtuple_0(X5)|~v1_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_3362, plain, (v1_xreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|v1_xboole_0(X4)|~v3_membered(k2_relset_1(X4,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X4)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2311, negated_conjecture, (r1_xxreal_0(esk9_3(esk2_0,X1,X2),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),esk9_3(esk2_0,X1,X2))|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_2348, negated_conjecture, (r2_hidden(esk6_3(k2_relset_1(X1,esk2_0),X2,X3),k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|v1_xboole_0(k2_relset_1(X1,esk2_0))|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|X2!=k3_seq_4(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1247, plain, (m1_subset_1(k3_funct_2(X1,X2,X3,X4),k1_numbers)|v1_xboole_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_4959, negated_conjecture, (r2_hidden(esk6_3(k10_xtuple_0(esk2_0),X1,X2),k10_xtuple_0(esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4498, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k1_numbers))).#trainneg
cnf(i_0_4499, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k5_numbers))).#trainneg
cnf(i_0_4420, negated_conjecture, (r1_xxreal_0(X2,X2)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_4749, plain, (v3_valued_0(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,k1_numbers),X2))))).#trainneg
cnf(i_0_4425, negated_conjecture, (v1_xcmplx_0(esk1_0)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_4835, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(k2_zfmisc_1(X1,X2))!=k1_numbers|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_741, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,k2_relset_1(X1,X2)),X1)|k2_relset_1(X1,X2)=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,k2_relset_1(X1,X2)),k9_xtuple_0(esk2_0))|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4438, negated_conjecture, (v1_xcmplx_0(k5_numbers)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(esk2_0)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_4831, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4898, negated_conjecture, (v1_xcmplx_0(k10_xtuple_0(esk2_0))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(esk2_0)!=k1_zfmisc_1(k1_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_641, plain, (m1_subset_1(k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4),X2)|X2=k6_numbers|~v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X2)|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~r2_hidden(X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3)))).#trainneg
cnf(i_0_646, plain, (X1=k6_numbers|r2_hidden(X5,k2_relset_1(X1,k1_relset_1(k2_zfmisc_1(X2,X1),X3)))|~v4_relat_1(X3,k2_zfmisc_1(X2,X1))|~v1_relat_1(X3)|~r2_hidden(esk9_3(k1_relset_1(k2_zfmisc_1(X2,X1),X3),X4,X5),X2)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X2,X1),X3),X2,X1)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X1),X3))|X4!=k10_xtuple_0(k1_relset_1(k2_zfmisc_1(X2,X1),X3))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X2,X1),X3))|~r2_hidden(X5,X4))).#trainneg
cnf(i_0_4771, negated_conjecture, (r1_xxreal_0(esk7_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|r1_xxreal_0(k6_numbers,esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4827, negated_conjecture, (k2_xcmplx_0(X1,k6_numbers)=X1|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4649, plain, (k2_xcmplx_0(k6_numbers,esk7_2(k2_relset_1(k1_numbers,X1),X2))=esk7_2(k2_relset_1(k1_numbers,X1),X2)|v1_xboole_0(k2_relset_1(k1_numbers,X1))|X2=k3_seq_4(k2_relset_1(k1_numbers,X1))|v1_xreal_0(esk8_2(k2_relset_1(k1_numbers,X1),X2))|~v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_membered(k2_relset_1(k1_numbers,X1))|~v2_seq_2(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4775, negated_conjecture, (r1_xxreal_0(k3_funct_2(k5_numbers,X1,esk2_0,X2),k6_numbers)|r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,X1,esk2_0,X2))|v1_xboole_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_4717, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(X2,esk9_3(X2,X3,X4)))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|X1!=k2_rinfsup1(esk2_0)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4932, negated_conjecture, (m1_subset_1(esk6_3(k2_relset_1(X1,esk2_0),X2,X3),k10_xtuple_0(esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4585, negated_conjecture, (k5_seq_4(k6_numbers)=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_4908, negated_conjecture, (k2_xcmplx_0(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(k2_zfmisc_1(X1,X2))!=k1_numbers|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_743, negated_conjecture, (r2_hidden(X1,X2)|k9_xtuple_0(X3)=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,k9_xtuple_0(X3)),k9_xtuple_0(esk2_0))|X1!=k1_funct_1(X3,esk10_2(esk2_0,k9_xtuple_0(X3)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2890, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk2_0,X1),esk1_0)|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1)))).#trainneg
cnf(i_0_4418, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_4628, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k2_rinfsup1(k2_zfmisc_1(k5_numbers,k1_numbers)))|r1_xxreal_0(k2_rinfsup1(k2_zfmisc_1(k5_numbers,k1_numbers)),k2_rinfsup1(esk2_0))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_4459, plain, (r1_xxreal_0(k2_rinfsup1(k2_zfmisc_1(k5_numbers,k1_numbers)),X1)|r1_xxreal_0(X1,k2_rinfsup1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_648, plain, (v1_xboole_0(X1)|v1_funct_2(k1_funct_1(X4,X5),X2,X3)|~v1_partfun1(k1_funct_1(X4,X5),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_652, plain, (v1_partfun1(k1_funct_1(X4,X5),X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(k1_funct_1(X4,X5),X2,X3)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4774, negated_conjecture, (r1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4819, plain, (r1_xxreal_0(k2_rinfsup1(X1),X3)|~v3_xxreal_2(k2_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k2_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk9_3(X1,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X3)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_4783, negated_conjecture, (k2_rinfsup1(esk2_0)=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk8_2(X1,k2_rinfsup1(esk2_0)))|v1_xboole_0(k10_xtuple_0(esk2_0))|~v3_xxreal_2(X1)|~v3_membered(X1)|~m1_subset_1(esk7_2(X1,k2_rinfsup1(esk2_0)),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4928, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xreal_0(k3_funct_2(k5_numbers,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|~m1_subset_1(k2_relset_1(X1,X3),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4976, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X3,X4),k1_numbers)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_4731, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk2_0,X1),k1_funct_1(esk2_0,X2))|r1_xxreal_0(k1_funct_1(esk2_0,X2),k1_funct_1(esk2_0,X1)))).#trainneg
cnf(i_0_1266, plain, (v4_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),X1))).#trainneg
cnf(i_0_4510, negated_conjecture, (k6_numbers=esk2_0|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_4832, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))!=k1_numbers|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_745, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,X1),X1)|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,X1),k9_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_2891, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk2_0,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_funct_1(esk2_0,X1)))).#trainneg
cnf(i_0_4696, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k2_rinfsup1(k2_zfmisc_1(k5_numbers,k1_numbers)))|r1_xxreal_0(k2_rinfsup1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_seq_1(esk2_0,X1))|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_3442, plain, (v7_ordinal1(k1_funct_1(k2_zfmisc_1(X1,X2),X3))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_663, plain, (X1=k6_numbers|r2_hidden(k1_funct_1(k1_funct_1(X4,X5),X6),k2_relset_1(X1,k1_funct_1(X4,X5)))|v1_xboole_0(X2)|~r2_hidden(X6,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))))|~m1_subset_1(X5,X2)|~v1_funct_2(k1_funct_1(X4,X5),X3,X1)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_1(k1_funct_1(X4,X5))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4419, negated_conjecture, (r1_xxreal_0(X1,X1)|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_4876, negated_conjecture, (r1_xxreal_0(esk6_3(X1,X2,X3),esk1_0)|r1_xxreal_0(esk1_0,esk6_3(X1,X2,X3))|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4640, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|r1_xxreal_0(X2,k6_numbers)|v1_xboole_0(X1)|X3!=k3_seq_4(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|k9_binop_2(X2,X3)!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k9_binop_2(X2,X3)))).#trainneg
cnf(i_0_4665, plain, (r1_xxreal_0(k3_seq_4(k10_xtuple_0(X1)),k6_numbers)|k10_xtuple_0(X1)!=k5_numbers|~v1_relat_1(X1)|~v3_xxreal_2(k10_xtuple_0(X1))|~v3_membered(k10_xtuple_0(X1))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4875, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),esk6_3(X1,X2,X3))|r1_xxreal_0(esk6_3(X1,X2,X3),k2_rinfsup1(esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4641, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|r1_xxreal_0(X2,k6_numbers)|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=k6_numbers|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X2)|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_1497, plain, (v5_relat_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)),X2))).#trainneg
cnf(i_0_4511, negated_conjecture, (v3_membered(k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_4834, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))!=k1_numbers|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_746, negated_conjecture, (v1_xreal_0(esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,X1),k9_xtuple_0(esk2_0))|X1!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1710, plain, (r1_xxreal_0(k6_numbers,k3_seq_4(X1))|r1_xxreal_0(k3_seq_4(X1),k6_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4439, negated_conjecture, (k2_xcmplx_0(k2_rinfsup1(esk2_0),k6_numbers)=k2_rinfsup1(esk2_0)|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4904, negated_conjecture, (k2_xcmplx_0(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4921, plain, (v1_xxreal_0(k1_funct_1(k2_zfmisc_1(X1,X2),X3))|v1_xboole_0(X1)|~v6_membered(X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3706, plain, (v1_xboole_0(X1)|v3_membered(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2)),X3))|~v3_membered(X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(X2)),X1,k1_zfmisc_1(X2))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_740, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk10_2(esk2_0,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4)),X2)|k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4)=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4)),k9_xtuple_0(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_747, negated_conjecture, (v1_xxreal_0(esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,X1),k9_xtuple_0(esk2_0))|X1!=k10_xtuple_0(X2)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4777, negated_conjecture, (r1_xxreal_0(k1_funct_1(X1,esk7_2(k9_xtuple_0(X1),X2)),k6_numbers)|r1_xxreal_0(k6_numbers,k1_funct_1(X1,esk7_2(k9_xtuple_0(X1),X2)))|X2=k3_seq_4(k9_xtuple_0(X1))|v1_xboole_0(k9_xtuple_0(X1))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X1),X2),k6_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4782, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_4643, negated_conjecture, (v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X2,k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X3,X4)))|r1_xxreal_0(X4,k6_numbers)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|X3!=k3_seq_4(k9_xtuple_0(X1))|X2!=k2_rinfsup1(esk2_0)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3900, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk6_3(k10_xtuple_0(esk2_0),X2,X3),k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4078, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),esk5_1(X1))|r1_xxreal_0(esk5_1(X1),esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(X1,k6_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2297, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|k9_binop_2(esk1_0,X1)!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk5_1(X1),k5_numbers)|~v1_xreal_0(k9_binop_2(esk1_0,X1)))).#trainneg
cnf(i_0_1682, negated_conjecture, (v1_xreal_0(X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_2555, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk5_1(X1),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3272, plain, (r1_xxreal_0(k9_binop_2(k6_numbers,X1),k9_binop_2(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4905, negated_conjecture, (k2_xcmplx_0(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3))!=k1_numbers|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_748, negated_conjecture, (v1_xcmplx_0(esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,X1),k9_xtuple_0(esk2_0))|X1!=k10_xtuple_0(X2)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3722, negated_conjecture, (k6_numbers=k1_numbers|m1_subset_1(X2,k1_numbers)|~r2_hidden(esk9_3(esk2_0,X1,X2),k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_4443, negated_conjecture, (k2_xcmplx_0(k9_xtuple_0(esk2_0),k6_numbers)=k9_xtuple_0(esk2_0)|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4601, negated_conjecture, (k2_xcmplx_0(k10_xtuple_0(esk2_0),k6_numbers)=k10_xtuple_0(esk2_0)|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3740, plain, (v3_valued_0(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2)),X3))|v1_xboole_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(X2)),X1,k1_zfmisc_1(X2))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_4363, plain, (v3_valued_0(k1_funct_1(X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_751, negated_conjecture, (X1=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,X1),k9_xtuple_0(esk2_0))|esk10_2(esk2_0,X1)!=k1_funct_1(esk2_0,X2)|~r2_hidden(X2,k9_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_753, negated_conjecture, (m1_subset_1(esk11_2(esk2_0,X1),k9_xtuple_0(esk2_0))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1))).#trainneg
cnf(i_0_5277, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(X2,X2)|v1_xboole_0(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k2_relset_1(X1,esk2_0)))).#trainneg
cnf(i_0_5009, negated_conjecture, (r1_xxreal_0(X1,X1)|~r2_hidden(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_5279, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X2,k6_numbers)|k9_binop_2(X2,k3_seq_4(k2_relset_1(X1,esk2_0)))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2)|~v1_xreal_0(k3_seq_4(k2_relset_1(X1,esk2_0))))).#trainneg
cnf(i_0_5400, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|k9_binop_2(esk1_0,X1)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5013, negated_conjecture, (k2_xcmplx_0(esk1_0,k6_numbers)=esk1_0|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_4907, negated_conjecture, (k2_xcmplx_0(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3))!=k1_numbers|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v1_membered(X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_754, negated_conjecture, (v1_xreal_0(esk11_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|k9_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5023, negated_conjecture, (k2_xcmplx_0(k5_numbers,k6_numbers)=k5_numbers|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(esk2_0)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_5283, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|k10_xtuple_0(esk2_0)!=k5_numbers)).#trainneg
cnf(i_0_5027, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_5141, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(k2_zfmisc_1(X1,X2),X3),k1_numbers)|~v3_membered(X2)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4108, plain, (v1_relat_1(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_755, negated_conjecture, (v1_xxreal_0(esk11_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|k9_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_756, negated_conjecture, (v1_xcmplx_0(esk11_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|k9_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4979, negated_conjecture, (m1_subset_1(esk6_3(k10_xtuple_0(esk2_0),X1,X2),k10_xtuple_0(esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5121, plain, (r1_xxreal_0(k2_rinfsup1(X1),X2)|X3!=k10_xtuple_0(X1)|~r2_hidden(X2,X3)|~v3_xxreal_2(k2_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k2_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5128, plain, (v1_xboole_0(X1)|v1_xreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~v5_relat_1(X2,k1_numbers)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1352, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk9_3(X3,X4,X5))|X4!=k10_xtuple_0(X3)|~v4_relat_1(X3,k1_zfmisc_1(X2))|~v1_relat_1(X3)|~r2_hidden(X5,X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3325, plain, (r1_xxreal_0(k9_binop_2(X1,k6_numbers),k9_binop_2(X1,k6_numbers))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5015, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=k2_xcmplx_0(esk1_0,k6_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_759, negated_conjecture, (r2_hidden(X1,X2)|X3=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X3),X3)|X1!=k1_funct_1(esk2_0,esk11_2(esk2_0,X3))|X2!=k10_xtuple_0(esk2_0))).#trainneg
cnf(i_0_5025, negated_conjecture, (k2_xcmplx_0(k6_numbers,k5_numbers)=k2_xcmplx_0(k5_numbers,k6_numbers)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|k10_xtuple_0(esk2_0)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_5295, negated_conjecture, (k9_binop_2(k2_rinfsup1(esk2_0),k6_numbers)=k2_rinfsup1(esk2_0)|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_5475, plain, (v3_valued_0(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers))),X3))|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers))),X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))))).#trainneg
cnf(i_0_4275, plain, (v5_relat_1(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4),X3)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_777, plain, (r1_xxreal_0(esk12_3(X1,X2,X3),X4)|r1_xxreal_0(X4,esk12_3(X1,X2,X3))|v1_xboole_0(X2)|~v1_xreal_0(X4)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_780, plain, (r1_xxreal_0(k3_funct_2(X1,k5_numbers,X2,X3),X4)|r1_xxreal_0(X4,k3_funct_2(X1,k5_numbers,X2,X3))|v1_xboole_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k5_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_2(X2,X1,k5_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5402, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xreal_0(esk9_3(X1,X2,X3))|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,k10_xtuple_0(esk2_0))|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5351, negated_conjecture, (X1=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(esk8_2(k2_relset_1(X2,esk2_0),X1),k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X1))|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(esk8_2(k2_relset_1(X2,esk2_0),X1)))).#trainneg
cnf(i_0_4586, negated_conjecture, (v3_xxreal_2(k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_820, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,k1_relset_1(X1,X2)),X1)|k1_funct_1(esk2_0,esk11_2(esk2_0,k1_relset_1(X1,X2)))=esk10_2(esk2_0,k1_relset_1(X1,X2))|k1_relset_1(X1,X2)=k10_xtuple_0(esk2_0)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5298, negated_conjecture, (k9_binop_2(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(esk2_0)!=k1_numbers|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_4437, plain, (v4_relat_1(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4),X2)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_4120, plain, (v3_valued_0(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4))|v1_xboole_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_821, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk10_2(esk2_0,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4)),X2)|k1_funct_1(esk2_0,esk11_2(esk2_0,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4)))=esk10_2(esk2_0,k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4))|k3_funct_2(X1,k1_zfmisc_1(X2),X3,X4)=k10_xtuple_0(esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,k1_zfmisc_1(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_823, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,k2_relset_1(X1,X2)),X1)|k1_funct_1(esk2_0,esk11_2(esk2_0,k2_relset_1(X1,X2)))=esk10_2(esk2_0,k2_relset_1(X1,X2))|k2_relset_1(X1,X2)=k10_xtuple_0(esk2_0)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2114, plain, (v1_xreal_0(esk9_3(X1,X2,X3))|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)|~m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3498, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,esk9_3(X2,X3,X4))|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|X3!=k10_xtuple_0(X2)|~v4_relat_1(X2,k10_xtuple_0(esk2_0))|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_999, plain, (v1_xreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|~v3_membered(k2_relset_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3197, plain, (r1_xxreal_0(k3_seq_4(X1),k3_seq_4(X1))|~v6_membered(k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_4566, negated_conjecture, (k2_rinfsup1(esk2_0)=k3_seq_4(k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_827, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,X1),X1)|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0))).#trainneg
cnf(i_0_4766, plain, (v1_valued_0(k1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X4))|v1_xboole_0(X1)|~v1_membered(X3)|~m1_subset_1(X4,X1)|~v1_funct_2(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))),X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_920, plain, (r1_xxreal_0(X1,k3_seq_4(X2))|r1_xxreal_0(k3_seq_4(X2),X1)|~v1_xreal_0(X1)|~v3_membered(X2))).#trainneg
cnf(i_0_825, negated_conjecture, (r2_hidden(X1,X2)|k1_funct_1(esk2_0,esk11_2(esk2_0,k9_xtuple_0(X3)))=esk10_2(esk2_0,k9_xtuple_0(X3))|k9_xtuple_0(X3)=k10_xtuple_0(esk2_0)|X1!=k1_funct_1(X3,esk10_2(esk2_0,k9_xtuple_0(X3)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_828, negated_conjecture, (v1_xreal_0(esk10_2(esk2_0,X1))|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|X1!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_907, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,k3_funct_2(k5_numbers,X1,X3,X4))|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1360, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,esk9_3(X3,X4,X5))|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(X2)|X4!=k10_xtuple_0(X3)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~r2_hidden(X5,X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_2560, plain, (m1_subset_1(k3_funct_2(k5_numbers,X1,X2,X3),k1_numbers)|v1_xboole_0(X1)|~v3_valued_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1596, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(X3)|X3!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_xreal_0(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_4048, negated_conjecture, (r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X1,esk2_0),X2))|r1_xxreal_0(esk7_2(k2_relset_1(X1,esk2_0),X2),esk1_0)|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3008, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k1_funct_1(X1,X2))|r1_xxreal_0(k1_funct_1(X1,X2),k2_rinfsup1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2639, plain, (v1_xboole_0(k1_relset_1(X1,X2))|m1_subset_1(k1_funct_1(X3,X4),X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_relset_1(X1,X2))))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_relset_1(X1,X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1136, plain, (v1_xboole_0(k1_zfmisc_1(X1))|v1_xreal_0(X1)|k1_zfmisc_1(X1)!=k10_xtuple_0(X2)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4582, negated_conjecture, (v3_valued_0(k6_numbers)|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_5069, negated_conjecture, (r1_xxreal_0(X1,X4)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|X1!=k2_rinfsup1(esk2_0)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_829, negated_conjecture, (v1_xxreal_0(esk10_2(esk2_0,X1))|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|X1!=k10_xtuple_0(X2)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2841, plain, (r1_xxreal_0(X1,k9_binop_2(k6_numbers,X2))|r1_xxreal_0(k9_binop_2(k6_numbers,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_830, negated_conjecture, (v1_xcmplx_0(esk10_2(esk2_0,X1))|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|X1!=k10_xtuple_0(X2)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_833, negated_conjecture, (X1=k10_xtuple_0(esk2_0)|k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|esk10_2(esk2_0,X1)!=k1_funct_1(esk2_0,X2)|~r2_hidden(X2,k9_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_6047, negated_conjecture, (r1_xxreal_0(X3,k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_3(X1,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_5737, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk7_2(k2_relset_1(X1,esk2_0),k2_rinfsup1(esk2_0)),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_5626, plain, (r1_xxreal_0(k2_rinfsup1(X1),X2)|X3!=k10_xtuple_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k2_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_seq_2(X1))).#trainneg
cnf(i_0_5738, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),k2_rinfsup1(esk2_0)),k6_numbers))).#trainneg
cnf(i_0_5569, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X2,k6_numbers)|k9_binop_2(k3_seq_4(k2_relset_1(X1,esk2_0)),X2)!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2)|~v1_xreal_0(k3_seq_4(k2_relset_1(X1,esk2_0))))).#trainneg
cnf(i_0_5535, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_5574, negated_conjecture, (k9_binop_2(esk1_0,k6_numbers)=esk1_0|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_5572, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|k9_binop_2(X1,esk1_0)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_834, negated_conjecture, (v1_xreal_0(esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1))).#trainneg
cnf(i_0_5944, plain, (r1_xxreal_0(k3_seq_4(k1_numbers),X1)|r1_xxreal_0(X1,k3_seq_4(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5945, plain, (r1_xxreal_0(k3_seq_4(k5_numbers),X1)|r1_xxreal_0(X1,k3_seq_4(k5_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_835, negated_conjecture, (v1_xxreal_0(esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|~v2_valued_0(esk2_0))).#trainneg
cnf(i_0_836, negated_conjecture, (v1_xcmplx_0(esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|~v1_valued_0(esk2_0))).#trainneg
cnf(i_0_5560, negated_conjecture, (r1_xxreal_0(k1_funct_1(X1,esk7_2(k9_xtuple_0(X1),X2)),k1_funct_1(X1,esk7_2(k9_xtuple_0(X1),X2)))|X2=k3_seq_4(k9_xtuple_0(X1))|v1_xboole_0(k9_xtuple_0(X1))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X1),X2),k6_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5554, negated_conjecture, (r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X1,X2),esk6_3(k10_xtuple_0(esk2_0),X1,X2))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X2,k6_numbers)|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5802, plain, (v1_xboole_0(X1)|v1_xreal_0(k3_funct_2(k5_numbers,X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5951, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,X4)|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X4,X1,X3),k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3)|~v1_xreal_0(X2)|~r2_hidden(X4,k2_relset_1(X1,X3)))).#trainneg
cnf(i_0_5557, negated_conjecture, (r1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1),k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_972, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~m1_subset_1(X1,k9_xtuple_0(X3)))).#trainneg
cnf(i_0_5575, negated_conjecture, (k9_binop_2(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_837, negated_conjecture, (k2_xcmplx_0(esk10_2(esk2_0,X1),k6_numbers)=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|~v1_valued_0(esk2_0))).#trainneg
cnf(i_0_2857, plain, (r1_xxreal_0(X1,k9_binop_2(X2,k6_numbers))|r1_xxreal_0(k9_binop_2(X2,k6_numbers),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_838, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,X1),k1_numbers)|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1))).#trainneg
cnf(i_0_841, negated_conjecture, (k2_xcmplx_0(X1,esk10_2(esk2_0,X2))=k2_xcmplx_0(esk10_2(esk2_0,X2),X1)|X2=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X2),X2)|~v1_valued_0(esk2_0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5736, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_xcmplx_0(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|k2_xcmplx_0(esk1_0,esk4_0)!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k2_xcmplx_0(esk1_0,esk4_0))|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k2_xcmplx_0(esk1_0,esk4_0)),k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_5735, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_xcmplx_0(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_xcmplx_0(esk1_0,esk4_0)!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k2_xcmplx_0(esk1_0,esk4_0))|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k2_xcmplx_0(esk1_0,esk4_0)),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_5631, plain, (m1_subset_1(k3_funct_2(k5_numbers,X1,X2,X3),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5820, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_5811, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(X1,k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_2175, negated_conjecture, (m1_subset_1(X3,k1_numbers)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3818, negated_conjecture, (r1_xxreal_0(k3_seq_4(X1),k10_xtuple_0(esk2_0))|X1!=k10_xtuple_0(X2)|k10_xtuple_0(X2)!=k1_numbers|~v1_relat_1(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X2)|~v1_xreal_0(k10_xtuple_0(esk2_0))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_845, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|~m2_subset_1(esk11_2(esk2_0,X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2328, negated_conjecture, (r1_xxreal_0(k3_seq_4(X1),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),k3_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_842, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(esk11_2(esk2_0,X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_843, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),esk10_2(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk11_2(esk2_0,X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5638, plain, (v1_xboole_0(X1)|v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X3,X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~r2_hidden(X3,k2_relset_1(X1,X2)))).#trainneg
cnf(i_0_5732, negated_conjecture, (k9_binop_2(X1,esk8_2(X2,X1))=k3_seq_4(k2_relset_1(X3,esk2_0))|v1_xboole_0(k2_relset_1(X3,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X3)|X1=k3_seq_4(X2)|k9_binop_2(X1,esk8_2(X2,X1))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X3,esk2_0))|~v3_membered(k2_relset_1(X3,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X3)|~v1_xreal_0(k9_binop_2(X1,esk8_2(X2,X1)))|~r2_hidden(esk7_2(k2_relset_1(X3,esk2_0),k9_binop_2(X1,esk8_2(X2,X1))),X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(X1,esk7_2(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(k2_relset_1(X3,esk2_0),k9_binop_2(X1,esk8_2(X2,X1)))))).#trainneg
cnf(i_0_5610, negated_conjecture, (X1=k2_rinfsup1(esk2_0)|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1)),esk6_3(k10_xtuple_0(esk2_0),X2,X3))|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X1)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5552, negated_conjecture, (r1_xxreal_0(esk7_2(k2_relset_1(X1,esk2_0),X2),esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2641, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X3))|X1!=k1_funct_1(X3,k1_funct_1(X4,X5))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k9_xtuple_0(X3))))|~m1_subset_1(X5,k5_numbers)|~v1_funct_2(X4,k5_numbers,k9_xtuple_0(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_4953, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk6_3(k2_relset_1(X2,esk2_0),X3,X4),k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|X3!=k3_seq_4(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_4636, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X2,k1_funct_1(X3,X4))|X2!=k2_rinfsup1(esk2_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_relset_1(X1,esk2_0))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,k2_relset_1(X1,esk2_0))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_1(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1775, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(X2)|X2!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_6115, negated_conjecture, (k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|k1_funct_1(esk2_0,esk11_2(esk2_0,k9_xtuple_0(esk2_0)))=esk10_2(esk2_0,k9_xtuple_0(esk2_0))|k9_xtuple_0(esk2_0)=k10_xtuple_0(esk2_0)|esk10_2(esk2_0,X1)!=k1_funct_1(esk2_0,esk10_2(esk2_0,k9_xtuple_0(esk2_0))))).#trainneg
cnf(i_0_852, plain, (X1=k10_xtuple_0(k2_zfmisc_1(X2,X3))|r2_hidden(esk11_2(k2_zfmisc_1(X2,X3),X1),k9_xtuple_0(k2_zfmisc_1(X2,X3)))|r2_hidden(esk10_2(k2_zfmisc_1(X2,X3),X1),X1)|~v1_funct_1(k2_zfmisc_1(X2,X3)))).#trainneg
cnf(i_0_6447, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,k5_numbers),k1_numbers)|k10_xtuple_0(esk2_0)=k5_numbers)).#trainneg
cnf(i_0_6162, negated_conjecture, (v1_xxreal_0(esk1_0)|~v6_membered(k1_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_919, plain, (r1_xxreal_0(X1,k9_binop_2(X2,X3))|r1_xxreal_0(k9_binop_2(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_844, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk4_0),esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m2_subset_1(esk11_2(esk2_0,X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_846, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,esk4_0),esk10_2(esk2_0,X1))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|k2_rinfsup1(esk2_0)!=esk1_0|~m2_subset_1(esk11_2(esk2_0,X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6134, plain, (r1_xxreal_0(k2_rinfsup1(X1),X2)|X3!=k10_xtuple_0(X1)|~r2_hidden(X2,X3)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v5_relat_1(X1,k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6503, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|v1_xboole_0(X1)|r1_xxreal_0(esk4_0,k6_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|k3_seq_4(k2_relset_1(X1,esk2_0))!=esk1_0|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_6225, negated_conjecture, (r1_xxreal_0(esk1_0,k3_seq_4(k1_numbers))|r1_xxreal_0(k3_seq_4(k1_numbers),esk1_0))).#trainneg
cnf(i_0_853, plain, (k1_funct_1(k2_zfmisc_1(X1,X2),esk11_2(k2_zfmisc_1(X1,X2),X3))=esk10_2(k2_zfmisc_1(X1,X2),X3)|X3=k10_xtuple_0(k2_zfmisc_1(X1,X2))|r2_hidden(esk10_2(k2_zfmisc_1(X1,X2),X3),X3)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_6252, plain, (r1_xxreal_0(k6_numbers,k3_seq_4(k1_numbers))|r1_xxreal_0(k3_seq_4(k1_numbers),k6_numbers))).#trainneg
cnf(i_0_2844, negated_conjecture, (r1_xxreal_0(k9_binop_2(k6_numbers,X1),k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),k9_binop_2(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6595, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(esk4_0,k6_numbers)|k2_rinfsup1(esk2_0)!=esk1_0|k3_seq_4(k2_relset_1(X1,esk2_0))!=esk1_0|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_6314, plain, (r1_xxreal_0(k6_numbers,k3_funct_2(k5_numbers,X1,X2,X3))|r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k6_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6310, plain, (m1_subset_1(k3_funct_2(k5_numbers,X1,X2,X3),k1_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6587, negated_conjecture, (k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))=k3_seq_4(k2_relset_1(X3,esk2_0))|X1=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X3,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X3)|v1_xboole_0(X2)|k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X3,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X3)|~v1_xreal_0(esk7_2(k2_relset_1(X3,esk2_0),k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))))|~v1_xreal_0(k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1)))|~v1_xreal_0(X1)|~r2_hidden(esk7_2(k2_relset_1(X3,esk2_0),k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))),k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_6267, negated_conjecture, (r1_xxreal_0(esk1_0,k3_seq_4(k5_numbers))|r1_xxreal_0(k3_seq_4(k5_numbers),esk1_0))).#trainneg
cnf(i_0_6294, plain, (r1_xxreal_0(k6_numbers,k3_seq_4(k5_numbers))|r1_xxreal_0(k3_seq_4(k5_numbers),k6_numbers))).#trainneg
cnf(i_0_2860, negated_conjecture, (r1_xxreal_0(k9_binop_2(X1,k6_numbers),k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),k9_binop_2(X1,k6_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_875, negated_conjecture, (k2_xcmplx_0(esk10_2(esk2_0,X1),X2)=k2_xcmplx_0(X2,esk10_2(esk2_0,X1))|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|~v1_valued_0(esk2_0)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_6470, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k9_binop_2(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k9_binop_2(esk1_0,esk4_0)!=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k9_binop_2(esk1_0,esk4_0)),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k9_binop_2(esk1_0,esk4_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_6136, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k6_numbers)|k9_binop_2(k2_rinfsup1(esk2_0),X1)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6464, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k9_binop_2(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|k9_binop_2(esk1_0,esk4_0)!=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k9_binop_2(esk1_0,esk4_0)),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k9_binop_2(esk1_0,esk4_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_6256, plain, (r1_xxreal_0(X1,k3_seq_4(k1_numbers))|r1_xxreal_0(k3_seq_4(k1_numbers),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6117, negated_conjecture, (k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|esk10_2(esk2_0,X1)!=k1_funct_1(esk2_0,esk9_3(esk2_0,X2,X3))|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_911, plain, (r1_xxreal_0(X1,esk8_2(X2,X3))|r1_xxreal_0(esk8_2(X2,X3),X1)|X3=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk7_2(X2,X3))|~v1_xreal_0(X1)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_6257, plain, (r1_xxreal_0(X1,k3_seq_4(k1_numbers))|r1_xxreal_0(k3_seq_4(k1_numbers),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_912, plain, (r1_xxreal_0(X1,esk6_3(X2,X3,X4))|r1_xxreal_0(esk6_3(X2,X3,X4),X1)|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X1)|X3!=k3_seq_4(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_6304, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk5_1(X1),esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k2_relset_1(X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6474, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_xcmplx_0(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_xcmplx_0(esk1_0,esk4_0)!=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)!=esk1_0|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k2_xcmplx_0(esk1_0,esk4_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),k2_xcmplx_0(esk1_0,esk4_0)),k6_numbers))).#trainneg
cnf(i_0_6327, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,esk7_2(k2_relset_1(X1,X3),X4))|X4=k3_seq_4(k2_relset_1(X1,X3))|v1_xboole_0(k2_relset_1(X1,X3))|X2!=k3_seq_4(X1)|~v3_xxreal_2(X1)|~v3_membered(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3)|~v1_xreal_0(X2)|~v3_xxreal_2(k2_relset_1(X1,X3))|~v3_membered(k2_relset_1(X1,X3))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,X3),X4),k6_numbers)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_6468, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_xcmplx_0(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|k2_xcmplx_0(esk1_0,esk4_0)!=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)!=esk1_0|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k2_xcmplx_0(esk1_0,esk4_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),k2_xcmplx_0(esk1_0,esk4_0)),k6_numbers))).#trainneg
cnf(i_0_6317, plain, (v1_xboole_0(X1)|v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk12_3(X3,X1,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~r2_hidden(X3,k2_relset_1(X1,X2)))).#trainneg
cnf(i_0_6471, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k9_binop_2(esk4_0,esk1_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k9_binop_2(esk4_0,esk1_0)!=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k9_binop_2(esk4_0,esk1_0)),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k9_binop_2(esk4_0,esk1_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_6298, plain, (r1_xxreal_0(X1,k3_seq_4(k5_numbers))|r1_xxreal_0(k3_seq_4(k5_numbers),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_914, plain, (r1_xxreal_0(X1,k5_seq_4(X2))|r1_xxreal_0(k5_seq_4(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6299, plain, (r1_xxreal_0(X1,k3_seq_4(k5_numbers))|r1_xxreal_0(k3_seq_4(k5_numbers),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2488, negated_conjecture, (r1_xxreal_0(k3_seq_4(X1),k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),k3_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_915, plain, (r1_xxreal_0(X1,k1_relset_1(X2,X3))|r1_xxreal_0(k1_relset_1(X2,X3),X1)|~v1_xreal_0(X1)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v3_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_916, plain, (r1_xxreal_0(X1,k3_funct_2(X2,X3,X4,X5))|r1_xxreal_0(k3_funct_2(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_6636, negated_conjecture, (X1=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X3,esk2_0))|v1_xboole_0(X3)|v1_xboole_0(X2)|r1_xxreal_0(esk8_2(k2_relset_1(X2,esk2_0),X1),k6_numbers)|X1!=k3_seq_4(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X3,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X1))|~v1_xreal_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_funct_2(esk2_0,k5_numbers,X3)|~v1_xreal_0(esk8_2(k2_relset_1(X2,esk2_0),X1)))).#trainneg
cnf(i_0_6465, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k9_binop_2(esk4_0,esk1_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|k9_binop_2(esk4_0,esk1_0)!=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k9_binop_2(esk4_0,esk1_0)),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(k9_binop_2(esk4_0,esk1_0))|~v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_6585, negated_conjecture, (esk8_2(X1,k6_numbers)=k3_seq_4(k2_relset_1(X2,esk2_0))|k3_seq_4(X1)=k6_numbers|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|esk8_2(X1,k6_numbers)!=k2_rinfsup1(esk2_0)|~r2_hidden(esk7_2(k2_relset_1(X2,esk2_0),esk8_2(X1,k6_numbers)),X1)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(X1)|~v3_membered(k2_relset_1(X2,esk2_0))|~v3_membered(X1)|~r1_xxreal_0(k6_numbers,esk7_2(X1,k6_numbers))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(esk7_2(k2_relset_1(X2,esk2_0),esk8_2(X1,k6_numbers)))|~v1_xreal_0(esk8_2(X1,k6_numbers)))).#trainneg
cnf(i_0_6126, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),X1)|r1_xxreal_0(X1,k2_rinfsup1(esk2_0))|X2!=k10_xtuple_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X3))).#trainneg
cnf(i_0_6312, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X1,X2,X3)),esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),esk5_1(k3_funct_2(k5_numbers,X1,X2,X3)))|r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k6_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2354, negated_conjecture, (r2_hidden(X1,k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|v1_xboole_0(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~m1_subset_1(X1,k2_relset_1(X2,esk2_0)))).#trainneg
cnf(i_0_1302, negated_conjecture, (r1_xxreal_0(esk1_0,k5_seq_4(X1))|r1_xxreal_0(k5_seq_4(X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_917, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6722, negated_conjecture, (v1_xreal_0(esk10_2(esk2_0,k5_numbers))|k10_xtuple_0(esk2_0)=k5_numbers)).#trainneg
cnf(i_0_921, plain, (r1_xxreal_0(X1,k3_seq_4(X2))|r1_xxreal_0(k3_seq_4(X2),X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_927, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~m1_subset_1(X1,k5_numbers)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_6797, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),esk6_3(k10_xtuple_0(esk2_0),X2,X3))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk6_3(k10_xtuple_0(esk2_0),X2,X3))|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_6801, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_6809, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),k1_funct_1(esk2_0,X2))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_7167, negated_conjecture, (X1=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|X1!=k3_seq_4(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(X1,esk7_2(k10_xtuple_0(esk2_0),X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6815, plain, (r1_xxreal_0(k2_rinfsup1(X1),X2)|v1_xboole_0(X3)|X3!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_6999, negated_conjecture, (k1_funct_1(esk2_0,esk11_2(esk2_0,X1))=esk10_2(esk2_0,X1)|X1=k10_xtuple_0(esk2_0)|esk10_2(esk2_0,X1)!=X3|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_6723, negated_conjecture, (v7_ordinal1(esk10_2(esk2_0,k5_numbers))|k10_xtuple_0(esk2_0)=k5_numbers|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_950, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xboole_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_958, plain, (r1_xxreal_0(X1,k5_seq_4(X2))|r1_xxreal_0(k5_seq_4(X2),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6795, plain, (r1_xxreal_0(k2_rinfsup1(X1),esk7_2(X2,X3))|X3=k3_seq_4(X2)|v1_xboole_0(X2)|X2!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk7_2(X2,X3))|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(esk8_2(X2,X3),k6_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_6798, plain, (r1_xxreal_0(k2_rinfsup1(X1),esk6_3(X2,X3,X4))|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|X2!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk6_3(X2,X3,X4))|X3!=k3_seq_4(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_6898, negated_conjecture, (k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))=k3_seq_4(k2_relset_1(X3,esk2_0))|X1=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X3,esk2_0))|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(k10_xtuple_0(esk2_0))|k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X3,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X3)|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1)))|~v1_xreal_0(X1)|~m1_subset_1(esk7_2(k2_relset_1(X3,esk2_0),k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_6816, negated_conjecture, (k10_xtuple_0(esk2_0)=k6_numbers|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0)).#trainneg
cnf(i_0_6817, negated_conjecture, (v1_xreal_0(k10_xtuple_0(esk2_0))|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k2_rinfsup1(esk2_0)!=esk1_0)).#trainneg
cnf(i_0_6792, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),esk7_2(k2_relset_1(X2,esk2_0),X3))|X3=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk7_2(k2_relset_1(X2,esk2_0),X3))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(esk8_2(k2_relset_1(X2,esk2_0),X3),k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_7204, negated_conjecture, (r1_xxreal_0(esk7_2(X1,X2),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),esk7_2(X1,X2))|X2=k3_seq_4(X1)|v1_xboole_0(X1)|X1!=k10_xtuple_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(esk8_2(X1,X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1061, plain, (v1_xreal_0(k2_relset_1(X1,k2_zfmisc_1(X2,X1)))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5664, negated_conjecture, (k2_xcmplx_0(esk1_0,k6_numbers)=k9_binop_2(k6_numbers,esk1_0)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_5665, negated_conjecture, (k2_xcmplx_0(esk1_0,k6_numbers)=k9_binop_2(esk1_0,k6_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_7127, plain, (r1_xxreal_0(k5_seq_4(k5_numbers),X1)|r1_xxreal_0(X1,k5_seq_4(k5_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_960, plain, (r1_xxreal_0(X1,k5_seq_4(X2))|r1_xxreal_0(k5_seq_4(X2),X1)|~v1_xboole_0(X1)|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_961, plain, (r1_xxreal_0(X1,k5_seq_4(X2))|r1_xxreal_0(k5_seq_4(X2),X1)|~m1_subset_1(X1,k5_numbers)|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6811, plain, (r1_xxreal_0(k2_rinfsup1(X1),k1_funct_1(X2,esk7_2(k9_xtuple_0(X2),X3)))|X3=k3_seq_4(k9_xtuple_0(X2))|v1_xboole_0(k9_xtuple_0(X2))|X4!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k1_funct_1(X2,esk7_2(k9_xtuple_0(X2),X3)))|X4!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X2),X3),k6_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_7176, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_xcmplx_0(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xboole_0(X2)|m2_subset_1(esk3_0,k1_numbers,k5_numbers)|k3_seq_4(k2_relset_1(X2,esk2_0))!=k2_xcmplx_0(esk1_0,esk4_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k2_xcmplx_0(esk1_0,esk4_0))|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k2_xcmplx_0(esk1_0,esk4_0)),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_6950, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k6_numbers)|k9_binop_2(X1,k2_rinfsup1(esk2_0))!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7179, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_seq_4(k2_relset_1(X2,esk2_0))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_2(esk2_0,k5_numbers,X2)|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),k2_rinfsup1(esk2_0)),k6_numbers))).#trainneg
cnf(i_0_6796, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),esk6_3(k2_relset_1(X2,esk2_0),X3,X4))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk6_3(k2_relset_1(X2,esk2_0),X3,X4))|X3!=k3_seq_4(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_6802, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),k3_funct_2(k5_numbers,X2,esk2_0,X3))|v1_xboole_0(X2)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(k3_funct_2(k5_numbers,X2,esk2_0,X3))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_6899, negated_conjecture, (k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))=k3_seq_4(k2_relset_1(X3,esk2_0))|X1=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X3,esk2_0))|v1_xboole_0(X2)|v1_xboole_0(X3)|k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X3,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X3)|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1)))|~v1_xreal_0(X1)|~r1_xxreal_0(esk8_2(k2_relset_1(X3,esk2_0),k9_binop_2(X1,esk8_2(k2_relset_1(X2,esk2_0),X1))),k6_numbers))).#trainneg
cnf(i_0_7128, plain, (r1_xxreal_0(k5_seq_4(k1_numbers),X1)|r1_xxreal_0(X1,k5_seq_4(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2755, plain, (m1_subset_1(k9_xtuple_0(k2_zfmisc_1(X1,X2)),k1_numbers)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_965, negated_conjecture, (m1_subset_1(esk11_2(esk2_0,X1),X2)|X1=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X1),X1)|~v4_relat_1(esk2_0,X2))).#trainneg
cnf(i_0_7250, negated_conjecture, (r1_xxreal_0(k1_funct_1(esk2_0,X1),X2)|r1_xxreal_0(X2,k1_funct_1(esk2_0,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_966, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,k9_xtuple_0(X1)),X2)|k9_xtuple_0(X1)=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,k9_xtuple_0(X1)),k9_xtuple_0(esk2_0))|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_967, negated_conjecture, (m1_subset_1(esk10_2(esk2_0,k9_xtuple_0(X1)),X2)|k1_funct_1(esk2_0,esk11_2(esk2_0,k9_xtuple_0(X1)))=esk10_2(esk2_0,k9_xtuple_0(X1))|k9_xtuple_0(X1)=k10_xtuple_0(esk2_0)|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6814, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),X2)|v1_xboole_0(X3)|k2_relset_1(X3,esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~r2_hidden(X2,k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X3))).#trainneg
cnf(i_0_6794, plain, (r1_xxreal_0(k2_rinfsup1(X1),esk7_2(X2,X3))|X3=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk8_2(X2,X3))|X2!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk7_2(X2,X3))|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_7188, negated_conjecture, (esk8_2(k2_relset_1(X1,esk2_0),k6_numbers)=k3_seq_4(k2_relset_1(X2,esk2_0))|k3_seq_4(k2_relset_1(X1,esk2_0))=k6_numbers|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|v1_xboole_0(X1)|esk8_2(k2_relset_1(X1,esk2_0),k6_numbers)!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~r1_xxreal_0(k6_numbers,esk7_2(k2_relset_1(X1,esk2_0),k6_numbers))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(esk7_2(k2_relset_1(X2,esk2_0),esk8_2(k2_relset_1(X1,esk2_0),k6_numbers)))|~v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),k6_numbers))|~r2_hidden(esk7_2(k2_relset_1(X2,esk2_0),esk8_2(k2_relset_1(X1,esk2_0),k6_numbers)),k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_7232, negated_conjecture, (r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|X1=k2_rinfsup1(esk2_0)|v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7177, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_xcmplx_0(esk1_0,esk4_0)|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_seq_4(k2_relset_1(X2,esk2_0))!=k2_xcmplx_0(esk1_0,esk4_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k2_xcmplx_0(esk1_0,esk4_0))|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k2_xcmplx_0(esk1_0,esk4_0)),k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_6895, negated_conjecture, (k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1))=k3_seq_4(k2_relset_1(X2,esk2_0))|X1=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1))!=k2_rinfsup1(esk2_0)|~r2_hidden(esk7_2(k2_relset_1(X2,esk2_0),k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1))),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~r1_xxreal_0(X1,esk7_2(k10_xtuple_0(esk2_0),X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2895, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),esk12_3(X1,X2,X3))|r1_xxreal_0(esk12_3(X1,X2,X3),k2_rinfsup1(esk2_0))|v1_xboole_0(X2)|~r2_hidden(X1,k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6509, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(esk1_0,k6_numbers)|k3_seq_4(X1)!=esk4_0|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(esk4_0)|k2_rinfsup1(esk2_0)!=esk1_0|~r2_hidden(esk6_3(X1,esk4_0,esk1_0),k10_xtuple_0(esk2_0))|~r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_3148, plain, (r1_xxreal_0(k3_seq_4(X1),esk7_2(X2,X3))|X3=k3_seq_4(X2)|v1_xboole_0(X2)|X1!=k10_xtuple_0(X4)|X2!=k10_xtuple_0(X4)|~v1_relat_1(X4)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X4)|~v1_xreal_0(esk7_2(X2,X3))|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(esk8_2(X2,X3),k6_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3808, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(k3_seq_4(X1),k1_funct_1(esk2_0,X2))|X1!=k10_xtuple_0(X3)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_funct_1(X3)|~v1_xreal_0(k1_funct_1(esk2_0,X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1310, negated_conjecture, (r1_xxreal_0(esk1_0,k3_seq_4(X1))|r1_xxreal_0(k3_seq_4(X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1201, negated_conjecture, (m1_subset_1(k1_funct_1(esk2_0,X1),k10_xtuple_0(esk2_0))|k6_numbers=k1_numbers|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1037, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(X3)|X3!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_xcmplx_0(X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_969, plain, (m1_subset_1(esk7_2(k9_xtuple_0(X1),X2),X3)|X2=k3_seq_4(k9_xtuple_0(X1))|v1_xboole_0(k9_xtuple_0(X1))|v1_xreal_0(esk8_2(k9_xtuple_0(X1),X2))|~v4_relat_1(X1,X3)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_970, plain, (m1_subset_1(esk7_2(k9_xtuple_0(X1),X2),X3)|X2=k3_seq_4(k9_xtuple_0(X1))|v1_xboole_0(k9_xtuple_0(X1))|~v4_relat_1(X1,X3)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~r1_xxreal_0(esk8_2(k9_xtuple_0(X1),X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7306, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),X3)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_3(esk2_0,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|X2!=k10_xtuple_0(esk2_0)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_7304, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),X2)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk12_3(X2,k1_numbers,esk2_0),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r2_hidden(X2,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_7305, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),X2)|v1_xboole_0(X3)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk12_3(X2,X3,esk2_0),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r2_hidden(X2,k2_relset_1(X3,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X3))).#trainneg
cnf(i_0_7346, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|r1_xxreal_0(k2_rinfsup1(X2),esk6_3(k10_xtuple_0(esk2_0),X3,X4))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(X4,k6_numbers)|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_xreal_0(esk6_3(k10_xtuple_0(esk2_0),X3,X4))|X3!=k2_rinfsup1(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_7397, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),k2_rinfsup1(esk2_0))|k10_xtuple_0(X1)!=k1_numbers|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1041, negated_conjecture, (r1_xxreal_0(k1_funct_1(X1,X2),k1_funct_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_971, plain, (m1_subset_1(esk6_3(k9_xtuple_0(X1),X2,X3),X4)|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|X2!=k3_seq_4(k9_xtuple_0(X1))|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_976, plain, (r1_xxreal_0(X1,k3_funct_2(X2,X3,X4,X5))|r1_xxreal_0(k3_funct_2(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v6_membered(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_7395, plain, (r1_xxreal_0(k2_rinfsup1(X1),k6_numbers)|k10_xtuple_0(X1)!=k5_numbers|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7409, plain, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(X2),k6_numbers)|k3_seq_4(X1)!=k6_numbers|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(k2_rinfsup1(X2))|X1!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7682, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(X2),X3)|v1_xboole_0(k10_xtuple_0(esk2_0))|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_1(X2)|~m1_subset_1(X3,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_7439, negated_conjecture, (X1=k3_seq_4(k10_xtuple_0(X2))|v1_xboole_0(k10_xtuple_0(X2))|r1_xxreal_0(k2_rinfsup1(esk2_0),esk7_2(k10_xtuple_0(X2),X1))|r1_xxreal_0(esk7_2(k10_xtuple_0(X2),X1),k2_rinfsup1(esk2_0))|~v3_xxreal_2(k10_xtuple_0(X2))|~v3_membered(k10_xtuple_0(X2))|~r1_xxreal_0(esk8_2(k10_xtuple_0(X2),X1),k6_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7761, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk1_0,k1_funct_1(esk2_0,X1))|k6_numbers=k1_numbers|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_7789, negated_conjecture, (k2_xcmplx_0(X1,esk6_3(k10_xtuple_0(esk2_0),X2,X3))=k2_xcmplx_0(esk6_3(k10_xtuple_0(esk2_0),X2,X3),X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_xcmplx_0(X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7483, negated_conjecture, (r1_xxreal_0(esk1_0,k5_seq_4(k5_numbers))|r1_xxreal_0(k5_seq_4(k5_numbers),esk1_0))).#trainneg
cnf(i_0_1410, plain, (v3_valued_0(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))|~v3_membered(X2))).#trainneg
cnf(i_0_1060, plain, (k2_relset_1(X1,k2_zfmisc_1(X2,X1))=k10_xtuple_0(k2_zfmisc_1(X2,X1)))).#trainneg
cnf(i_0_7404, negated_conjecture, (v1_xxreal_0(esk10_2(esk2_0,k5_numbers))|k10_xtuple_0(esk2_0)=k5_numbers|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_7458, negated_conjecture, (k9_binop_2(k6_numbers,esk1_0)=k9_binop_2(esk1_0,k6_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_7517, plain, (r1_xxreal_0(X1,k5_seq_4(k5_numbers))|r1_xxreal_0(k5_seq_4(k5_numbers),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7518, plain, (r1_xxreal_0(X1,k5_seq_4(k5_numbers))|r1_xxreal_0(k5_seq_4(k5_numbers),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_981, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|~v3_membered(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_983, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X6))|v1_xboole_0(X2)|~v3_membered(X3)|~m1_subset_1(X6,X1)|~v1_funct_2(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5),X1,X3)|~v1_funct_1(k3_funct_2(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)),X4,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))|~m1_subset_1(X5,X2)|~v1_funct_2(X4,X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_7671, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(X2),esk6_3(k10_xtuple_0(esk2_0),X3,X4))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X4,k6_numbers)|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_1(X2)|X3!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_7384, plain, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(X2),k1_funct_1(X3,X4))|X1!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_xreal_0(k1_funct_1(X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_7759, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|k6_numbers=k1_numbers|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~r2_hidden(X2,k5_numbers))).#trainneg
cnf(i_0_7345, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|r1_xxreal_0(k2_rinfsup1(X2),esk6_3(k2_relset_1(X3,esk2_0),X4,X5))|v1_xboole_0(k2_relset_1(X3,esk2_0))|v1_xboole_0(X3)|v1_xboole_0(X1)|r1_xxreal_0(X5,k6_numbers)|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_xreal_0(esk6_3(k2_relset_1(X3,esk2_0),X4,X5))|X4!=k3_seq_4(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X3,esk2_0))|~v3_membered(k2_relset_1(X3,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_2(esk2_0,k5_numbers,X3)|~v1_xreal_0(X5)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_7651, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|r1_xxreal_0(k1_funct_1(esk2_0,X2),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7686, plain, (k2_rinfsup1(X1)=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk8_2(X2,k2_rinfsup1(X1)))|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(k2_rinfsup1(X1))|X2!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7394, plain, (r1_xxreal_0(k2_rinfsup1(X1),k6_numbers)|k10_xtuple_0(X1)!=k1_numbers|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7349, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),esk6_3(k10_xtuple_0(esk2_0),X2,X3))|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|k10_xtuple_0(X1)!=k1_numbers|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk6_3(k10_xtuple_0(esk2_0),X2,X3))|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7650, negated_conjecture, (r1_xxreal_0(X1,k1_funct_1(esk2_0,X2))|r1_xxreal_0(k1_funct_1(esk2_0,X2),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7355, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),esk5_1(X2))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X2,k6_numbers)|k10_xtuple_0(X1)!=k5_numbers|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk5_1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_7549, negated_conjecture, (r1_xxreal_0(esk1_0,k5_seq_4(k1_numbers))|r1_xxreal_0(k5_seq_4(k1_numbers),esk1_0))).#trainneg
cnf(i_0_1491, plain, (v1_valued_0(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))|~v1_membered(X2))).#trainneg
cnf(i_0_1063, plain, (v1_xreal_0(k9_xtuple_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3)))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(k1_zfmisc_1(X1))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_7513, plain, (r1_xxreal_0(k6_numbers,k5_seq_4(k5_numbers))|r1_xxreal_0(k5_seq_4(k5_numbers),k6_numbers))).#trainneg
cnf(i_0_7583, plain, (r1_xxreal_0(X1,k5_seq_4(k1_numbers))|r1_xxreal_0(k5_seq_4(k1_numbers),X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7584, plain, (r1_xxreal_0(X1,k5_seq_4(k1_numbers))|r1_xxreal_0(k5_seq_4(k1_numbers),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_984, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X4))|~v3_membered(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_985, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_funct_1(k9_xtuple_0(X2),X3))|~v3_membered(X4)|~m1_subset_1(X3,X1)|~v1_funct_2(k9_xtuple_0(X2),X1,X4)|~v1_funct_1(k9_xtuple_0(X2))|~v4_relat_1(X2,k2_zfmisc_1(X1,X4))|~v1_relat_1(X2))).#trainneg
cnf(i_0_3907, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(X2,k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(X2,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_5628, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk5_1(X1),esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_7847, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),X2)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|X3!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X3)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7985, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(k2_rinfsup1(X1),X2)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k10_xtuple_0(esk2_0))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7849, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(X2),X3)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~r2_hidden(X3,k2_relset_1(X1,esk2_0))|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_8210, plain, (r1_xxreal_0(k6_numbers,esk8_2(X1,k2_rinfsup1(X2)))|r1_xxreal_0(esk8_2(X1,k2_rinfsup1(X2)),k6_numbers)|k2_rinfsup1(X2)=k3_seq_4(X1)|v1_xboole_0(X1)|X1!=k10_xtuple_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_8364, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(X1)|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X1,esk2_0,X2)),esk5_1(X3))|r1_xxreal_0(esk5_1(X3),esk5_1(k3_funct_2(k5_numbers,X1,esk2_0,X2)))|r1_xxreal_0(k3_funct_2(k5_numbers,X1,esk2_0,X2),k6_numbers)|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_7983, plain, (v1_xboole_0(k10_xtuple_0(X1))|r1_xxreal_0(k2_rinfsup1(X1),k6_numbers)|k3_seq_4(k10_xtuple_0(X1))!=k6_numbers|~v3_xxreal_2(k10_xtuple_0(X1))|~v3_membered(k10_xtuple_0(X1))|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7848, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),X2)|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~r2_hidden(X2,k10_xtuple_0(esk2_0))|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8053, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k10_xtuple_0(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_8035, plain, (m1_subset_1(k10_xtuple_0(k2_zfmisc_1(X2,X1)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1061, plain, (v1_xreal_0(k10_xtuple_0(k2_zfmisc_1(X2,X1)))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1065, plain, (v1_xreal_0(k9_xtuple_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3)))|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(k1_zfmisc_1(X1))|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_7579, plain, (r1_xxreal_0(k6_numbers,k5_seq_4(k1_numbers))|r1_xxreal_0(k5_seq_4(k1_numbers),k6_numbers))).#trainneg
cnf(i_0_8079, negated_conjecture, (v1_xreal_0(k9_binop_2(esk1_0,k6_numbers))|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_2487, negated_conjecture, (r1_xxreal_0(k9_binop_2(X1,X2),k1_seq_1(esk2_0,X3))|r1_xxreal_0(k1_seq_1(esk2_0,X3),k9_binop_2(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1000, plain, (v7_ordinal1(k3_funct_2(k5_numbers,X1,X2,X3))|v1_xboole_0(X1)|~v6_membered(k2_relset_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7989, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k6_numbers=k1_numbers|r1_xxreal_0(esk1_0,X2)|~r2_hidden(esk9_3(esk2_0,X1,X2),k5_numbers)|X1!=k10_xtuple_0(esk2_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_7988, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|k6_numbers=k1_numbers|r1_xxreal_0(esk1_0,X1)|v1_xboole_0(X2)|~r2_hidden(esk12_3(X1,X2,esk2_0),k5_numbers)|~r2_hidden(X1,k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_2907, negated_conjecture, (r1_xxreal_0(X1,esk7_2(k10_xtuple_0(esk2_0),X2))|X2=k2_rinfsup1(esk2_0)|v1_xboole_0(k10_xtuple_0(esk2_0))|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),X2),k6_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1530, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_valued_0(esk2_0)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_3309, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|k9_binop_2(esk1_0,X1)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(k9_binop_2(esk1_0,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2146, plain, (v1_xboole_0(X1)|r1_xxreal_0(k3_seq_4(k2_relset_1(X1,X2)),X3)|v1_xboole_0(X4)|~v3_xxreal_2(k2_relset_1(X1,X2))|~v3_membered(k2_relset_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk12_3(X3,X4,X2),k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2)|~v1_xreal_0(X3)|~r2_hidden(X3,k2_relset_1(X4,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X4)))|~v1_funct_2(X2,k5_numbers,X4))).#trainneg
cnf(i_0_4248, negated_conjecture, (k2_xcmplx_0(k6_numbers,k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1))=k3_funct_2(k5_numbers,k1_numbers,esk2_0,X1)|~v1_valued_0(esk2_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7854, negated_conjecture, (r1_xxreal_0(X3,X3)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_8002, plain, (v3_valued_0(k10_xtuple_0(k2_zfmisc_1(X2,k2_zfmisc_1(X1,k1_numbers)))))).#trainneg
cnf(i_0_375, plain, (X1=k6_numbers|r2_hidden(k1_funct_1(k2_zfmisc_1(X2,X1),X3),k10_xtuple_0(k2_zfmisc_1(X2,X1)))|~r2_hidden(X3,X2)|~v1_funct_2(k2_zfmisc_1(X2,X1),X2,X1)|~v1_funct_1(k2_zfmisc_1(X2,X1)))).#trainneg
cnf(i_0_1099, plain, (v3_membered(k9_xtuple_0(k1_relset_1(k2_zfmisc_1(X1,X2),X3)))|~v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X1)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_8212, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_3199, plain, (r1_xxreal_0(X1,k3_seq_4(X2))|r1_xxreal_0(k3_seq_4(X2),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers)|~v3_membered(X2))).#trainneg
cnf(i_0_1025, negated_conjecture, (k2_xcmplx_0(X1,esk11_2(esk2_0,X2))=k2_xcmplx_0(esk11_2(esk2_0,X2),X1)|X2=k10_xtuple_0(esk2_0)|r2_hidden(esk10_2(esk2_0,X2),X2)|k9_xtuple_0(esk2_0)!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_xcmplx_0(X1)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1026, negated_conjecture, (k2_xcmplx_0(X1,esk10_2(esk2_0,X2))=k2_xcmplx_0(esk10_2(esk2_0,X2),X1)|X2=k10_xtuple_0(esk2_0)|r2_hidden(esk11_2(esk2_0,X2),k9_xtuple_0(esk2_0))|X2!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_xcmplx_0(X1)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_7373, plain, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(X2),k3_funct_2(k5_numbers,X1,X3,X4))|X1!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_xreal_0(k3_funct_2(k5_numbers,X1,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_4929, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xreal_0(k3_funct_2(k5_numbers,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|~v5_relat_1(X3,X1)|~v1_relat_1(X3)|~v3_membered(X1))).#trainneg
cnf(i_0_391, negated_conjecture, (v1_xreal_0(esk5_1(X1))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2640, plain, (v1_xboole_0(k2_relset_1(X1,X2))|m1_subset_1(k1_funct_1(X3,X4),X1)|~v5_relat_1(X2,X1)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_relset_1(X1,X2))))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,k2_relset_1(X1,X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_2921, negated_conjecture, (X1=k2_rinfsup1(esk2_0)|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),X1))|~r1_xxreal_0(esk8_2(k10_xtuple_0(esk2_0),X1),k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8003, plain, (v1_relat_1(k10_xtuple_0(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_8011, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(k10_xtuple_0(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2))))!=k1_numbers|~v1_membered(X2)|~v1_funct_1(k10_xtuple_0(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1101, plain, (v3_membered(k9_xtuple_0(k2_relset_1(k2_zfmisc_1(X1,X2),X3)))|~v1_relat_1(k2_relset_1(k2_zfmisc_1(X1,X2),X3))|~v3_membered(X1)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_8080, negated_conjecture, (v7_ordinal1(k9_binop_2(esk1_0,k6_numbers))|~v6_membered(k1_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_8081, negated_conjecture, (v1_xxreal_0(k9_binop_2(esk1_0,k6_numbers))|~v6_membered(k1_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_3201, plain, (r1_xxreal_0(X1,k3_seq_4(X2))|r1_xxreal_0(k3_seq_4(X2),X1)|~v1_xboole_0(X1)|~v6_membered(k1_numbers)|~v3_membered(X2))).#trainneg
cnf(i_0_1027, negated_conjecture, (k2_xcmplx_0(X1,esk10_2(esk2_0,X2))=k2_xcmplx_0(esk10_2(esk2_0,X2),X1)|k1_funct_1(esk2_0,esk11_2(esk2_0,X2))=esk10_2(esk2_0,X2)|X2=k10_xtuple_0(esk2_0)|X2!=k10_xtuple_0(X3)|~v1_valued_0(X3)|~v1_xcmplx_0(X1)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1028, plain, (k2_xcmplx_0(X1,esk9_3(X2,X3,X4))=k2_xcmplx_0(esk9_3(X2,X3,X4),X1)|k9_xtuple_0(X2)!=k10_xtuple_0(X5)|~v1_valued_0(X5)|~v1_xcmplx_0(X1)|~v1_relat_1(X5)|~v1_funct_1(X5)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_2346, negated_conjecture, (r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),X2),k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1096, plain, (r1_xxreal_0(X1,k1_funct_1(X2,esk6_3(k9_xtuple_0(X2),X3,X4)))|v1_xboole_0(k9_xtuple_0(X2))|r1_xxreal_0(X4,k6_numbers)|X1!=k3_seq_4(X5)|~v3_xxreal_2(X5)|~v3_membered(X5)|~v1_xreal_0(X1)|~v1_xreal_0(k1_funct_1(X2,esk6_3(k9_xtuple_0(X2),X3,X4)))|X3!=k3_seq_4(k9_xtuple_0(X2))|X5!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2019, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k6_numbers)|r1_xxreal_0(k6_numbers,k1_funct_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8787, negated_conjecture, (r1_xxreal_0(esk5_1(X1),k2_rinfsup1(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8783, negated_conjecture, (r1_xxreal_0(esk1_0,esk5_1(X1))|r1_xxreal_0(esk5_1(X1),esk1_0)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8788, negated_conjecture, (r1_xxreal_0(k6_numbers,esk5_1(X1))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3726, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_1122, plain, (v1_xboole_0(k1_zfmisc_1(X1))|v1_xreal_0(k1_relset_1(X1,X2))|k1_zfmisc_1(X1)!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8588, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k9_binop_2(X2,esk1_0))|r1_xxreal_0(k9_binop_2(X2,esk1_0),k1_seq_1(esk2_0,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1029, plain, (k2_xcmplx_0(X1,esk7_2(X2,X3))=k2_xcmplx_0(esk7_2(X2,X3),X1)|X3=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk8_2(X2,X3))|X2!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_xcmplx_0(X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1030, plain, (k2_xcmplx_0(X1,esk7_2(X2,X3))=k2_xcmplx_0(esk7_2(X2,X3),X1)|X3=k3_seq_4(X2)|v1_xboole_0(X2)|X2!=k10_xtuple_0(X4)|~v1_valued_0(X4)|~v1_xcmplx_0(X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(esk8_2(X2,X3),k6_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_8784, negated_conjecture, (m1_subset_1(esk5_1(X1),k1_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8777, negated_conjecture, (r1_xxreal_0(X1,esk5_1(X2))|r1_xxreal_0(esk5_1(X2),X1)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8400, negated_conjecture, (k2_rinfsup1(X1)=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(esk7_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)),esk1_0)|r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k2_rinfsup1(X1))|k2_relset_1(X2,esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8401, negated_conjecture, (k2_rinfsup1(X1)=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,esk7_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|r1_xxreal_0(esk7_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)),k6_numbers)|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k2_rinfsup1(X1))|k2_relset_1(X2,esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8813, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_8865, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8397, negated_conjecture, (k2_rinfsup1(X1)=k3_seq_4(k2_relset_1(X2,esk2_0))|k2_rinfsup1(esk2_0)=esk1_0|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(esk1_0,esk7_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k2_rinfsup1(X1))|k2_relset_1(X2,esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8651, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k3_seq_4(k2_relset_1(X1,esk2_0)),X2)|~r2_hidden(X2,k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8037, plain, (v3_membered(k10_xtuple_0(k2_zfmisc_1(X2,X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_4074, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|m1_subset_1(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k1_numbers))).#trainneg
cnf(i_0_1126, plain, (v1_xboole_0(k1_zfmisc_1(X1))|v1_xreal_0(k2_relset_1(X1,X2))|k1_zfmisc_1(X1)!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v5_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8622, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k9_binop_2(X2,k2_rinfsup1(esk2_0)))|r1_xxreal_0(k9_binop_2(X2,k2_rinfsup1(esk2_0)),k1_seq_1(esk2_0,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1031, plain, (k2_xcmplx_0(X1,esk6_3(X2,X3,X4))=k2_xcmplx_0(esk6_3(X2,X3,X4),X1)|v1_xboole_0(X2)|r1_xxreal_0(X4,k6_numbers)|X2!=k10_xtuple_0(X5)|~v1_valued_0(X5)|~v1_xcmplx_0(X1)|~v1_relat_1(X5)|~v1_funct_1(X5)|X3!=k3_seq_4(X2)|~v3_xxreal_2(X2)|~v3_membered(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1032, plain, (k2_xcmplx_0(X1,k3_funct_2(k5_numbers,X2,X3,X4))=k2_xcmplx_0(k3_funct_2(k5_numbers,X2,X3,X4),X1)|v1_xboole_0(X2)|k2_relset_1(X2,X3)!=k10_xtuple_0(X5)|~v1_valued_0(X5)|~v1_xcmplx_0(X1)|~v1_relat_1(X5)|~v1_funct_1(X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_8402, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X1,esk2_0),k2_rinfsup1(esk2_0)))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(esk2_0))).#trainneg
cnf(i_0_8772, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(esk5_1(X2)))|r1_xxreal_0(esk5_1(esk5_1(X2)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk5_1(X2),k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8641, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|k9_binop_2(esk1_0,X1)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~m1_subset_1(k9_binop_2(esk1_0,X1),k5_numbers))).#trainneg
cnf(i_0_8635, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k10_xtuple_0(esk2_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_8399, negated_conjecture, (k2_rinfsup1(X1)=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(X3,esk7_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|X3!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X3)|~v1_xreal_0(k2_rinfsup1(X1))|k2_relset_1(X2,esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8859, negated_conjecture, (r1_xxreal_0(X1,esk7_2(k2_relset_1(X2,esk2_0),X3))|X3=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|v1_xreal_0(esk8_2(k2_relset_1(X2,esk2_0),X3))|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(k2_relset_1(X2,esk2_0),X3))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_4073, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),esk1_0)|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_1129, plain, (v1_xboole_0(k1_zfmisc_1(X1))|v1_xreal_0(k9_xtuple_0(X2))|k1_zfmisc_1(X1)!=k10_xtuple_0(X3)|~v3_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_8077, negated_conjecture, (m1_subset_1(k9_binop_2(esk1_0,k6_numbers),k1_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_1621, plain, (r1_xxreal_0(k6_numbers,k1_funct_1(X1,X2))|r1_xxreal_0(k1_funct_1(X1,X2),k6_numbers)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1033, plain, (k2_xcmplx_0(X1,k1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X3),X4),X5))=k2_xcmplx_0(k1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X3),X4),X5),X1)|X3=k6_numbers|k2_relset_1(X3,k1_relset_1(k2_zfmisc_1(X2,X3),X4))!=k10_xtuple_0(X6)|~v1_valued_0(X6)|~v1_xcmplx_0(X1)|~v1_relat_1(X6)|~v1_funct_1(X6)|~v4_relat_1(X4,k2_zfmisc_1(X2,X3))|~v1_relat_1(X4)|~r2_hidden(X5,X2)|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X2,X3),X4),X2,X3)|~v1_funct_1(k1_relset_1(k2_zfmisc_1(X2,X3),X4)))).#trainneg
cnf(i_0_1035, plain, (k2_xcmplx_0(X1,k1_funct_1(k2_relset_1(k2_zfmisc_1(X2,X3),X4),X5))=k2_xcmplx_0(k1_funct_1(k2_relset_1(k2_zfmisc_1(X2,X3),X4),X5),X1)|X3=k6_numbers|k2_relset_1(X3,k2_relset_1(k2_zfmisc_1(X2,X3),X4))!=k10_xtuple_0(X6)|~v1_valued_0(X6)|~v1_xcmplx_0(X1)|~v1_relat_1(X6)|~v1_funct_1(X6)|~v5_relat_1(X4,k2_zfmisc_1(X2,X3))|~v1_relat_1(X4)|~r2_hidden(X5,X2)|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X2,X3),X4),X2,X3)|~v1_funct_1(k2_relset_1(k2_zfmisc_1(X2,X3),X4)))).#trainneg
cnf(i_0_8867, negated_conjecture, (r1_xxreal_0(esk7_2(k2_relset_1(X1,esk2_0),X2),k6_numbers)|r1_xxreal_0(k6_numbers,esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8847, negated_conjecture, (X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~v1_xboole_0(k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8409, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk5_1(X1),esk5_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2)))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2)),esk5_1(X1))|r1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X2,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8786, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk5_1(esk5_1(X1)),esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8848, negated_conjecture, (m1_subset_1(esk7_2(k2_relset_1(X1,esk2_0),X2),k10_xtuple_0(esk2_0))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8713, negated_conjecture, (r1_xxreal_0(X1,X1)|v1_xboole_0(X2)|X2!=k10_xtuple_0(X3)|~v2_valued_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_8004, plain, (v4_relat_1(k10_xtuple_0(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_4077, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k6_numbers)|r1_xxreal_0(k6_numbers,esk7_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_1137, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers))|v1_xreal_0(X1)|k10_xtuple_0(X2)!=k2_zfmisc_1(k5_numbers,k1_numbers)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_3202, plain, (r1_xxreal_0(X1,k3_seq_4(X2))|r1_xxreal_0(k3_seq_4(X2),X1)|~m1_subset_1(X1,k5_numbers)|~v6_membered(k1_numbers)|~v3_membered(X2))).#trainneg
cnf(i_0_3274, plain, (r1_xxreal_0(X1,k9_binop_2(k6_numbers,X2))|r1_xxreal_0(k9_binop_2(k6_numbers,X2),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1064, plain, (v1_xreal_0(k9_xtuple_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5)))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|~v3_membered(k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_1079, plain, (m1_subset_1(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)),X4)|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|~v4_relat_1(X5,X4)|~v1_relat_1(X5)|X2!=k3_seq_4(k9_xtuple_0(X1))|k1_relset_1(X4,X5)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9208, negated_conjecture, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_8993, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9049, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k10_xtuple_0(esk2_0))))|~v1_funct_2(esk2_0,k5_numbers,k10_xtuple_0(esk2_0))|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_9044, negated_conjecture, (r1_xxreal_0(esk5_1(X1),X2)|r1_xxreal_0(X2,esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_8007, plain, (v5_relat_1(k10_xtuple_0(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_4076, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(k2_rinfsup1(esk2_0),esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),k2_rinfsup1(esk2_0)))).#trainneg
cnf(i_0_1144, plain, (m1_subset_1(k1_funct_1(X1,esk9_3(X1,X2,X3)),X4)|~v4_relat_1(X5,X4)|~v1_relat_1(X5)|k1_relset_1(X4,X5)!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3276, plain, (r1_xxreal_0(X1,k9_binop_2(k6_numbers,X2))|r1_xxreal_0(k9_binop_2(k6_numbers,X2),X1)|~v1_xboole_0(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3327, plain, (r1_xxreal_0(X1,k9_binop_2(X2,k6_numbers))|r1_xxreal_0(k9_binop_2(X2,k6_numbers),X1)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1082, plain, (m1_subset_1(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)),X4)|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|~v5_relat_1(X5,X4)|~v1_relat_1(X5)|X2!=k3_seq_4(k9_xtuple_0(X1))|k2_relset_1(X4,X5)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1087, plain, (r2_hidden(X1,X2)|v1_xboole_0(k9_xtuple_0(X4))|r1_xxreal_0(X6,k6_numbers)|X1!=k1_funct_1(X3,k1_funct_1(X4,esk6_3(k9_xtuple_0(X4),X5,X6)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|X5!=k3_seq_4(k9_xtuple_0(X4))|k9_xtuple_0(X3)!=k10_xtuple_0(X4)|~v1_relat_1(X4)|~v3_xxreal_2(k9_xtuple_0(X4))|~v3_membered(k9_xtuple_0(X4))|~v1_funct_1(X4)|~v1_xreal_0(X6)|~v1_xreal_0(X5))).#trainneg
cnf(i_0_9060, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k3_seq_4(k2_relset_1(X2,esk2_0)),k6_numbers)|v1_xboole_0(X2)|k3_seq_4(X1)!=k6_numbers|~v3_xxreal_2(X1)|~v3_membered(X1)|~v1_xreal_0(k3_seq_4(k2_relset_1(X2,esk2_0)))|~r2_hidden(esk6_3(X1,k6_numbers,k3_seq_4(k2_relset_1(X2,esk2_0))),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2))).#trainneg
cnf(i_0_9056, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k3_seq_4(k2_relset_1(X1,esk2_0)))|~r2_hidden(esk7_2(k2_relset_1(X2,esk2_0),k3_seq_4(k2_relset_1(X1,esk2_0))),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_9048, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=esk1_0|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_8643, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|k9_binop_2(esk1_0,X1)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8634, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X1,esk7_2(k10_xtuple_0(esk2_0),esk1_0))|X1!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1)|~v1_xreal_0(esk8_2(k10_xtuple_0(esk2_0),esk1_0)))).#trainneg
cnf(i_0_9058, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=k3_seq_4(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k3_seq_4(k2_relset_1(X1,esk2_0)))|~r2_hidden(esk7_2(k10_xtuple_0(esk2_0),k3_seq_4(k2_relset_1(X1,esk2_0))),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_8398, negated_conjecture, (k2_rinfsup1(X1)=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X2,esk2_0),k2_rinfsup1(X1)))|~v1_xboole_0(k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k2_rinfsup1(X1))|k2_relset_1(X2,esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9163, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk6_3(k10_xtuple_0(esk2_0),X2,X3)))|r1_xxreal_0(esk5_1(esk6_3(k10_xtuple_0(esk2_0),X2,X3)),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk6_3(k10_xtuple_0(esk2_0),X2,X3),k6_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X1)|X2!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8382, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(X2),esk7_2(k2_relset_1(X1,esk2_0),X3))|X3=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X3))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X2)|~v2_seq_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_1(X2)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_8266, negated_conjecture, (k2_xcmplx_0(k6_numbers,esk1_0)=esk1_0|k10_xtuple_0(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))!=k1_numbers|~v1_funct_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))|~v1_membered(X2))).#trainneg
cnf(i_0_1147, plain, (m1_subset_1(k1_funct_1(X1,esk9_3(X1,X2,X3)),X4)|~v5_relat_1(X5,X4)|~v1_relat_1(X5)|k2_relset_1(X4,X5)!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3329, plain, (r1_xxreal_0(X1,k9_binop_2(X2,k6_numbers))|r1_xxreal_0(k9_binop_2(X2,k6_numbers),X1)|~v1_xboole_0(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1088, plain, (m1_subset_1(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)),X4)|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|~v4_relat_1(X5,X4)|~v1_relat_1(X5)|X2!=k3_seq_4(k9_xtuple_0(X1))|k9_xtuple_0(X5)!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1089, plain, (m1_subset_1(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)),k1_numbers)|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(k9_xtuple_0(X1))|k10_xtuple_0(X1)!=k5_numbers|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8812, negated_conjecture, (k2_rinfsup1(X1)=k2_rinfsup1(esk2_0)|v1_xboole_0(k10_xtuple_0(esk2_0))|r1_xxreal_0(esk1_0,esk7_2(k10_xtuple_0(esk2_0),k2_rinfsup1(X1)))|r1_xxreal_0(k6_numbers,esk8_2(k10_xtuple_0(esk2_0),k2_rinfsup1(X1)))|~v1_xreal_0(k2_rinfsup1(X1))|k10_xtuple_0(esk2_0)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8664, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k3_seq_4(k2_relset_1(X2,X3)),k3_funct_2(k5_numbers,X1,X3,X4))|~v3_xxreal_2(k2_relset_1(X2,X3))|~v3_membered(k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3)|~v1_xreal_0(k3_funct_2(k5_numbers,X1,X3,X4))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_9043, negated_conjecture, (r1_xxreal_0(esk5_1(X1),X2)|r1_xxreal_0(X2,esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8896, plain, (r1_xxreal_0(k6_numbers,X3)|r1_xxreal_0(X3,k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(esk9_3(X1,X2,X3),k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers)|~v1_funct_1(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_8395, plain, (k2_rinfsup1(X1)=k3_seq_4(k2_relset_1(X2,X3))|v1_xboole_0(k2_relset_1(X2,X3))|v1_xboole_0(X2)|r1_xxreal_0(X4,esk7_2(k2_relset_1(X2,X3),k2_rinfsup1(X1)))|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X2,X3),k2_rinfsup1(X1)))|X4!=k3_seq_4(X2)|~v3_xxreal_2(k2_relset_1(X2,X3))|~v3_xxreal_2(X2)|~v3_membered(k2_relset_1(X2,X3))|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3)|~v1_xreal_0(X4)|~v1_xreal_0(k2_rinfsup1(X1))|k2_relset_1(X2,X3)!=k10_xtuple_0(X1)|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8403, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X1,esk2_0),k2_rinfsup1(esk2_0)))|k3_seq_4(k2_relset_1(X2,esk2_0))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_funct_2(esk2_0,k5_numbers,X2)|k2_relset_1(X1,esk2_0)!=k10_xtuple_0(esk2_0))).#trainneg
cnf(i_0_9202, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8012, negated_conjecture, (k2_xcmplx_0(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(k10_xtuple_0(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2))))!=k1_numbers|~v1_membered(X2)|~v1_funct_1(k10_xtuple_0(k2_zfmisc_1(X3,k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1153, plain, (m1_subset_1(k1_funct_1(X1,esk9_3(X1,X2,X3)),X4)|~v4_relat_1(X5,X4)|~v1_relat_1(X5)|k9_xtuple_0(X5)!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_7750, negated_conjecture, (v7_ordinal1(k1_funct_1(esk2_0,X1))|k6_numbers=k1_numbers|~v6_membered(k10_xtuple_0(esk2_0))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_6109, plain, (r1_xxreal_0(k9_binop_2(k6_numbers,X1),X2)|r1_xxreal_0(X2,k9_binop_2(k6_numbers,X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6110, plain, (r1_xxreal_0(k9_binop_2(k6_numbers,X1),X2)|r1_xxreal_0(X2,k9_binop_2(k6_numbers,X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1091, plain, (m1_subset_1(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)),X4)|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|X2!=k3_seq_4(k9_xtuple_0(X1))|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1092, plain, (v1_xreal_0(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)))|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|X4!=k10_xtuple_0(X5)|~v3_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X2!=k3_seq_4(k9_xtuple_0(X1))|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9165, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(X2)))|r1_xxreal_0(esk5_1(esk5_1(X2)),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk5_1(X2),k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9547, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~r2_hidden(esk7_2(k2_relset_1(X1,esk2_0),k2_rinfsup1(esk2_0)),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_9544, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(k2_relset_1(X2,esk2_0))|r1_xxreal_0(k3_seq_4(k2_relset_1(X1,esk2_0)),k6_numbers)|v1_xboole_0(X2)|k3_seq_4(k2_relset_1(X2,esk2_0))!=k6_numbers|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(k3_seq_4(k2_relset_1(X1,esk2_0))))).#trainneg
cnf(i_0_9505, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_9504, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8036, plain, (v7_ordinal1(k10_xtuple_0(k2_zfmisc_1(X2,X1)))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8267, negated_conjecture, (k2_xcmplx_0(k6_numbers,k2_rinfsup1(esk2_0))=k2_rinfsup1(esk2_0)|k10_xtuple_0(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))!=k1_numbers|~v1_funct_1(k9_xtuple_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)))|~v1_membered(X2))).#trainneg
cnf(i_0_1154, plain, (m1_subset_1(k1_funct_1(X1,esk9_3(X1,X2,X3)),k1_numbers)|k10_xtuple_0(X1)!=k5_numbers|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9507, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k10_xtuple_0(esk2_0))))|~v1_funct_2(esk2_0,k5_numbers,k10_xtuple_0(esk2_0))|~m1_subset_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k1_numbers))).#trainneg
cnf(i_0_9508, negated_conjecture, (v1_xboole_0(k10_xtuple_0(esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k10_xtuple_0(esk2_0))))|~v1_funct_2(esk2_0,k5_numbers,k10_xtuple_0(esk2_0))|~m1_subset_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k5_numbers))).#trainneg
cnf(i_0_6420, plain, (r1_xxreal_0(k9_binop_2(X1,k6_numbers),X2)|r1_xxreal_0(X2,k9_binop_2(X1,k6_numbers))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6421, plain, (r1_xxreal_0(k9_binop_2(X1,k6_numbers),X2)|r1_xxreal_0(X2,k9_binop_2(X1,k6_numbers))|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1093, plain, (v1_xxreal_0(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)))|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|X4!=k10_xtuple_0(X5)|~v2_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X2!=k3_seq_4(k9_xtuple_0(X1))|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1094, plain, (v1_xcmplx_0(k1_funct_1(X1,esk6_3(k9_xtuple_0(X1),X2,X3)))|v1_xboole_0(k9_xtuple_0(X1))|r1_xxreal_0(X3,k6_numbers)|X4!=k10_xtuple_0(X5)|~v1_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X2!=k3_seq_4(k9_xtuple_0(X1))|X4!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v3_xxreal_2(k9_xtuple_0(X1))|~v3_membered(k9_xtuple_0(X1))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9540, negated_conjecture, (v1_xboole_0(X1)|r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|k3_seq_4(X1)!=k6_numbers|~r2_hidden(esk6_3(X1,k6_numbers,k2_rinfsup1(esk2_0)),k10_xtuple_0(esk2_0))|~v3_xxreal_2(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_9977, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))))|r1_xxreal_0(esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_9975, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))))|r1_xxreal_0(esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_9542, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k3_seq_4(k2_relset_1(X1,esk2_0)),k6_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|k3_seq_4(X2)!=k6_numbers|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_xxreal_2(X2)|~v3_membered(k2_relset_1(X1,esk2_0))|~v3_membered(X2)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk6_3(X2,k6_numbers,k3_seq_4(k2_relset_1(X1,esk2_0))),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_9550, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(k10_xtuple_0(esk2_0))|k3_seq_4(k2_relset_1(X1,esk2_0))!=k2_rinfsup1(esk2_0)|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk7_2(k2_relset_1(X2,esk2_0),k3_seq_4(k2_relset_1(X1,esk2_0))),k10_xtuple_0(esk2_0)))).#trainneg
cnf(i_0_9976, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))))|r1_xxreal_0(esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_valued_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_9475, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_valued_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_8290, plain, (v1_xreal_0(k9_xtuple_0(k9_xtuple_0(X3)))|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8041, plain, (v1_xxreal_0(k10_xtuple_0(k2_zfmisc_1(X2,X1)))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4569, negated_conjecture, (m1_subset_1(k6_numbers,k1_zfmisc_1(k1_numbers))|k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_1156, plain, (m1_subset_1(k1_funct_1(X1,esk9_3(X1,X2,X3)),X4)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5026, negated_conjecture, (k2_xcmplx_0(X1,k5_numbers)=k2_xcmplx_0(k5_numbers,X1)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v1_xcmplx_0(X1)|k10_xtuple_0(esk2_0)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_1097, plain, (k2_xcmplx_0(X1,k1_funct_1(X2,esk6_3(k9_xtuple_0(X2),X3,X4)))=k2_xcmplx_0(k1_funct_1(X2,esk6_3(k9_xtuple_0(X2),X3,X4)),X1)|v1_xboole_0(k9_xtuple_0(X2))|r1_xxreal_0(X4,k6_numbers)|X5!=k10_xtuple_0(X6)|~v1_valued_0(X6)|~v1_xcmplx_0(X1)|~v1_relat_1(X6)|~v1_funct_1(X6)|X3!=k3_seq_4(k9_xtuple_0(X2))|X5!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v3_xxreal_2(k9_xtuple_0(X2))|~v3_membered(k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_1100, plain, (v3_membered(k9_xtuple_0(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5)))|v1_xboole_0(X1)|~v1_relat_1(k3_funct_2(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)),X4,X5))|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))|~m1_subset_1(X5,X1)|~v1_funct_2(X4,X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4))).#trainneg
cnf(i_0_9554, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=esk1_0|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k1_numbers))).#trainneg
cnf(i_0_9476, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_9658, negated_conjecture, (r1_xxreal_0(esk7_2(k10_xtuple_0(esk2_0),esk1_0),esk5_1(X1))|r1_xxreal_0(esk5_1(X1),esk7_2(k10_xtuple_0(esk2_0),esk1_0))|r1_xxreal_0(X1,k6_numbers)|v1_xboole_0(k10_xtuple_0(esk2_0))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9555, negated_conjecture, (v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(k10_xtuple_0(esk2_0))|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=esk1_0|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~m1_subset_1(esk8_2(k10_xtuple_0(esk2_0),esk1_0),k5_numbers))).#trainneg
cnf(i_0_9971, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))))|r1_xxreal_0(esk5_1(esk5_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,X2),k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_10132, plain, (m1_subset_1(X3,k1_numbers)|k10_xtuple_0(X1)!=k5_numbers|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10074, plain, (v1_xreal_0(k10_xtuple_0(k2_zfmisc_1(X1,X2)))|~v6_membered(k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1157, plain, (v1_xreal_0(k1_funct_1(X1,esk9_3(X1,X2,X3)))|X4!=k10_xtuple_0(X5)|~v3_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10134, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k6_numbers)))|~m1_subset_1(esk8_2(k6_numbers,esk1_0),k5_numbers)|~v1_funct_2(esk2_0,k5_numbers,k6_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_10011, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k2_rinfsup1(esk2_0)|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),k2_rinfsup1(esk2_0)),k6_numbers))).#trainneg
cnf(i_0_10072, negated_conjecture, (r1_xxreal_0(esk5_1(k2_rinfsup1(esk2_0)),esk5_1(X1))|r1_xxreal_0(esk5_1(X1),esk5_1(k2_rinfsup1(esk2_0)))|r1_xxreal_0(k2_rinfsup1(esk2_0),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9474, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)))|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_9579, negated_conjecture, (r1_xxreal_0(k6_numbers,esk8_2(k2_relset_1(X1,esk2_0),esk1_0))|r1_xxreal_0(esk8_2(k2_relset_1(X1,esk2_0),esk1_0),k6_numbers)|k3_seq_4(k2_relset_1(X1,esk2_0))=esk1_0|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_9694, negated_conjecture, (r1_xxreal_0(k3_seq_4(X1),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),k3_seq_4(X1))|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_8780, negated_conjecture, (r1_xxreal_0(esk5_1(X1),k5_seq_4(k5_numbers))|r1_xxreal_0(k5_seq_4(k5_numbers),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9203, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_8872, negated_conjecture, (k2_xcmplx_0(esk7_2(k2_relset_1(X1,esk2_0),X2),X3)=k2_xcmplx_0(X3,esk7_2(k2_relset_1(X1,esk2_0),X2))|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~v1_valued_0(esk2_0)|~v1_xreal_0(X3)|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10433, plain, (m1_subset_1(X3,X4)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6254, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k3_seq_4(k1_numbers))|r1_xxreal_0(k3_seq_4(k1_numbers),k2_rinfsup1(esk2_0)))).#trainneg
cnf(i_0_9582, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|k9_binop_2(X1,esk1_0)!=k2_rinfsup1(esk2_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1158, plain, (v1_xxreal_0(k1_funct_1(X1,esk9_3(X1,X2,X3)))|X4!=k10_xtuple_0(X5)|~v2_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_8092, negated_conjecture, (r1_xxreal_0(k9_binop_2(esk1_0,k6_numbers),k9_binop_2(esk1_0,k6_numbers))|~v6_membered(k1_numbers)|k10_xtuple_0(esk2_0)!=k1_numbers)).#trainneg
cnf(i_0_1379, negated_conjecture, (r1_xxreal_0(X1,k9_xtuple_0(esk2_0))|r1_xxreal_0(k9_xtuple_0(esk2_0),X1)|~v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1125, plain, (v1_xboole_0(X1)|v1_xreal_0(k3_funct_2(X2,X1,X3,X4))|v1_xboole_0(X2)|X1!=k10_xtuple_0(X5)|~v3_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(X4,X2)|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_9955, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(esk12_3(X2,X3,X4))))|r1_xxreal_0(esk5_1(esk5_1(esk12_3(X2,X3,X4))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(esk12_3(X2,X3,X4)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk12_3(X2,X3,X4),k6_numbers)|v1_xboole_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X2,k2_relset_1(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X4,k5_numbers,X3)|~v1_funct_1(X4))).#trainneg
cnf(i_0_9501, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(k1_seq_1(X2,X3)))|r1_xxreal_0(esk5_1(k1_seq_1(X2,X3)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k1_seq_1(X2,X3),k6_numbers)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_9967, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(esk5_1(X2))))|r1_xxreal_0(esk5_1(esk5_1(esk5_1(X2))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(esk5_1(X2)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(esk5_1(X2),k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10004, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(k2_rinfsup1(X2))))|r1_xxreal_0(esk5_1(esk5_1(k2_rinfsup1(X2))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(k2_rinfsup1(X2)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k2_rinfsup1(X2),k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_9057, negated_conjecture, (k3_seq_4(k2_relset_1(X1,esk2_0))=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X3,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_seq_4(k2_relset_1(X1,esk2_0))!=k3_seq_4(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X3,esk2_0))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X3,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_funct_2(esk2_0,k5_numbers,X3)|~v1_xreal_0(k3_seq_4(k2_relset_1(X1,esk2_0)))|~r2_hidden(esk7_2(k2_relset_1(X2,esk2_0),k3_seq_4(k2_relset_1(X1,esk2_0))),k10_xtuple_0(esk2_0))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1))).#trainneg
cnf(i_0_9974, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))))|r1_xxreal_0(esk5_1(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X2,X3,X4)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X2,X3,X4),k6_numbers)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(k2_relset_1(X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k5_numbers)|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_10002, negated_conjecture, (r1_xxreal_0(esk5_1(esk5_1(X1)),esk5_1(esk5_1(k1_seq_1(X2,X3))))|r1_xxreal_0(esk5_1(esk5_1(k1_seq_1(X2,X3))),esk5_1(esk5_1(X1)))|r1_xxreal_0(esk5_1(X1),k6_numbers)|r1_xxreal_0(esk5_1(k1_seq_1(X2,X3)),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k1_seq_1(X2,X3),k6_numbers)|~v1_xreal_0(X1)|~v3_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_10385, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_funct_1(X3,esk9_3(X3,X4,X5)))|k1_zfmisc_1(X2)!=k10_xtuple_0(X3)|X4!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~r2_hidden(X5,X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_6296, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(esk2_0),k3_seq_4(k5_numbers))|r1_xxreal_0(k3_seq_4(k5_numbers),k2_rinfsup1(esk2_0)))).#trainneg
cnf(i_0_1159, plain, (v1_xcmplx_0(k1_funct_1(X1,esk9_3(X1,X2,X3)))|X4!=k10_xtuple_0(X5)|~v1_valued_0(X5)|~v1_relat_1(X5)|~v1_funct_1(X5)|X4!=k10_xtuple_0(X1)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_6734, negated_conjecture, (r1_xxreal_0(esk10_2(esk2_0,k5_numbers),esk1_0)|r1_xxreal_0(esk1_0,esk10_2(esk2_0,k5_numbers))|k10_xtuple_0(esk2_0)=k5_numbers)).#trainneg
cnf(i_0_5008, negated_conjecture, (k2_rinfsup1(esk2_0)=esk1_0|r1_xxreal_0(esk1_0,k6_numbers)|~m1_subset_1(esk8_2(k6_numbers,esk1_0),k5_numbers)|k2_rinfsup1(esk2_0)!=k6_numbers)).#trainneg
cnf(i_0_1807, negated_conjecture, (r1_xxreal_0(X1,k10_xtuple_0(esk2_0))|r1_xxreal_0(k10_xtuple_0(esk2_0),X1)|~v1_xreal_0(X1)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1145, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X2,esk9_3(X2,X3,X4)),X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k1_zfmisc_1(X5))))|~m1_subset_1(X7,X1)|~v1_funct_2(X6,X1,k1_zfmisc_1(X5))|~v1_funct_1(X6)|k3_funct_2(X1,k1_zfmisc_1(X5),X6,X7)!=k10_xtuple_0(X2)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~r2_hidden(X4,X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1152, plain, (r2_hidden(X1,X2)|X1!=k1_funct_1(X3,k1_funct_1(X4,esk9_3(X4,X5,X6)))|X2!=k10_xtuple_0(X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|k9_xtuple_0(X3)!=k10_xtuple_0(X4)|X5!=k10_xtuple_0(X4)|~v1_relat_1(X4)|~r2_hidden(X6,X5)|~v1_funct_1(X4))).#trainneg
cnf(i_0_9503, negated_conjecture, (r1_xxreal_0(esk5_1(X1),esk5_1(k2_rinfsup1(X2)))|r1_xxreal_0(esk5_1(k2_rinfsup1(X2)),esk5_1(X1))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k2_rinfsup1(X2),k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_8871, negated_conjecture, (m1_subset_1(esk7_2(k2_relset_1(X1,esk2_0),X2),k1_numbers)|X2=k3_seq_4(k2_relset_1(X1,esk2_0))|v1_xboole_0(k2_relset_1(X1,esk2_0))|v1_xboole_0(X1)|v1_xreal_0(esk8_2(k2_relset_1(X1,esk2_0),X2))|~v3_xxreal_2(k2_relset_1(X1,esk2_0))|~v3_membered(k2_relset_1(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(esk2_0,k5_numbers,X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8983, negated_conjecture, (r1_xxreal_0(k2_rinfsup1(X1),esk5_1(X2))|r1_xxreal_0(X2,k6_numbers)|k10_xtuple_0(X1)!=k1_numbers|~v2_seq_2(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(esk5_1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_9286, negated_conjecture, (r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,esk5_1(X1))),esk5_1(X2))|r1_xxreal_0(esk5_1(X2),esk5_1(k3_funct_2(k5_numbers,k1_numbers,esk2_0,esk5_1(X1))))|r1_xxreal_0(k3_funct_2(k5_numbers,k1_numbers,esk2_0,esk5_1(X1)),k6_numbers)|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8662, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k3_seq_4(k2_relset_1(X2,X3)),esk7_2(k2_relset_1(X1,X3),X4))|X4=k3_seq_4(k2_relset_1(X1,X3))|v1_xboole_0(k2_relset_1(X1,X3))|~v3_xxreal_2(k2_relset_1(X2,X3))|~v3_membered(k2_relset_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X3,k5_numbers,X2)|~v1_funct_2(X3,k5_numbers,X1)|~v1_funct_1(X3)|~v1_xreal_0(esk7_2(k2_relset_1(X1,X3),X4))|~v3_xxreal_2(k2_relset_1(X1,X3))|~v3_membered(k2_relset_1(X1,X3))|~r1_xxreal_0(esk8_2(k2_relset_1(X1,X3),X4),k6_numbers)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_8857, negated_conjecture, (X1=k2_rinfsup1(esk2_0)|r2_hidden(esk7_2(k10_xtuple_0(esk2_0),X1),k10_xtuple_0(esk2_0))|r1_xxreal_0(k9_binop_2(X1,esk8_2(k10_xtuple_0(esk2_0),X1)),esk7_2(k2_relset_1(X2,esk2_0),X3))|X3=k3_seq_4(k2_relset_1(X2,esk2_0))|v1_xboole_0(k2_relset_1(X2,esk2_0))|v1_xboole_0(X2)|v1_xreal_0(esk8_2(k2_relset_1(X2,esk2_0),X3))|~v1_xreal_0(X1)|~v1_xreal_0(esk7_2(k2_relset_1(X2,esk2_0),X3))|~v3_xxreal_2(k2_relset_1(X2,esk2_0))|~v3_membered(k2_relset_1(X2,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(esk2_0,k5_numbers,X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_9738, negated_conjecture, (r1_xxreal_0(esk5_1(k3_funct_2(k5_numbers,X1,X2,X3)),esk5_1(esk5_1(X4)))|r1_xxreal_0(esk5_1(esk5_1(X4)),esk5_1(k3_funct_2(k5_numbers,X1,X2,X3)))|r1_xxreal_0(esk5_1(X4),k6_numbers)|r1_xxreal_0(k3_funct_2(k5_numbers,X1,X2,X3),k6_numbers)|r1_xxreal_0(X4,k6_numbers)|v1_xboole_0(X1)|~v1_xreal_0(X4)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 75
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 120
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 117
# Processed clauses                    : 2458
# ...of these trivial                  : 45
# ...subsumed                          : 950
# ...remaining for further processing  : 1463
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 151
# Backward-rewritten                   : 106
# Generated clauses                    : 10972
# ...of the previous two non-trivial   : 10477
# Contextual simplify-reflections      : 196
# Paramodulations                      : 10873
# Factorizations                       : 4
# Equation resolutions                 : 89
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1200
#    Positive orientable unit clauses  : 60
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 1135
# Current number of unprocessed clauses: 7892
# ...number of literals in the above   : 72078
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1726
# Proof object given clauses           : 152
# Proof search given clauses           : 1463
# Clause-clause subsumption calls (NU) : 733595
# Rec. Clause-clause subsumption calls : 26237
# Non-unit clause-clause subsumptions  : 1292
# Unit Clause-clause subsumption calls : 9430
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 32
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 353040

# -------------------------------------------------
# User time                : 30.389 s
# System time              : 1.057 s
# Total time               : 31.446 s
# Maximum resident set size: 409888 pages
30.38user 1.15system 0:31.69elapsed 99%CPU (0avgtext+0avgdata 409888maxresident)k
0inputs+832outputs (0major+304048minor)pagefaults 0swaps
