# ENIGMA: Connected to the TF server '127.0.0.1:8888'.
# Preprocessing time       : 0.020 s
#TF#SERVER: Sending data to the TF server (conj=17, context=73, query=56, total=146).
#TF#SERVER: Sent 94321 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (520 bytes expected).
#TF#SERVER: Received 520 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=107, query=132, total=256).
#TF#SERVER: Sent 234306 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (960 bytes expected).
#TF#SERVER: Received 960 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=155, query=128, total=300).
#TF#SERVER: Sent 277423 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1136 bytes expected).
#TF#SERVER: Received 1136 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=173, query=138, total=328).
#TF#SERVER: Sent 366703 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1248 bytes expected).
#TF#SERVER: Received 1248 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=202, query=133, total=352).
#TF#SERVER: Sent 347753 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1344 bytes expected).
#TF#SERVER: Received 1344 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=216, query=128, total=361).
#TF#SERVER: Sent 414139 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1380 bytes expected).
#TF#SERVER: Received 1380 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=216, query=76, total=309).
#TF#SERVER: Sent 314162 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1172 bytes expected).
#TF#SERVER: Received 1172 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=230, query=180, total=427).
#TF#SERVER: Sent 591025 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1644 bytes expected).
#TF#SERVER: Received 1644 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=247, query=128, total=392).
#TF#SERVER: Sent 475229 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1504 bytes expected).
#TF#SERVER: Received 1504 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=250, query=130, total=397).
#TF#SERVER: Sent 542328 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1524 bytes expected).
#TF#SERVER: Received 1524 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=265, query=143, total=425).
#TF#SERVER: Sent 419245 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1636 bytes expected).
#TF#SERVER: Received 1636 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=286, query=130, total=433).
#TF#SERVER: Sent 426915 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1668 bytes expected).
#TF#SERVER: Received 1668 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=300, query=130, total=447).
#TF#SERVER: Sent 495346 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1724 bytes expected).
#TF#SERVER: Received 1724 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=312, query=128, total=457).
#TF#SERVER: Sent 570203 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1764 bytes expected).
#TF#SERVER: Received 1764 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=321, query=128, total=466).
#TF#SERVER: Sent 600822 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1800 bytes expected).
#TF#SERVER: Received 1800 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=330, query=135, total=482).
#TF#SERVER: Sent 624825 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1864 bytes expected).
#TF#SERVER: Received 1864 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=340, query=141, total=498).
#TF#SERVER: Sent 646979 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1928 bytes expected).
#TF#SERVER: Received 1928 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=349, query=128, total=494).
#TF#SERVER: Sent 577876 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1912 bytes expected).
#TF#SERVER: Received 1912 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=349, query=128, total=494).
#TF#SERVER: Sent 653189 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1912 bytes expected).
#TF#SERVER: Received 1912 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=349, query=19, total=385).
#TF#SERVER: Sent 342879 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1476 bytes expected).
#TF#SERVER: Received 1476 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=353, query=128, total=498).
#TF#SERVER: Sent 642020 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1928 bytes expected).
#TF#SERVER: Received 1928 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=353, query=96, total=466).
#TF#SERVER: Sent 583054 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1800 bytes expected).
#TF#SERVER: Received 1800 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=356, query=128, total=501).
#TF#SERVER: Sent 674317 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1940 bytes expected).
#TF#SERVER: Received 1940 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=356, query=128, total=501).
#TF#SERVER: Sent 667704 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1940 bytes expected).
#TF#SERVER: Received 1940 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=356, query=23, total=396).
#TF#SERVER: Sent 366469 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1520 bytes expected).
#TF#SERVER: Received 1520 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=360, query=128, total=505).
#TF#SERVER: Sent 694707 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1956 bytes expected).
#TF#SERVER: Received 1956 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=360, query=128, total=505).
#TF#SERVER: Sent 687303 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1956 bytes expected).
#TF#SERVER: Received 1956 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=360, query=29, total=406).
#TF#SERVER: Sent 394110 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1560 bytes expected).
#TF#SERVER: Received 1560 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=373, query=159, total=549).
#TF#SERVER: Sent 773027 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2132 bytes expected).
#TF#SERVER: Received 2132 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=379, query=139, total=535).
#TF#SERVER: Sent 694903 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2076 bytes expected).
#TF#SERVER: Received 2076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=389, query=139, total=545).
#TF#SERVER: Sent 701166 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2116 bytes expected).
#TF#SERVER: Received 2116 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=401, query=151, total=569).
#TF#SERVER: Sent 757914 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2212 bytes expected).
#TF#SERVER: Received 2212 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=412, query=176, total=605).
#TF#SERVER: Sent 910470 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2356 bytes expected).
#TF#SERVER: Received 2356 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=417, query=128, total=562).
#TF#SERVER: Sent 744104 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2184 bytes expected).
#TF#SERVER: Received 2184 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=417, query=128, total=562).
#TF#SERVER: Sent 784357 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2184 bytes expected).
#TF#SERVER: Received 2184 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=417, query=107, total=541).
#TF#SERVER: Sent 721060 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2100 bytes expected).
#TF#SERVER: Received 2100 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=421, query=128, total=566).
#TF#SERVER: Sent 791838 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2200 bytes expected).
#TF#SERVER: Received 2200 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=421, query=128, total=566).
#TF#SERVER: Sent 792725 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2200 bytes expected).
#TF#SERVER: Received 2200 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=421, query=53, total=491).
#TF#SERVER: Sent 561080 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1900 bytes expected).
#TF#SERVER: Received 1900 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=434, query=130, total=581).
#TF#SERVER: Sent 794774 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2260 bytes expected).
#TF#SERVER: Received 2260 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=446, query=186, total=649).
#TF#SERVER: Sent 988409 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2532 bytes expected).
#TF#SERVER: Received 2532 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=451, query=185, total=653).
#TF#SERVER: Sent 1074725 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2548 bytes expected).
#TF#SERVER: Received 2548 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=457, query=139, total=613).
#TF#SERVER: Sent 892225 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2388 bytes expected).
#TF#SERVER: Received 2388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=463, query=162, total=642).
#TF#SERVER: Sent 1002487 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2504 bytes expected).
#TF#SERVER: Received 2504 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=476, query=131, total=624).
#TF#SERVER: Sent 798250 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2432 bytes expected).
#TF#SERVER: Received 2432 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=487, query=157, total=661).
#TF#SERVER: Sent 954788 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2580 bytes expected).
#TF#SERVER: Received 2580 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=497, query=135, total=649).
#TF#SERVER: Sent 902484 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2532 bytes expected).
#TF#SERVER: Received 2532 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=511, query=155, total=683).
#TF#SERVER: Sent 1019900 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2668 bytes expected).
#TF#SERVER: Received 2668 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=142, total=671).
#TF#SERVER: Sent 993762 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2620 bytes expected).
#TF#SERVER: Received 2620 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 997150 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=79, total=608).
#TF#SERVER: Sent 816037 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2368 bytes expected).
#TF#SERVER: Received 2368 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=136, total=665).
#TF#SERVER: Sent 1024522 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2596 bytes expected).
#TF#SERVER: Received 2596 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=158, total=687).
#TF#SERVER: Sent 1116595 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2684 bytes expected).
#TF#SERVER: Received 2684 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=129, total=658).
#TF#SERVER: Sent 940744 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=134, total=663).
#TF#SERVER: Sent 946158 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2588 bytes expected).
#TF#SERVER: Received 2588 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 977069 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=76, total=605).
#TF#SERVER: Sent 794120 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2356 bytes expected).
#TF#SERVER: Received 2356 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=187, total=716).
#TF#SERVER: Sent 1193882 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2800 bytes expected).
#TF#SERVER: Received 2800 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=159, total=688).
#TF#SERVER: Sent 963480 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2688 bytes expected).
#TF#SERVER: Received 2688 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 909098 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 981604 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=6, total=535).
#TF#SERVER: Sent 562908 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2076 bytes expected).
#TF#SERVER: Received 2076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 996720 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 1082737 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=15, total=544).
#TF#SERVER: Sent 595400 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2112 bytes expected).
#TF#SERVER: Received 2112 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 808953 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 977373 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=27, total=556).
#TF#SERVER: Sent 632743 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2160 bytes expected).
#TF#SERVER: Received 2160 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 944008 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=98, total=627).
#TF#SERVER: Sent 917377 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2444 bytes expected).
#TF#SERVER: Received 2444 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=128, total=657).
#TF#SERVER: Sent 921662 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=121, total=650).
#TF#SERVER: Sent 1014661 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2536 bytes expected).
#TF#SERVER: Received 2536 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=17, context=512, query=146, total=675).
#TF#SERVER: Sent 1135867 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2636 bytes expected).
#TF#SERVER: Received 2636 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t92_tsep_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(r3_tsep_1(X1,X2,X3)<=>?[X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))&?[X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))&(r4_tsep_1(X1,X4,X5)&(m1_pre_topc(X2,X4)&(m1_pre_topc(X3,X5)&(r1_tsep_1(X4,X5)|r1_tsep_1(k2_tsep_1(X1,X4,X5),k1_tsep_1(X1,X2,X3)))))))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t92_tsep_1)).
fof(t77_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(r3_tsep_1(X1,X2,X3)<=>?[X4]:((~(v2_struct_0(X4))&(v1_tsep_1(X4,X1)&m1_pre_topc(X4,X1)))&?[X5]:((~(v2_struct_0(X5))&(v1_tsep_1(X5,X1)&m1_pre_topc(X5,X1)))&(m1_pre_topc(X2,X4)&(m1_pre_topc(X3,X5)&(r1_tsep_1(X4,X5)|r1_tsep_1(k2_tsep_1(X1,X4,X5),k1_tsep_1(X1,X2,X3))))))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t77_tsep_1)).
fof(t81_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))=>![X3]:((v1_tsep_1(X3,X1)&m1_pre_topc(X3,X1))=>r4_tsep_1(X1,X2,X3)))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t81_tsep_1)).
fof(d4_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k2_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', d4_tsep_1)).
fof(dt_k2_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k2_tsep_1(X1,X2,X3)))&(v1_pre_topc(k2_tsep_1(X1,X2,X3))&m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', dt_k2_tsep_1)).
fof(d3_tsep_1, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>(r1_tsep_1(X1,X2)<=>r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', d3_tsep_1)).
fof(symmetry_r1_tsep_1, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_struct_0(X2))=>(r1_tsep_1(X1,X2)=>r1_tsep_1(X2,X1))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', symmetry_r1_tsep_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', symmetry_r1_xboole_0)).
fof(t62_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_connsp_1(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&?[X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))&(r2_tsep_1(X1,X4,X5)&(r1_tarski(X2,X4)&(r1_tarski(X3,X5)&r1_xboole_0(k9_subset_1(u1_struct_0(X1),X4,X5),k4_subset_1(u1_struct_0(X1),X2,X3))))))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t62_tsep_1)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', redefinition_k4_subset_1)).
fof(d2_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k1_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', d2_tsep_1)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', redefinition_k9_subset_1)).
fof(dt_k1_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k1_tsep_1(X1,X2,X3)))&(v1_pre_topc(k1_tsep_1(X1,X2,X3))&m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', dt_k1_tsep_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', dt_l1_pre_topc)).
fof(d7_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r4_tsep_1(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>((X4=u1_struct_0(X2)&X5=u1_struct_0(X3))=>r2_tsep_1(X1,X4,X5)))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', d7_tsep_1)).
fof(d6_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r3_tsep_1(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>((X4=u1_struct_0(X2)&X5=u1_struct_0(X3))=>r1_connsp_1(X1,X4,X5)))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', d6_tsep_1)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t4_tsep_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t1_tsep_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', dt_m1_pre_topc)).
fof(t78_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>((r1_tsep_1(X2,X3)&r4_tsep_1(X1,X2,X3))<=>r3_tsep_1(X1,X2,X3))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t78_tsep_1)).
fof(t7_connsp_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>((r1_connsp_1(X1,X2,X3)&(r1_tarski(X4,X2)&r1_tarski(X5,X3)))=>r1_connsp_1(X1,X4,X5))))))), file('/scratch/mptp/predsb1rnn_cl/GeneralTopology/t92_tsep_1__3', t7_connsp_1)).
fof(c_0_21, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(r3_tsep_1(X1,X2,X3)<=>?[X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))&?[X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))&(r4_tsep_1(X1,X4,X5)&(m1_pre_topc(X2,X4)&(m1_pre_topc(X3,X5)&(r1_tsep_1(X4,X5)|r1_tsep_1(k2_tsep_1(X1,X4,X5),k1_tsep_1(X1,X2,X3))))))))))))), inference(assume_negation,[status(cth)],[t92_tsep_1])).
fof(c_0_22, negated_conjecture, ![X9, X10]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&(((~r1_tsep_1(X9,X10)|~m1_pre_topc(esk3_0,X10)|~m1_pre_topc(esk2_0,X9)|~r4_tsep_1(esk1_0,X9,X10)|(v2_struct_0(X10)|~m1_pre_topc(X10,esk1_0))|(v2_struct_0(X9)|~m1_pre_topc(X9,esk1_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0))&(~r1_tsep_1(k2_tsep_1(esk1_0,X9,X10),k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X10)|~m1_pre_topc(esk2_0,X9)|~r4_tsep_1(esk1_0,X9,X10)|(v2_struct_0(X10)|~m1_pre_topc(X10,esk1_0))|(v2_struct_0(X9)|~m1_pre_topc(X9,esk1_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)))&(((~v2_struct_0(esk4_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))&(m1_pre_topc(esk4_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)))&(((~v2_struct_0(esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))&(m1_pre_topc(esk5_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)))&((r4_tsep_1(esk1_0,esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))&((m1_pre_topc(esk2_0,esk4_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))&((m1_pre_topc(esk3_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))&(r1_tsep_1(esk4_0,esk5_0)|r1_tsep_1(k2_tsep_1(esk1_0,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0))|r3_tsep_1(esk1_0,esk2_0,esk3_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_21])])])])])])).
fof(c_0_23, plain, ![X68, X69, X70, X73, X74]:((((~v2_struct_0(esk12_3(X68,X69,X70))|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&((v1_tsep_1(esk12_3(X68,X69,X70),X68)|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&(m1_pre_topc(esk12_3(X68,X69,X70),X68)|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))))&(((~v2_struct_0(esk13_3(X68,X69,X70))|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&((v1_tsep_1(esk13_3(X68,X69,X70),X68)|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&(m1_pre_topc(esk13_3(X68,X69,X70),X68)|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))))&((m1_pre_topc(X69,esk12_3(X68,X69,X70))|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&((m1_pre_topc(X70,esk13_3(X68,X69,X70))|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&(r1_tsep_1(esk12_3(X68,X69,X70),esk13_3(X68,X69,X70))|r1_tsep_1(k2_tsep_1(X68,esk12_3(X68,X69,X70),esk13_3(X68,X69,X70)),k1_tsep_1(X68,X69,X70))|~r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))))))&((~r1_tsep_1(X73,X74)|~m1_pre_topc(X70,X74)|~m1_pre_topc(X69,X73)|(v2_struct_0(X74)|(~v1_tsep_1(X74,X68)|~m1_pre_topc(X74,X68)))|(v2_struct_0(X73)|(~v1_tsep_1(X73,X68)|~m1_pre_topc(X73,X68)))|r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68))))&(~r1_tsep_1(k2_tsep_1(X68,X73,X74),k1_tsep_1(X68,X69,X70))|~m1_pre_topc(X70,X74)|~m1_pre_topc(X69,X73)|(v2_struct_0(X74)|(~v1_tsep_1(X74,X68)|~m1_pre_topc(X74,X68)))|(v2_struct_0(X73)|(~v1_tsep_1(X73,X68)|~m1_pre_topc(X73,X68)))|r3_tsep_1(X68,X69,X70)|(v2_struct_0(X70)|~m1_pre_topc(X70,X68))|(v2_struct_0(X69)|~m1_pre_topc(X69,X68))|(v2_struct_0(X68)|(~v2_pre_topc(X68)|~l1_pre_topc(X68)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t77_tsep_1])])])])])])).
fof(c_0_24, plain, ![X83, X84, X85]:(v2_struct_0(X83)|(~v2_pre_topc(X83)|~l1_pre_topc(X83))|(~v1_tsep_1(X84,X83)|~m1_pre_topc(X84,X83)|(~v1_tsep_1(X85,X83)|~m1_pre_topc(X85,X83)|r4_tsep_1(X83,X84,X85)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t81_tsep_1])])])])).
cnf(c_0_25, negated_conjecture, (v2_struct_0(X2)|v2_struct_0(X1)|~r1_tsep_1(k2_tsep_1(esk1_0,X1,X2),k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X2)|~m1_pre_topc(esk2_0,X1)|~r4_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|r1_tsep_1(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)),k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, plain, (v2_struct_0(X1)|r4_tsep_1(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tsep_1(X2,X1)|~m1_pre_topc(X2,X1)|~v1_tsep_1(X3,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (v1_tsep_1(esk13_3(X1,X2,X3),X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, negated_conjecture, (v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X2)|~r4_tsep_1(esk1_0,X1,X2)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r1_tsep_1(k2_tsep_1(esk1_0,X1,X2),k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|r1_tsep_1(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)),k1_tsep_1(X1,X2,X3))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_26])).
cnf(c_0_38, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_27])).
cnf(c_0_39, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_28])).
cnf(c_0_40, negated_conjecture, (l1_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_29])).
cnf(c_0_41, negated_conjecture, (v2_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_30])).
cnf(c_0_42, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_31])).
cnf(c_0_43, negated_conjecture, (~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_32])).
cnf(c_0_44, negated_conjecture, (~v2_struct_0(esk3_0)), inference(evalgc,[status(thm)],[c_0_33])).
cnf(c_0_45, negated_conjecture, (v2_struct_0(X2)|v2_struct_0(X1)|~r1_tsep_1(X1,X2)|~m1_pre_topc(esk3_0,X2)|~m1_pre_topc(esk2_0,X1)|~r4_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_46, plain, (v2_struct_0(X1)|r4_tsep_1(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v1_tsep_1(X3,X1)|~v1_tsep_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_tsep_1(esk13_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (m1_pre_topc(esk13_3(X1,X2,X3),X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_49, plain, ![X19, X20, X21, X22]:((X22!=k2_tsep_1(X19,X20,X21)|u1_struct_0(X22)=k3_xboole_0(u1_struct_0(X20),u1_struct_0(X21))|(v2_struct_0(X22)|(~v1_pre_topc(X22)|~m1_pre_topc(X22,X19)))|r1_tsep_1(X20,X21)|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|~m1_pre_topc(X20,X19))|(v2_struct_0(X19)|~l1_pre_topc(X19)))&(u1_struct_0(X22)!=k3_xboole_0(u1_struct_0(X20),u1_struct_0(X21))|X22=k2_tsep_1(X19,X20,X21)|(v2_struct_0(X22)|(~v1_pre_topc(X22)|~m1_pre_topc(X22,X19)))|r1_tsep_1(X20,X21)|(v2_struct_0(X21)|~m1_pre_topc(X21,X19))|(v2_struct_0(X20)|~m1_pre_topc(X20,X19))|(v2_struct_0(X19)|~l1_pre_topc(X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tsep_1])])])])])).
cnf(c_0_50, negated_conjecture, (v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk12_3(esk1_0,esk2_0,esk3_0),esk13_3(esk1_0,esk2_0,esk3_0))|~r4_tsep_1(esk1_0,esk12_3(esk1_0,esk2_0,esk3_0),esk13_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_51, negated_conjecture, (v2_struct_0(X2)|v2_struct_0(X1)|~r1_tsep_1(X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X2)|~r4_tsep_1(esk1_0,X1,X2)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (r4_tsep_1(X1,X2,esk13_3(X1,X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_tsep_1(X2,X1)|~m1_pre_topc(esk13_3(X1,X3,X4),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47])])).
cnf(c_0_53, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(esk13_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_54, plain, (u1_struct_0(X1)=k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X1)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k2_tsep_1(X2,X3,X4)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_55, plain, ![X40, X41, X42]:((~v2_struct_0(k2_tsep_1(X40,X41,X42))|(v2_struct_0(X40)|~l1_pre_topc(X40)|(v2_struct_0(X41)|~m1_pre_topc(X41,X40)|(v2_struct_0(X42)|~m1_pre_topc(X42,X40)))))&((v1_pre_topc(k2_tsep_1(X40,X41,X42))|(v2_struct_0(X40)|~l1_pre_topc(X40)|(v2_struct_0(X41)|~m1_pre_topc(X41,X40)|(v2_struct_0(X42)|~m1_pre_topc(X42,X40)))))&(m1_pre_topc(k2_tsep_1(X40,X41,X42),X40)|(v2_struct_0(X40)|~l1_pre_topc(X40)|(v2_struct_0(X41)|~m1_pre_topc(X41,X40)|(v2_struct_0(X42)|~m1_pre_topc(X42,X40))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tsep_1])])])])).
cnf(c_0_56, negated_conjecture, (v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|~r4_tsep_1(esk1_0,esk12_3(esk1_0,esk2_0,esk3_0),esk13_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))), inference(csr,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57, plain, (r4_tsep_1(X1,X2,esk13_3(X1,X3,X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X1)|~v1_tsep_1(X2,X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, plain, (v1_tsep_1(esk12_3(X1,X2,X3),X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_59, plain, ![X17, X18]:((~r1_tsep_1(X17,X18)|r1_xboole_0(u1_struct_0(X17),u1_struct_0(X18))|~l1_struct_0(X18)|~l1_struct_0(X17))&(~r1_xboole_0(u1_struct_0(X17),u1_struct_0(X18))|r1_tsep_1(X17,X18)|~l1_struct_0(X18)|~l1_struct_0(X17))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tsep_1])])])])).
cnf(c_0_60, plain, (u1_struct_0(X1)=k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X3,X4)|X1!=k2_tsep_1(X2,X3,X4)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_61, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62, plain, (v1_pre_topc(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k2_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_64, plain, ![X52, X53]:(~l1_struct_0(X52)|~l1_struct_0(X53)|(~r1_tsep_1(X52,X53)|r1_tsep_1(X53,X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_tsep_1])])).
cnf(c_0_65, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk13_3(X1,X2,X3))|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_66, negated_conjecture, (v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))|~v1_tsep_1(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_44]), c_0_43]), c_0_42])])).
cnf(c_0_67, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_tsep_1(esk12_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|r1_tsep_1(X2,X3)|v2_struct_0(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_60])])).
cnf(c_0_70, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_71, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_pre_topc(k2_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_72, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_struct_0(k2_tsep_1(X1,X2,X3))), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_73, plain, (r1_tsep_1(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~r1_tsep_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_74, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|r1_tsep_1(k2_tsep_1(esk1_0,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0))|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_75, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk12_3(X1,X2,X3))|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_76, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)|~v2_struct_0(esk13_3(X1,X2,X3))), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_77, negated_conjecture, (v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_78, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_79, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_69, c_0_70]), c_0_71]), c_0_72])).
cnf(c_0_80, plain, (r1_tsep_1(X2,X1)|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_81, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|r1_tsep_1(k2_tsep_1(esk1_0,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_82, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)|~v2_struct_0(esk12_3(X1,X2,X3))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_83, negated_conjecture, (v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_42]), c_0_43]), c_0_44])])).
fof(c_0_84, plain, ![X54, X55]:(~r1_xboole_0(X54,X55)|r1_xboole_0(X55,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_85, plain, (r1_xboole_0(u1_struct_0(X1),k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79])])).
cnf(c_0_86, negated_conjecture, (r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),k2_tsep_1(esk1_0,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81])])).
cnf(c_0_87, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_88, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_89, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_90, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_91, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_92, plain, (m1_pre_topc(esk12_3(X1,X2,X3),X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_93, plain, ![X61, X62, X63, X66, X67]:(((m1_subset_1(esk10_3(X61,X62,X63),k1_zfmisc_1(u1_struct_0(X61)))|~r1_connsp_1(X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X61)))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))&((m1_subset_1(esk11_3(X61,X62,X63),k1_zfmisc_1(u1_struct_0(X61)))|~r1_connsp_1(X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X61)))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))&((r2_tsep_1(X61,esk10_3(X61,X62,X63),esk11_3(X61,X62,X63))|~r1_connsp_1(X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X61)))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))&((r1_tarski(X62,esk10_3(X61,X62,X63))|~r1_connsp_1(X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X61)))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))&((r1_tarski(X63,esk11_3(X61,X62,X63))|~r1_connsp_1(X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X61)))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))&(r1_xboole_0(k9_subset_1(u1_struct_0(X61),esk10_3(X61,X62,X63),esk11_3(X61,X62,X63)),k4_subset_1(u1_struct_0(X61),X62,X63))|~r1_connsp_1(X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X61)))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61)))))))))&(~m1_subset_1(X66,k1_zfmisc_1(u1_struct_0(X61)))|(~m1_subset_1(X67,k1_zfmisc_1(u1_struct_0(X61)))|(~r2_tsep_1(X61,X66,X67)|(~r1_tarski(X62,X66)|(~r1_tarski(X63,X67)|~r1_xboole_0(k9_subset_1(u1_struct_0(X61),X66,X67),k4_subset_1(u1_struct_0(X61),X62,X63))))))|r1_connsp_1(X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(u1_struct_0(X61)))|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))|(v2_struct_0(X61)|(~v2_pre_topc(X61)|~l1_pre_topc(X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t62_tsep_1])])])])])])).
fof(c_0_94, plain, ![X46, X47, X48]:(~m1_subset_1(X47,k1_zfmisc_1(X46))|~m1_subset_1(X48,k1_zfmisc_1(X46))|k4_subset_1(X46,X47,X48)=k2_xboole_0(X47,X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
fof(c_0_95, plain, ![X13, X14, X15, X16]:((X16!=k1_tsep_1(X13,X14,X15)|u1_struct_0(X16)=k2_xboole_0(u1_struct_0(X14),u1_struct_0(X15))|(v2_struct_0(X16)|(~v1_pre_topc(X16)|~m1_pre_topc(X16,X13)))|(v2_struct_0(X15)|~m1_pre_topc(X15,X13))|(v2_struct_0(X14)|~m1_pre_topc(X14,X13))|(v2_struct_0(X13)|~l1_pre_topc(X13)))&(u1_struct_0(X16)!=k2_xboole_0(u1_struct_0(X14),u1_struct_0(X15))|X16=k1_tsep_1(X13,X14,X15)|(v2_struct_0(X16)|(~v1_pre_topc(X16)|~m1_pre_topc(X16,X13)))|(v2_struct_0(X15)|~m1_pre_topc(X15,X13))|(v2_struct_0(X14)|~m1_pre_topc(X14,X13))|(v2_struct_0(X13)|~l1_pre_topc(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_tsep_1])])])])])).
cnf(c_0_96, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_97, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_40])]), c_0_42])])).
cnf(c_0_98, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_99, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_100, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_101, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|~v2_struct_0(esk4_0)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_102, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_53]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_103, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(esk12_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_104, plain, (m1_pre_topc(X1,esk13_3(X2,X3,X1))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~r3_tsep_1(X2,X3,X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_105, plain, (r1_connsp_1(X2,X4,X5)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_tsep_1(X2,X1,X3)|~r1_tarski(X4,X1)|~r1_tarski(X5,X3)|~r1_xboole_0(k9_subset_1(u1_struct_0(X2),X1,X3),k4_subset_1(u1_struct_0(X2),X4,X5))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_106, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_94])).
fof(c_0_107, plain, ![X49, X50, X51]:(~m1_subset_1(X51,k1_zfmisc_1(X49))|k9_subset_1(X49,X50,X51)=k3_xboole_0(X50,X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
cnf(c_0_108, plain, (u1_struct_0(X1)=k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k1_tsep_1(X2,X3,X4)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
fof(c_0_109, plain, ![X37, X38, X39]:((~v2_struct_0(k1_tsep_1(X37,X38,X39))|(v2_struct_0(X37)|~l1_pre_topc(X37)|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|(v2_struct_0(X39)|~m1_pre_topc(X39,X37)))))&((v1_pre_topc(k1_tsep_1(X37,X38,X39))|(v2_struct_0(X37)|~l1_pre_topc(X37)|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|(v2_struct_0(X39)|~m1_pre_topc(X39,X37)))))&(m1_pre_topc(k1_tsep_1(X37,X38,X39),X37)|(v2_struct_0(X37)|~l1_pre_topc(X37)|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|(v2_struct_0(X39)|~m1_pre_topc(X39,X37))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tsep_1])])])])).
cnf(c_0_110, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_111, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_97, c_0_98]), c_0_99]), c_0_100]), c_0_101])).
cnf(c_0_112, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_113, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(X1,esk13_3(X2,X3,X1))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~r3_tsep_1(X2,X3,X1)), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_114, plain, (m1_pre_topc(X1,esk12_3(X2,X1,X3))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~r3_tsep_1(X2,X1,X3)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_115, plain, (v2_struct_0(X2)|r1_connsp_1(X2,X4,X5)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~r1_tarski(X5,X3)|~r1_tarski(X4,X1)|~r2_tsep_1(X2,X1,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r1_xboole_0(k9_subset_1(u1_struct_0(X2),X1,X3),k4_subset_1(u1_struct_0(X2),X4,X5))), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_116, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_117, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_118, plain, (u1_struct_0(X1)=k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|X1!=k1_tsep_1(X2,X3,X4)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_119, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_120, plain, (v1_pre_topc(k1_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_121, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k1_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_122, negated_conjecture, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)),u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_123, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_39]), c_0_38]), c_0_40]), c_0_41])]), c_0_44]), c_0_42]), c_0_43])])).
cnf(c_0_124, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(X1,esk12_3(X2,X1,X3))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~r3_tsep_1(X2,X1,X3)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_125, plain, (r1_connsp_1(X1,X2,X3)|v2_struct_0(X1)|~r1_tarski(X3,X4)|~r1_tarski(X2,X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k9_subset_1(u1_struct_0(X1),X5,X4),k2_xboole_0(X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116])])).
cnf(c_0_126, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_127, plain, (u1_struct_0(k1_tsep_1(X1,X2,X3))=k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_118])])).
cnf(c_0_128, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_129, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_pre_topc(k1_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_130, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_struct_0(k1_tsep_1(X1,X2,X3))), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_131, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_122])])).
cnf(c_0_132, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_38]), c_0_39]), c_0_40]), c_0_41])]), c_0_43]), c_0_42]), c_0_44])])).
cnf(c_0_133, plain, (r1_connsp_1(X1,X2,X3)|v2_struct_0(X1)|~r1_tarski(X3,X4)|~r1_tarski(X2,X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k3_xboole_0(X5,X4),k2_xboole_0(X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_134, plain, (u1_struct_0(k1_tsep_1(X1,X2,X3))=k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_127, c_0_128]), c_0_129]), c_0_130])).
cnf(c_0_135, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[c_0_131, c_0_132])).
fof(c_0_136, plain, ![X43]:(~l1_pre_topc(X43)|l1_struct_0(X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_137, plain, (r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~r1_tarski(u1_struct_0(X3),X4)|~r1_tarski(u1_struct_0(X2),X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k3_xboole_0(X5,X4),u1_struct_0(k1_tsep_1(X6,X2,X3)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134])])).
cnf(c_0_138, negated_conjecture, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)),u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_135])])).
cnf(c_0_139, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_136])).
fof(c_0_140, plain, ![X30, X31, X32, X33, X34]:((~r4_tsep_1(X30,X31,X32)|(~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X30)))|(~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(X30)))|(X33!=u1_struct_0(X31)|X34!=u1_struct_0(X32)|r2_tsep_1(X30,X33,X34))))|~m1_pre_topc(X32,X30)|~m1_pre_topc(X31,X30)|~l1_pre_topc(X30))&((m1_subset_1(esk8_3(X30,X31,X32),k1_zfmisc_1(u1_struct_0(X30)))|r4_tsep_1(X30,X31,X32)|~m1_pre_topc(X32,X30)|~m1_pre_topc(X31,X30)|~l1_pre_topc(X30))&((m1_subset_1(esk9_3(X30,X31,X32),k1_zfmisc_1(u1_struct_0(X30)))|r4_tsep_1(X30,X31,X32)|~m1_pre_topc(X32,X30)|~m1_pre_topc(X31,X30)|~l1_pre_topc(X30))&(((esk8_3(X30,X31,X32)=u1_struct_0(X31)|r4_tsep_1(X30,X31,X32)|~m1_pre_topc(X32,X30)|~m1_pre_topc(X31,X30)|~l1_pre_topc(X30))&(esk9_3(X30,X31,X32)=u1_struct_0(X32)|r4_tsep_1(X30,X31,X32)|~m1_pre_topc(X32,X30)|~m1_pre_topc(X31,X30)|~l1_pre_topc(X30)))&(~r2_tsep_1(X30,esk8_3(X30,X31,X32),esk9_3(X30,X31,X32))|r4_tsep_1(X30,X31,X32)|~m1_pre_topc(X32,X30)|~m1_pre_topc(X31,X30)|~l1_pre_topc(X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_tsep_1])])])])])).
cnf(c_0_141, negated_conjecture, (r4_tsep_1(esk1_0,esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_142, plain, ![X23, X24, X25, X26, X27]:((~r3_tsep_1(X23,X24,X25)|(~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X23)))|(~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X23)))|(X26!=u1_struct_0(X24)|X27!=u1_struct_0(X25)|r1_connsp_1(X23,X26,X27))))|~m1_pre_topc(X25,X23)|~m1_pre_topc(X24,X23)|~l1_pre_topc(X23))&((m1_subset_1(esk6_3(X23,X24,X25),k1_zfmisc_1(u1_struct_0(X23)))|r3_tsep_1(X23,X24,X25)|~m1_pre_topc(X25,X23)|~m1_pre_topc(X24,X23)|~l1_pre_topc(X23))&((m1_subset_1(esk7_3(X23,X24,X25),k1_zfmisc_1(u1_struct_0(X23)))|r3_tsep_1(X23,X24,X25)|~m1_pre_topc(X25,X23)|~m1_pre_topc(X24,X23)|~l1_pre_topc(X23))&(((esk6_3(X23,X24,X25)=u1_struct_0(X24)|r3_tsep_1(X23,X24,X25)|~m1_pre_topc(X25,X23)|~m1_pre_topc(X24,X23)|~l1_pre_topc(X23))&(esk7_3(X23,X24,X25)=u1_struct_0(X25)|r3_tsep_1(X23,X24,X25)|~m1_pre_topc(X25,X23)|~m1_pre_topc(X24,X23)|~l1_pre_topc(X23)))&(~r1_connsp_1(X23,esk6_3(X23,X24,X25),esk7_3(X23,X24,X25))|r3_tsep_1(X23,X24,X25)|~m1_pre_topc(X25,X23)|~m1_pre_topc(X24,X23)|~l1_pre_topc(X23)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_tsep_1])])])])])).
cnf(c_0_143, negated_conjecture, (r1_connsp_1(X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|v2_struct_0(X1)|r1_tsep_1(esk4_0,esk5_0)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~r2_tsep_1(X1,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_38]), c_0_39]), c_0_40])]), c_0_42]), c_0_43]), c_0_44])])).
cnf(c_0_144, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_145, plain, (r2_tsep_1(X1,X4,X5)|~r4_tsep_1(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|X4!=u1_struct_0(X2)|X5!=u1_struct_0(X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_146, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|r4_tsep_1(esk1_0,esk4_0,esk5_0)), inference(evalgc,[status(thm)],[c_0_141])).
fof(c_0_147, plain, ![X58, X59, X60]:((~r1_tarski(u1_struct_0(X59),u1_struct_0(X60))|m1_pre_topc(X59,X60)|~m1_pre_topc(X60,X58)|~m1_pre_topc(X59,X58)|(~v2_pre_topc(X58)|~l1_pre_topc(X58)))&(~m1_pre_topc(X59,X60)|r1_tarski(u1_struct_0(X59),u1_struct_0(X60))|~m1_pre_topc(X60,X58)|~m1_pre_topc(X59,X58)|(~v2_pre_topc(X58)|~l1_pre_topc(X58)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_148, plain, (r3_tsep_1(X1,X2,X3)|~r1_connsp_1(X1,esk6_3(X1,X2,X3),esk7_3(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_149, plain, (esk7_3(X1,X2,X3)=u1_struct_0(X3)|r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_150, plain, (r1_connsp_1(X1,X4,X5)|~r3_tsep_1(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|X4!=u1_struct_0(X2)|X5!=u1_struct_0(X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_151, negated_conjecture, (r1_connsp_1(X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~r2_tsep_1(X1,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144])])).
cnf(c_0_152, plain, (r2_tsep_1(X1,X4,X5)|X5!=u1_struct_0(X3)|X4!=u1_struct_0(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r4_tsep_1(X1,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_153, negated_conjecture, (r4_tsep_1(esk1_0,esk4_0,esk5_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_146, c_0_132])])).
cnf(c_0_154, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_99, c_0_132])])).
cnf(c_0_155, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_98, c_0_132])])).
cnf(c_0_156, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_157, negated_conjecture, (m1_pre_topc(esk3_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_158, plain, (r3_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r1_connsp_1(X1,esk6_3(X1,X2,X3),esk7_3(X1,X2,X3))), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_159, plain, (esk7_3(X1,X2,X3)=u1_struct_0(X3)|r3_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_149])).
cnf(c_0_160, plain, (r1_connsp_1(X1,X4,X5)|X5!=u1_struct_0(X3)|X4!=u1_struct_0(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_161, plain, (esk6_3(X1,X2,X3)=u1_struct_0(X2)|r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_162, negated_conjecture, (r1_connsp_1(X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~r2_tsep_1(X1,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_144])])).
cnf(c_0_163, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk5_0)|X1!=u1_struct_0(esk4_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_154]), c_0_155]), c_0_40])])])).
cnf(c_0_164, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_165, negated_conjecture, (m1_pre_topc(esk3_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_166, negated_conjecture, (m1_pre_topc(esk2_0,esk4_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_167, plain, (r3_tsep_1(X1,X2,X3)|~r1_connsp_1(X1,esk6_3(X1,X2,X3),u1_struct_0(X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159])])).
cnf(c_0_168, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_100]), c_0_38]), c_0_39]), c_0_40])])])).
cnf(c_0_169, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_170, plain, (esk6_3(X1,X2,X3)=u1_struct_0(X2)|r3_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_161])).
fof(c_0_171, plain, ![X56, X57]:(~l1_pre_topc(X56)|(~m1_pre_topc(X57,X56)|m1_subset_1(u1_struct_0(X57),k1_zfmisc_1(u1_struct_0(X56))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
cnf(c_0_172, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163]), c_0_40]), c_0_41])]), c_0_42])])).
cnf(c_0_173, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk5_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_154]), c_0_40]), c_0_41])])])).
cnf(c_0_174, negated_conjecture, (m1_pre_topc(esk3_0,esk5_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_165, c_0_132])])).
cnf(c_0_175, negated_conjecture, (m1_pre_topc(esk2_0,esk4_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_166])).
fof(c_0_176, plain, ![X44, X45]:(~l1_pre_topc(X44)|(~m1_pre_topc(X45,X44)|l1_pre_topc(X45))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_177, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_40])])])).
cnf(c_0_178, plain, (r3_tsep_1(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_179, plain, (r3_tsep_1(X1,X2,X3)|~r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_170])])).
cnf(c_0_180, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_101]), c_0_38]), c_0_39]), c_0_40])])])).
cnf(c_0_181, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_171])).
cnf(c_0_182, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_38]), c_0_174])])])).
cnf(c_0_183, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk4_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_155]), c_0_40]), c_0_41])])])).
cnf(c_0_184, negated_conjecture, (m1_pre_topc(esk2_0,esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_175, c_0_132])])).
cnf(c_0_185, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_186, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_40])])])).
cnf(c_0_187, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_40])])])).
cnf(c_0_188, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_189, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183]), c_0_39]), c_0_184])])])).
cnf(c_0_190, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_191, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_170]), c_0_40])])])).
cnf(c_0_192, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_40])])])).
cnf(c_0_193, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_189]), c_0_38]), c_0_39]), c_0_40])]), c_0_132])])).
cnf(c_0_194, plain, (l1_pre_topc(k1_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_128])])).
cnf(c_0_195, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_188]), c_0_40])])])).
cnf(c_0_196, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_188]), c_0_40])])])).
cnf(c_0_197, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194]), c_0_38]), c_0_39]), c_0_40])]), c_0_44]), c_0_43]), c_0_42])])).
cnf(c_0_198, plain, (l1_pre_topc(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_70])])).
cnf(c_0_199, negated_conjecture, (~v2_struct_0(esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_195]), c_0_38]), c_0_39])])])).
cnf(c_0_200, negated_conjecture, (~v2_struct_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_196]), c_0_38]), c_0_39])])])).
cnf(c_0_201, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198]), c_0_154]), c_0_155]), c_0_40])]), c_0_199]), c_0_200]), c_0_42])])).
fof(c_0_202, plain, ![X75, X76, X77]:((~r1_tsep_1(X76,X77)|~r4_tsep_1(X75,X76,X77)|r3_tsep_1(X75,X76,X77)|(v2_struct_0(X77)|~m1_pre_topc(X77,X75))|(v2_struct_0(X76)|~m1_pre_topc(X76,X75))|(v2_struct_0(X75)|(~v2_pre_topc(X75)|~l1_pre_topc(X75))))&((r1_tsep_1(X76,X77)|~r3_tsep_1(X75,X76,X77)|(v2_struct_0(X77)|~m1_pre_topc(X77,X75))|(v2_struct_0(X76)|~m1_pre_topc(X76,X75))|(v2_struct_0(X75)|(~v2_pre_topc(X75)|~l1_pre_topc(X75))))&(r4_tsep_1(X75,X76,X77)|~r3_tsep_1(X75,X76,X77)|(v2_struct_0(X77)|~m1_pre_topc(X77,X75))|(v2_struct_0(X76)|~m1_pre_topc(X76,X75))|(v2_struct_0(X75)|(~v2_pre_topc(X75)|~l1_pre_topc(X75)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t78_tsep_1])])])])])).
cnf(c_0_203, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_188]), c_0_154]), c_0_40])])])).
cnf(c_0_204, plain, (r3_tsep_1(X3,X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~r1_tsep_1(X1,X2)|~r4_tsep_1(X3,X1,X2)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_205, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_188]), c_0_155]), c_0_40])])])).
cnf(c_0_206, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r3_tsep_1(X3,X1,X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~r1_tsep_1(X1,X2)|~r4_tsep_1(X3,X1,X2)), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_207, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_188]), c_0_38]), c_0_40])])])).
fof(c_0_208, plain, ![X78, X79, X80, X81, X82]:(~l1_pre_topc(X78)|(~m1_subset_1(X79,k1_zfmisc_1(u1_struct_0(X78)))|(~m1_subset_1(X80,k1_zfmisc_1(u1_struct_0(X78)))|(~m1_subset_1(X81,k1_zfmisc_1(u1_struct_0(X78)))|(~m1_subset_1(X82,k1_zfmisc_1(u1_struct_0(X78)))|(~r1_connsp_1(X78,X79,X80)|(~r1_tarski(X81,X79)|~r1_tarski(X82,X80))|r1_connsp_1(X78,X81,X82))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_connsp_1])])])).
cnf(c_0_209, negated_conjecture, (r3_tsep_1(esk1_0,esk4_0,esk5_0)|~r1_tsep_1(esk4_0,esk5_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_153]), c_0_154]), c_0_155]), c_0_40]), c_0_41])]), c_0_200]), c_0_199]), c_0_42])])).
cnf(c_0_210, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_188]), c_0_39]), c_0_40])])])).
cnf(c_0_211, plain, (r1_connsp_1(X1,X4,X5)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_connsp_1(X1,X2,X3)|~r1_tarski(X4,X2)|~r1_tarski(X5,X3)), inference(split_conjunct,[status(thm)],[c_0_208])).
cnf(c_0_212, negated_conjecture, (r3_tsep_1(esk1_0,esk4_0,esk5_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_209, c_0_210])])])).
cnf(c_0_213, plain, (r1_connsp_1(X1,X4,X5)|~l1_pre_topc(X1)|~r1_tarski(X5,X3)|~r1_tarski(X4,X2)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_214, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk5_0)|X1!=u1_struct_0(esk4_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_212]), c_0_154]), c_0_155]), c_0_40])])])).
cnf(c_0_215, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk5_0)|X4!=u1_struct_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214]), c_0_40])])])).
cnf(c_0_216, negated_conjecture, (r1_connsp_1(esk1_0,X1,u1_struct_0(X2))|X3!=u1_struct_0(esk4_0)|~r1_tarski(X1,X3)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,esk5_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_173])])).
cnf(c_0_217, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_183])])).
cnf(c_0_218, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_217]), c_0_40])])])).
cnf(c_0_219, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_188]), c_0_154]), c_0_40])])])).
cnf(c_0_220, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_188]), c_0_155]), c_0_40])])])).
cnf(c_0_221, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_188]), c_0_40])])])).
cnf(c_0_222, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_188]), c_0_40])])])).
cnf(c_0_223, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_222]), c_0_38]), c_0_39]), c_0_174]), c_0_184])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 224
# Proof object clause steps            : 181
# Proof object formula steps           : 43
# Proof object conjectures             : 98
# Proof object clause conjectures      : 95
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 53
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 212
# Training examples: 121 positive, 470 negative
# Training: Positive examples begin
cnf(c_0_41, negated_conjecture, (v2_pre_topc(esk1_0))).# trainpos
cnf(c_0_40, negated_conjecture, (l1_pre_topc(esk1_0))).# trainpos
cnf(c_0_42, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_43, negated_conjecture, (~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_44, negated_conjecture, (~v2_struct_0(esk3_0))).# trainpos
cnf(c_0_39, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0))).# trainpos
cnf(c_0_38, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0))).# trainpos
cnf(c_0_144, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_110, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2))).# trainpos
cnf(c_0_101, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|~v2_struct_0(esk4_0))).# trainpos
cnf(c_0_100, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_175, negated_conjecture, (m1_pre_topc(esk2_0,esk4_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))).# trainpos
cnf(c_0_165, negated_conjecture, (m1_pre_topc(esk3_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))).# trainpos
cnf(c_0_98, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))).# trainpos
cnf(c_0_99, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0))).# trainpos
cnf(c_0_51, negated_conjecture, (v2_struct_0(X2)|v2_struct_0(X1)|~r1_tsep_1(X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X2)|~r4_tsep_1(esk1_0,X1,X2)|~r3_tsep_1(esk1_0,esk2_0,esk3_0))).# trainpos
cnf(c_0_190, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_146, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|r4_tsep_1(esk1_0,esk4_0,esk5_0))).# trainpos
cnf(c_0_81, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|r1_tsep_1(k2_tsep_1(esk1_0,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_188, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_71, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_pre_topc(k2_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_129, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_pre_topc(k1_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_80, plain, (r1_tsep_1(X2,X1)|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(X1,X2))).# trainpos
cnf(c_0_78, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(X1,X2))).# trainpos
cnf(c_0_126, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_116, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_170, plain, (esk6_3(X1,X2,X3)=u1_struct_0(X2)|r3_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_159, plain, (esk7_3(X1,X2,X3)=u1_struct_0(X3)|r3_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_164, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_178, plain, (r3_tsep_1(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_36, negated_conjecture, (v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X2)|~r4_tsep_1(esk1_0,X1,X2)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r1_tsep_1(k2_tsep_1(esk1_0,X1,X2),k1_tsep_1(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_46, plain, (v2_struct_0(X1)|r4_tsep_1(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v1_tsep_1(X3,X1)|~v1_tsep_1(X2,X1))).# trainpos
cnf(c_0_70, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_72, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_struct_0(k2_tsep_1(X1,X2,X3)))).# trainpos
cnf(c_0_130, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_struct_0(k1_tsep_1(X1,X2,X3)))).# trainpos
cnf(c_0_158, plain, (r3_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r1_connsp_1(X1,esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)))).# trainpos
cnf(c_0_128, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_118, plain, (u1_struct_0(X1)=k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|X1!=k1_tsep_1(X2,X3,X4)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_82, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)|~v2_struct_0(esk12_3(X1,X2,X3)))).# trainpos
cnf(c_0_76, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)|~v2_struct_0(esk13_3(X1,X2,X3)))).# trainpos
cnf(c_0_124, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(X1,esk12_3(X2,X1,X3))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~r3_tsep_1(X2,X1,X3))).# trainpos
cnf(c_0_113, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(X1,esk13_3(X2,X3,X1))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~r3_tsep_1(X2,X3,X1))).# trainpos
cnf(c_0_103, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(esk12_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3))).# trainpos
cnf(c_0_53, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|m1_pre_topc(esk13_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3))).# trainpos
cnf(c_0_67, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_tsep_1(esk12_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3))).# trainpos
cnf(c_0_47, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v1_tsep_1(esk13_3(X1,X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3))).# trainpos
cnf(c_0_206, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r3_tsep_1(X3,X1,X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~r1_tsep_1(X1,X2)|~r4_tsep_1(X3,X1,X2))).# trainpos
cnf(c_0_160, plain, (r1_connsp_1(X1,X4,X5)|X5!=u1_struct_0(X3)|X4!=u1_struct_0(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_115, plain, (v2_struct_0(X2)|r1_connsp_1(X2,X4,X5)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~r1_tarski(X5,X3)|~r1_tarski(X4,X1)|~r2_tsep_1(X2,X1,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r1_xboole_0(k9_subset_1(u1_struct_0(X2),X1,X3),k4_subset_1(u1_struct_0(X2),X4,X5)))).# trainpos
cnf(c_0_152, plain, (r2_tsep_1(X1,X4,X5)|X5!=u1_struct_0(X3)|X4!=u1_struct_0(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r4_tsep_1(X1,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_213, plain, (r1_connsp_1(X1,X4,X5)|~l1_pre_topc(X1)|~r1_tarski(X5,X3)|~r1_tarski(X4,X2)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_60, plain, (u1_struct_0(X1)=k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X3,X4)|X1!=k2_tsep_1(X2,X3,X4)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_37, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|r1_tsep_1(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)),k1_tsep_1(X1,X2,X3))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3))).# trainpos
cnf(c_0_50, negated_conjecture, (v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk12_3(esk1_0,esk2_0,esk3_0),esk13_3(esk1_0,esk2_0,esk3_0))|~r4_tsep_1(esk1_0,esk12_3(esk1_0,esk2_0,esk3_0),esk13_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_86, negated_conjecture, (r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),k2_tsep_1(esk1_0,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).# trainpos
cnf(c_0_69, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|r1_tsep_1(X2,X3)|v2_struct_0(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_127, plain, (u1_struct_0(k1_tsep_1(X1,X2,X3))=k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_167, plain, (r3_tsep_1(X1,X2,X3)|~r1_connsp_1(X1,esk6_3(X1,X2,X3),u1_struct_0(X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_198, plain, (l1_pre_topc(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_194, plain, (l1_pre_topc(k1_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_52, plain, (r4_tsep_1(X1,X2,esk13_3(X1,X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_tsep_1(X2,X1)|~m1_pre_topc(esk13_3(X1,X3,X4),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1))).# trainpos
cnf(c_0_168, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_180, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk4_0))).# trainpos
cnf(c_0_66, negated_conjecture, (v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))|~v1_tsep_1(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0))).# trainpos
cnf(c_0_125, plain, (r1_connsp_1(X1,X2,X3)|v2_struct_0(X1)|~r1_tarski(X3,X4)|~r1_tarski(X2,X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k9_subset_1(u1_struct_0(X1),X5,X4),k2_xboole_0(X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_179, plain, (r3_tsep_1(X1,X2,X3)|~r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_177, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_77, negated_conjecture, (v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_133, plain, (r1_connsp_1(X1,X2,X3)|v2_struct_0(X1)|~r1_tarski(X3,X4)|~r1_tarski(X2,X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k3_xboole_0(X5,X4),k2_xboole_0(X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_187, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk4_0))).# trainpos
cnf(c_0_186, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_83, negated_conjecture, (v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_192, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk4_0))).# trainpos
cnf(c_0_191, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_85, plain, (r1_xboole_0(u1_struct_0(X1),k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_91, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_196, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk4_0))).# trainpos
cnf(c_0_195, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_102, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_97, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0))).# trainpos
cnf(c_0_112, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_122, negated_conjecture, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)),u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_137, plain, (r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~r1_tarski(u1_struct_0(X3),X4)|~r1_tarski(u1_struct_0(X2),X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k3_xboole_0(X5,X4),u1_struct_0(k1_tsep_1(X6,X2,X3)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).# trainpos
cnf(c_0_123, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_132, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0))).# trainpos
cnf(c_0_174, negated_conjecture, (m1_pre_topc(esk3_0,esk5_0))).# trainpos
cnf(c_0_155, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0))).# trainpos
cnf(c_0_184, negated_conjecture, (m1_pre_topc(esk2_0,esk4_0))).# trainpos
cnf(c_0_154, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0))).# trainpos
cnf(c_0_200, negated_conjecture, (~v2_struct_0(esk4_0))).# trainpos
cnf(c_0_199, negated_conjecture, (~v2_struct_0(esk5_0))).# trainpos
cnf(c_0_153, negated_conjecture, (r4_tsep_1(esk1_0,esk4_0,esk5_0))).# trainpos
cnf(c_0_131, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_163, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk5_0)|X1!=u1_struct_0(esk4_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_183, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk4_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk4_0))).# trainpos
cnf(c_0_173, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk5_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk5_0))).# trainpos
cnf(c_0_209, negated_conjecture, (r3_tsep_1(esk1_0,esk4_0,esk5_0)|~r1_tsep_1(esk4_0,esk5_0))).# trainpos
cnf(c_0_138, negated_conjecture, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)),u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_143, negated_conjecture, (r1_connsp_1(X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|v2_struct_0(X1)|r1_tsep_1(esk4_0,esk5_0)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~r2_tsep_1(X1,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_151, negated_conjecture, (r1_connsp_1(X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~r2_tsep_1(X1,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_162, negated_conjecture, (r1_connsp_1(X1,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~r2_tsep_1(X1,u1_struct_0(esk4_0),u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).# trainpos
cnf(c_0_172, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|~r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk5_0))|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).# trainpos
cnf(c_0_182, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|~r1_tarski(u1_struct_0(esk2_0),u1_struct_0(esk4_0))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).# trainpos
cnf(c_0_189, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(esk2_0),u1_struct_0(esk3_0))|r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).# trainpos
cnf(c_0_193, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).# trainpos
cnf(c_0_197, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).# trainpos
cnf(c_0_201, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_203, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_205, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_207, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_210, negated_conjecture, (r1_tsep_1(esk4_0,esk5_0))).# trainpos
cnf(c_0_212, negated_conjecture, (r3_tsep_1(esk1_0,esk4_0,esk5_0))).# trainpos
cnf(c_0_214, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk5_0)|X1!=u1_struct_0(esk4_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_215, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk5_0)|X4!=u1_struct_0(esk4_0))).# trainpos
cnf(c_0_216, negated_conjecture, (r1_connsp_1(esk1_0,X1,u1_struct_0(X2))|X3!=u1_struct_0(esk4_0)|~r1_tarski(X1,X3)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,esk5_0))).# trainpos
cnf(c_0_217, negated_conjecture, (r1_connsp_1(esk1_0,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk4_0))).# trainpos
cnf(c_0_218, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0))).# trainpos
cnf(c_0_219, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0))).# trainpos
cnf(c_0_220, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0))).# trainpos
cnf(c_0_221, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0))).# trainpos
cnf(c_0_222, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk5_0)|~m1_pre_topc(X1,esk4_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_20, plain, (r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_32, plain, (esk8_3(X1,X2,X3)=u1_struct_0(X2)|r4_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).#trainneg
cnf(i_0_31, plain, (esk9_3(X1,X2,X3)=u1_struct_0(X3)|r4_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).#trainneg
cnf(i_0_50, plain, (m1_pre_topc(X1,X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2)))).#trainneg
cnf(i_0_27, plain, (r3_tsep_1(X1,X2,X3)|m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).#trainneg
cnf(i_0_34, plain, (r4_tsep_1(X1,X2,X3)|m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).#trainneg
cnf(i_0_70, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X1,X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~r3_tsep_1(X3,X1,X2))).#trainneg
cnf(i_0_30, plain, (r4_tsep_1(X1,X2,X3)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r2_tsep_1(X1,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)))).#trainneg
cnf(i_0_69, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r4_tsep_1(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~r3_tsep_1(X1,X2,X3))).#trainneg
cnf(i_0_18, plain, (X1=k1_tsep_1(X4,X2,X3)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(X1)!=k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|~l1_pre_topc(X4)|~v1_pre_topc(X1)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4))).#trainneg
cnf(i_0_33, plain, (r4_tsep_1(X1,X2,X3)|m1_subset_1(esk9_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1))).#trainneg
cnf(i_0_54, plain, (v2_struct_0(X2)|r1_tarski(X1,esk10_3(X2,X1,X3))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~r1_connsp_1(X2,X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_53, plain, (v2_struct_0(X2)|r1_tarski(X1,esk11_3(X2,X3,X1))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~r1_connsp_1(X2,X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_57, plain, (v2_struct_0(X1)|m1_subset_1(esk10_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_56, plain, (v2_struct_0(X1)|m1_subset_1(esk11_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_22, plain, (X1=k2_tsep_1(X4,X2,X3)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X2,X3)|u1_struct_0(X1)!=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))|~l1_pre_topc(X4)|~v1_pre_topc(X1)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4))).#trainneg
cnf(i_0_55, plain, (v2_struct_0(X1)|r2_tsep_1(X1,esk10_3(X1,X2,X3),esk11_3(X1,X2,X3))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_58, plain, (v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r3_tsep_1(X1,X4,X5)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v1_tsep_1(X3,X1)|~v1_tsep_1(X2,X1)|~r1_tsep_1(k2_tsep_1(X1,X2,X3),k1_tsep_1(X1,X4,X5)))).#trainneg
cnf(i_0_59, plain, (v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r3_tsep_1(X5,X4,X3)|~v2_pre_topc(X5)|~l1_pre_topc(X5)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X1,X5)|~r1_tsep_1(X1,X2)|~v1_tsep_1(X2,X5)|~v1_tsep_1(X1,X5))).#trainneg
cnf(i_0_52, plain, (v2_struct_0(X1)|r1_xboole_0(k9_subset_1(u1_struct_0(X1),esk10_3(X1,X2,X3),esk11_3(X1,X2,X3)),k4_subset_1(u1_struct_0(X1),X2,X3))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_100, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(X2)|~r1_tsep_1(X1,X2)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X2)|~r3_tsep_1(esk1_0,X1,X2))).#trainneg
cnf(i_0_78, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_struct_0(X2)|~r1_tsep_1(X2,X1))).#trainneg
cnf(i_0_125, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r1_tsep_1(esk4_0,esk5_0)|~v1_tsep_1(esk5_0,esk1_0)|~v1_tsep_1(esk4_0,esk1_0)|~m1_pre_topc(esk3_0,esk5_0)|~m1_pre_topc(esk2_0,esk4_0)|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0))).#trainneg
cnf(i_0_113, negated_conjecture, (r3_tsep_1(esk1_0,esk4_0,esk5_0)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r1_tsep_1(esk4_0,esk5_0)|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0))).#trainneg
cnf(i_0_91, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)|m1_pre_topc(esk3_0,esk5_0))).#trainneg
cnf(i_0_97, plain, (r4_tsep_1(X1,X2,X3)|~r2_tsep_1(X1,u1_struct_0(X2),esk9_3(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_90, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)|m1_pre_topc(esk2_0,esk4_0))).#trainneg
cnf(i_0_74, negated_conjecture, (l1_pre_topc(esk3_0))).#trainneg
cnf(i_0_75, negated_conjecture, (l1_pre_topc(esk2_0))).#trainneg
cnf(i_0_88, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_80, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk3_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk3_0))).#trainneg
cnf(i_0_81, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_89, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_92, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)|m1_pre_topc(esk4_0,esk1_0))).#trainneg
cnf(i_0_96, plain, (r4_tsep_1(X1,X2,X3)|~r2_tsep_1(X1,esk8_3(X1,X2,X3),u1_struct_0(X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_93, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)|m1_pre_topc(esk5_0,esk1_0))).#trainneg
cnf(i_0_123, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|X2!=u1_struct_0(esk5_0)|X1!=u1_struct_0(esk4_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0))).#trainneg
cnf(i_0_111, plain, (r4_tsep_1(X1,X2,esk12_3(X1,X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~v1_tsep_1(X2,X1)|~m1_pre_topc(esk12_3(X1,X3,X4),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1))).#trainneg
cnf(i_0_131, plain, (r1_xboole_0(k9_subset_1(u1_struct_0(X1),esk10_3(X1,X2,X3),esk11_3(X1,X2,X3)),k2_xboole_0(X2,X3))|v2_struct_0(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_120, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk5_0,esk1_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_119, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk4_0,esk1_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_117, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk2_0,esk4_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_118, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk3_0,esk5_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_95, plain, (r3_tsep_1(X1,X2,X3)|~r1_connsp_1(X1,u1_struct_0(X2),esk7_3(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_210, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk2_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_211, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_222, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_221, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk2_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_108, plain, (l1_pre_topc(esk12_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_86, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(k2_tsep_1(X2,X3,X4)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X1,k2_tsep_1(X2,X3,X4))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2))).#trainneg
cnf(i_0_98, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(k1_tsep_1(X2,X3,X4)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2))).#trainneg
cnf(i_0_140, negated_conjecture, (v2_struct_0(esk13_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk12_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk12_3(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(esk2_0,esk12_3(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,esk13_3(esk1_0,esk2_0,esk3_0))|~r3_tsep_1(esk1_0,esk12_3(esk1_0,esk2_0,esk3_0),esk13_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_223, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk4_0,esk1_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_203, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk3_0,esk5_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_259, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk2_0,esk4_0)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk3_0)|X4!=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_250, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk4_0,esk1_0)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk3_0)|X4!=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_208, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk2_0,esk4_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_228, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk5_0,esk1_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_262, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk3_0,esk5_0)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk3_0)|X4!=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_245, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(esk13_3(esk1_0,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tsep_1(X1,esk13_3(esk1_0,X2,X3))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk13_3(esk1_0,X2,X3),esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,esk13_3(esk1_0,X2,X3))|~v1_tsep_1(X1,esk1_0)|~r3_tsep_1(esk1_0,X2,X3)|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X2,esk1_0))).#trainneg
cnf(i_0_205, plain, (r4_tsep_1(X1,X2,X3)|~r2_tsep_1(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_204, negated_conjecture, (m1_pre_topc(esk3_0,esk5_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_209, negated_conjecture, (m1_pre_topc(esk2_0,esk4_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_103, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,esk12_3(X3,X2,X4))|~m1_pre_topc(X1,X2)|~l1_pre_topc(esk12_3(X3,X2,X4))|~v2_pre_topc(esk12_3(X3,X2,X4))|~r3_tsep_1(X3,X2,X4)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X2,X3)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_224, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_229, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_110, plain, (l1_pre_topc(esk13_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_216, negated_conjecture, (r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(esk3_0))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk3_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_105, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,esk13_3(X3,X4,X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(esk13_3(X3,X4,X2))|~v2_pre_topc(esk13_3(X3,X4,X2))|~r3_tsep_1(X3,X4,X2)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X2,X3)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_233, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(esk12_3(esk1_0,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tsep_1(X1,esk12_3(esk1_0,X2,X3))|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(esk12_3(esk1_0,X2,X3),esk1_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,esk12_3(esk1_0,X2,X3))|~v1_tsep_1(X1,esk1_0)|~r3_tsep_1(esk1_0,X2,X3)|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X2,esk1_0))).#trainneg
cnf(i_0_129, plain, (r1_xboole_0(k4_subset_1(u1_struct_0(X1),X2,X3),k9_subset_1(u1_struct_0(X1),esk10_3(X1,X2,X3),esk11_3(X1,X2,X3)))|v2_struct_0(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_133, plain, (r1_tsep_1(k1_tsep_1(X1,X2,X3),k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~l1_struct_0(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_243, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|m1_pre_topc(esk5_0,esk1_0)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk3_0)|X4!=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_246, plain, (r3_tsep_1(X1,X2,esk13_3(X1,X3,X4))|v2_struct_0(X2)|v2_struct_0(esk13_3(X1,X3,X4))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tsep_1(X2,esk13_3(X1,X3,X4))|~m1_pre_topc(esk13_3(X1,X3,X4),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_tsep_1(X2,X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1))).#trainneg
cnf(i_0_256, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk3_0)|X4!=u1_struct_0(esk2_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_253, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|~r1_tarski(X2,X3)|~r1_tarski(X1,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(esk1_0)))|X3!=u1_struct_0(esk3_0)|X4!=u1_struct_0(esk2_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_234, plain, (r3_tsep_1(X1,X2,esk12_3(X1,X3,X4))|v2_struct_0(X2)|v2_struct_0(esk12_3(X1,X3,X4))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tsep_1(X2,esk12_3(X1,X3,X4))|~m1_pre_topc(esk12_3(X1,X3,X4),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_tsep_1(X2,X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1))).#trainneg
cnf(i_0_137, negated_conjecture, (v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r3_tsep_1(esk1_0,k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)),k1_tsep_1(X1,X2,X3))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)),esk1_0)|~m1_pre_topc(esk2_0,k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|~m1_pre_topc(esk3_0,k1_tsep_1(X1,X2,X3))|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_124, plain, (r2_tsep_1(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X4)|X3!=u1_struct_0(X4)|X2!=u1_struct_0(X5)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~l1_pre_topc(X1)|~r3_tsep_1(X1,X5,X4)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_122, plain, (r1_connsp_1(X1,X2,X3)|v2_struct_0(X1)|~r1_tarski(X3,X4)|~r1_tarski(X2,X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k3_xboole_0(X5,X4),k4_subset_1(u1_struct_0(X1),X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_132, plain, (r1_xboole_0(k3_xboole_0(esk10_3(X1,X2,X3),esk11_3(X1,X2,X3)),k4_subset_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk11_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_255, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_107, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk12_3(X2,X3,X4)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X1,esk12_3(X2,X3,X4))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r3_tsep_1(X2,X3,X4)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2))).#trainneg
cnf(i_0_230, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(esk9_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk8_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_266, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(esk7_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_267, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(esk7_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_219, negated_conjecture, (r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(esk2_0))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk2_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_175, plain, (r1_xboole_0(u1_struct_0(X1),k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~r1_tsep_1(X1,k1_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_109, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk13_3(X2,X3,X4)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X1,esk13_3(X2,X3,X4))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r3_tsep_1(X2,X3,X4)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2))).#trainneg
cnf(i_0_126, plain, (r3_tsep_1(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(esk13_3(X1,X5,X6))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X6)|~v1_tsep_1(X4,X1)|~r1_tsep_1(X4,esk13_3(X1,X5,X6))|~m1_pre_topc(X2,X1)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X3,esk13_3(X1,X5,X6))|~m1_pre_topc(esk13_3(X1,X5,X6),X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_tsep_1(X1,X5,X6)|~m1_pre_topc(X6,X1)|~m1_pre_topc(X5,X1))).#trainneg
cnf(i_0_283, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk4_0,esk1_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_297, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk2_0,esk4_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_284, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk3_0,esk5_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_298, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk5_0,esk1_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_272, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|~l1_struct_0(esk2_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_404, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)),k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_273, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|~l1_struct_0(esk2_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_373, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|m1_pre_topc(esk4_0,esk1_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_371, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|m1_pre_topc(esk2_0,esk4_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_333, negated_conjecture, (l1_pre_topc(esk4_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_387, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_143, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=u1_struct_0(k2_tsep_1(X4,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_400, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_335, negated_conjecture, (l1_pre_topc(esk5_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_420, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_127, plain, (r3_tsep_1(X1,X2,X3)|v2_struct_0(X4)|v2_struct_0(esk12_3(X1,X5,X6))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X6)|~v1_tsep_1(X4,X1)|~r1_tsep_1(X4,esk12_3(X1,X5,X6))|~m1_pre_topc(X2,X1)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X3,esk12_3(X1,X5,X6))|~m1_pre_topc(esk12_3(X1,X5,X6),X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r3_tsep_1(X1,X5,X6)|~m1_pre_topc(X6,X1)|~m1_pre_topc(X5,X1))).#trainneg
cnf(i_0_144, plain, (X1=k2_tsep_1(X2,X3,X4)|r1_tsep_1(X3,X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X5)|u1_struct_0(X1)!=u1_struct_0(k2_tsep_1(X5,X3,X4))|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_418, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk3_0,esk5_0))).#trainneg
cnf(i_0_419, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk5_0,esk1_0))).#trainneg
cnf(i_0_417, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk2_0,esk4_0))).#trainneg
cnf(i_0_416, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|m1_pre_topc(esk4_0,esk1_0))).#trainneg
cnf(i_0_423, negated_conjecture, (r1_xboole_0(k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)),u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_487, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0)),k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_425, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_374, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|m1_pre_topc(esk5_0,esk1_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_350, negated_conjecture, (v2_struct_0(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|v2_struct_0(k1_tsep_1(X1,X2,X3))|r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r3_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|~m1_pre_topc(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)),esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(esk2_0,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(esk3_0,k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|~l1_struct_0(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_372, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|m1_pre_topc(esk3_0,esk5_0)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_332, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk4_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk4_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_402, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_398, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(esk8_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_334, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk5_0))|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk5_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_325, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk3_0))|~m1_pre_topc(X1,esk5_0)|~m1_pre_topc(X1,esk3_0)|~l1_pre_topc(esk5_0)|~v2_pre_topc(esk5_0)|~m1_pre_topc(esk2_0,esk2_0)|~m1_pre_topc(esk3_0,esk3_0))).#trainneg
cnf(i_0_145, plain, (X1=k1_tsep_1(X2,X3,k2_tsep_1(X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k2_tsep_1(X4,X5,X6))|v2_struct_0(X2)|r1_tsep_1(X5,X6)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|u1_struct_0(X1)!=k2_xboole_0(u1_struct_0(X3),k3_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|~v1_pre_topc(X1)|~m1_pre_topc(k2_tsep_1(X4,X5,X6),X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_146, plain, (X1=k2_tsep_1(X2,X3,k2_tsep_1(X4,X5,X6))|r1_tsep_1(X3,k2_tsep_1(X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k2_tsep_1(X4,X5,X6))|v2_struct_0(X2)|r1_tsep_1(X5,X6)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|u1_struct_0(X1)!=k3_xboole_0(u1_struct_0(X3),k3_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|~v1_pre_topc(X1)|~m1_pre_topc(k2_tsep_1(X4,X5,X6),X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_370, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_369, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk3_0)|X1!=u1_struct_0(esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_236, plain, (r1_xboole_0(k2_xboole_0(X1,X2),k9_subset_1(u1_struct_0(X3),esk10_3(X3,X1,X2),esk11_3(X3,X1,X2)))|v2_struct_0(X3)|~r1_connsp_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_547, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r3_tsep_1(esk1_0,esk3_0,esk2_0)|~m1_pre_topc(esk2_0,esk3_0)|~m1_pre_topc(esk3_0,esk2_0)|~v2_struct_0(esk4_0))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4))))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X1,k2_tsep_1(X2,X3,X4))|~l1_pre_topc(k2_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_609, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_575, negated_conjecture, (~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r3_tsep_1(esk1_0,esk3_0,esk2_0)|~m1_pre_topc(esk2_0,esk3_0)|~m1_pre_topc(esk3_0,esk2_0)|~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_558, negated_conjecture, (m1_pre_topc(esk3_0,esk5_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r3_tsep_1(esk1_0,esk3_0,esk2_0)|~m1_pre_topc(esk2_0,esk3_0)|~m1_pre_topc(esk3_0,esk2_0))).#trainneg
cnf(i_0_548, plain, (k2_tsep_1(X1,X2,X3)=k2_tsep_1(X4,X2,X3)|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(k2_tsep_1(X1,X2,X3))|~v1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_488, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(k2_tsep_1(X5,X3,X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|r1_tsep_1(X3,X4)|v2_struct_0(X5)|~m1_pre_topc(X1,k2_tsep_1(X2,X3,X4))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_149, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_xboole_0(u1_struct_0(X1),k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_150, plain, (m1_pre_topc(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tarski(u1_struct_0(X1),k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X5)|~m1_pre_topc(X1,X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_299, plain, (r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~r1_tarski(u1_struct_0(X3),X4)|~r1_tarski(u1_struct_0(X2),X5)|~r2_tsep_1(X1,X5,X4)|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k9_subset_1(u1_struct_0(X1),X5,X4),u1_struct_0(k1_tsep_1(X6,X2,X3)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_171, plain, (u1_struct_0(k1_tsep_1(X1,X2,X3))=u1_struct_0(k1_tsep_1(X4,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_312, plain, (r4_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_tsep_1(X1,k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_568, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0)),u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_383, plain, (r1_xboole_0(k3_xboole_0(esk10_3(X1,X2,X3),esk11_3(X1,X2,X3)),k2_xboole_0(X2,X3))|v2_struct_0(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_841, plain, (r1_xboole_0(u1_struct_0(k1_tsep_1(X5,X2,X3)),u1_struct_0(X4))|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(X4)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~l1_pre_topc(X1)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_783, plain, (r1_xboole_0(u1_struct_0(k1_tsep_1(X5,X2,X3)),u1_struct_0(X4))|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~r1_tsep_1(X4,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_231, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk8_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_313, plain, (r4_tsep_1(X1,k2_tsep_1(X2,X3,X4),X5)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_tsep_1(X1,k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_698, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_692, negated_conjecture, (~m1_pre_topc(esk2_0,esk3_0)|~m1_pre_topc(esk3_0,esk2_0)|~v2_struct_0(esk4_0)|u1_struct_0(esk3_0)!=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_700, negated_conjecture, (~m1_pre_topc(esk2_0,esk3_0)|~m1_pre_topc(esk3_0,esk2_0)|~v2_struct_0(esk5_0)|u1_struct_0(esk3_0)!=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_151, plain, (r1_xboole_0(u1_struct_0(X1),k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~r1_tsep_1(k2_tsep_1(X2,X3,X4),X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_152, plain, (u1_struct_0(k2_tsep_1(X1,k2_tsep_1(X2,X3,X4),X5))=k3_xboole_0(k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(X5))|r1_tsep_1(k2_tsep_1(X2,X3,X4),X5)|v2_struct_0(X5)|v2_struct_0(k2_tsep_1(X2,X3,X4))|v2_struct_0(X1)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X5,X1)|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_153, plain, (X1=k1_tsep_1(X2,k2_tsep_1(X3,X4,X5),X6)|v2_struct_0(X1)|v2_struct_0(k2_tsep_1(X3,X4,X5))|v2_struct_0(X6)|v2_struct_0(X2)|r1_tsep_1(X4,X5)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|u1_struct_0(X1)!=k2_xboole_0(k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5)),u1_struct_0(X6))|~v1_pre_topc(X1)|~m1_pre_topc(X6,X2)|~m1_pre_topc(k2_tsep_1(X3,X4,X5),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_913, negated_conjecture, (r1_tsep_1(k1_tsep_1(X1,esk2_0,esk3_0),k2_tsep_1(esk1_0,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_920, plain, (r1_connsp_1(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|~r1_tarski(X3,esk11_3(X4,X2,X3))|~r1_tarski(X2,esk10_3(X4,X2,X3))|~r2_tsep_1(X1,esk10_3(X4,X2,X3),esk11_3(X4,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk11_3(X4,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk10_3(X4,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r1_connsp_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_955, negated_conjecture, (r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),k2_tsep_1(X1,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_919, plain, (r1_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(esk10_3(X3,X1,X2),esk11_3(X3,X1,X2)))|v2_struct_0(X3)|~r1_connsp_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_424, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_178, plain, (m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tarski(u1_struct_0(X1),k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X5)|~m1_pre_topc(X1,X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_179, plain, (r1_xboole_0(u1_struct_0(X1),k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~r1_tsep_1(k1_tsep_1(X2,X3,X4),X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_959, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_156, plain, (r1_xboole_0(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X4)|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_241, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk3_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|~m1_subset_1(esk7_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_560, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r3_tsep_1(esk1_0,esk3_0,esk2_0)|~m1_pre_topc(esk2_0,esk3_0)|~m1_pre_topc(esk3_0,esk2_0))).#trainneg
cnf(i_0_562, negated_conjecture, (m1_pre_topc(esk2_0,esk4_0)|~r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r3_tsep_1(esk1_0,esk3_0,esk2_0)|~m1_pre_topc(esk2_0,esk3_0)|~m1_pre_topc(esk3_0,esk2_0))).#trainneg
cnf(i_0_154, plain, (X1=k2_tsep_1(X2,k2_tsep_1(X3,X4,X5),X6)|r1_tsep_1(k2_tsep_1(X3,X4,X5),X6)|v2_struct_0(X1)|v2_struct_0(k2_tsep_1(X3,X4,X5))|v2_struct_0(X6)|v2_struct_0(X2)|r1_tsep_1(X4,X5)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|u1_struct_0(X1)!=k3_xboole_0(k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5)),u1_struct_0(X6))|~v1_pre_topc(X1)|~m1_pre_topc(X6,X2)|~m1_pre_topc(k2_tsep_1(X3,X4,X5),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_155, plain, (u1_struct_0(k2_tsep_1(X1,X2,k2_tsep_1(X3,X4,X5)))=k3_xboole_0(u1_struct_0(X2),k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|r1_tsep_1(X2,k2_tsep_1(X3,X4,X5))|v2_struct_0(k2_tsep_1(X3,X4,X5))|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X4,X5)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_1107, negated_conjecture, (r1_tsep_1(k2_tsep_1(X1,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1095, negated_conjecture, (r1_tsep_1(k2_tsep_1(esk1_0,esk4_0,esk5_0),k1_tsep_1(X1,esk2_0,esk3_0))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1328, negated_conjecture, (r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r1_tsep_1(esk4_0,esk5_0)|~v1_tsep_1(esk5_0,esk1_0)|~v1_tsep_1(esk4_0,esk1_0)|~m1_pre_topc(esk3_0,esk5_0)|~m1_pre_topc(esk2_0,esk4_0)|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_1156, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~r1_tsep_1(X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0))).#trainneg
cnf(i_0_1117, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0)),u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1152, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)),k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_1321, negated_conjecture, (r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),k2_tsep_1(X1,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1144, plain, (m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X6)|~r1_tarski(u1_struct_0(X1),u1_struct_0(k1_tsep_1(X6,X3,X4)))|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X5)|~m1_pre_topc(X1,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2)|~v2_pre_topc(X5)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X3,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_1116, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0)),u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(esk5_0)|v2_struct_0(esk4_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_321, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|m1_pre_topc(esk2_0,esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk3_0)|u1_struct_0(X1)!=u1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_157, plain, (m1_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k1_zfmisc_1(u1_struct_0(X4)))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1162, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|m1_pre_topc(esk5_0,esk1_0)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk3_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk2_0)|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_808, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(k1_tsep_1(X5,X3,X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X5)|~m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_158, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_xboole_0(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|~l1_struct_0(X4)|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_159, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|~m1_pre_topc(X4,X5)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1, negated_conjecture, (r1_tsep_1(k2_tsep_1(esk1_0,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,esk5_0))).#trainneg
cnf(i_0_1339, negated_conjecture, (v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r1_tsep_1(esk4_0,esk5_0)|~v1_tsep_1(esk5_0,esk1_0)|~v1_tsep_1(esk4_0,esk1_0)|~m1_pre_topc(esk3_0,esk5_0)|~m1_pre_topc(esk2_0,esk4_0)|~m1_pre_topc(esk5_0,esk1_0)|~m1_pre_topc(esk4_0,esk1_0)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_314, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1408, plain, (m1_subset_1(k3_xboole_0(u1_struct_0(X1),u1_struct_0(X2)),k1_zfmisc_1(u1_struct_0(X3)))|r1_tsep_1(X1,X2)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_162, plain, (m1_subset_1(esk6_3(k2_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r3_tsep_1(k2_tsep_1(X1,X2,X3),X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k2_tsep_1(X1,X2,X3))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1157, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|X2!=u1_struct_0(X3)|X1!=u1_struct_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X4,esk1_0)|u1_struct_0(X4)!=u1_struct_0(esk4_0)|u1_struct_0(X3)!=u1_struct_0(esk5_0))).#trainneg
cnf(i_0_1389, plain, (m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X6)|~m1_pre_topc(X1,X6)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X6)|~l1_pre_topc(X2)|~l1_pre_topc(X5)|~v2_pre_topc(X6)|~m1_pre_topc(X1,k1_tsep_1(X5,X3,X4))|~m1_pre_topc(X1,X5)|~v2_pre_topc(X5))).#trainneg
cnf(i_0_160, plain, (k2_tsep_1(X1,X2,X3)=k1_tsep_1(X4,X5,X6)|v2_struct_0(k2_tsep_1(X1,X2,X3))|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6))|~v1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_161, plain, (k2_tsep_1(X1,X2,X3)=k2_tsep_1(X4,X5,X6)|r1_tsep_1(X5,X6)|v2_struct_0(k2_tsep_1(X1,X2,X3))|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=k3_xboole_0(u1_struct_0(X5),u1_struct_0(X6))|~v1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1391, negated_conjecture, (r1_tsep_1(k2_tsep_1(X1,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1368, negated_conjecture, (r1_xboole_0(k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)),u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_1333, negated_conjecture, (r1_tsep_1(k1_tsep_1(X1,esk2_0,esk3_0),k2_tsep_1(esk1_0,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_946, plain, (r1_tsep_1(X1,X2)|r4_tsep_1(X3,k2_tsep_1(X4,X1,X2),X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X6)|~r2_tsep_1(X3,u1_struct_0(k2_tsep_1(X6,X1,X2)),u1_struct_0(X5))|~m1_pre_topc(k2_tsep_1(X4,X1,X2),X3)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_177, plain, (r1_tsep_1(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_xboole_0(u1_struct_0(X1),k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_379, plain, (r1_xboole_0(k4_subset_1(u1_struct_0(X1),X2,X3),k3_xboole_0(esk10_3(X1,X2,X3),esk11_3(X1,X2,X3)))|v2_struct_0(X1)|~r1_connsp_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1469, negated_conjecture, (l1_pre_topc(esk4_0))).#trainneg
cnf(i_0_163, plain, (m1_subset_1(esk7_3(k2_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r3_tsep_1(k2_tsep_1(X1,X2,X3),X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k2_tsep_1(X1,X2,X3))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1482, negated_conjecture, (l1_pre_topc(esk5_0))).#trainneg
cnf(i_0_1550, plain, (m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(X1,k1_tsep_1(X5,X3,X4))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X1,X5)|~l1_pre_topc(X2)|~l1_pre_topc(X5)|~v2_pre_topc(X2)|~v2_pre_topc(X5))).#trainneg
cnf(i_0_784, negated_conjecture, (r1_tarski(u1_struct_0(k1_tsep_1(X4,X2,X3)),u1_struct_0(esk2_0))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk2_0)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_164, plain, (m1_subset_1(esk8_3(k2_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r4_tsep_1(k2_tsep_1(X1,X2,X3),X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k2_tsep_1(X1,X2,X3))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_165, plain, (m1_subset_1(esk9_3(k2_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r4_tsep_1(k2_tsep_1(X1,X2,X3),X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k2_tsep_1(X1,X2,X3))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1473, negated_conjecture, (r1_tsep_1(k1_tsep_1(esk1_0,esk2_0,esk3_0),k2_tsep_1(esk1_0,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_1479, negated_conjecture, (v2_struct_0(esk4_0)|v2_struct_0(esk5_0)|r1_tsep_1(esk4_0,esk5_0)|~v1_tsep_1(esk5_0,esk1_0)|~v1_tsep_1(esk4_0,esk1_0)|~m1_pre_topc(esk5_0,esk1_0))).#trainneg
cnf(i_0_1700, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(esk9_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk8_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1520, plain, (m1_subset_1(u1_struct_0(k2_tsep_1(X4,X1,X2)),k1_zfmisc_1(u1_struct_0(X3)))|r1_tsep_1(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_170, plain, (r1_xboole_0(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~r1_tsep_1(X4,k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1654, plain, (m1_pre_topc(k1_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5),k1_tsep_1(X6,X2,X3))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(k1_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5),X6)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(k1_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5),X1)|~l1_pre_topc(X6)|~l1_pre_topc(X1)|~v2_pre_topc(X6)|~v2_pre_topc(X1)|~m1_pre_topc(X5,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k1_tsep_1(X1,X2,X3))|~l1_pre_topc(k1_tsep_1(X1,X2,X3)))).#trainneg
cnf(i_0_1655, plain, (m1_pre_topc(k2_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5),k1_tsep_1(X6,X2,X3))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(k2_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5),X6)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(k2_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5),X1)|~l1_pre_topc(X6)|~l1_pre_topc(X1)|~v2_pre_topc(X6)|~v2_pre_topc(X1)|~m1_pre_topc(X5,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k1_tsep_1(X1,X2,X3))|~l1_pre_topc(k1_tsep_1(X1,X2,X3)))).#trainneg
cnf(i_0_166, plain, (m1_subset_1(esk10_3(k2_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|v2_struct_0(k2_tsep_1(X1,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k2_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~v2_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_167, plain, (m1_subset_1(esk11_3(k2_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|v2_struct_0(k2_tsep_1(X1,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k2_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~v2_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1699, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|v2_struct_0(X3)|~r1_connsp_1(X3,X1,X2)|~m1_subset_1(esk11_3(X3,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk10_3(X3,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|esk11_3(X3,X1,X2)!=u1_struct_0(esk5_0)|esk10_3(X3,X1,X2)!=u1_struct_0(esk4_0))).#trainneg
cnf(i_0_1553, plain, (k2_tsep_1(X1,X2,X3)=k1_tsep_1(X1,X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X4)|k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1556, plain, (k2_tsep_1(X1,X2,X3)=k2_tsep_1(X1,X4,X5)|r1_tsep_1(X2,X3)|r1_tsep_1(X4,X5)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X4)|k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_310, plain, (r4_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r2_tsep_1(X1,u1_struct_0(X2),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_311, plain, (r4_tsep_1(X1,X2,k2_tsep_1(X3,X4,X5))|r1_tsep_1(X4,X5)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r2_tsep_1(X1,u1_struct_0(X2),k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|~m1_pre_topc(k2_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_497, plain, (m1_pre_topc(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_tarski(u1_struct_0(X1),u1_struct_0(k2_tsep_1(X6,X3,X4)))|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X5)|~m1_pre_topc(X1,X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X3,X6)|~l1_pre_topc(X2)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_331, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~r3_tsep_1(X3,X1,X2)|~m1_pre_topc(X1,X1)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)|~v2_pre_topc(esk12_3(X3,X1,X2))|~v2_pre_topc(X3))).#trainneg
cnf(i_0_1539, negated_conjecture, (r2_tsep_1(esk1_0,X1,u1_struct_0(X2))|u1_struct_0(X3)!=u1_struct_0(esk4_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|X1!=u1_struct_0(X3)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X3,esk1_0))).#trainneg
cnf(i_0_172, plain, (X1=k1_tsep_1(X2,X3,X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X5)|u1_struct_0(X1)!=u1_struct_0(k1_tsep_1(X5,X3,X4))|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_341, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~r3_tsep_1(X3,X2,X1)|~m1_pre_topc(X1,X1)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)|~v2_pre_topc(esk13_3(X3,X2,X1))|~v2_pre_topc(X3))).#trainneg
cnf(i_0_785, negated_conjecture, (r1_tarski(u1_struct_0(k1_tsep_1(X4,X2,X3)),u1_struct_0(esk3_0))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk3_0)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_168, plain, (r1_connsp_1(k2_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(k2_tsep_1(X1,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(X5,X6)|~r1_tarski(X4,X7)|~r2_tsep_1(k2_tsep_1(X1,X2,X3),X7,X6)|~m1_subset_1(X5,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X7,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~r1_xboole_0(k9_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X7,X6),k4_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~v2_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_169, plain, (r1_xboole_0(k9_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),esk10_3(k2_tsep_1(X1,X2,X3),X4,X5),esk11_3(k2_tsep_1(X1,X2,X3),X4,X5)),k4_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))|v2_struct_0(k2_tsep_1(X1,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k2_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~v2_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1938, plain, (k1_tsep_1(X1,X2,X3)=k1_tsep_1(X4,X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(k1_tsep_1(X1,X2,X3))|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1705, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(esk9_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_378, plain, (r1_connsp_1(X1,X2,X3)|v2_struct_0(X1)|r1_tsep_1(X5,X4)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~r1_tarski(X3,u1_struct_0(X4))|~r1_tarski(X2,u1_struct_0(X5))|~r2_tsep_1(X1,u1_struct_0(X5),u1_struct_0(X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X5),k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(u1_struct_0(k2_tsep_1(X6,X5,X4)),k4_subset_1(u1_struct_0(X1),X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X5,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_1806, plain, (k2_tsep_1(X1,X2,X3)=k1_tsep_1(X1,X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(k1_tsep_1(X6,X4,X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_1623, negated_conjecture, (r1_tsep_1(k2_tsep_1(X1,esk4_0,esk5_0),k1_tsep_1(X2,esk2_0,esk3_0))|v2_struct_0(X2)|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(X2,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0))|~m1_pre_topc(esk3_0,X2)|~m1_pre_topc(esk2_0,X2)|~l1_pre_topc(X2)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_186, plain, (r1_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X4)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_189, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|~m1_pre_topc(X4,X5)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_206, plain, (r4_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_tsep_1(X1,k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),esk9_3(X1,k1_tsep_1(X2,X3,X4),X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_746, plain, (r1_tsep_1(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_xboole_0(u1_struct_0(X1),u1_struct_0(k1_tsep_1(X5,X3,X4)))|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2028, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_176, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4))))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|~l1_pre_topc(k1_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1928, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(X3,k2_tsep_1(X4,X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(k2_tsep_1(X4,X1,X2),X6)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X1,X5)|~l1_pre_topc(X6)|~l1_pre_topc(X4)|~l1_pre_topc(X5)|~v2_pre_topc(X6)|~m1_pre_topc(X3,k2_tsep_1(X5,X1,X2))|~m1_pre_topc(X3,X5)|~v2_pre_topc(X5))).#trainneg
cnf(i_0_173, plain, (X1=k1_tsep_1(X2,X3,k1_tsep_1(X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k1_tsep_1(X4,X5,X6))|v2_struct_0(X2)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|u1_struct_0(X1)!=k2_xboole_0(u1_struct_0(X3),k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|~v1_pre_topc(X1)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_174, plain, (X1=k2_tsep_1(X2,X3,k1_tsep_1(X4,X5,X6))|r1_tsep_1(X3,k1_tsep_1(X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(k1_tsep_1(X4,X5,X6))|v2_struct_0(X2)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|u1_struct_0(X1)!=k3_xboole_0(u1_struct_0(X3),k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|~v1_pre_topc(X1)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_475, plain, (r1_xboole_0(u1_struct_0(k2_tsep_1(X5,X2,X3)),u1_struct_0(X4))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~r1_tsep_1(X4,k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1439, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X5)|~r1_xboole_0(u1_struct_0(k2_tsep_1(X5,X2,X3)),u1_struct_0(X4))|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_1470, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk2_0))|~m1_pre_topc(X1,esk4_0)|~m1_pre_topc(X1,esk2_0)|~l1_pre_topc(esk4_0)|~v2_pre_topc(esk4_0))).#trainneg
cnf(i_0_207, plain, (r4_tsep_1(X1,k2_tsep_1(X2,X3,X4),X5)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_tsep_1(X1,k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),esk9_3(X1,k2_tsep_1(X2,X3,X4),X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2144, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(X3,k2_tsep_1(X4,X1,X2))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X3,k2_tsep_1(X5,X1,X2))|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X1,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X4)|~l1_pre_topc(X5)|~v2_pre_topc(X4)|~v2_pre_topc(X5))).#trainneg
cnf(i_0_2211, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0)),u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2135, plain, (r1_tsep_1(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X2)|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2)|~l1_struct_0(k1_tsep_1(X5,X3,X4))|~r1_tsep_1(X1,k1_tsep_1(X5,X3,X4)))).#trainneg
cnf(i_0_275, plain, (r1_tarski(u1_struct_0(X1),k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|r1_tsep_1(X3,X4)|~m1_pre_topc(X1,k2_tsep_1(X2,X3,X4))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1466, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk3_0))|~m1_pre_topc(X1,esk5_0)|~m1_pre_topc(X1,esk3_0)|~l1_pre_topc(esk5_0)|~v2_pre_topc(esk5_0))).#trainneg
cnf(i_0_772, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X3)|~r1_tarski(u1_struct_0(k1_tsep_1(X6,X2,X3)),u1_struct_0(X4))|~m1_pre_topc(X4,X5)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X5)|~l1_pre_topc(X5)|~v2_pre_topc(X5)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X6)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1472, negated_conjecture, (r3_tsep_1(esk1_0,k2_tsep_1(esk1_0,esk4_0,esk5_0),k1_tsep_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk4_0)|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0))).#trainneg
cnf(i_0_1495, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_180, plain, (u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5))=k2_xboole_0(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(X5))|v2_struct_0(X5)|v2_struct_0(k1_tsep_1(X2,X3,X4))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1693, negated_conjecture, (r3_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk2_0,esk3_0),k2_tsep_1(esk1_0,esk4_0,esk5_0))|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|r1_tsep_1(esk4_0,esk5_0)|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk4_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk5_0)|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_2136, plain, (r1_tsep_1(X1,k1_tsep_1(X2,X3,X4))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X2)|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2)|~l1_struct_0(k1_tsep_1(X5,X3,X4))|~r1_tsep_1(k1_tsep_1(X5,X3,X4),X1))).#trainneg
cnf(i_0_2228, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~l1_pre_topc(X1)|~l1_pre_topc(X5)|~l1_struct_0(k2_tsep_1(X5,X2,X3))|~r1_tsep_1(k2_tsep_1(X5,X2,X3),X4))).#trainneg
cnf(i_0_181, plain, (u1_struct_0(k2_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5))=k3_xboole_0(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(X5))|r1_tsep_1(k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X5)|v2_struct_0(k1_tsep_1(X2,X3,X4))|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_182, plain, (X1=k1_tsep_1(X2,k1_tsep_1(X3,X4,X5),X6)|v2_struct_0(X1)|v2_struct_0(k1_tsep_1(X3,X4,X5))|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|u1_struct_0(X1)!=k2_xboole_0(k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)),u1_struct_0(X6))|~v1_pre_topc(X1)|~m1_pre_topc(X6,X2)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_2549, plain, (r1_tsep_1(k2_tsep_1(X1,esk12_3(X2,X3,X4),esk13_3(X2,X3,X4)),k1_tsep_1(X2,X3,X4))|r1_tsep_1(esk12_3(X2,X3,X4),esk13_3(X2,X3,X4))|v2_struct_0(X1)|v2_struct_0(esk12_3(X2,X3,X4))|v2_struct_0(esk13_3(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~l1_struct_0(k2_tsep_1(X1,esk12_3(X2,X3,X4),esk13_3(X2,X3,X4)))|~l1_struct_0(k2_tsep_1(X2,esk12_3(X2,X3,X4),esk13_3(X2,X3,X4)))|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~m1_pre_topc(esk13_3(X2,X3,X4),X1)|~m1_pre_topc(esk12_3(X2,X3,X4),X1)|~m1_pre_topc(esk13_3(X2,X3,X4),X2)|~m1_pre_topc(esk12_3(X2,X3,X4),X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~r3_tsep_1(X2,X3,X4)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_184, plain, (u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5)))=k2_xboole_0(u1_struct_0(X2),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|v2_struct_0(k1_tsep_1(X3,X4,X5))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_2229, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~l1_pre_topc(X1)|~l1_pre_topc(X5)|~l1_struct_0(k2_tsep_1(X5,X2,X3))|~r1_tsep_1(X4,k2_tsep_1(X5,X2,X3)))).#trainneg
cnf(i_0_2312, negated_conjecture, (m1_pre_topc(k1_tsep_1(X1,X2,X3),esk3_0)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X5)|~m1_pre_topc(esk3_0,X5)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X5)|~m1_pre_topc(k1_tsep_1(X4,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X4,X2,X3),esk3_0))).#trainneg
cnf(i_0_183, plain, (X1=k2_tsep_1(X2,k1_tsep_1(X3,X4,X5),X6)|r1_tsep_1(k1_tsep_1(X3,X4,X5),X6)|v2_struct_0(X1)|v2_struct_0(k1_tsep_1(X3,X4,X5))|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|u1_struct_0(X1)!=k3_xboole_0(k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)),u1_struct_0(X6))|~v1_pre_topc(X1)|~m1_pre_topc(X6,X2)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_185, plain, (u1_struct_0(k2_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5)))=k3_xboole_0(u1_struct_0(X2),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|r1_tsep_1(X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(k1_tsep_1(X3,X4,X5))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_1730, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(esk8_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_1465, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),X4)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X6)|~r1_tarski(u1_struct_0(k2_tsep_1(X6,X1,X2)),u1_struct_0(X4))|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X5)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X5)|~l1_pre_topc(X3)|~v2_pre_topc(X5)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_2298, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X7)|v2_struct_0(X1)|v2_struct_0(X9)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~r1_tarski(u1_struct_0(k1_tsep_1(X7,X2,X3)),u1_struct_0(k1_tsep_1(X9,X5,X6)))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X8)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X8)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X8)|~l1_pre_topc(X7)|~l1_pre_topc(X1)|~v2_pre_topc(X8)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(X6,X9)|~m1_pre_topc(X5,X9)|~l1_pre_topc(X4)|~l1_pre_topc(X9))).#trainneg
cnf(i_0_188, plain, (r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|~l1_struct_0(X4)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2319, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|X2!=u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|X1!=u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk4_0)|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk5_0))).#trainneg
cnf(i_0_187, plain, (m1_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k1_zfmisc_1(u1_struct_0(X4)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2314, negated_conjecture, (m1_pre_topc(k1_tsep_1(X1,X2,X3),esk2_0)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X5)|~m1_pre_topc(esk2_0,X5)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X5)|~m1_pre_topc(k1_tsep_1(X4,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X4,X2,X3),esk2_0))).#trainneg
cnf(i_0_2316, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X7)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X8)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X8)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X8)|~l1_pre_topc(X7)|~l1_pre_topc(X1)|~v2_pre_topc(X8)|~m1_pre_topc(k1_tsep_1(X7,X2,X3),k1_tsep_1(X4,X5,X6))|~m1_pre_topc(k1_tsep_1(X7,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_190, plain, (k1_tsep_1(X1,X2,X3)=k1_tsep_1(X4,X5,X6)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6))|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_191, plain, (k1_tsep_1(X1,X2,X3)=k2_tsep_1(X4,X5,X6)|r1_tsep_1(X5,X6)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=k3_xboole_0(u1_struct_0(X5),u1_struct_0(X6))|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1701, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk8_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1823, plain, (k2_tsep_1(X1,X2,X3)=k1_tsep_1(X1,X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|u1_struct_0(k2_tsep_1(X6,X2,X3))!=k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_1858, plain, (k2_tsep_1(X1,X2,X3)=k2_tsep_1(X1,X4,X5)|r1_tsep_1(X4,X5)|r1_tsep_1(X2,X3)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|u1_struct_0(k2_tsep_1(X6,X2,X3))!=k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_3413, plain, (m1_subset_1(k2_xboole_0(u1_struct_0(X1),u1_struct_0(X2)),k1_zfmisc_1(u1_struct_0(X3)))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_3418, plain, (k1_tsep_1(X1,X2,X3)=k1_tsep_1(X1,X4,X5)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X4)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3330, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|esk9_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_750, plain, (r4_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X2)|v2_struct_0(X6)|v2_struct_0(X3)|v2_struct_0(X4)|~r2_tsep_1(X1,u1_struct_0(k1_tsep_1(X6,X3,X4)),esk9_3(X1,k1_tsep_1(X2,X3,X4),X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X6)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1942, plain, (k1_tsep_1(X1,X2,X3)=k1_tsep_1(X4,X5,X6)|v2_struct_0(X7)|v2_struct_0(X4)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(k1_tsep_1(X7,X5,X6))|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X7)|~m1_pre_topc(X5,X7)|~l1_pre_topc(X4)|~l1_pre_topc(X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_315, plain, (r3_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_connsp_1(X1,u1_struct_0(X2),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_3524, plain, (m1_subset_1(u1_struct_0(k1_tsep_1(X4,X1,X2)),k1_zfmisc_1(u1_struct_0(X3)))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_2541, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|X2!=u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|X1!=u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk4_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk5_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_3576, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|esk8_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_192, plain, (m1_subset_1(esk6_3(k1_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r3_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k1_tsep_1(X1,X2,X3))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_571, negated_conjecture, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)),k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|v2_struct_0(esk5_0)|v2_struct_0(esk4_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3354, negated_conjecture, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),esk3_0)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X5)|~m1_pre_topc(esk3_0,X5)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~v2_pre_topc(X5)|~m1_pre_topc(k2_tsep_1(X4,X1,X2),esk1_0)|~m1_pre_topc(k2_tsep_1(X4,X1,X2),esk3_0))).#trainneg
cnf(i_0_3356, negated_conjecture, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),esk2_0)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X5)|~m1_pre_topc(esk2_0,X5)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~v2_pre_topc(X5)|~m1_pre_topc(k2_tsep_1(X4,X1,X2),esk1_0)|~m1_pre_topc(k2_tsep_1(X4,X1,X2),esk2_0))).#trainneg
cnf(i_0_193, plain, (m1_subset_1(esk7_3(k1_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r3_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k1_tsep_1(X1,X2,X3))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_194, plain, (m1_subset_1(esk8_3(k1_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r4_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k1_tsep_1(X1,X2,X3))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3575, plain, (k1_tsep_1(X1,X2,X3)=k1_tsep_1(X1,X4,X5)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|u1_struct_0(k1_tsep_1(X6,X2,X3))!=k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_3628, negated_conjecture, (r4_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|u1_struct_0(X4)!=u1_struct_0(esk5_0)|u1_struct_0(k1_tsep_1(X1,X2,X3))!=u1_struct_0(esk4_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_820, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(k1_tsep_1(X5,X3,X4)))|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~l1_struct_0(X1)|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~r1_tsep_1(k1_tsep_1(X2,X3,X4),X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X2)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_744, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(k1_tsep_1(X5,X3,X4)))|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~l1_struct_0(k1_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~r1_tsep_1(X1,k1_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_961, plain, (r1_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k3_xboole_0(u1_struct_0(esk12_3(X1,X2,X3)),u1_struct_0(esk13_3(X1,X2,X3))))|r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(esk12_3(X1,X2,X3))|v2_struct_0(esk13_3(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(esk13_3(X1,X2,X3),X1)|~m1_pre_topc(esk12_3(X1,X2,X3),X1)|~l1_pre_topc(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3680, negated_conjecture, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)),k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_766, plain, (r4_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X4)|v2_struct_0(X5)|~r2_tsep_1(X1,u1_struct_0(X2),u1_struct_0(k1_tsep_1(X6,X4,X5)))|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X6)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_212, plain, (r3_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_connsp_1(X1,esk6_3(X1,X2,k1_tsep_1(X3,X4,X5)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_195, plain, (m1_subset_1(esk9_3(k1_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|r4_tsep_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X5,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X4,k1_tsep_1(X1,X2,X3))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2317, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k2_tsep_1(X4,X5,X6))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X7)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X8)|~m1_pre_topc(k2_tsep_1(X4,X5,X6),X8)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X8)|~l1_pre_topc(X7)|~l1_pre_topc(X1)|~v2_pre_topc(X8)|~m1_pre_topc(k1_tsep_1(X7,X2,X3),k2_tsep_1(X4,X5,X6))|~m1_pre_topc(k1_tsep_1(X7,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_3358, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),k1_tsep_1(X4,X5,X6))|v2_struct_0(X7)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X8)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X8)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X1,X7)|~l1_pre_topc(X8)|~l1_pre_topc(X3)|~l1_pre_topc(X7)|~v2_pre_topc(X8)|~m1_pre_topc(k2_tsep_1(X7,X1,X2),k1_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X7,X1,X2),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_196, plain, (m1_subset_1(esk10_3(k1_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~v2_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_197, plain, (m1_subset_1(esk11_3(k1_tsep_1(X1,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~v2_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3788, negated_conjecture, (r1_xboole_0(k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_1128, plain, (r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(X3))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|r1_tsep_1(X6,X5)|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X7)|~r1_tarski(u1_struct_0(X3),u1_struct_0(X5))|~r1_tarski(u1_struct_0(X2),u1_struct_0(X6))|~r2_tsep_1(X1,u1_struct_0(X6),u1_struct_0(X5))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X5),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X6),k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(u1_struct_0(k2_tsep_1(X7,X6,X5)),u1_struct_0(k1_tsep_1(X4,X2,X3)))|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~v2_pre_topc(X1)|~m1_pre_topc(X5,X7)|~m1_pre_topc(X6,X7)|~l1_pre_topc(X7))).#trainneg
cnf(i_0_2330, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_225, plain, (r4_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r2_tsep_1(X1,esk8_3(X1,X2,k1_tsep_1(X3,X4,X5)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_570, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0)),u1_struct_0(k1_tsep_1(X2,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|r3_tsep_1(esk1_0,esk2_0,esk3_0)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~m1_pre_topc(esk3_0,X2)|~m1_pre_topc(esk2_0,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3646, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk4_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk5_0)|X2!=u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|X1!=u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_200, plain, (r1_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(X4))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~r1_tsep_1(X4,k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3359, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),k2_tsep_1(X4,X5,X6))|v2_struct_0(X7)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X8)|~m1_pre_topc(k2_tsep_1(X4,X5,X6),X8)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X1,X7)|~l1_pre_topc(X8)|~l1_pre_topc(X3)|~l1_pre_topc(X7)|~v2_pre_topc(X8)|~m1_pre_topc(k2_tsep_1(X7,X1,X2),k2_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X7,X1,X2),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_2266, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k1_tsep_1(k2_tsep_1(X3,X1,X2),X4,X5),k2_tsep_1(X6,X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X3)|v2_struct_0(k2_tsep_1(X3,X1,X2))|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(k1_tsep_1(k2_tsep_1(X3,X1,X2),X4,X5),X6)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(k1_tsep_1(k2_tsep_1(X3,X1,X2),X4,X5),X3)|~l1_pre_topc(X6)|~l1_pre_topc(X3)|~v2_pre_topc(X6)|~v2_pre_topc(X3)|~m1_pre_topc(X5,k2_tsep_1(X3,X1,X2))|~m1_pre_topc(X4,k2_tsep_1(X3,X1,X2))|~l1_pre_topc(k2_tsep_1(X3,X1,X2)))).#trainneg
cnf(i_0_198, plain, (r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(X5,X6)|~r1_tarski(X4,X7)|~r2_tsep_1(k1_tsep_1(X1,X2,X3),X7,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X7,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~r1_xboole_0(k9_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X7,X6),k4_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~v2_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_199, plain, (r1_xboole_0(k9_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),esk10_3(k1_tsep_1(X1,X2,X3),X4,X5),esk11_3(k1_tsep_1(X1,X2,X3),X4,X5)),k4_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X4,X5))|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~v2_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_317, plain, (r3_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_connsp_1(X1,k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3791, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0)),k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3577, plain, (r4_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X2)|~r2_tsep_1(X1,u1_struct_0(k1_tsep_1(X6,X3,X4)),u1_struct_0(X5))|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X6)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3980, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk4_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk5_0)|X2!=u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|X1!=u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)|~l1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_201, plain, (u1_struct_0(k1_tsep_1(X1,X2,k2_tsep_1(X3,X4,X5)))=k2_xboole_0(u1_struct_0(X2),k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|v2_struct_0(k2_tsep_1(X3,X4,X5))|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X4,X5)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_279, plain, (r1_tarski(u1_struct_0(X1),k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_2267, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(k2_tsep_1(X3,X1,X2),X4,X5),k2_tsep_1(X6,X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X3)|v2_struct_0(k2_tsep_1(X3,X1,X2))|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(k2_tsep_1(k2_tsep_1(X3,X1,X2),X4,X5),X6)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(k2_tsep_1(k2_tsep_1(X3,X1,X2),X4,X5),X3)|~l1_pre_topc(X6)|~l1_pre_topc(X3)|~v2_pre_topc(X6)|~v2_pre_topc(X3)|~m1_pre_topc(X5,k2_tsep_1(X3,X1,X2))|~m1_pre_topc(X4,k2_tsep_1(X3,X1,X2))|~l1_pre_topc(k2_tsep_1(X3,X1,X2)))).#trainneg
cnf(i_0_3370, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X7)|v2_struct_0(X1)|v2_struct_0(X8)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X9)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X9)|~m1_pre_topc(X3,X8)|~m1_pre_topc(X2,X8)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(X6,X7)|~m1_pre_topc(X5,X7)|~l1_pre_topc(X9)|~l1_pre_topc(X8)|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~l1_pre_topc(X7)|~v2_pre_topc(X9)|~m1_pre_topc(k1_tsep_1(X8,X2,X3),k1_tsep_1(X7,X5,X6))|~m1_pre_topc(k1_tsep_1(X8,X2,X3),X7)|~v2_pre_topc(X7))).#trainneg
cnf(i_0_202, plain, (u1_struct_0(k1_tsep_1(X1,k2_tsep_1(X2,X3,X4),X5))=k2_xboole_0(k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(X5))|v2_struct_0(X5)|v2_struct_0(k2_tsep_1(X2,X3,X4))|v2_struct_0(X1)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X5,X1)|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_213, plain, (r3_tsep_1(X1,X2,k2_tsep_1(X3,X4,X5))|r1_tsep_1(X4,X5)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_connsp_1(X1,esk6_3(X1,X2,k2_tsep_1(X3,X4,X5)),k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|~m1_pre_topc(k2_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_2064, negated_conjecture, (r1_tsep_1(k1_tsep_1(X1,esk2_0,esk3_0),k2_tsep_1(X2,esk4_0,esk5_0))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(X2,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~m1_pre_topc(esk5_0,X2)|~m1_pre_topc(esk4_0,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1710, negated_conjecture, (r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(esk4_0))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk4_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_552, plain, (k1_tsep_1(X1,X2,X3)=k2_tsep_1(X4,X5,X6)|r1_tsep_1(X5,X6)|v2_struct_0(X7)|v2_struct_0(X4)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(k2_tsep_1(X7,X5,X6))|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X7)|~m1_pre_topc(X5,X7)|~l1_pre_topc(X4)|~l1_pre_topc(X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4159, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk4_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk5_0)|X2!=u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|X1!=u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0)|~l1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0)))).#trainneg
cnf(i_0_217, negated_conjecture, (r1_tarski(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(esk3_0))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk3_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_220, negated_conjecture, (r1_tarski(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(esk2_0))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk2_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3415, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X1)|v2_struct_0(k1_tsep_1(X4,X5,X6))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X7)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X7)|~m1_pre_topc(k1_tsep_1(k1_tsep_1(X4,X5,X6),X2,X3),X4)|~m1_pre_topc(X3,k1_tsep_1(X4,X5,X6))|~m1_pre_topc(X2,k1_tsep_1(X4,X5,X6))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X7)|~l1_pre_topc(k1_tsep_1(X4,X5,X6))|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~v2_pre_topc(X7)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_3866, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k2_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X1)|v2_struct_0(k2_tsep_1(X4,X5,X6))|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X7)|~m1_pre_topc(k2_tsep_1(X4,X5,X6),X7)|~m1_pre_topc(k1_tsep_1(k2_tsep_1(X4,X5,X6),X2,X3),X4)|~m1_pre_topc(X3,k2_tsep_1(X4,X5,X6))|~m1_pre_topc(X2,k2_tsep_1(X4,X5,X6))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X7)|~l1_pre_topc(k2_tsep_1(X4,X5,X6))|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~v2_pre_topc(X7)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_226, plain, (r4_tsep_1(X1,X2,k2_tsep_1(X3,X4,X5))|r1_tsep_1(X4,X5)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r2_tsep_1(X1,esk8_3(X1,X2,k2_tsep_1(X3,X4,X5)),k3_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|~m1_pre_topc(k2_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_235, plain, (r2_tsep_1(X1,X2,X3)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|X3!=u1_struct_0(esk12_3(X1,X4,X5))|X2!=u1_struct_0(X6)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(esk12_3(X1,X4,X5),X1)|~m1_pre_topc(X6,X1)|~l1_pre_topc(X1)|~v1_tsep_1(X6,X1)|~r3_tsep_1(X1,X4,X5)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1704, negated_conjecture, (r1_tsep_1(X1,X2)|r4_tsep_1(esk1_0,k2_tsep_1(X3,X1,X2),X4)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X3,X1,X2),esk1_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X1,X5)|~l1_pre_topc(X3)|~l1_pre_topc(X5)|u1_struct_0(X4)!=u1_struct_0(esk5_0)|u1_struct_0(k2_tsep_1(X5,X1,X2))!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(k2_tsep_1(X5,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2297, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X7)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|~r1_tarski(u1_struct_0(k1_tsep_1(X7,X2,X3)),k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X8)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X8)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X8)|~l1_pre_topc(X7)|~l1_pre_topc(X1)|~v2_pre_topc(X8)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_4850, plain, (r2_tsep_1(X1,X2,u1_struct_0(esk12_3(X1,X3,X4)))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|X2!=u1_struct_0(X5)|~v1_tsep_1(X5,X1)|~m1_subset_1(u1_struct_0(esk12_3(X1,X3,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1748, plain, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk12_3(X1,X2,X3)),u1_struct_0(esk13_3(X1,X2,X3))),u1_struct_0(k1_tsep_1(X1,X2,X3)))|r1_tsep_1(esk12_3(X1,X2,X3),esk13_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(esk12_3(X1,X2,X3))|v2_struct_0(esk13_3(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|~l1_struct_0(k2_tsep_1(X1,esk12_3(X1,X2,X3),esk13_3(X1,X2,X3)))|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(esk13_3(X1,X2,X3),X1)|~m1_pre_topc(esk12_3(X1,X2,X3),X1)|~l1_pre_topc(X1)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1943, plain, (k1_tsep_1(X1,X2,X3)=k1_tsep_1(X4,X5,X6)|v2_struct_0(X7)|v2_struct_0(X4)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X8)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|u1_struct_0(k1_tsep_1(X8,X2,X3))!=u1_struct_0(k1_tsep_1(X7,X5,X6))|~v1_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X7)|~m1_pre_topc(X5,X7)|~l1_pre_topc(X4)|~l1_pre_topc(X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X8)|~m1_pre_topc(X2,X8)|~l1_pre_topc(X1)|~l1_pre_topc(X8))).#trainneg
cnf(i_0_1721, negated_conjecture, (r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(esk5_0))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk5_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4795, plain, (r1_tsep_1(X1,X2)|r3_tsep_1(X3,X4,k2_tsep_1(X5,X1,X2))|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|~r1_connsp_1(X3,u1_struct_0(X4),k3_xboole_0(u1_struct_0(X1),u1_struct_0(X2)))|~m1_pre_topc(k2_tsep_1(X5,X1,X2),X3)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X1,X5)|~l1_pre_topc(X3)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_4814, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk4_0)|u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))!=u1_struct_0(esk5_0)|X2!=u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|X1!=u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k2_tsep_1(esk1_0,esk4_0,esk5_0),esk1_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk2_0,esk3_0),esk1_0))).#trainneg
cnf(i_0_264, plain, (r3_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_connsp_1(X1,k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),esk7_3(X1,k1_tsep_1(X2,X3,X4),X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1361, negated_conjecture, (r1_xboole_0(u1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3867, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),k1_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k1_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X7)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X7)|~m1_pre_topc(k2_tsep_1(k1_tsep_1(X4,X5,X6),X1,X2),X4)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,k1_tsep_1(X4,X5,X6))|~m1_pre_topc(X1,k1_tsep_1(X4,X5,X6))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X7)|~l1_pre_topc(X3)|~l1_pre_topc(k1_tsep_1(X4,X5,X6))|~l1_pre_topc(X4)|~v2_pre_topc(X7)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_3987, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),k2_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k2_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X7)|~m1_pre_topc(k2_tsep_1(X4,X5,X6),X7)|~m1_pre_topc(k2_tsep_1(k2_tsep_1(X4,X5,X6),X1,X2),X4)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,k2_tsep_1(X4,X5,X6))|~m1_pre_topc(X1,k2_tsep_1(X4,X5,X6))|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X7)|~l1_pre_topc(X3)|~l1_pre_topc(k2_tsep_1(X4,X5,X6))|~l1_pre_topc(X4)|~v2_pre_topc(X7)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_237, plain, (r1_xboole_0(k9_subset_1(u1_struct_0(X1),esk10_3(X1,u1_struct_0(X2),u1_struct_0(X3)),esk11_3(X1,u1_struct_0(X2),u1_struct_0(X3))),u1_struct_0(k1_tsep_1(X4,X2,X3)))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_238, plain, (r1_xboole_0(k9_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),esk10_3(k1_tsep_1(X1,X2,X3),X4,X5),esk11_3(k1_tsep_1(X1,X2,X3),X4,X5)),k2_xboole_0(X4,X5))|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~v2_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5003, plain, (r1_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k9_subset_1(u1_struct_0(X4),esk10_3(X4,u1_struct_0(X2),u1_struct_0(X3)),esk11_3(X4,u1_struct_0(X2),u1_struct_0(X3))))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_connsp_1(X4,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X4)))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_759, plain, (r4_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X4)|v2_struct_0(X5)|~r2_tsep_1(X1,esk8_3(X1,X2,k1_tsep_1(X3,X4,X5)),u1_struct_0(k1_tsep_1(X6,X4,X5)))|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X6)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_1145, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X7)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X7)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X7)|~l1_pre_topc(X4)|~v2_pre_topc(X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_265, plain, (r3_tsep_1(X1,k2_tsep_1(X2,X3,X4),X5)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_connsp_1(X1,k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),esk7_3(X1,k2_tsep_1(X2,X3,X4),X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1567, negated_conjecture, (r1_xboole_0(k2_xboole_0(u1_struct_0(esk2_0),u1_struct_0(esk3_0)),u1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4461, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(k1_tsep_1(X7,X5,X6))|v2_struct_0(X1)|v2_struct_0(X7)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X8)|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X8)|~m1_pre_topc(k1_tsep_1(k1_tsep_1(X7,X5,X6),X2,X3),X7)|~m1_pre_topc(X3,k1_tsep_1(X7,X5,X6))|~m1_pre_topc(X2,k1_tsep_1(X7,X5,X6))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(X6,X7)|~m1_pre_topc(X5,X7)|~l1_pre_topc(X8)|~l1_pre_topc(k1_tsep_1(X7,X5,X6))|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~l1_pre_topc(X7)|~v2_pre_topc(X8)|~v2_pre_topc(X7))).#trainneg
cnf(i_0_239, plain, (r1_xboole_0(k9_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),esk10_3(k2_tsep_1(X1,X2,X3),X4,X5),esk11_3(k2_tsep_1(X1,X2,X3),X4,X5)),k2_xboole_0(X4,X5))|v2_struct_0(k2_tsep_1(X1,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k2_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~v2_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_247, plain, (r2_tsep_1(X1,X2,X3)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|X3!=u1_struct_0(esk13_3(X1,X4,X5))|X2!=u1_struct_0(X6)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_pre_topc(esk13_3(X1,X4,X5),X1)|~m1_pre_topc(X6,X1)|~l1_pre_topc(X1)|~v1_tsep_1(X6,X1)|~r3_tsep_1(X1,X4,X5)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5190, plain, (r1_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k3_xboole_0(esk10_3(X4,u1_struct_0(X2),u1_struct_0(X3)),esk11_3(X4,u1_struct_0(X2),u1_struct_0(X3))))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(X4,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X4)))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X4)|~m1_subset_1(esk11_3(X4,u1_struct_0(X2),u1_struct_0(X3)),k1_zfmisc_1(u1_struct_0(X4))))).#trainneg
cnf(i_0_5238, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X8)|~r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X8,X5,X6)))|~m1_pre_topc(k1_tsep_1(X4,X5,X6),X7)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X7)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X7)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(X7)|~m1_pre_topc(X6,X8)|~m1_pre_topc(X5,X8)|~l1_pre_topc(X8))).#trainneg
cnf(i_0_5083, plain, (r1_xboole_0(k2_xboole_0(X1,X2),k9_subset_1(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),esk10_3(k1_tsep_1(X5,X3,X4),X1,X2),esk11_3(k1_tsep_1(X5,X3,X4),X1,X2)))|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_connsp_1(k1_tsep_1(X5,X3,X4),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X1,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4))))|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X5)|~v2_pre_topc(k1_tsep_1(X5,X3,X4)))).#trainneg
cnf(i_0_4013, plain, (r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~r1_tarski(X5,X6)|~r1_tarski(X4,X7)|~r2_tsep_1(k1_tsep_1(X1,X2,X3),X7,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X7,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~r1_xboole_0(k9_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X7,X6),k2_xboole_0(X4,X5))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(k1_tsep_1(X1,X2,X3)))).#trainneg
cnf(i_0_461, plain, (r1_xboole_0(u1_struct_0(k2_tsep_1(X5,X2,X3)),u1_struct_0(X4))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(X4)|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5268, plain, (r1_tsep_1(X1,X2)|r3_tsep_1(X3,k2_tsep_1(X4,X1,X2),X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|~r1_connsp_1(X3,k3_xboole_0(u1_struct_0(X1),u1_struct_0(X2)),u1_struct_0(X5))|~m1_pre_topc(k2_tsep_1(X4,X1,X2),X3)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~l1_pre_topc(X3)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_476, negated_conjecture, (r1_tarski(u1_struct_0(k2_tsep_1(X4,X2,X3)),u1_struct_0(esk2_0))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk2_0)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_477, negated_conjecture, (r1_tarski(u1_struct_0(k2_tsep_1(X4,X2,X3)),u1_struct_0(esk3_0))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),esk3_0)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_276, plain, (r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k2_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),k2_tsep_1(X4,X5,X6))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_277, plain, (r1_tarski(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k2_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),k2_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5342, plain, (r2_tsep_1(X1,X2,u1_struct_0(esk13_3(X1,X3,X4)))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|X2!=u1_struct_0(X5)|~v1_tsep_1(X5,X1)|~m1_subset_1(u1_struct_0(esk13_3(X1,X3,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3613, plain, (r3_tsep_1(X1,X2,k1_tsep_1(X3,X4,X5))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~r1_connsp_1(X1,u1_struct_0(X2),u1_struct_0(k1_tsep_1(X6,X4,X5)))|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_2121, negated_conjecture, (r4_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|esk9_3(esk1_0,k1_tsep_1(X1,X2,X3),X4)!=u1_struct_0(esk5_0)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(esk4_0)|~m1_subset_1(esk9_3(esk1_0,k1_tsep_1(X1,X2,X3),X4),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_336, negated_conjecture, (r1_tarski(k2_xboole_0(u1_struct_0(X1),u1_struct_0(X2)),u1_struct_0(esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(k1_tsep_1(esk3_0,X1,X2),esk1_0)|~m1_pre_topc(X2,esk3_0)|~m1_pre_topc(X1,esk3_0))).#trainneg
cnf(i_0_1703, negated_conjecture, (r4_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|u1_struct_0(X4)!=u1_struct_0(esk5_0)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_441, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(k2_tsep_1(X5,X3,X4)))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_771, plain, (r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|~r1_xboole_0(u1_struct_0(k1_tsep_1(X5,X2,X3)),u1_struct_0(X4))|~l1_struct_0(X4)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1578, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0)),u1_struct_0(k2_tsep_1(X1,esk4_0,esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk5_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5689, plain, (r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1)|~l1_struct_0(k1_tsep_1(X5,X2,X3))|~r1_tsep_1(k1_tsep_1(X5,X2,X3),X4))).#trainneg
cnf(i_0_5690, plain, (r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X1)|~l1_struct_0(k1_tsep_1(X1,X2,X3))|~l1_struct_0(X4)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1)|~l1_struct_0(k1_tsep_1(X5,X2,X3))|~r1_tsep_1(X4,k1_tsep_1(X5,X2,X3)))).#trainneg
cnf(i_0_280, plain, (r1_tarski(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_281, plain, (r1_tarski(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5600, negated_conjecture, (r4_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|esk9_3(esk1_0,k1_tsep_1(X1,X2,X3),X4)!=u1_struct_0(esk5_0)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(esk4_0)|~m1_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4810, negated_conjecture, (r1_tarski(k2_xboole_0(u1_struct_0(X1),u1_struct_0(X2)),u1_struct_0(esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(k1_tsep_1(esk4_0,X1,X2),esk1_0)|~m1_pre_topc(X2,esk4_0)|~m1_pre_topc(X1,esk4_0))).#trainneg
cnf(i_0_5516, plain, (r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(X5,X6)|~r1_tarski(X4,X7)|~r2_tsep_1(k1_tsep_1(X1,X2,X3),X7,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X7,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~r1_xboole_0(k3_xboole_0(X7,X6),k2_xboole_0(X4,X5))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(k1_tsep_1(X1,X2,X3)))).#trainneg
cnf(i_0_442, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k2_tsep_1(X5,X3,X4))))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,k2_tsep_1(X2,X3,X4))|~l1_pre_topc(k2_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1579, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0)),k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_301, plain, (r1_connsp_1(k2_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(k2_tsep_1(X1,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_tarski(X5,X6)|~r1_tarski(X4,X7)|~r2_tsep_1(k2_tsep_1(X1,X2,X3),X7,X6)|~m1_subset_1(X5,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X7,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~r1_xboole_0(k9_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X7,X6),k2_xboole_0(X4,X5))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~v2_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5743, negated_conjecture, (r4_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|esk9_3(esk1_0,k1_tsep_1(X1,X2,X3),X4)!=u1_struct_0(esk5_0)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(esk4_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X2,esk1_0))).#trainneg
cnf(i_0_5775, negated_conjecture, (r1_tarski(u1_struct_0(k1_tsep_1(X3,X1,X2)),u1_struct_0(esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(esk4_0,X1,X2),esk1_0)|~m1_pre_topc(X2,esk4_0)|~m1_pre_topc(X1,esk4_0)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_3001, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X2,X3,X4)),k2_xboole_0(u1_struct_0(X6),u1_struct_0(X7)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(k1_tsep_1(X5,X6,X7)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(k1_tsep_1(X5,X6,X7))|v2_struct_0(X5)|v2_struct_0(X6)|v2_struct_0(X7)|v2_struct_0(k1_tsep_1(X2,X3,X4))|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~m1_pre_topc(k1_tsep_1(X5,X6,X7),X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_pre_topc(X7,X5)|~m1_pre_topc(X6,X5)|~l1_pre_topc(X5))).#trainneg
cnf(i_0_1415, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|~r1_xboole_0(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|~l1_struct_0(k2_tsep_1(X1,X2,X3))|~l1_struct_0(k1_tsep_1(X4,X5,X6))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_443, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_xboole_0(u1_struct_0(X1),u1_struct_0(k2_tsep_1(X5,X3,X4)))|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_770, plain, (m1_subset_1(u1_struct_0(k1_tsep_1(X5,X2,X3)),k1_zfmisc_1(u1_struct_0(X4)))|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1752, negated_conjecture, (r1_xboole_0(k3_xboole_0(u1_struct_0(esk4_0),u1_struct_0(esk5_0)),u1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_445, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(k2_tsep_1(X5,X3,X4)))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~l1_struct_0(X1)|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~r1_tsep_1(k2_tsep_1(X2,X3,X4),X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_327, plain, (r1_tarski(u1_struct_0(esk13_3(esk12_3(X1,X2,X3),X4,X5)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(esk12_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(esk13_3(esk12_3(X1,X2,X3),X4,X5),X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk12_3(X1,X2,X3))|~v2_pre_topc(X1)|~r3_tsep_1(esk12_3(X1,X2,X3),X4,X5)|~m1_pre_topc(X5,esk12_3(X1,X2,X3))|~m1_pre_topc(X4,esk12_3(X1,X2,X3))|~l1_pre_topc(esk12_3(X1,X2,X3)))).#trainneg
cnf(i_0_328, plain, (r1_tarski(u1_struct_0(esk12_3(esk12_3(X1,X2,X3),X4,X5)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(esk12_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(esk12_3(esk12_3(X1,X2,X3),X4,X5),X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk12_3(X1,X2,X3))|~v2_pre_topc(X1)|~r3_tsep_1(esk12_3(X1,X2,X3),X4,X5)|~m1_pre_topc(X5,esk12_3(X1,X2,X3))|~m1_pre_topc(X4,esk12_3(X1,X2,X3))|~l1_pre_topc(esk12_3(X1,X2,X3)))).#trainneg
cnf(i_0_5956, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X2)|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2)|~l1_struct_0(k2_tsep_1(X5,X3,X4))|~r1_tsep_1(X1,k2_tsep_1(X5,X3,X4)))).#trainneg
cnf(i_0_5859, negated_conjecture, (r4_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(X4)!=u1_struct_0(esk5_0)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(esk4_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5873, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X6,X4,X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X6)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(k1_tsep_1(X6,X4,X5),X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~l1_pre_topc(X1)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_5517, plain, (r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|~r1_tarski(X5,esk11_3(k1_tsep_1(X6,X2,X3),X4,X5))|~r1_tarski(X4,esk10_3(k1_tsep_1(X6,X2,X3),X4,X5))|~r2_tsep_1(k1_tsep_1(X1,X2,X3),esk10_3(k1_tsep_1(X6,X2,X3),X4,X5),esk11_3(k1_tsep_1(X6,X2,X3),X4,X5))|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(esk11_3(k1_tsep_1(X6,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(esk10_3(k1_tsep_1(X6,X2,X3),X4,X5),k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(k1_tsep_1(X1,X2,X3))|~r1_connsp_1(k1_tsep_1(X6,X2,X3),X4,X5)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6)|~v2_pre_topc(k1_tsep_1(X6,X2,X3)))).#trainneg
cnf(i_0_1921, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k1_tsep_1(X3,X4,X5),k2_tsep_1(X6,X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X7)|v2_struct_0(X9)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|~r1_tarski(u1_struct_0(k1_tsep_1(X9,X4,X5)),u1_struct_0(k2_tsep_1(X7,X1,X2)))|~m1_pre_topc(k2_tsep_1(X6,X1,X2),X8)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X8)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X1,X7)|~l1_pre_topc(X8)|~l1_pre_topc(X6)|~l1_pre_topc(X7)|~v2_pre_topc(X8)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X5,X9)|~m1_pre_topc(X4,X9)|~l1_pre_topc(X3)|~l1_pre_topc(X9))).#trainneg
cnf(i_0_462, plain, (m1_subset_1(u1_struct_0(k2_tsep_1(X5,X2,X3)),k1_zfmisc_1(u1_struct_0(X4)))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5957, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X2)|~l1_struct_0(k2_tsep_1(X2,X3,X4))|~l1_struct_0(X1)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2)|~l1_struct_0(k2_tsep_1(X5,X3,X4))|~r1_tsep_1(k2_tsep_1(X5,X3,X4),X1))).#trainneg
cnf(i_0_786, plain, (r1_tarski(u1_struct_0(k1_tsep_1(X7,X2,X3)),u1_struct_0(k2_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X7)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),k2_tsep_1(X4,X5,X6))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X7)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_329, plain, (r1_tarski(u1_struct_0(k1_tsep_1(esk12_3(X1,X2,X3),X4,X5)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(esk12_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(k1_tsep_1(esk12_3(X1,X2,X3),X4,X5),X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk12_3(X1,X2,X3))|~v2_pre_topc(X1)|~m1_pre_topc(X5,esk12_3(X1,X2,X3))|~m1_pre_topc(X4,esk12_3(X1,X2,X3))|~l1_pre_topc(esk12_3(X1,X2,X3)))).#trainneg
cnf(i_0_330, plain, (r1_tarski(u1_struct_0(k2_tsep_1(esk12_3(X1,X2,X3),X4,X5)),u1_struct_0(X2))|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(esk12_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(k2_tsep_1(esk12_3(X1,X2,X3),X4,X5),X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk12_3(X1,X2,X3))|~v2_pre_topc(X1)|~m1_pre_topc(X5,esk12_3(X1,X2,X3))|~m1_pre_topc(X4,esk12_3(X1,X2,X3))|~l1_pre_topc(esk12_3(X1,X2,X3)))).#trainneg
cnf(i_0_5817, negated_conjecture, (r1_xboole_0(u1_struct_0(k1_tsep_1(X1,esk2_0,esk3_0)),u1_struct_0(k2_tsep_1(X2,esk4_0,esk5_0)))|r1_tsep_1(esk4_0,esk5_0)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(k2_tsep_1(esk1_0,esk4_0,esk5_0))|~l1_struct_0(k1_tsep_1(esk1_0,esk2_0,esk3_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~m1_pre_topc(esk5_0,X2)|~m1_pre_topc(esk4_0,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1920, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k1_tsep_1(X3,X4,X5),k2_tsep_1(X6,X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X7)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tarski(k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)),u1_struct_0(k2_tsep_1(X7,X1,X2)))|~m1_pre_topc(k2_tsep_1(X6,X1,X2),X8)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X8)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X1,X7)|~l1_pre_topc(X8)|~l1_pre_topc(X6)|~l1_pre_topc(X7)|~v2_pre_topc(X8)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_6023, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6018, negated_conjecture, (r3_tsep_1(esk1_0,k1_tsep_1(X1,X2,X3),X4)|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~r1_tsep_1(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),esk1_0)|k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))!=u1_struct_0(esk4_0)|u1_struct_0(X4)!=u1_struct_0(esk5_0)|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6039, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k1_tsep_1(X3,X4,X5),k2_tsep_1(X6,X1,X2))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X7)|v2_struct_0(X8)|v2_struct_0(X6)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(k2_tsep_1(X6,X1,X2),X9)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X9)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~m1_pre_topc(X2,X8)|~m1_pre_topc(X1,X8)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X5,X7)|~m1_pre_topc(X4,X7)|~l1_pre_topc(X9)|~l1_pre_topc(X6)|~l1_pre_topc(X8)|~l1_pre_topc(X3)|~l1_pre_topc(X7)|~v2_pre_topc(X9)|~m1_pre_topc(k1_tsep_1(X7,X4,X5),k2_tsep_1(X8,X1,X2))|~m1_pre_topc(k1_tsep_1(X7,X4,X5),X8)|~v2_pre_topc(X8))).#trainneg
cnf(i_0_787, plain, (r1_tarski(u1_struct_0(k1_tsep_1(X7,X2,X3)),u1_struct_0(k1_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X7)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k1_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|~m1_pre_topc(k1_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X7)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_337, plain, (r1_tarski(u1_struct_0(esk13_3(esk13_3(X1,X2,X3),X4,X5)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(esk13_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(esk13_3(esk13_3(X1,X2,X3),X4,X5),X3)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk13_3(X1,X2,X3))|~v2_pre_topc(X1)|~r3_tsep_1(esk13_3(X1,X2,X3),X4,X5)|~m1_pre_topc(X5,esk13_3(X1,X2,X3))|~m1_pre_topc(X4,esk13_3(X1,X2,X3))|~l1_pre_topc(esk13_3(X1,X2,X3)))).#trainneg
cnf(i_0_338, plain, (r1_tarski(u1_struct_0(esk12_3(esk13_3(X1,X2,X3),X4,X5)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(esk13_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(esk12_3(esk13_3(X1,X2,X3),X4,X5),X3)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk13_3(X1,X2,X3))|~v2_pre_topc(X1)|~r3_tsep_1(esk13_3(X1,X2,X3),X4,X5)|~m1_pre_topc(X5,esk13_3(X1,X2,X3))|~m1_pre_topc(X4,esk13_3(X1,X2,X3))|~l1_pre_topc(esk13_3(X1,X2,X3)))).#trainneg
cnf(i_0_6182, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X6,X4,X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_6247, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X6,X2,X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_6243, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(k1_tsep_1(X6,X4,X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_6591, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(k1_tsep_1(X4,X5,X6)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X6,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X1)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_6181, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6246, plain, (k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6927, negated_conjecture, (r1_tsep_1(esk5_0,esk4_0)|~l1_struct_0(esk5_0)|~l1_struct_0(esk4_0))).#trainneg
cnf(i_0_6256, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k1_tsep_1(X3,X4,X5),k2_tsep_1(X6,X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X7)|v2_struct_0(k2_tsep_1(X7,X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|~m1_pre_topc(k2_tsep_1(X6,X1,X2),X8)|~m1_pre_topc(k1_tsep_1(X3,X4,X5),X8)|~m1_pre_topc(k1_tsep_1(k2_tsep_1(X7,X1,X2),X4,X5),X7)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X1,X7)|~m1_pre_topc(X5,X3)|~m1_pre_topc(X4,X3)|~m1_pre_topc(X5,k2_tsep_1(X7,X1,X2))|~m1_pre_topc(X4,k2_tsep_1(X7,X1,X2))|~l1_pre_topc(X8)|~l1_pre_topc(X6)|~l1_pre_topc(X7)|~l1_pre_topc(X3)|~l1_pre_topc(k2_tsep_1(X7,X1,X2))|~v2_pre_topc(X8)|~v2_pre_topc(X7))).#trainneg
cnf(i_0_339, plain, (r1_tarski(u1_struct_0(k1_tsep_1(esk13_3(X1,X2,X3),X4,X5)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(esk13_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(k1_tsep_1(esk13_3(X1,X2,X3),X4,X5),X3)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk13_3(X1,X2,X3))|~v2_pre_topc(X1)|~m1_pre_topc(X5,esk13_3(X1,X2,X3))|~m1_pre_topc(X4,esk13_3(X1,X2,X3))|~l1_pre_topc(esk13_3(X1,X2,X3)))).#trainneg
cnf(i_0_340, plain, (r1_tarski(u1_struct_0(k2_tsep_1(esk13_3(X1,X2,X3),X4,X5)),u1_struct_0(X3))|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(esk13_3(X1,X2,X3))|v2_struct_0(X4)|v2_struct_0(X5)|~r3_tsep_1(X1,X2,X3)|~m1_pre_topc(k2_tsep_1(esk13_3(X1,X2,X3),X4,X5),X3)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(esk13_3(X1,X2,X3))|~v2_pre_topc(X1)|~m1_pre_topc(X5,esk13_3(X1,X2,X3))|~m1_pre_topc(X4,esk13_3(X1,X2,X3))|~l1_pre_topc(esk13_3(X1,X2,X3)))).#trainneg
cnf(i_0_6214, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))=k2_xboole_0(u1_struct_0(k1_tsep_1(X6,X2,X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_7062, negated_conjecture, (r1_tsep_1(esk5_0,esk4_0)|~l1_struct_0(esk4_0))).#trainneg
cnf(i_0_7210, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X6,X2,X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),k2_xboole_0(u1_struct_0(X4),u1_struct_0(X5)))|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_7280, negated_conjecture, (r1_tsep_1(esk5_0,esk4_0))).#trainneg
cnf(i_0_7266, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7274, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7275, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_403, negated_conjecture, (r1_tarski(k2_xboole_0(u1_struct_0(X1),u1_struct_0(X2)),u1_struct_0(esk2_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(k1_tsep_1(esk2_0,X1,X2),esk1_0)|~m1_pre_topc(X2,esk2_0)|~m1_pre_topc(X1,esk2_0))).#trainneg
cnf(i_0_7267, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_745, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k1_tsep_1(X5,X3,X4))))|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_pre_topc(X1,k1_tsep_1(X2,X3,X4))|~l1_pre_topc(k1_tsep_1(X2,X3,X4))|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X5)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_478, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X7,X2,X3)),u1_struct_0(k2_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X7)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),k2_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X7)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_343, plain, (r1_xboole_0(k4_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X4,X5),k9_subset_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),esk10_3(k1_tsep_1(X1,X2,X3),X4,X5),esk11_3(k1_tsep_1(X1,X2,X3),X4,X5)))|v2_struct_0(k1_tsep_1(X1,X2,X3))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k1_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k1_tsep_1(X1,X2,X3))|~v2_pre_topc(k1_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_344, plain, (r1_xboole_0(k4_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),X4,X5),k9_subset_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),esk10_3(k2_tsep_1(X1,X2,X3),X4,X5),esk11_3(k2_tsep_1(X1,X2,X3),X4,X5)))|v2_struct_0(k2_tsep_1(X1,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_connsp_1(k2_tsep_1(X1,X2,X3),X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3))))|~l1_pre_topc(k2_tsep_1(X1,X2,X3))|~v2_pre_topc(k2_tsep_1(X1,X2,X3))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_7480, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_7482, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_7499, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_7538, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_7209, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X6,X2,X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_7699, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|esk7_3(esk1_0,X1,X2)!=u1_struct_0(esk5_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_7453, negated_conjecture, (r3_tsep_1(esk1_0,esk5_0,esk4_0)|u1_struct_0(esk5_0)!=u1_struct_0(esk4_0))).#trainneg
cnf(i_0_7717, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_479, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X7,X2,X3)),u1_struct_0(k1_tsep_1(X4,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X4)|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X7)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),k1_tsep_1(X4,X5,X6))|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X7)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6166, plain, (r1_connsp_1(X1,u1_struct_0(k1_tsep_1(X2,X3,X4)),k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)))|v2_struct_0(X1)|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tarski(k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)),X7)|~r1_tarski(u1_struct_0(k1_tsep_1(X2,X3,X4)),X8)|~r2_tsep_1(X1,X8,X7)|~m1_subset_1(k2_xboole_0(u1_struct_0(X5),u1_struct_0(X6)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(k1_tsep_1(X2,X3,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k9_subset_1(u1_struct_0(X1),X8,X7),k2_xboole_0(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),u1_struct_0(k1_tsep_1(X2,X5,X6))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X6,X2)|~m1_pre_topc(X5,X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_6958, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(k1_tsep_1(X4,X5,X6)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X7,X5,X6)))|v2_struct_0(X6)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X6,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~m1_pre_topc(X6,X7)|~m1_pre_topc(X5,X7)|~l1_pre_topc(X7))).#trainneg
cnf(i_0_7877, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk4_0)|X1!=u1_struct_0(esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk5_0)!=u1_struct_0(esk4_0))).#trainneg
cnf(i_0_7879, negated_conjecture, (r1_tsep_1(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X1)!=u1_struct_0(esk4_0))).#trainneg
cnf(i_0_7878, negated_conjecture, (r2_tsep_1(esk1_0,X1,X2)|X2!=u1_struct_0(esk4_0)|X1!=u1_struct_0(esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk5_0)!=u1_struct_0(esk4_0))).#trainneg
cnf(i_0_7888, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|~r1_tarski(X2,u1_struct_0(X3))|~r1_tarski(X1,u1_struct_0(X4))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X4),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X3,esk5_0)|~m1_pre_topc(X4,esk1_0)|~m1_pre_topc(X4,esk4_0))).#trainneg
cnf(i_0_7875, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|u1_struct_0(X2)!=u1_struct_0(esk5_0)|esk6_3(esk1_0,X1,X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0))).#trainneg
cnf(i_0_7880, negated_conjecture, (r1_connsp_1(esk1_0,X1,X2)|X2!=u1_struct_0(X3)|X1!=u1_struct_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X4,esk1_0)|u1_struct_0(X3)!=u1_struct_0(esk5_0)|u1_struct_0(X4)!=u1_struct_0(esk4_0))).#trainneg
cnf(i_0_7884, negated_conjecture, (u1_struct_0(esk5_0)!=u1_struct_0(esk3_0)|u1_struct_0(esk4_0)!=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_6996, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X6,X1,X2)),k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)))=k2_xboole_0(u1_struct_0(k1_tsep_1(X6,X1,X2)),u1_struct_0(k1_tsep_1(X5,X3,X4)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X6)|~m1_pre_topc(X2,X5)|~m1_pre_topc(X1,X5)|~m1_pre_topc(X4,X5)|~m1_pre_topc(X3,X5)|~l1_pre_topc(X5)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X1,X6)|~l1_pre_topc(X6))).#trainneg
cnf(i_0_6049, plain, (r1_tarski(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(k2_tsep_1(X4,X5,X6)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(k2_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~m1_pre_topc(k1_tsep_1(k2_tsep_1(X4,X5,X6),X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,k2_tsep_1(X4,X5,X6))|~m1_pre_topc(X2,k2_tsep_1(X4,X5,X6))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~l1_pre_topc(k2_tsep_1(X4,X5,X6))|~v2_pre_topc(X4))).#trainneg
cnf(i_0_8251, negated_conjecture, (r1_connsp_1(esk1_0,X1,u1_struct_0(X2))|u1_struct_0(X2)!=u1_struct_0(esk5_0)|u1_struct_0(X3)!=u1_struct_0(esk4_0)|X1!=u1_struct_0(X3)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X3,esk1_0))).#trainneg
cnf(i_0_8194, negated_conjecture, (r1_tsep_1(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_struct_0(X1)|~l1_struct_0(X2)|u1_struct_0(X1)!=u1_struct_0(esk5_0)|u1_struct_0(X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk1_0))).#trainneg
cnf(i_0_8177, negated_conjecture, (r3_tsep_1(X1,X2,esk13_3(X1,X3,X4))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(esk13_3(X1,X3,X4))|~v1_tsep_1(X2,X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|u1_struct_0(esk13_3(X1,X3,X4))!=u1_struct_0(esk5_0)|u1_struct_0(X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(esk13_3(X1,X3,X4),esk1_0)|~m1_pre_topc(X2,esk1_0))).#trainneg
cnf(i_0_8180, negated_conjecture, (r3_tsep_1(X1,X2,esk12_3(X1,X3,X4))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(esk12_3(X1,X3,X4))|~v1_tsep_1(X2,X1)|~r3_tsep_1(X1,X3,X4)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|u1_struct_0(esk12_3(X1,X3,X4))!=u1_struct_0(esk5_0)|u1_struct_0(X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(esk12_3(X1,X3,X4),esk1_0)|~m1_pre_topc(X2,esk1_0))).#trainneg
cnf(i_0_8244, negated_conjecture, (r1_connsp_1(esk1_0,X1,u1_struct_0(X2))|~r1_tarski(X1,u1_struct_0(X3))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X3,esk1_0)|~m1_pre_topc(X3,esk4_0)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X2,esk3_0))).#trainneg
cnf(i_0_6495, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(k1_tsep_1(X7,X4,X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X6,X4,X5)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X6)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X7)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X6)|~l1_pre_topc(X1)|~m1_pre_topc(X5,X7)|~m1_pre_topc(X4,X7)|~l1_pre_topc(X7))).#trainneg
cnf(i_0_8156, negated_conjecture, (r3_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(esk5_0)!=u1_struct_0(esk4_0)|u1_struct_0(X2)!=u1_struct_0(esk4_0)|u1_struct_0(X1)!=u1_struct_0(esk5_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_755, plain, (r3_tsep_1(X1,k1_tsep_1(X2,X3,X4),X5)|v2_struct_0(X2)|v2_struct_0(X6)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_connsp_1(X1,u1_struct_0(k1_tsep_1(X6,X3,X4)),u1_struct_0(X5))|~m1_pre_topc(X5,X1)|~m1_pre_topc(k1_tsep_1(X2,X3,X4),X1)|~l1_pre_topc(X1)|~m1_pre_topc(X4,X6)|~m1_pre_topc(X3,X6)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X6)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8210, negated_conjecture, (r4_tsep_1(esk1_0,X1,X2)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)|u1_struct_0(esk5_0)!=u1_struct_0(esk4_0)|u1_struct_0(X2)!=u1_struct_0(esk4_0)|u1_struct_0(X1)!=u1_struct_0(esk5_0)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6257, plain, (r1_tarski(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(k1_tsep_1(X4,X5,X6)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(k1_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~m1_pre_topc(k1_tsep_1(k1_tsep_1(X4,X5,X6),X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,k1_tsep_1(X4,X5,X6))|~m1_pre_topc(X2,k1_tsep_1(X4,X5,X6))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~l1_pre_topc(k1_tsep_1(X4,X5,X6))|~v2_pre_topc(X4))).#trainneg
cnf(i_0_7543, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X1,X2,X3)),u1_struct_0(k2_tsep_1(X4,X5,X6)))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(k2_tsep_1(X4,X5,X6))|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X6)|~m1_pre_topc(k2_tsep_1(k2_tsep_1(X4,X5,X6),X2,X3),X4)|~m1_pre_topc(X6,X4)|~m1_pre_topc(X5,X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,k2_tsep_1(X4,X5,X6))|~m1_pre_topc(X2,k2_tsep_1(X4,X5,X6))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~l1_pre_topc(k2_tsep_1(X4,X5,X6))|~v2_pre_topc(X4))).#trainneg
cnf(i_0_8446, negated_conjecture, (r1_tsep_1(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|u1_struct_0(X1)!=u1_struct_0(esk5_0)|u1_struct_0(X2)!=u1_struct_0(esk4_0)|~l1_struct_0(X2)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_7211, plain, (k2_xboole_0(u1_struct_0(k1_tsep_1(X7,X2,X3)),u1_struct_0(k1_tsep_1(X1,X4,X5)))=k2_xboole_0(k2_xboole_0(u1_struct_0(X2),u1_struct_0(X3)),u1_struct_0(k1_tsep_1(X6,X4,X5)))|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X7)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X6)|~m1_pre_topc(X4,X6)|~l1_pre_topc(X1)|~l1_pre_topc(X6)|~m1_pre_topc(X3,X7)|~m1_pre_topc(X2,X7)|~l1_pre_topc(X7))).#trainneg
cnf(i_0_8733, negated_conjecture, (r1_tsep_1(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(X1)!=u1_struct_0(esk5_0)|u1_struct_0(X2)!=u1_struct_0(esk4_0)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk1_0)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_6652, plain, (r1_connsp_1(X1,u1_struct_0(k1_tsep_1(X2,X3,X4)),u1_struct_0(k1_tsep_1(X5,X6,X7)))|v2_struct_0(X1)|v2_struct_0(X7)|v2_struct_0(X6)|v2_struct_0(X2)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X5)|~r1_tarski(u1_struct_0(k1_tsep_1(X5,X6,X7)),X8)|~r1_tarski(u1_struct_0(k1_tsep_1(X2,X3,X4)),X9)|~r2_tsep_1(X1,X9,X8)|~m1_subset_1(u1_struct_0(k1_tsep_1(X5,X6,X7)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(u1_struct_0(k1_tsep_1(X2,X3,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X1)))|~r1_xboole_0(k9_subset_1(u1_struct_0(X1),X9,X8),k2_xboole_0(k2_xboole_0(u1_struct_0(X3),u1_struct_0(X4)),k2_xboole_0(u1_struct_0(X6),u1_struct_0(X7))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X7,X2)|~m1_pre_topc(X6,X2)|~m1_pre_topc(X7,X5)|~m1_pre_topc(X6,X5)|~l1_pre_topc(X2)|~l1_pre_topc(X5))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 73
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 73
# Processed clauses                    : 945
# ...of these trivial                  : 1
# ...subsumed                          : 353
# ...remaining for further processing  : 591
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 93
# Backward-rewritten                   : 62
# Generated clauses                    : 8991
# ...of the previous two non-trivial   : 8941
# Contextual simplify-reflections      : 194
# Paramodulations                      : 8973
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 430
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 408
# Current number of unprocessed clauses: 7919
# ...number of literals in the above   : 184802
# Current number of archived formulas  : 0
# Current number of archived clauses   : 752
# Proof object given clauses           : 121
# Proof search given clauses           : 591
# Clause-clause subsumption calls (NU) : 184887
# Rec. Clause-clause subsumption calls : 4831
# Non-unit clause-clause subsumptions  : 572
# Unit Clause-clause subsumption calls : 980
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 519534

# -------------------------------------------------
# User time                : 8.981 s
# System time              : 0.279 s
# Total time               : 9.261 s
# Maximum resident set size: 101556 pages
